Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvlSem
% (103511)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 190
% Passive clauses: 3506
% Generated clauses: 24831
% Final active clauses: 130
% Final passive clauses: 3017
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 927
% Trivial inequalities: 168
% Fw subsumption resolutions: 2309
% Fw demodulations: 4511
% Bw demodulations: 50
% 
% Simple tautologies: 220
% Equational tautologies: 52
% Forward subsumptions: 12392
% 
% Binary resolution: 542
% Factoring: 3
% Forward superposition: 9745
% Backward superposition: 6038
% Self superposition: 209
% Equality factoring: 106
% Equality resolution: 19
% 
% Split clauses: 275
% Split components: 550
% SAT solver clauses: 874
% SAT solver unit clauses: 41
% SAT solver binary clauses: 130
% 
% Sat splits: 275
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 872
% TWLsolver calls for satisfiability: 159
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvlSem
% (103541)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 960
% Generated clauses: 15186
% Final active clauses: 91
% Final passive clauses: 684
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 517
% Trivial inequalities: 41
% Fw subsumption resolutions: 760
% Fw demodulations: 1516
% Bw demodulations: 171
% Condensations: 1221
% 
% Simple tautologies: 60
% Equational tautologies: 41
% Forward subsumptions: 8319
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 268
% Factoring: 1
% Forward superposition: 4944
% Backward superposition: 4275
% Self superposition: 212
% Equality factoring: 82
% Equality resolution: 3
% 
% Split clauses: 81
% Split components: 162
% SAT solver clauses: 1725
% SAT solver unit clauses: 131
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 81
% Sat splitting refutations: 33
% 
% TWLsolver clauses: 1724
% TWLsolver calls for satisfiability: 252
% 
% Memory used [KB]: 22387
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvlSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvlSem
% SZS output start Proof for thm.bvlSem
fof(f30507,plain,(
  $false),
  inference(subsumption_resolution,[],[f30506,f213])).
fof(f213,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','thm.bool.TRUTH')).
fof(f30506,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f30504,f3176])).
fof(f3176,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2984,f294])).
fof(f294,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f215,f216])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7',aHL_BOOL_CASES)).
fof(f215,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7',aHL_FALSITY)).
fof(f2984,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f2876,f604])).
fof(f604,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X0,X0),'_i_0_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f570,f266])).
fof(f266,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X6,X15] : ! [X10,X9,X17] : s(X14,'const.combin.o_3'(s(fun(X6,X14),X10),s(fun(X15,X6),X9),s(X15,X17))) = s(X14,happ(s(fun(X6,X14),X10),s(X6,happ(s(fun(X15,X6),X9),s(X15,X17)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','thm.combin.o_THM')).
fof(f570,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f217,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X44,X45,X46] : (p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X44,fun(X44,bool)),'=_0'),s(X44,X45))),s(X44,X46)))) <=> s(X44,X45) = s(X44,X46))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7',arity_0_HL_EQ)).
fof(f217,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X34] : ! [X36] : s(X34,X36) = s(X34,happ(s(fun(X34,X34),'_i_0_0'),s(X34,X36)))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','_i0_ax')).
fof(f2876,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f266,f217])).
fof(f30504,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),fun(fun(sK0,'type.bvlSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f30441,f10076])).
fof(f10076,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.bvlSem.state'(sK1)),fun(fun(sK0,'type.bvlSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f212,f263])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f211])).
fof(f212,plain,(
  s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2))) != s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK4))),s(fun(sK0,'type.bvlSem.state'(sK1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(sK1),'type.bvlSem.state'(sK1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),sK3))),s(fun(sK0,'type.bvlSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) != s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) = s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))))),s(fun(X0,'type.bvlSem.state'(X1)),X2))) = s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X4))),s(fun(X0,'type.bvlSem.state'(X1)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X1),'type.bvlSem.state'(X1)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X3))),s(fun(X0,'type.bvlSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9))))),s(fun(X6,'type.bvlSem.state'(X7)),X8))) = s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10))),s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9))),s(fun(X6,'type.bvlSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9))))),s(fun(X6,'type.bvlSem.state'(X7)),X8))) = s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10))),s(fun(X6,'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9))),s(fun(X6,'type.bvlSem.state'(X7)),X8)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','thm.bvlSem.state_fupdfupds_comp|split|7')).
fof(f30441,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2))))),s(fun(X3,'type.bvlSem.state'(X0)),X4))) = s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1))),s(fun(X3,'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2))),s(fun(X3,'type.bvlSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f270,f228])).
fof(f228,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1))))) = s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X2))),s(fun('type.bvlSem.state'(X0),'type.bvlSem.state'(X0)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),'const.combin.o_2'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10),s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9))))) = s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.combin.o_2'(s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X10))),s(fun('type.bvlSem.state'(X7),'type.bvlSem.state'(X7)),'const.bvlSem.state_code_fupd_1'(s(fun('type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp')),'type.sptree.spt'('type.pair.prod'('type.num.num','type.bvl.exp'))),X9)))))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','thm.bvlSem.state_fupdfupds_comp|split|6')).
fof(f270,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14,X6,X19,X15] : ! [X10,X9,X8] : s(fun(X19,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X19,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X19,X15),X8))))) = s(fun(X19,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X19,X15),X8)))),
  file('thm.bvlSem.state_fupdfupds_comp|split|7','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.bvlSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 218
% Passive clauses: 4787
% Generated clauses: 33996
% Final active clauses: 99
% Final passive clauses: 1113
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 884
% Trivial inequalities: 84
% Fw subsumption resolutions: 1253
% Fw demodulations: 4446
% Bw demodulations: 3256
% Condensations: 1358
% Global subsumptions: 1
% 
% Simple tautologies: 82
% Equational tautologies: 203
% Forward subsumptions: 20640
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 124
% 
% Unit resulting resolution: 986
% Factoring: 1
% Forward superposition: 9075
% Backward superposition: 8277
% Self superposition: 293
% Equality factoring: 110
% Equality resolution: 3
% 
% Split clauses: 36
% Split components: 72
% SAT solver clauses: 5019
% SAT solver unit clauses: 706
% SAT solver binary clauses: 2827
% 
% Sat splits: 36
% Sat splitting refutations: 157
% 
% Memory used [KB]: 37739
% Time elapsed: 2.117 s
% ------------------------------
% ------------------------------
% Success in time 3.46 s
