Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (51536)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 3121
% Generated clauses: 13333
% Final active clauses: 123
% Final passive clauses: 2853
% Input formulas: 37
% Initial clauses: 61
% 
% Duplicate literals: 424
% Trivial inequalities: 313
% Fw subsumption resolutions: 1141
% Fw demodulations: 1644
% 
% Simple tautologies: 253
% Equational tautologies: 42
% Forward subsumptions: 5372
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 289
% Factoring: 1
% Forward superposition: 5159
% Backward superposition: 3949
% Self superposition: 169
% Equality factoring: 100
% Equality resolution: 30
% 
% Split clauses: 154
% Split components: 317
% SAT solver clauses: 345
% SAT solver unit clauses: 72
% SAT solver binary clauses: 85
% 
% Sat splits: 154
% 
% TWLsolver clauses: 344
% TWLsolver calls for satisfiability: 129
% 
% Memory used [KB]: 15735
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (51582)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 1005
% Generated clauses: 8473
% Final active clauses: 92
% Final passive clauses: 900
% Input formulas: 37
% Initial clauses: 61
% 
% Duplicate literals: 215
% Trivial inequalities: 33
% Fw subsumption resolutions: 236
% Fw demodulations: 1496
% Bw demodulations: 7
% Condensations: 327
% 
% Simple tautologies: 69
% Equational tautologies: 32
% Forward subsumptions: 3521
% 
% Binary resolution: 84
% Factoring: 1
% Forward superposition: 3034
% Backward superposition: 1794
% Self superposition: 106
% Equality factoring: 54
% Equality resolution: 14
% 
% Split clauses: 23
% Split components: 50
% SAT solver clauses: 1582
% SAT solver unit clauses: 77
% SAT solver binary clauses: 37
% 
% Sat splits: 23
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1582
% TWLsolver calls for satisfiability: 704
% 
% Memory used [KB]: 20468
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% (51607)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 645
% Passive clauses: 15278
% Generated clauses: 125628
% Final active clauses: 163
% Final passive clauses: 1310
% Input formulas: 37
% Initial clauses: 61
% Discarded non-redundant clauses: 6838
% 
% Duplicate literals: 2398
% Trivial inequalities: 1136
% Fw subsumption resolutions: 5547
% Fw demodulations: 23266
% Bw demodulations: 6974
% Condensations: 4209
% Global subsumptions: 160
% 
% Simple tautologies: 548
% Equational tautologies: 444
% Forward subsumptions: 71349
% Fw demodulations to eq. taut.: 42
% Bw demodulations to eq. taut.: 812
% 
% Unit resulting resolution: 6289
% Factoring: 14
% Forward superposition: 39361
% Backward superposition: 25620
% Self superposition: 431
% Equality factoring: 128
% Equality resolution: 46
% 
% Split clauses: 137
% Split components: 283
% SAT solver clauses: 16010
% SAT solver unit clauses: 235
% SAT solver binary clauses: 4283
% 
% Sat splits: 137
% Sat splitting refutations: 296
% 
% Memory used [KB]: 108996
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f22617,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f72,f89,f71,f90,f88,f70,f69,f94,f68,f92,f67,f98,f96,f87,f311,f475,f478,f479,f482,f484,f485,f486,f489,f488,f490,f491,f481,f492,f495,f494,f536,f558,f706,f707,f928,f929,f1083,f1313,f1315,f1318,f1319,f1320,f1321,f477,f1595,f1596,f1621,f1624,f1651,f1656,f1691,f1700,f1702,f1740,f2018,f2026,f2037,f2047,f2701,f2707,f2717,f2725,f2780,f2782,f3538,f3542,f3551,f3558,f2702,f2048,f3613,f3615,f3614,f3872,f3873,f1314,f3913,f3915,f3908,f3909,f3916,f4262,f4264,f3921,f3923,f3928,f4267,f4268,f3945,f3947,f3948,f3949,f3950,f3951,f4271,f4272,f3967,f3970,f3974,f4004,f4007,f4275,f4276,f4012,f4013,f4014,f4016,f2000,f4017,f4018,f4277,f4279,f4048,f2002,f4049,f4266,f4050,f4269,f4280,f4281,f4282,f4283,f4059,f4060,f476,f4061,f4284,f4285,f4286,f4287,f4288,f4102,f4103,f4289,f4290,f4294,f4295,f4180,f4263,f4181,f4298,f4299,f4300,f4301,f4302,f4191,f2003,f4192,f4270,f4193,f493,f4303,f4194,f4195,f4196,f2001,f4304,f4305,f4278,f487,f4306,f4307,f480,f4308,f4265,f2038,f4738,f4739,f2781,f2019,f5406,f6878,f6939,f6938,f6911,f6948,f6950,f6957,f1084,f5670,f6959,f6958,f7133,f13784,f13788,f13754,f13759,f13760,f13785,f13761,f13786,f13764,f13765,f13789,f13790,f16423,f16426,f16429,f16431,f16435,f16111,f16114,f16121,f16531,f16425,f16532,f16427,f16534,f16430,f16535,f16432,f16538,f16436,f16424,f16349,f16452,f16352,f16451,f16453,f16359,f16467,f16466,f16566,f16570,f16568,f16573,f16577,f16574,f16582,f16586,f16583,f16589,f16593,f16590,f16598,f16602,f16599,f16605,f16609,f16606,f16612,f16616,f16613,f16567,f16619,f16623,f16620,f3914,f16794,f1085,f16795,f19272,f19273,f19312,f19311,f19274,f19275,f19314,f19313,f19787,f19788,f19820,f19819,f19789,f19790,f19822,f19821,f20386,f20389,f20388,f20387,f21502,f21517,f21516,f21575,f21578,f21576,f22044,f22098,f22101,f22099,f22184,f22187,f22185,f22616])).
fof(f22616,plain,(
  $false | ($spl0 | $spl4 | $spl157)),
  inference(subsumption_resolution,[],[f22615,f22046])).
fof(f22046,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) | $spl157),
  inference(trivial_inequality_removal,[],[f22045])).
fof(f22045,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),sK2) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) | $spl157),
  inference(superposition,[],[f21515,f78])).
fof(f78,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK10(X1,X0))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK10(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f46])).
fof(f46,plain,(
  ! [X0,X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X43] : (? [X12] : s('type.semanticPrimitives.error_result'(X6),X43) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X12))) | ? [X12] : s('type.semanticPrimitives.error_result'(X6),X43) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X12))))),
  file('thm.closProps.eqs|split|7','thm.semanticPrimitives.error_result_nchotomy')).
fof(f21515,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) | $spl157),
  inference(cnf_transformation,[],[f21515_D])).
fof(f21515_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f22615,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) | ($spl0 | $spl4 | $spl157)),
  inference(trivial_inequality_removal,[],[f22613])).
fof(f22613,plain,(
  s(sK1,sK3) != s(sK1,sK3) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) | ($spl0 | $spl4 | $spl157)),
  inference(superposition,[],[f90,f22251])).
fof(f22251,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK10(sK2,sK0)))) | ($spl0 | $spl157)),
  inference(superposition,[],[f22055,f87])).
fof(f22055,plain,(
  ( ! [X12,X10,X11] : (s(X10,happ(s(fun(sK0,X10),X11),s(sK0,sK10(sK2,sK0)))) = s(X10,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,X10),X11),s(fun('type.semanticPrimitives.abort',X10),X12)))) ) | $spl157),
  inference(superposition,[],[f80,f22046])).
fof(f80,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X2))),s(fun(X1,X0),X3),s(fun('type.semanticPrimitives.abort',X0),X4)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X12,X11,X10] : s(X7,happ(s(fun(X6,X7),X11),s(X6,X12))) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X12))),s(fun(X6,X7),X11),s(fun('type.semanticPrimitives.abort',X7),X10)))),
  file('thm.closProps.eqs|split|7','thm.semanticPrimitives.error_result_case_def|split|0')).
fof(f22185,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) | $spl168),
  inference(cnf_transformation,[],[f22185_D])).
fof(f22185_D,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f22187,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) | $spl170),
  inference(cnf_transformation,[],[f22187_D])).
fof(f22187_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f22184,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0)))) ) | ($spl0 | $spl157)),
  inference(forward_demodulation,[],[f22183,f22046])).
fof(f22183,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))))) ) | ($spl0 | $spl157)),
  inference(trivial_inequality_removal,[],[f22153])).
fof(f22153,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),sK2) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))))) ) | ($spl0 | $spl157)),
  inference(superposition,[],[f1328,f22046])).
fof(f1328,plain,(
  ( ! [X39,X41,X40] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X39) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X40,X41,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X39,sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X40,X41,sK2,'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X39,sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X39,sK0)))),s(fun(sK0,sK1),X41),s(fun('type.semanticPrimitives.abort',sK1),X40))) | s('type.semanticPrimitives.error_result'(sK0),X39) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(X39,sK0))))) ) | $spl0),
  inference(superposition,[],[f688,f78])).
fof(f688,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X3) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X5,X4,sK2,X3,sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X5,X4,sK2,X3,sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X3),s(fun(sK0,sK1),X4),s(fun('type.semanticPrimitives.abort',sK1),X5)))) ) | $spl0),
  inference(superposition,[],[f87,f73])).
fof(f73,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))) | s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5) | (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X6),s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0))))) | (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0))))) | s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5) | ? [X9] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X9))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,X9))) != s(X1,happ(s(fun(X0,X1),X6),s(X0,X9)))) | ? [X8] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) & s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',X8))) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',X8)))) | s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5) | ? [X9] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X9))) & s(X1,happ(s(fun(X0,X1),X3),s(X0,X9))) != s(X1,happ(s(fun(X0,X1),X6),s(X0,X9)))) | ? [X8] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) & s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',X8))) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',X8))))) | s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.semanticPrimitives.error_result'(X0),X4) = s('type.semanticPrimitives.error_result'(X0),X5) & ! [X9] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X9))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X9))) = s(X1,happ(s(fun(X0,X1),X6),s(X0,X9)))) & ! [X8] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) => s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',X8))) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',X8))))) => s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((s('type.semanticPrimitives.error_result'(X0),X4) = s('type.semanticPrimitives.error_result'(X0),X5) & ! [X9] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X9))) => s(X1,happ(s(fun(X0,X1),X3),s(X0,X9))) = s(X1,happ(s(fun(X0,X1),X6),s(X0,X9)))) & ! [X8] : (s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) => s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',X8))) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',X8))))) => s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X16,X17,X18,X19,X11,X10] : ((s('type.semanticPrimitives.error_result'(X6),X18) = s('type.semanticPrimitives.error_result'(X6),X19) & ! [X12] : (s('type.semanticPrimitives.error_result'(X6),X19) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X12))) => s(X7,happ(s(fun(X6,X7),X11),s(X6,X12))) = s(X7,happ(s(fun(X6,X7),X17),s(X6,X12)))) & ! [X12] : (s('type.semanticPrimitives.error_result'(X6),X19) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X12))) => s(X7,happ(s(fun('type.semanticPrimitives.abort',X7),X10),s('type.semanticPrimitives.abort',X12))) = s(X7,happ(s(fun('type.semanticPrimitives.abort',X7),X16),s('type.semanticPrimitives.abort',X12))))) => s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X18),s(fun(X6,X7),X11),s(fun('type.semanticPrimitives.abort',X7),X10))) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X19),s(fun(X6,X7),X17),s(fun('type.semanticPrimitives.abort',X7),X16))))),
  file('thm.closProps.eqs|split|7','thm.semanticPrimitives.error_result_case_cong')).
fof(f22099,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X5),s(fun('type.semanticPrimitives.abort',sK1),X4))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) | $spl164),
  inference(cnf_transformation,[],[f22099_D])).
fof(f22099_D,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X5),s(fun('type.semanticPrimitives.abort',sK1),X4))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f22101,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) | $spl166),
  inference(cnf_transformation,[],[f22101_D])).
fof(f22101_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f22098,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X5),s(fun('type.semanticPrimitives.abort',sK1),X4)))) ) | ($spl0 | $spl157)),
  inference(forward_demodulation,[],[f22091,f22046])).
fof(f22091,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X5),s(fun('type.semanticPrimitives.abort',sK1),X4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))))) ) | ($spl0 | $spl157)),
  inference(trivial_inequality_removal,[],[f22049])).
fof(f22049,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),sK2) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X4,X5,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X5),s(fun('type.semanticPrimitives.abort',sK1),X4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK0))))) ) | ($spl0 | $spl157)),
  inference(superposition,[],[f1325,f22046])).
fof(f1325,plain,(
  ( ! [X19,X20,X18] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X18) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X19,X20,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(X18,sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X19,X20,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(X18,sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(X18,sK0)))),s(fun(sK0,sK1),X20),s(fun('type.semanticPrimitives.abort',sK1),X19))) | s('type.semanticPrimitives.error_result'(sK0),X18) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X18,sK0))))) ) | $spl0),
  inference(superposition,[],[f688,f78])).
fof(f22044,plain,(
  $false | ($spl0 | $spl2 | $spl156)),
  inference(subsumption_resolution,[],[f22043,f21516])).
fof(f22043,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) | ($spl0 | $spl2 | $spl156)),
  inference(trivial_inequality_removal,[],[f22041])).
fof(f22041,plain,(
  s(sK1,sK3) != s(sK1,sK3) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) | ($spl0 | $spl2 | $spl156)),
  inference(superposition,[],[f89,f21681])).
fof(f21681,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) | ($spl0 | $spl156)),
  inference(superposition,[],[f21528,f87])).
fof(f21528,plain,(
  ( ! [X14,X15,X16] : (s(X14,happ(s(fun('type.semanticPrimitives.abort',X14),X15),s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) = s(X14,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,X14),X16),s(fun('type.semanticPrimitives.abort',X14),X15)))) ) | $spl156),
  inference(superposition,[],[f77,f21516])).
fof(f77,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X2))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X12,X11,X10] : s(X7,happ(s(fun('type.semanticPrimitives.abort',X7),X10),s('type.semanticPrimitives.abort',X12))) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X12))),s(fun(X6,X7),X11),s(fun('type.semanticPrimitives.abort',X7),X10)))),
  file('thm.closProps.eqs|split|7','thm.semanticPrimitives.error_result_case_def|split|1')).
fof(f21576,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X9),s(fun('type.semanticPrimitives.abort',sK1),X8))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) | $spl160),
  inference(cnf_transformation,[],[f21576_D])).
fof(f21576_D,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X9),s(fun('type.semanticPrimitives.abort',sK1),X8))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f21578,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))) | $spl162),
  inference(cnf_transformation,[],[f21578_D])).
fof(f21578_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f21575,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X9),s(fun('type.semanticPrimitives.abort',sK1),X8)))) ) | ($spl0 | $spl156)),
  inference(forward_demodulation,[],[f21564,f21516])).
fof(f21564,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X9),s(fun('type.semanticPrimitives.abort',sK1),X8))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))))) ) | ($spl0 | $spl156)),
  inference(trivial_inequality_removal,[],[f21522])).
fof(f21522,plain,(
  ( ! [X8,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),sK2) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X8,X9,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0)))),s(fun(sK0,sK1),X9),s(fun('type.semanticPrimitives.abort',sK1),X8))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))),sK0))))) ) | ($spl0 | $spl156)),
  inference(superposition,[],[f1325,f21516])).
fof(f21516,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) | $spl156),
  inference(cnf_transformation,[],[f21516_D])).
fof(f21516_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f21517,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0))))) ) | $spl158),
  inference(cnf_transformation,[],[f21517_D])).
fof(f21517_D,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f21502,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X0,X1,sK2,'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0))),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK10(sK2,sK0)))),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(sK2,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f1325])).
fof(f20387,plain,(
  ( ! [X8,X11] : (p(s(bool,X11)) | s(bool,X8) = s(bool,X11)) ) | $spl152),
  inference(cnf_transformation,[],[f20387_D])).
fof(f20387_D,plain,(
  ( ! [X8,X11] : (p(s(bool,X11)) | s(bool,X8) = s(bool,X11)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f20388,plain,(
  ( ! [X10,X9] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X10),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool))))) | ~p(s(bool,X9))) ) | $spl154),
  inference(cnf_transformation,[],[f20388_D])).
fof(f20388_D,plain,(
  ( ! [X10,X9] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X10),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool))))) | ~p(s(bool,X9))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f20389,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X13)),bool))))) | p(s(bool,X15)) | s(bool,X12) = s(bool,X15) | ~p(s(bool,X13))) ) | ($spl80 | $spl134 | $spl142)),
  inference(subsumption_resolution,[],[f20336,f81])).
fof(f81,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X20,X21] : ((p(s(bool,X20)) => p(s(bool,X21))) => ((p(s(bool,X21)) => p(s(bool,X20))) => s(bool,X20) = s(bool,X21)))),
  file('thm.closProps.eqs|split|7','thm.bool.IMP_ANTISYM_AX')).
fof(f20336,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X13)),bool))))) | p(s(bool,X15)) | s(bool,X12) = s(bool,X15) | ~p(s(bool,X13)) | ~p(s(bool,X12))) ) | ($spl80 | $spl134 | $spl142)),
  inference(superposition,[],[f20050,f84])).
fof(f84,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f60])).
fof(f20050,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X26),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X25)),bool))))) | p(s(bool,X27)) | s(bool,X25) = s(bool,X27)) ) | ($spl80 | $spl134 | $spl142)),
  inference(superposition,[],[f19829,f17561])).
fof(f17561,plain,(
  ( ! [X66,X64,X67,X65] : (s(X64,happ(s(fun(bool,X64),X66),s(bool,X65))) = s(X64,happ(s(fun('type.semanticPrimitives.abort',X64),X67),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X65)),bool))))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f17395,f80])).
fof(f17395,plain,(
  ( ! [X10,X8,X7,X9] : (s(X8,happ(s(fun('type.semanticPrimitives.abort',X8),X9),s('type.semanticPrimitives.abort',sK11(X7,bool)))) = s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X7),s(fun(bool,X8),X10),s(fun('type.semanticPrimitives.abort',X8),X9)))) ) | ($spl80 | $spl134)),
  inference(forward_demodulation,[],[f17372,f16846])).
fof(f16846,plain,(
  ( ! [X6,X4,X5] : (s(X5,happ(s(fun('type.semanticPrimitives.abort',X5),X6),s('type.semanticPrimitives.abort',sK11(X4,bool)))) = s(X5,happ(s(fun('type.semanticPrimitives.abort',X5),X6),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool))),bool))))) ) | $spl134),
  inference(forward_demodulation,[],[f16823,f77])).
fof(f16823,plain,(
  ( ! [X6,X4,X7,X5] : (s(X5,happ(s(fun('type.semanticPrimitives.abort',X5),X6),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool))),bool)))) = s(X5,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool)))),s(fun(bool,X5),X7),s(fun('type.semanticPrimitives.abort',X5),X6)))) ) | $spl134),
  inference(superposition,[],[f77,f16795])).
fof(f17372,plain,(
  ( ! [X10,X8,X7,X9] : (s(X8,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X7),s(fun(bool,X8),X10),s(fun('type.semanticPrimitives.abort',X8),X9))) = s(X8,happ(s(fun('type.semanticPrimitives.abort',X8),X9),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X7,bool))),bool))))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f77,f17305])).
fof(f17305,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool))))) ) | ($spl80 | $spl134)),
  inference(backward_demodulation,[],[f17294,f16795])).
fof(f17294,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | $spl80),
  inference(duplicate_literal_removal,[],[f17293])).
fof(f17293,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | $spl80),
  inference(equality_resolution,[],[f6948])).
fof(f19829,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2)) ) | $spl142),
  inference(resolution,[],[f19314,f81])).
fof(f20386,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X10),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool))))) | p(s(bool,X11)) | s(bool,X8) = s(bool,X11) | ~p(s(bool,X9))) ) | ($spl80 | $spl134 | $spl142)),
  inference(subsumption_resolution,[],[f20335,f81])).
fof(f20335,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X10),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool))))) | p(s(bool,X11)) | s(bool,X8) = s(bool,X11) | ~p(s(bool,X8)) | ~p(s(bool,X9))) ) | ($spl80 | $spl134 | $spl142)),
  inference(superposition,[],[f20050,f84])).
fof(f19821,plain,(
  ( ! [X14,X10,X13] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10))) | ~p(s(bool,X10))) ) | $spl148),
  inference(cnf_transformation,[],[f19821_D])).
fof(f19821_D,plain,(
  ( ! [X14,X10,X13] : (s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10))) | ~p(s(bool,X10))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f19822,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | ~p(s(bool,X11))) ) | $spl150),
  inference(cnf_transformation,[],[f19822_D])).
fof(f19822_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | ~p(s(bool,X11))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f19790,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X15))) ) | ($spl80 | $spl134 | $spl136 | $spl140)),
  inference(superposition,[],[f19760,f84])).
fof(f19760,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X26),s(bool,X24)))) ) | ($spl80 | $spl134 | $spl136 | $spl140)),
  inference(subsumption_resolution,[],[f18632,f19443])).
fof(f19443,plain,(
  ( ! [X14,X10,X13] : (p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10)))) ) | ($spl136 | $spl140)),
  inference(subsumption_resolution,[],[f19313,f19311])).
fof(f18632,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X25)) | ~p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),X26),s(bool,X24)))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18310,f84])).
fof(f18310,plain,(
  ( ! [X4,X2,X0,X3] : (s(X0,happ(s(fun(bool,X0),X3),s(bool,X2))) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X2)))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f17561,f17561])).
fof(f19789,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10))) | ~p(s(bool,X10)) | ~p(s(bool,X11))) ) | ($spl80 | $spl134 | $spl136 | $spl140)),
  inference(superposition,[],[f19760,f84])).
fof(f19819,plain,(
  ( ! [X4,X0,X3] : (s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0))) | p(s(bool,X0))) ) | $spl144),
  inference(cnf_transformation,[],[f19819_D])).
fof(f19819_D,plain,(
  ( ! [X4,X0,X3] : (s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0))) | p(s(bool,X0))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f19820,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | p(s(bool,X1))) ) | $spl146),
  inference(cnf_transformation,[],[f19820_D])).
fof(f19820_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | p(s(bool,X1))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f19788,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,X8) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X5))) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl80 | $spl134 | $spl136 | $spl140)),
  inference(superposition,[],[f19760,f81])).
fof(f19787,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl80 | $spl134 | $spl136 | $spl140)),
  inference(superposition,[],[f19760,f81])).
fof(f19313,plain,(
  ( ! [X14,X10,X13] : (~p(s(bool,X10)) | p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10)))) ) | $spl140),
  inference(cnf_transformation,[],[f19313_D])).
fof(f19313_D,plain,(
  ( ! [X14,X10,X13] : (~p(s(bool,X10)) | p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f19314,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11))))) ) | $spl142),
  inference(cnf_transformation,[],[f19314_D])).
fof(f19314_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f19275,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | p(s(bool,X18)) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X15))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18630,f84])).
fof(f18630,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,X17)) | p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X16)))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X16)))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18310,f81])).
fof(f19274,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X14),s(bool,X10))) | ~p(s(bool,X10)) | ~p(s(bool,X11))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18630,f84])).
fof(f19311,plain,(
  ( ! [X4,X0,X3] : (p(s(bool,X3)) | p(s(bool,X0)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0)))) ) | $spl136),
  inference(cnf_transformation,[],[f19311_D])).
fof(f19311_D,plain,(
  ( ! [X4,X0,X3] : (p(s(bool,X3)) | p(s(bool,X0)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f19312,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | p(s(bool,X1))) ) | $spl138),
  inference(cnf_transformation,[],[f19312_D])).
fof(f19312_D,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | p(s(bool,X1))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f19273,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | p(s(bool,X8)) | s(bool,X8) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X5))) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18630,f81])).
fof(f19272,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | p(s(bool,X3)) | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X0))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl80 | $spl134)),
  inference(superposition,[],[f18630,f81])).
fof(f16795,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool))))) ) | $spl134),
  inference(cnf_transformation,[],[f16795_D])).
fof(f16795_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f1085,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2)))) ) | $spl32),
  inference(cnf_transformation,[],[f1085_D])).
fof(f1085_D,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f16794,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool))))) ) | $spl82),
  inference(duplicate_literal_removal,[],[f16644])).
fof(f16644,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))),bool))))) ) | $spl82),
  inference(superposition,[],[f6958,f6958])).
fof(f3914,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl56),
  inference(cnf_transformation,[],[f3914_D])).
fof(f3914_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f16620,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool))))))) ) | $spl132),
  inference(cnf_transformation,[],[f16620_D])).
fof(f16620_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f16623,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16622,f7013])).
fof(f7013,plain,(
  ( ! [X24,X23,X25,X22] : (s(X23,happ(s(fun('type.semanticPrimitives.abort',X23),X24),s('type.semanticPrimitives.abort',sK11(X22,bool)))) = s(X23,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X22),s(fun(bool,X23),X25),s(fun('type.semanticPrimitives.abort',X23),X24)))) ) | $spl30),
  inference(superposition,[],[f77,f1084])).
fof(f16622,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,'type.semanticPrimitives.abort'),X4),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16621,f7013])).
fof(f16621,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,'type.semanticPrimitives.abort'),X1),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,'type.semanticPrimitives.abort'),X4),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6982,f9302])).
fof(f9302,plain,(
  ( ! [X4,X2,X0,X3] : (s(X0,happ(s(fun(bool,X0),X3),s(bool,X2))) = s(X0,happ(s(fun(bool,X0),X4),s(bool,X2)))) ) | $spl30),
  inference(superposition,[],[f8531,f8531])).
fof(f8531,plain,(
  ( ! [X66,X64,X67,X65] : (s(X64,happ(s(fun(bool,X64),X66),s(bool,X65))) = s(X64,happ(s(fun('type.semanticPrimitives.abort',X64),X67),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X65)),bool))))) ) | $spl30),
  inference(superposition,[],[f7013,f80])).
fof(f6982,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1),s(bool,X6))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X4),s(bool,X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,'type.semanticPrimitives.abort'),X1),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,'type.semanticPrimitives.abort'),X4),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | $spl2),
  inference(superposition,[],[f912,f81])).
fof(f912,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.abort',happ(s(fun(X14,'type.semanticPrimitives.abort'),X16),s(X14,sK8(X17,X16,X18,X15,X19,X20,'type.semanticPrimitives.abort',X14)))) != s('type.semanticPrimitives.abort',happ(s(fun(X14,'type.semanticPrimitives.abort'),X19),s(X14,sK8(X17,X16,X18,X15,X19,X20,'type.semanticPrimitives.abort',X14)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X14),X15),s(fun(X14,'type.semanticPrimitives.abort'),X16),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X17))))) | s('type.semanticPrimitives.error_result'(X14),X18) = s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X17,X16,X18,X15,X19,X20,'type.semanticPrimitives.abort',X14)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X14),X18),s(fun(X14,'type.semanticPrimitives.abort'),X19),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X20))))) | s('type.semanticPrimitives.error_result'(X14),X15) != s('type.semanticPrimitives.error_result'(X14),X18)) ) | $spl2),
  inference(superposition,[],[f89,f74])).
fof(f74,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))) | s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X1,happ(s(fun(X0,X1),X3),s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun(X0,X1),X6),s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5)) )),
  inference(cnf_transformation,[],[f65])).
fof(f16619,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16618,f7013])).
fof(f16618,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,'type.semanticPrimitives.abort'),X11),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16617,f7013])).
fof(f16617,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,'type.semanticPrimitives.abort'),X8),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,'type.semanticPrimitives.abort'),X11),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6983,f9302])).
fof(f6983,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X8),s(bool,X13))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X11),s(bool,X13))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,'type.semanticPrimitives.abort'),X8),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,'type.semanticPrimitives.abort'),X11),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | $spl2),
  inference(superposition,[],[f912,f81])).
fof(f16567,plain,(
  ( ! [X27] : (~p(s(bool,X27))) ) | $spl116),
  inference(cnf_transformation,[],[f16567_D])).
fof(f16567_D,plain,(
  ( ! [X27] : (~p(s(bool,X27))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f16613,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool))))))) ) | $spl130),
  inference(cnf_transformation,[],[f16613_D])).
fof(f16613_D,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool))))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f16616,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16615,f7013])).
fof(f16615,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,'type.semanticPrimitives.abort'),X18),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16614,f7013])).
fof(f16614,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,'type.semanticPrimitives.abort'),X15),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,'type.semanticPrimitives.abort'),X18),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6984,f9302])).
fof(f6984,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X15),s(bool,X20))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X18),s(bool,X20))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,'type.semanticPrimitives.abort'),X15),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,'type.semanticPrimitives.abort'),X18),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | $spl2),
  inference(superposition,[],[f912,f84])).
fof(f16612,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16611,f7013])).
fof(f16611,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,'type.semanticPrimitives.abort'),X25),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16610,f7013])).
fof(f16610,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,'type.semanticPrimitives.abort'),X22),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,'type.semanticPrimitives.abort'),X25),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6985,f9302])).
fof(f6985,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X22),s(bool,X27))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X25),s(bool,X27))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,'type.semanticPrimitives.abort'),X22),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,'type.semanticPrimitives.abort'),X25),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | $spl2),
  inference(superposition,[],[f912,f84])).
fof(f16606,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool))))))) ) | $spl128),
  inference(cnf_transformation,[],[f16606_D])).
fof(f16606_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool))))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f16609,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16608,f7013])).
fof(f16608,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,'type.semanticPrimitives.abort'),X1),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16607,f7013])).
fof(f16607,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,'type.semanticPrimitives.abort'),X4),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,'type.semanticPrimitives.abort'),X1),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f7367,f9302])).
fof(f7367,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1),s(bool,X6))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X4),s(bool,X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,'type.semanticPrimitives.abort'),X4),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,'type.semanticPrimitives.abort'),X1),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,'type.semanticPrimitives.abort',bool)))) ) | $spl2),
  inference(superposition,[],[f913,f81])).
fof(f913,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.abort',happ(s(fun(X21,'type.semanticPrimitives.abort'),X23),s(X21,sK8(X27,X26,X22,X25,X23,X24,'type.semanticPrimitives.abort',X21)))) != s('type.semanticPrimitives.abort',happ(s(fun(X21,'type.semanticPrimitives.abort'),X26),s(X21,sK8(X27,X26,X22,X25,X23,X24,'type.semanticPrimitives.abort',X21)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X21),X22),s(fun(X21,'type.semanticPrimitives.abort'),X23),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X24))))) | s('type.semanticPrimitives.error_result'(X21),X22) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X27,X26,X22,X25,X23,X24,'type.semanticPrimitives.abort',X21)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X21),X25),s(fun(X21,'type.semanticPrimitives.abort'),X26),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X27))))) | s('type.semanticPrimitives.error_result'(X21),X22) != s('type.semanticPrimitives.error_result'(X21),X25)) ) | $spl2),
  inference(superposition,[],[f89,f74])).
fof(f16605,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16604,f7013])).
fof(f16604,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,'type.semanticPrimitives.abort'),X8),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16603,f7013])).
fof(f16603,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,'type.semanticPrimitives.abort'),X11),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,'type.semanticPrimitives.abort'),X8),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f7368,f9302])).
fof(f7368,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X8),s(bool,X13))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X11),s(bool,X13))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,'type.semanticPrimitives.abort'),X11),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,'type.semanticPrimitives.abort'),X8),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,'type.semanticPrimitives.abort',bool))) | p(s(bool,X13))) ) | $spl2),
  inference(superposition,[],[f913,f81])).
fof(f16599,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool))))))) ) | $spl126),
  inference(cnf_transformation,[],[f16599_D])).
fof(f16599_D,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool))))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f16602,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16601,f7013])).
fof(f16601,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,'type.semanticPrimitives.abort'),X15),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16600,f7013])).
fof(f16600,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,'type.semanticPrimitives.abort'),X18),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,'type.semanticPrimitives.abort'),X15),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f7369,f9302])).
fof(f7369,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X15),s(bool,X20))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X18),s(bool,X20))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,'type.semanticPrimitives.abort'),X18),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,'type.semanticPrimitives.abort'),X15),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,'type.semanticPrimitives.abort',bool))) | ~p(s(bool,X20))) ) | $spl2),
  inference(superposition,[],[f913,f84])).
fof(f16598,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16597,f7013])).
fof(f16597,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,'type.semanticPrimitives.abort'),X22),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f16596,f7013])).
fof(f16596,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,'type.semanticPrimitives.abort'),X25),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,'type.semanticPrimitives.abort'),X22),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f7370,f9302])).
fof(f7370,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X22),s(bool,X27))) != s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X25),s(bool,X27))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,'type.semanticPrimitives.abort'),X25),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort','const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,'type.semanticPrimitives.abort'),X22),s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.abort'),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,'type.semanticPrimitives.abort',bool)))) ) | $spl2),
  inference(superposition,[],[f913,f84])).
fof(f16590,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool))))))) ) | $spl124),
  inference(cnf_transformation,[],[f16590_D])).
fof(f16590_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f16593,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16592,f7013])).
fof(f16592,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,sK0),X4),s(fun('type.semanticPrimitives.abort',sK0),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16591,f7013])).
fof(f16591,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,sK0),X1),s(fun('type.semanticPrimitives.abort',sK0),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,sK0),X4),s(fun('type.semanticPrimitives.abort',sK0),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8832,f9302])).
fof(f8832,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK0,happ(s(fun(bool,sK0),X1),s(bool,X6))) != s(sK0,happ(s(fun(bool,sK0),X4),s(bool,X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,sK0),X1),s(fun('type.semanticPrimitives.abort',sK0),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,sK0),X4),s(fun('type.semanticPrimitives.abort',sK0),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | $spl4),
  inference(superposition,[],[f932,f81])).
fof(f932,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK0,happ(s(fun(X14,sK0),X16),s(X14,sK8(X17,X16,X18,X15,X19,X20,sK0,X14)))) != s(sK0,happ(s(fun(X14,sK0),X19),s(X14,sK8(X17,X16,X18,X15,X19,X20,sK0,X14)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X14),X15),s(fun(X14,sK0),X16),s(fun('type.semanticPrimitives.abort',sK0),X17))))) | s('type.semanticPrimitives.error_result'(X14),X18) = s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X17,X16,X18,X15,X19,X20,sK0,X14)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X14),X18),s(fun(X14,sK0),X19),s(fun('type.semanticPrimitives.abort',sK0),X20))))) | s('type.semanticPrimitives.error_result'(X14),X15) != s('type.semanticPrimitives.error_result'(X14),X18)) ) | $spl4),
  inference(superposition,[],[f90,f74])).
fof(f16589,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16588,f7013])).
fof(f16588,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,sK0),X11),s(fun('type.semanticPrimitives.abort',sK0),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16587,f7013])).
fof(f16587,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,sK0),X8),s(fun('type.semanticPrimitives.abort',sK0),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,sK0),X11),s(fun('type.semanticPrimitives.abort',sK0),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8833,f9302])).
fof(f8833,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK0,happ(s(fun(bool,sK0),X8),s(bool,X13))) != s(sK0,happ(s(fun(bool,sK0),X11),s(bool,X13))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,sK0),X8),s(fun('type.semanticPrimitives.abort',sK0),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,sK0),X11),s(fun('type.semanticPrimitives.abort',sK0),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | $spl4),
  inference(superposition,[],[f932,f81])).
fof(f16583,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool))))))) ) | $spl122),
  inference(cnf_transformation,[],[f16583_D])).
fof(f16583_D,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool))))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f16586,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16585,f7013])).
fof(f16585,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,sK0),X18),s(fun('type.semanticPrimitives.abort',sK0),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16584,f7013])).
fof(f16584,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,sK0),X15),s(fun('type.semanticPrimitives.abort',sK0),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,sK0),X18),s(fun('type.semanticPrimitives.abort',sK0),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8834,f9302])).
fof(f8834,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK0,happ(s(fun(bool,sK0),X15),s(bool,X20))) != s(sK0,happ(s(fun(bool,sK0),X18),s(bool,X20))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,sK0),X15),s(fun('type.semanticPrimitives.abort',sK0),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,sK0),X18),s(fun('type.semanticPrimitives.abort',sK0),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | $spl4),
  inference(superposition,[],[f932,f84])).
fof(f16582,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16581,f7013])).
fof(f16581,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,sK0),X25),s(fun('type.semanticPrimitives.abort',sK0),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16580,f7013])).
fof(f16580,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,sK0),X22),s(fun('type.semanticPrimitives.abort',sK0),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,sK0),X25),s(fun('type.semanticPrimitives.abort',sK0),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8835,f9302])).
fof(f8835,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK0,happ(s(fun(bool,sK0),X22),s(bool,X27))) != s(sK0,happ(s(fun(bool,sK0),X25),s(bool,X27))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,sK0),X22),s(fun('type.semanticPrimitives.abort',sK0),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,sK0),X25),s(fun('type.semanticPrimitives.abort',sK0),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | $spl4),
  inference(superposition,[],[f932,f84])).
fof(f16574,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool))))))) ) | $spl120),
  inference(cnf_transformation,[],[f16574_D])).
fof(f16574_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f16577,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X0),s('type.semanticPrimitives.abort',sK11(X3,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16576,f7013])).
fof(f16576,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X5),s('type.semanticPrimitives.abort',sK11(X2,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,sK0),X1),s(fun('type.semanticPrimitives.abort',sK0),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16575,f7013])).
fof(f16575,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,sK0),X4),s(fun('type.semanticPrimitives.abort',sK0),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,sK0),X1),s(fun('type.semanticPrimitives.abort',sK0),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f10140,f9302])).
fof(f10140,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(sK0,happ(s(fun(bool,sK0),X1),s(bool,X6))) != s(sK0,happ(s(fun(bool,sK0),X4),s(bool,X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X2),s(fun(bool,sK0),X4),s(fun('type.semanticPrimitives.abort',sK0),X5))))) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,X1,X2,X3,X4,X5,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X3),s(fun(bool,sK0),X1),s(fun('type.semanticPrimitives.abort',sK0),X0))))) | s('type.semanticPrimitives.error_result'(bool),X2) != s('type.semanticPrimitives.error_result'(bool),X3) | p(s(bool,X6)) | p(s(bool,sK8(X0,X1,X2,X3,X4,X5,sK0,bool)))) ) | $spl4),
  inference(superposition,[],[f933,f81])).
fof(f933,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK0,happ(s(fun(X21,sK0),X23),s(X21,sK8(X27,X26,X22,X25,X23,X24,sK0,X21)))) != s(sK0,happ(s(fun(X21,sK0),X26),s(X21,sK8(X27,X26,X22,X25,X23,X24,sK0,X21)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X21),X22),s(fun(X21,sK0),X23),s(fun('type.semanticPrimitives.abort',sK0),X24))))) | s('type.semanticPrimitives.error_result'(X21),X22) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X27,X26,X22,X25,X23,X24,sK0,X21)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X21),X25),s(fun(X21,sK0),X26),s(fun('type.semanticPrimitives.abort',sK0),X27))))) | s('type.semanticPrimitives.error_result'(X21),X22) != s('type.semanticPrimitives.error_result'(X21),X25)) ) | $spl4),
  inference(superposition,[],[f90,f74])).
fof(f16573,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X7),s('type.semanticPrimitives.abort',sK11(X10,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16572,f7013])).
fof(f16572,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X12),s('type.semanticPrimitives.abort',sK11(X9,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,sK0),X8),s(fun('type.semanticPrimitives.abort',sK0),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16571,f7013])).
fof(f16571,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,sK0),X11),s(fun('type.semanticPrimitives.abort',sK0),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,sK0),X8),s(fun('type.semanticPrimitives.abort',sK0),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f10141,f9302])).
fof(f10141,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(sK0,happ(s(fun(bool,sK0),X8),s(bool,X13))) != s(sK0,happ(s(fun(bool,sK0),X11),s(bool,X13))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X9),s(fun(bool,sK0),X11),s(fun('type.semanticPrimitives.abort',sK0),X12))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X7,X8,X9,X10,X11,X12,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X10),s(fun(bool,sK0),X8),s(fun('type.semanticPrimitives.abort',sK0),X7))))) | s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | p(s(bool,sK8(X7,X8,X9,X10,X11,X12,sK0,bool))) | p(s(bool,X13))) ) | $spl4),
  inference(superposition,[],[f933,f81])).
fof(f16568,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool))))))) ) | $spl118),
  inference(cnf_transformation,[],[f16568_D])).
fof(f16568_D,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool))))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f16570,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X14),s('type.semanticPrimitives.abort',sK11(X17,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16569,f7013])).
fof(f16569,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X19),s('type.semanticPrimitives.abort',sK11(X16,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,sK0),X15),s(fun('type.semanticPrimitives.abort',sK0),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f10175,f7013])).
fof(f10175,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,sK0),X18),s(fun('type.semanticPrimitives.abort',sK0),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,sK0),X15),s(fun('type.semanticPrimitives.abort',sK0),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f10142,f9302])).
fof(f10142,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(sK0,happ(s(fun(bool,sK0),X15),s(bool,X20))) != s(sK0,happ(s(fun(bool,sK0),X18),s(bool,X20))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X16),s(fun(bool,sK0),X18),s(fun('type.semanticPrimitives.abort',sK0),X19))))) | s('type.semanticPrimitives.error_result'(bool),X16) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X14,X15,X16,X17,X18,X19,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X17),s(fun(bool,sK0),X15),s(fun('type.semanticPrimitives.abort',sK0),X14))))) | s('type.semanticPrimitives.error_result'(bool),X16) != s('type.semanticPrimitives.error_result'(bool),X17) | ~p(s(bool,sK8(X14,X15,X16,X17,X18,X19,sK0,bool))) | ~p(s(bool,X20))) ) | $spl4),
  inference(superposition,[],[f933,f84])).
fof(f16566,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X21),s('type.semanticPrimitives.abort',sK11(X24,bool)))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f16565,f7013])).
fof(f16565,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun('type.semanticPrimitives.abort',sK0),X26),s('type.semanticPrimitives.abort',sK11(X23,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,sK0),X22),s(fun('type.semanticPrimitives.abort',sK0),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(forward_demodulation,[],[f10180,f7013])).
fof(f10180,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,sK0),X25),s(fun('type.semanticPrimitives.abort',sK0),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,sK0),X22),s(fun('type.semanticPrimitives.abort',sK0),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f10143,f9302])).
fof(f10143,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(sK0,happ(s(fun(bool,sK0),X22),s(bool,X27))) != s(sK0,happ(s(fun(bool,sK0),X25),s(bool,X27))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X23),s(fun(bool,sK0),X25),s(fun('type.semanticPrimitives.abort',sK0),X26))))) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X21,X22,X23,X24,X25,X26,sK0,bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),X24),s(fun(bool,sK0),X22),s(fun('type.semanticPrimitives.abort',sK0),X21))))) | s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),X24) | ~p(s(bool,X27)) | ~p(s(bool,sK8(X21,X22,X23,X24,X25,X26,sK0,bool)))) ) | $spl4),
  inference(superposition,[],[f933,f84])).
fof(f16466,plain,(
  ( ! [X1810,X1809] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun(bool,sK0),X1809),s(bool,X1810)))))) ) | $spl112),
  inference(cnf_transformation,[],[f16466_D])).
fof(f16466_D,plain,(
  ( ! [X1810,X1809] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun(bool,sK0),X1809),s(bool,X1810)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f16467,plain,(
  ( ! [X1812,X1811] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun(bool,sK0),X1811),s(bool,X1812)))))) ) | $spl114),
  inference(cnf_transformation,[],[f16467_D])).
fof(f16467_D,plain,(
  ( ! [X1812,X1811] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun(bool,sK0),X1811),s(bool,X1812)))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f16359,plain,(
  ( ! [X1812,X1810,X1811,X1809] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun(bool,sK0),X1811),s(bool,X1812))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun(bool,sK0),X1809),s(bool,X1810)))))) ) | ($spl4 | $spl30 | $spl44)),
  inference(superposition,[],[f90,f14897])).
fof(f14897,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(X16,happ(s(fun(bool,X16),X17),s(bool,X20))) = s(X16,happ(s(fun(bool,X16),X19),s(bool,X18)))) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f14024,f9302])).
fof(f14024,plain,(
  ( ! [X4,X2,X0,X5] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X4))) = s(X0,happ(s(fun(bool,X0),X2),s(bool,X5)))) ) | $spl44),
  inference(superposition,[],[f7215,f7215])).
fof(f7215,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(X59,happ(s(fun(bool,X59),X60),s(bool,X57))) = s(X59,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X58))),s(fun(bool,X59),X60),s(fun('type.semanticPrimitives.abort',X59),X61)))) ) | $spl44),
  inference(superposition,[],[f80,f6951])).
fof(f6951,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) ) | $spl44),
  inference(subsumption_resolution,[],[f6798,f6801])).
fof(f6801,plain,(
  ( ! [X30,X31,X34] : (s('type.semanticPrimitives.error_result'(bool),X30) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X31))) | s('type.semanticPrimitives.error_result'(bool),X30) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34)))) ) | $spl44),
  inference(subsumption_resolution,[],[f6800,f5591])).
fof(f5591,plain,(
  ( ! [X37,X35,X36] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X36))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X37,bool)))) | s('type.semanticPrimitives.error_result'(bool),X37) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl44),
  inference(superposition,[],[f5405,f4842])).
fof(f4842,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2)) ) | $spl44),
  inference(superposition,[],[f2019,f2019])).
fof(f5405,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) ) | $spl44),
  inference(equality_factoring,[],[f5076])).
fof(f5076,plain,(
  ( ! [X273,X272] : (s('type.semanticPrimitives.error_result'(bool),X272) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X273))) | s('type.semanticPrimitives.error_result'(bool),X272) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X272,bool))))) ) | $spl44),
  inference(superposition,[],[f78,f4842])).
fof(f6800,plain,(
  ( ! [X30,X33,X31,X34] : (s('type.semanticPrimitives.error_result'(bool),X30) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X31))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X30,bool)))) | s('type.semanticPrimitives.error_result'(bool),X30) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34)))) ) | $spl44),
  inference(subsumption_resolution,[],[f6734,f5591])).
fof(f6734,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.semanticPrimitives.error_result'(bool),X30) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X31))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X32))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X30,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X30,bool)))) | s('type.semanticPrimitives.error_result'(bool),X30) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34)))) ) | $spl44),
  inference(superposition,[],[f5669,f5076])).
fof(f5669,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X3)))) ) | $spl44),
  inference(superposition,[],[f5591,f5076])).
fof(f6798,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2)))) ) | $spl44),
  inference(equality_resolution,[],[f5669])).
fof(f16453,plain,(
  ( ! [X1761,X1760] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1760),s(bool,X1761)))))) ) | $spl110),
  inference(cnf_transformation,[],[f16453_D])).
fof(f16453_D,plain,(
  ( ! [X1761,X1760] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1760),s(bool,X1761)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f16451,plain,(
  ( ! [X1745,X1746] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1745),s(bool,X1746)))))) ) | $spl106),
  inference(cnf_transformation,[],[f16451_D])).
fof(f16451_D,plain,(
  ( ! [X1745,X1746] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1745),s(bool,X1746)))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f16352,plain,(
  ( ! [X1761,X1763,X1760,X1762] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1762),s(bool,X1763))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1760),s(bool,X1761)))))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f89,f14897])).
fof(f16452,plain,(
  ( ! [X1747,X1748] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1747),s(bool,X1748)))),bool))))) ) | $spl108),
  inference(cnf_transformation,[],[f16452_D])).
fof(f16452_D,plain,(
  ( ! [X1747,X1748] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1747),s(bool,X1748)))),bool))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f16349,plain,(
  ( ! [X1745,X1747,X1746,X1748] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1747),s(bool,X1748)))),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1745),s(bool,X1746)))))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9046,f14897])).
fof(f9046,plain,(
  ( ! [X847] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X847)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',X847)))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f89,f8530])).
fof(f8530,plain,(
  ( ! [X61,X60,X63] : (s(X60,happ(s(fun('type.semanticPrimitives.abort',X60),X63),s('type.semanticPrimitives.abort',X61))) = s(X60,happ(s(fun('type.semanticPrimitives.abort',X60),X63),s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X61)),bool))))) ) | $spl30),
  inference(superposition,[],[f7013,f77])).
fof(f16424,plain,(
  ( ! [X1600,X1599] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600)))) ) | $spl94),
  inference(cnf_transformation,[],[f16424_D])).
fof(f16424_D,plain,(
  ( ! [X1600,X1599] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f16436,plain,(
  ( ! [X1619,X1623,X1618,X1622] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0))))) ) | $spl104),
  inference(cnf_transformation,[],[f16436_D])).
fof(f16436_D,plain,(
  ( ! [X1619,X1623,X1618,X1622] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f16538,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16537,f9314])).
fof(f9314,plain,(
  ( ! [X59,X57,X58] : (s('type.semanticPrimitives.error_result'(X57),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X58)),bool)))) = s('type.semanticPrimitives.error_result'(X57),happ(s(fun(bool,'type.semanticPrimitives.error_result'(X57)),X59),s(bool,X58)))) ) | $spl30),
  inference(superposition,[],[f8531,f79])).
fof(f79,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rabort_0'),s('type.semanticPrimitives.abort',X1)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rabort_0'),s('type.semanticPrimitives.abort',X1)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rabort_0'),s('type.semanticPrimitives.abort',X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X13] : s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X13))) = s('type.semanticPrimitives.error_result'(X6),happ(s(fun('type.semanticPrimitives.abort','type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rabort_0'),s('type.semanticPrimitives.abort',X13)))),
  file('thm.closProps.eqs|split|7','ARITY_0_1_const.semanticPrimitives.Rabort')).
fof(f16537,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16536,f9314])).
fof(f16536,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16329,f9314])).
fof(f16329,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f687,f14897])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X0) | s('type.semanticPrimitives.error_result'(sK0),X0) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,X0,sK2,X1,X2,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),X0) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,X0,sK2,X1,X2,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X0),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),X2)))) ) | $spl0),
  inference(superposition,[],[f87,f73])).
fof(f16432,plain,(
  ( ! [X1612,X1617,X1613,X1616] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1613)),bool)))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0))))) ) | $spl102),
  inference(cnf_transformation,[],[f16432_D])).
fof(f16432_D,plain,(
  ( ! [X1612,X1617,X1613,X1616] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1613)),bool)))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f16535,plain,(
  ( ! [X1612,X1617,X1614,X1613,X1615,X1616] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1613)),bool)))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1614),s(bool,X1615))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16328,f9314])).
fof(f16328,plain,(
  ( ! [X1612,X1614,X1617,X1613,X1615,X1616] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1614),s(bool,X1615))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f688,f14897])).
fof(f16430,plain,(
  ( ! [X1608,X1607,X1611] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611)))) ) | $spl100),
  inference(cnf_transformation,[],[f16430_D])).
fof(f16430_D,plain,(
  ( ! [X1608,X1607,X1611] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f16534,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16533,f9314])).
fof(f16533,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16327,f9314])).
fof(f16327,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f1586,f14897])).
fof(f1586,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X0) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X0),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1))) | s('type.semanticPrimitives.error_result'(sK0),X0) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,X0,sK2,sK5,X1,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f689])).
fof(f689,plain,(
  ( ! [X6,X8,X7] : (s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK8(sK4,sK5,X6,sK2,X7,X8,sK1,sK0)))) != s(sK1,happ(s(fun(sK0,sK1),X7),s(sK0,sK8(sK4,sK5,X6,sK2,X7,X8,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),X6) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,X6,sK2,X7,X8,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X6),s(fun(sK0,sK1),X7),s(fun('type.semanticPrimitives.abort',sK1),X8))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X6)) ) | $spl0),
  inference(superposition,[],[f87,f74])).
fof(f16427,plain,(
  ( ! [X1602,X1606,X1603] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1603)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) | $spl98),
  inference(cnf_transformation,[],[f16427_D])).
fof(f16427_D,plain,(
  ( ! [X1602,X1606,X1603] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1603)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f16532,plain,(
  ( ! [X1602,X1604,X1606,X1603,X1605] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1603)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1604),s(bool,X1605))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16326,f9314])).
fof(f16326,plain,(
  ( ! [X1602,X1604,X1606,X1603,X1605] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1604),s(bool,X1605))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f2096,f14897])).
fof(f2096,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X1) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X1),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X0))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X0,sK5,sK2,X1,sK5,sK4,sK1,sK0))))) ) | $spl0),
  inference(equality_resolution,[],[f690])).
fof(f690,plain,(
  ( ! [X10,X11,X9] : (s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK8(X11,X10,sK2,X9,sK5,sK4,sK1,sK0)))) != s(sK1,happ(s(fun(sK0,sK1),X10),s(sK0,sK8(X11,X10,sK2,X9,sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X11,X10,sK2,X9,sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X9),s(fun(sK0,sK1),X10),s(fun('type.semanticPrimitives.abort',sK1),X11))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X9)) ) | $spl0),
  inference(superposition,[],[f87,f74])).
fof(f16425,plain,(
  ( ! [X1598,X1601,X1597] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0))))) ) | $spl96),
  inference(cnf_transformation,[],[f16425_D])).
fof(f16425_D,plain,(
  ( ! [X1598,X1601,X1597] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f16531,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16530,f9314])).
fof(f16530,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16325,f9314])).
fof(f16325,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f2187,f14897])).
fof(f2187,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X0) | s('type.semanticPrimitives.error_result'(sK0),X0) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,X0,sK2,X1,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X0),s(fun(sK0,sK1),X1),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) ) | $spl0),
  inference(equality_resolution,[],[f691])).
fof(f691,plain,(
  ( ! [X14,X12,X13] : (s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK9(sK4,sK5,X12,sK2,X13,X14,sK1,sK0)))) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),X14),s('type.semanticPrimitives.abort',sK9(sK4,sK5,X12,sK2,X13,X14,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),X12),s(fun(sK0,sK1),X13),s(fun('type.semanticPrimitives.abort',sK1),X14))) | s('type.semanticPrimitives.error_result'(sK0),X12) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,X12,sK2,X13,X14,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),X12)) ) | $spl0),
  inference(superposition,[],[f87,f75])).
fof(f75,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X4),s(fun(X0,X1),X6),s(fun('type.semanticPrimitives.abort',X1),X7))) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X5),s(fun(X0,X1),X3),s(fun('type.semanticPrimitives.abort',X1),X2))) | s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X2),s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X7),s('type.semanticPrimitives.abort',sK9(X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X5) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK8(X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X4) != s('type.semanticPrimitives.error_result'(X0),X5)) )),
  inference(cnf_transformation,[],[f65])).
fof(f16121,plain,(
  ( ! [X1812,X1810,X1811,X1809] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,happ(s(fun(bool,sK0),X1811),s(bool,X1812))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,happ(s(fun(bool,sK0),X1809),s(bool,X1810)))))) ) | ($spl4 | $spl30 | $spl44)),
  inference(superposition,[],[f90,f14897])).
fof(f16114,plain,(
  ( ! [X1761,X1763,X1760,X1762] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1762),s(bool,X1763))))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1760),s(bool,X1761)))))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f89,f14897])).
fof(f16111,plain,(
  ( ! [X1745,X1747,X1746,X1748] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1747),s(bool,X1748)))),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',happ(s(fun(bool,'type.semanticPrimitives.abort'),X1745),s(bool,X1746)))))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9046,f14897])).
fof(f16435,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16434,f9314])).
fof(f16434,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16433,f9314])).
fof(f16433,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1619)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16091,f9314])).
fof(f16091,plain,(
  ( ! [X1619,X1621,X1623,X1618,X1620,X1622] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1620),s(bool,X1621))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619)),sK2,X1622,X1623,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1618),s(bool,X1619))),s(fun(sK0,sK1),X1622),s(fun('type.semanticPrimitives.abort',sK1),X1623)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f687,f14897])).
fof(f16431,plain,(
  ( ! [X1612,X1617,X1614,X1613,X1615,X1616] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1613)),bool)))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1614),s(bool,X1615))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16090,f9314])).
fof(f16090,plain,(
  ( ! [X1612,X1614,X1617,X1613,X1615,X1616] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1614),s(bool,X1615))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(X1616,X1617,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613)),sK5,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1612),s(bool,X1613))),s(fun(sK0,sK1),X1617),s(fun('type.semanticPrimitives.abort',sK1),X1616)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f688,f14897])).
fof(f16429,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16428,f9314])).
fof(f16428,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1608)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16089,f9314])).
fof(f16089,plain,(
  ( ! [X1608,X1610,X1607,X1609,X1611] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1609),s(bool,X1610))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1611))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1607),s(bool,X1608)),sK2,sK5,X1611,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f1586,f14897])).
fof(f16426,plain,(
  ( ! [X1602,X1604,X1606,X1603,X1605] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1603)),bool)))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1604),s(bool,X1605))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16088,f9314])).
fof(f16088,plain,(
  ( ! [X1602,X1604,X1606,X1603,X1605] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1604),s(bool,X1605))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603))),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),X1606))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK9(X1606,sK5,sK2,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1602),s(bool,X1603)),sK5,sK4,sK1,sK0))))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f2096,f14897])).
fof(f16423,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16422,f9314])).
fof(f16422,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1598)),bool)))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(forward_demodulation,[],[f16087,f9314])).
fof(f16087,plain,(
  ( ! [X1600,X1598,X1601,X1597,X1599] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1599),s(bool,X1600))) | s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK8(sK4,sK5,happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598)),sK2,X1601,sK4,sK1,sK0)))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),happ(s(fun(bool,'type.semanticPrimitives.error_result'(sK0)),X1597),s(bool,X1598))),s(fun(sK0,sK1),X1601),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) ) | ($spl0 | $spl30 | $spl44)),
  inference(superposition,[],[f2187,f14897])).
fof(f13790,plain,(
  ( ! [X43,X41] : (s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X43),s(bool,happ(s(fun(sK0,bool),X41),s(sK0,sK6)))))) ) | $spl92),
  inference(cnf_transformation,[],[f13790_D])).
fof(f13790_D,plain,(
  ( ! [X43,X41] : (s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X43),s(bool,happ(s(fun(sK0,bool),X41),s(sK0,sK6)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f13789,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6)))),bool))))) ) | $spl90),
  inference(cnf_transformation,[],[f13789_D])).
fof(f13789_D,plain,(
  ( ! [X13] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6)))),bool))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f13765,plain,(
  ( ! [X59,X57,X58] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6)))),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X59),s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK6)))))) ) | ($spl2 | $spl16 | $spl30)),
  inference(superposition,[],[f9611,f2797])).
fof(f2797,plain,(
  ( ! [X12,X13] : (s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK6))) = s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6)))) ) | $spl16),
  inference(resolution,[],[f1349,f477])).
fof(f1349,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK6)))) ) | $spl16),
  inference(resolution,[],[f477,f84])).
fof(f9611,plain,(
  ( ! [X1002,X1001] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X1001)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X1002),s(bool,X1001)))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f89,f8531])).
fof(f13764,plain,(
  ( ! [X54,X56,X55] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X56),s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6)))))) ) | ($spl2 | $spl16 | $spl30)),
  inference(superposition,[],[f9611,f2797])).
fof(f13786,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool))))) ) | $spl88),
  inference(cnf_transformation,[],[f13786_D])).
fof(f13786_D,plain,(
  ( ! [X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f13761,plain,(
  ( ! [X43,X41,X42] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X42)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X43),s(bool,happ(s(fun(sK0,bool),X41),s(sK0,sK6)))))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9611,f2019])).
fof(f13785,plain,(
  ( ! [X6,X8] : (s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X8),s(bool,X6)))) ) | $spl86),
  inference(cnf_transformation,[],[f13785_D])).
fof(f13785_D,plain,(
  ( ! [X6,X8] : (s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X8),s(bool,X6)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f13760,plain,(
  ( ! [X39,X38,X40] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X39)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X40),s(bool,X38)))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9611,f4842])).
fof(f13759,plain,(
  ( ! [X37,X35,X36] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X37),s(bool,X35)))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9611,f4842])).
fof(f13754,plain,(
  ( ! [X14,X12,X13] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6)))),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X14),s(bool,X12)))) ) | ($spl2 | $spl30 | $spl44)),
  inference(superposition,[],[f9611,f2019])).
fof(f13788,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X11),s(bool,X9)))) ) | ($spl2 | $spl16 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f13787,f4903])).
fof(f4903,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl16 | $spl44)),
  inference(superposition,[],[f477,f2019])).
fof(f13787,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X11),s(bool,X9))) | ~p(s(bool,X10))) ) | ($spl2 | $spl16 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f13753,f4903])).
fof(f13753,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X11),s(bool,X9))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f9611,f84])).
fof(f13784,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X8),s(bool,X6)))) ) | ($spl2 | $spl16 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f13783,f4903])).
fof(f13783,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X8),s(bool,X6))) | ~p(s(bool,X6))) ) | ($spl2 | $spl16 | $spl30 | $spl44)),
  inference(subsumption_resolution,[],[f13752,f4903])).
fof(f13752,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool)))) | s(sK1,sK3) != s(sK1,happ(s(fun(bool,sK1),X8),s(bool,X6))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl2 | $spl30)),
  inference(superposition,[],[f9611,f84])).
fof(f7133,plain,(
  $false | $spl84),
  inference(equality_resolution,[],[f6959])).
fof(f6958,plain,(
  ( ! [X4,X7] : (s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool))),bool))))) ) | $spl82),
  inference(cnf_transformation,[],[f6958_D])).
fof(f6958_D,plain,(
  ( ! [X4,X7] : (s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool))),bool))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f6959,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X6)))) ) | $spl84),
  inference(cnf_transformation,[],[f6959_D])).
fof(f6959_D,plain,(
  ( ! [X6,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X6)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5670,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X6))) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X4,bool))),bool))))) ) | $spl44),
  inference(superposition,[],[f5591,f5076])).
fof(f1084,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | $spl30),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6957,plain,(
  ( ! [X28,X26,X25] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s('type.semanticPrimitives.error_result'(bool),X26) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X26,bool))))) ) | $spl44),
  inference(subsumption_resolution,[],[f6956,f5076])).
fof(f6956,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.semanticPrimitives.error_result'(bool),X26) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s('type.semanticPrimitives.error_result'(bool),X26) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X26,bool))))) ) | $spl44),
  inference(subsumption_resolution,[],[f6733,f6801])).
fof(f6733,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s('type.semanticPrimitives.error_result'(bool),X26) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29))) | s('type.semanticPrimitives.error_result'(bool),X26) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X26,bool))))) ) | $spl44),
  inference(superposition,[],[f5669,f5076])).
fof(f6950,plain,(
  ( ! [X23,X21,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X23))) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X21,bool))))) ) | $spl44),
  inference(subsumption_resolution,[],[f6809,f5076])).
fof(f6809,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),X21) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X23))) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X21,bool))))) ) | $spl44),
  inference(superposition,[],[f6801,f5076])).
fof(f6948,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.error_result'(bool),X21) != s('type.semanticPrimitives.error_result'(bool),X22) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X21,bool)))) | s('type.semanticPrimitives.error_result'(bool),X22) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X22,bool))))) ) | $spl80),
  inference(cnf_transformation,[],[f6948_D])).
fof(f6948_D,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.error_result'(bool),X21) != s('type.semanticPrimitives.error_result'(bool),X22) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X21,bool)))) | s('type.semanticPrimitives.error_result'(bool),X22) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X22,bool))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6911,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),X21) != s('type.semanticPrimitives.error_result'(bool),X22) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X23))) | s('type.semanticPrimitives.error_result'(bool),X22) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X22,bool)))) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X21,bool))))) ) | $spl44),
  inference(superposition,[],[f6868,f5076])).
fof(f6868,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'(bool),X13) != s('type.semanticPrimitives.error_result'(bool),X14) | s('type.semanticPrimitives.error_result'(bool),X14) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) | s('type.semanticPrimitives.error_result'(bool),X13) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X13,bool))))) ) | $spl44),
  inference(superposition,[],[f6801,f5076])).
fof(f6938,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool))))) ) | $spl76),
  inference(cnf_transformation,[],[f6938_D])).
fof(f6938_D,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6939,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool))),bool)))) | s('type.semanticPrimitives.error_result'(bool),X10) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11)))) ) | $spl78),
  inference(cnf_transformation,[],[f6939_D])).
fof(f6939_D,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool))),bool)))) | s('type.semanticPrimitives.error_result'(bool),X10) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f6878,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),X9) != s('type.semanticPrimitives.error_result'(bool),X10) | s('type.semanticPrimitives.error_result'(bool),X10) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X9,bool))),bool))))) ) | ($spl32 | $spl44)),
  inference(superposition,[],[f6868,f5634])).
fof(f5634,plain,(
  ( ! [X5,X3] : (s('type.semanticPrimitives.error_result'(bool),X3) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X3,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X3,bool))),bool))))) ) | ($spl32 | $spl44)),
  inference(subsumption_resolution,[],[f5617,f1085])).
fof(f5617,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),X3) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X3,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X3,bool))),bool))))) ) | $spl44),
  inference(superposition,[],[f5405,f5076])).
fof(f5406,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | $spl44),
  inference(duplicate_literal_removal,[],[f5247])).
fof(f5247,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | $spl44),
  inference(superposition,[],[f5076,f5076])).
fof(f2019,plain,(
  ( ! [X55,X51] : (s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))) ) | $spl44),
  inference(cnf_transformation,[],[f2019_D])).
fof(f2019_D,plain,(
  ( ! [X55,X51] : (s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2781,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | p(s(bool,X1))) ) | $spl52),
  inference(cnf_transformation,[],[f2781_D])).
fof(f2781_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | p(s(bool,X1))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4739,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,X5))) ) | $spl16),
  inference(subsumption_resolution,[],[f4692,f1349])).
fof(f4692,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl16),
  inference(superposition,[],[f4657,f81])).
fof(f4657,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6)))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f4655])).
fof(f4655,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6)))) ) | $spl16),
  inference(equality_factoring,[],[f3043])).
fof(f3043,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6)))) ) | $spl16),
  inference(resolution,[],[f2794,f1349])).
fof(f2794,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X6)) | s(bool,X4) = s(bool,X6) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X5),s(sK0,sK6)))) ) | $spl16),
  inference(resolution,[],[f1349,f81])).
fof(f4738,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1))) ) | $spl16),
  inference(subsumption_resolution,[],[f4691,f1349])).
fof(f4691,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl16),
  inference(superposition,[],[f4657,f81])).
fof(f2038,plain,(
  ( ! [X87,X86,X82] : (s(bool,X82) = s(bool,X86) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK6)))) ) | $spl46),
  inference(cnf_transformation,[],[f2038_D])).
fof(f2038_D,plain,(
  ( ! [X87,X86,X82] : (s(bool,X82) = s(bool,X86) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK6)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4265,plain,(
  ( ! [X68,X65] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | $spl62),
  inference(cnf_transformation,[],[f4265_D])).
fof(f4265_D,plain,(
  ( ! [X68,X65] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4308,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X40),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK6))) | s(bool,X41) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(sK0,bool),X44),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4230,f2794])).
fof(f4230,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,X41)) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X40),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK6))) | s(bool,X41) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(sK0,bool),X44),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f1349,f2702])).
fof(f480,plain,(
  ( ! [X66,X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7)))) ) | $spl18),
  inference(cnf_transformation,[],[f480_D])).
fof(f480_D,plain,(
  ( ! [X66,X67] : (~p(s(bool,X67)) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4307,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X35),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X37)) | s(bool,X36) = s(bool,X38) | s(bool,X38) = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4229,f2794])).
fof(f4229,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,X36)) | s(bool,X37) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X35),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X37)) | s(bool,X36) = s(bool,X38) | s(bool,X38) = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4306,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X30),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X32)) | s(bool,X31) = s(bool,X33) | s(bool,X33) = s(bool,happ(s(fun(sK0,bool),X34),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4228,f2794])).
fof(f4228,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X30),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X32)) | s(bool,X31) = s(bool,X33) | s(bool,X33) = s(bool,happ(s(fun(sK0,bool),X34),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f487,plain,(
  ( ! [X69] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))))) ) | $spl22),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  ( ! [X69] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4278,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6))) | s(bool,X9) = s(bool,X10)) ) | $spl70),
  inference(cnf_transformation,[],[f4278_D])).
fof(f4278_D,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6))) | s(bool,X9) = s(bool,X10)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4305,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X6) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4223,f3046])).
fof(f3046,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X17)) | s(bool,X14) = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK6))) | s(bool,X15) = s(bool,X17) | s(bool,X14) = s(bool,X15)) ) | $spl16),
  inference(resolution,[],[f2794,f84])).
fof(f4223,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X7)) | s(bool,X6) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4304,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X1) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4222,f3046])).
fof(f4222,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X1) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),X4),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f2001,plain,(
  ( ! [X28,X31,X29] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK6)))) ) | $spl38),
  inference(cnf_transformation,[],[f2001_D])).
fof(f2001_D,plain,(
  ( ! [X28,X31,X29] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK6)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4196,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X68),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X69),s(sK0,sK6))) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2797,f2702])).
fof(f4195,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X63),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X64),s(sK0,sK6))) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK6))) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2797,f2702])).
fof(f4194,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X59),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X60) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6))) | s(bool,X60) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f477,f2702])).
fof(f4303,plain,(
  ( ! [X57,X54,X56,X55,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK6))) | s(bool,X51) = s(bool,X56) | s(bool,X53) = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) | $spl74),
  inference(cnf_transformation,[],[f4303_D])).
fof(f4303_D,plain,(
  ( ! [X57,X54,X56,X55,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK6))) | s(bool,X51) = s(bool,X56) | s(bool,X53) = s(bool,X55) | s(bool,X51) = s(bool,X55)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f493,plain,(
  ( ! [X81] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7))))) ) | $spl26),
  inference(cnf_transformation,[],[f493_D])).
fof(f493_D,plain,(
  ( ! [X81] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4193,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X52),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6))) | s(bool,X51) = s(bool,X55) | s(bool,X53) = s(bool,X55) | s(bool,X51) = s(bool,X56) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f3046,f2702])).
fof(f4270,plain,(
  ( ! [X70,X68,X66,X71,X67] : (s(bool,X66) = s(bool,X67) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6))) | s(bool,X66) = s(bool,X70) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | $spl68),
  inference(cnf_transformation,[],[f4270_D])).
fof(f4270_D,plain,(
  ( ! [X70,X68,X66,X71,X67] : (s(bool,X66) = s(bool,X67) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6))) | s(bool,X66) = s(bool,X70) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4192,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X46),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X45) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(sK0,bool),X48),s(sK0,sK6))) | s(bool,X45) = s(bool,X49) | s(bool,X49) = s(bool,happ(s(fun(sK0,bool),X50),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2794,f2702])).
fof(f2003,plain,(
  ( ! [X35,X38,X36,X40] : (s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK6))) | s(bool,X36) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6)))) ) | $spl42),
  inference(cnf_transformation,[],[f2003_D])).
fof(f2003_D,plain,(
  ( ! [X35,X38,X36,X40] : (s(bool,X35) = s(bool,X36) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK6))) | s(bool,X36) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4191,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X41),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X40) = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK6))) | s(bool,X40) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(sK0,bool),X44),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f1349,f2702])).
fof(f4302,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X26),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X27)) | s(bool,X25) = s(bool,X28) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4188,f2794])).
fof(f4188,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X26),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X27)) | ~p(s(bool,X25)) | s(bool,X25) = s(bool,X28) | s(bool,X28) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4301,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X22)) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4187,f2794])).
fof(f4187,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X20)) | ~p(s(bool,X22)) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4300,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X16),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X15) = s(bool,X17) | s(bool,X15) = s(bool,X18) | s(bool,X18) = s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4186,f3046])).
fof(f4186,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X16),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X15) = s(bool,X17) | p(s(bool,X17)) | s(bool,X15) = s(bool,X18) | s(bool,X18) = s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4299,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X11),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4185,f3046])).
fof(f4185,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X11),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X10) = s(bool,X12) | p(s(bool,X12)) | s(bool,X10) = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4298,plain,(
  ( ! [X70,X66,X71] : (~p(s(bool,X66)) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X71),s('type.semanticPrimitives.abort',sK7))) | s(bool,X66) = s(bool,X70)) ) | $spl72),
  inference(cnf_transformation,[],[f4298_D])).
fof(f4298_D,plain,(
  ( ! [X70,X66,X71] : (~p(s(bool,X66)) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X71),s('type.semanticPrimitives.abort',sK7))) | s(bool,X66) = s(bool,X70)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4181,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (~p(s(bool,X66)) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | s(bool,X69) = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6))) | s(bool,X67) = s(bool,X69) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X71),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f3046,f2702])).
fof(f4263,plain,(
  ( ! [X4,X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1)) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) | $spl60),
  inference(cnf_transformation,[],[f4263_D])).
fof(f4263_D,plain,(
  ( ! [X4,X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1)) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4180,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (p(s(bool,X60)) | s(bool,X61) = s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6))) | s(bool,X61) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK6))) | s(bool,X60) = s(bool,X63) | s(bool,X63) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X64),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2794,f2702])).
fof(f4295,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | s(bool,X31) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK6))) | s(bool,X30) = s(bool,X32) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4174,f1349])).
fof(f4174,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | s(bool,X31) = s(bool,happ(s(fun(sK0,bool),X29),s(sK0,sK6))) | p(s(bool,X31)) | s(bool,X30) = s(bool,X32) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4294,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | s(bool,X26) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X28),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4173,f1349])).
fof(f4173,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | s(bool,X26) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6))) | p(s(bool,X26)) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X28),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4290,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X7) = s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | p(s(bool,X5)) | s(bool,X5) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4130,f1349])).
fof(f4130,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X7) = s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X5) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X2) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6))) | p(s(bool,X0)) | s(bool,X0) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4129,f1349])).
fof(f4129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X2) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6))) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4103,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X68) = s(bool,happ(s(fun(sK0,bool),X69),s(sK0,sK6))) | s(bool,X68) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X71),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2797,f2702])).
fof(f4102,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X64),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2797,f2702])).
fof(f4288,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X59) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X61),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4101,f1349])).
fof(f4101,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X59)) | s(bool,X59) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X61),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f477,f2702])).
fof(f4287,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X26) = s(bool,X27) | ~p(s(bool,X25)) | s(bool,X26) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X25) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X29),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4095,f2794])).
fof(f4095,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,X26) = s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X25)) | s(bool,X26) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X25) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X29),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4286,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X21) = s(bool,X22) | ~p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK6))) | s(bool,X20) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X24),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4094,f2794])).
fof(f4094,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,X21) = s(bool,X22) | ~p(s(bool,X20)) | ~p(s(bool,X22)) | s(bool,X21) = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK6))) | s(bool,X20) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X24),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4285,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X17) | p(s(bool,X17)) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6))) | s(bool,X15) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X19),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4093,f1349])).
fof(f4093,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) = s(bool,X17) | p(s(bool,X17)) | s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6))) | s(bool,X15) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X19),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4284,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X10) = s(bool,X12) | p(s(bool,X12)) | s(bool,X11) = s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6))) | s(bool,X10) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4092,f1349])).
fof(f4092,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X11)) | s(bool,X10) = s(bool,X12) | p(s(bool,X12)) | s(bool,X11) = s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6))) | s(bool,X10) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4061,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,X52)) | s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6))) | s(bool,X51) = s(bool,X55) | s(bool,X53) = s(bool,X55) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK6))) | s(bool,X52) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X57),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f3046,f2702])).
fof(f476,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7)))) ) | $spl14),
  inference(cnf_transformation,[],[f476_D])).
fof(f476_D,plain,(
  ( ! [X61,X60] : (p(s(bool,X61)) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4060,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,X46)) | s(bool,X45) = s(bool,X47) | s(bool,X47) = s(bool,happ(s(fun(sK0,bool),X48),s(sK0,sK6))) | s(bool,X45) = s(bool,happ(s(fun(sK0,bool),X49),s(sK0,sK6))) | s(bool,X46) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X50),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2794,f2702])).
fof(f4059,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,X41)) | s(bool,X40) = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK6))) | s(bool,X40) = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK6))) | s(bool,X41) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X44),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f1349,f2702])).
fof(f4283,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,X36)) | s(bool,X35) = s(bool,X37) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6))) | s(bool,X36) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4058,f2794])).
fof(f4058,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,X36)) | s(bool,X35) = s(bool,X37) | ~p(s(bool,X37)) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6))) | s(bool,X36) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4282,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X32) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),X33),s(sK0,sK6))) | s(bool,X31) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X34),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4057,f2794])).
fof(f4057,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X32) | ~p(s(bool,X32)) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),X33),s(sK0,sK6))) | s(bool,X31) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X34),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f84,f2702])).
fof(f4281,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | p(s(bool,X7)) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK6))) | s(bool,X6) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4052,f1349])).
fof(f4052,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK6))) | s(bool,X6) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4280,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4051,f1349])).
fof(f4051,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X0)) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f81,f2702])).
fof(f4269,plain,(
  ( ! [X69,X65] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) ) | $spl66),
  inference(cnf_transformation,[],[f4269_D])).
fof(f4269_D,plain,(
  ( ! [X69,X65] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4050,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,X49) = s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X52),s(sK0,sK6))) | s(bool,X49) = s(bool,happ(s(fun(sK0,bool),X53),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X48),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X54),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4266,plain,(
  ( ! [X70,X66,X71,X69,X67] : (s(bool,X66) = s(bool,X67) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6))) | s(bool,X66) = s(bool,X70) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) ) | $spl64),
  inference(cnf_transformation,[],[f4266_D])).
fof(f4266_D,plain,(
  ( ! [X70,X66,X71,X69,X67] : (s(bool,X66) = s(bool,X67) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6))) | s(bool,X66) = s(bool,X70) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4049,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X41),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK6))) | s(bool,X42) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X47),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f2002,plain,(
  ( ! [X39,X37,X34] : (s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X37),s('type.semanticPrimitives.abort',sK7))) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7)))) ) | $spl40),
  inference(cnf_transformation,[],[f2002_D])).
fof(f2002_D,plain,(
  ( ! [X39,X37,X34] : (s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X37),s('type.semanticPrimitives.abort',sK7))) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4048,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(bool,X36) = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK6))) | s(bool,X36) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6))) | s(bool,X39) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X34),s('type.semanticPrimitives.abort',sK7))) | s(bool,X39) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X40),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4279,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,X13) = s(bool,X14) | s(bool,X14) = s(bool,X15) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X12),s('type.semanticPrimitives.abort',sK7))))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4044,f3046])).
fof(f4044,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,X13) = s(bool,X14) | s(bool,X14) = s(bool,X15) | s(bool,X15) = s(bool,happ(s(fun(sK0,bool),X16),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X12),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X13))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f4277,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4043,f3046])).
fof(f4043,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6))) | p(s(bool,X8)) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f4018,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X62),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X64),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2702,f2797])).
fof(f4017,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X57),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK6))) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6))) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2702,f2797])).
fof(f2000,plain,(
  ( ! [X30,X33,X27,X32] : (s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X30),s('type.semanticPrimitives.abort',sK7))) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7))) | s(bool,X27) = s(bool,X32)) ) | $spl36),
  inference(cnf_transformation,[],[f2000_D])).
fof(f2000_D,plain,(
  ( ! [X30,X33,X27,X32] : (s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X30),s('type.semanticPrimitives.abort',sK7))) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7))) | s(bool,X27) = s(bool,X32)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4016,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,X49) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X50),s('type.semanticPrimitives.abort',sK7))) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X48),s(sK0,sK6))) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X52),s(sK0,sK6))) | s(bool,X49) = s(bool,X53) | s(bool,X53) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X54),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4014,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(bool,X35) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X36),s('type.semanticPrimitives.abort',sK7))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6))) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK6))) | s(bool,X35) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X40),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4013,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X29),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X27) = s(bool,X30) | s(bool,X30) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK6))) | s(bool,X27) = s(bool,X32) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4012,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X22),s('type.semanticPrimitives.abort',sK7))) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6))) | s(bool,X20) = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4276,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X17),s('type.semanticPrimitives.abort',sK7))) | s(bool,X15) = s(bool,X18) | s(bool,X18) = s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK6))) | ~p(s(bool,X16))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4011,f2794])).
fof(f4011,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X17),s('type.semanticPrimitives.abort',sK7))) | s(bool,X15) = s(bool,X18) | s(bool,X18) = s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK6))) | ~p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl50),
  inference(superposition,[],[f2702,f84])).
fof(f4275,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X11) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X12),s('type.semanticPrimitives.abort',sK7))) | s(bool,X10) = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK6))) | ~p(s(bool,X11))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f4010,f2794])).
fof(f4010,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X11) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X12),s('type.semanticPrimitives.abort',sK7))) | s(bool,X10) = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK6))) | ~p(s(bool,X10)) | ~p(s(bool,X11))) ) | $spl50),
  inference(superposition,[],[f2702,f84])).
fof(f4007,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s(bool,X59) = s(bool,X60) | s(bool,X60) = s(bool,X61) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X62),s('type.semanticPrimitives.abort',sK7))) | s(bool,X59) = s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X64),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4004,plain,(
  ( ! [X39,X37,X43,X41,X38,X42,X40] : (s(bool,X39) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X38),s('type.semanticPrimitives.abort',sK7))) | s(bool,X39) = s(bool,X40) | s(bool,X40) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X41),s('type.semanticPrimitives.abort',sK7))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X37),s(sK0,sK6))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3974,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s(bool,X66) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X68) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) | s(bool,X68) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3970,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK6))) | s(bool,X41) = s(bool,X44) | s(bool,X44) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X45),s('type.semanticPrimitives.abort',sK7))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK6))) | s(bool,X41) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X47),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3967,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X22),s(sK0,sK6))) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X24),s('type.semanticPrimitives.abort',sK7))) | s(bool,X20) = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4272,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | s(bool,X5) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X5))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3964,f1349])).
fof(f3964,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | s(bool,X5) = s(bool,X8) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f4271,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X0))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3963,f1349])).
fof(f3963,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f3951,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3950,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,X62) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK6))) | s(bool,X62) = s(bool,happ(s(fun(sK0,bool),X63),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X64),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X60),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2702,f2797])).
fof(f3949,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X57) = s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK6))) | s(bool,X57) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X59),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))) ) | ($spl16 | $spl50)),
  inference(superposition,[],[f2702,f2797])).
fof(f3948,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,X49) = s(bool,X50) | s(bool,X50) = s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK6))) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X52),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X48),s(sK0,sK6))) | s(bool,X49) = s(bool,X53) | s(bool,X53) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X54),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3947,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(sK0,bool),X44),s(sK0,sK6))) | s(bool,X41) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X45),s('type.semanticPrimitives.abort',sK7))) | s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK6))) | s(bool,X41) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X47),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3945,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X30),s(sK0,sK6))) | s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X31),s('type.semanticPrimitives.abort',sK7))) | s(bool,X27) = s(bool,X32) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4268,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6))) | s(bool,X15) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X19),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X15))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3943,f2794])).
fof(f3943,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6))) | s(bool,X15) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X19),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl50),
  inference(superposition,[],[f2702,f84])).
fof(f4267,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X11) = s(bool,X12) | s(bool,X12) = s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6))) | s(bool,X10) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X10))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3942,f2794])).
fof(f3942,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,X11) = s(bool,X12) | s(bool,X12) = s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6))) | s(bool,X10) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X10)) | ~p(s(bool,X11))) ) | $spl50),
  inference(superposition,[],[f2702,f84])).
fof(f3928,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s(bool,X66) = s(bool,X67) | s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X65),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))) | s(bool,X66) = s(bool,X70) | s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3923,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(bool,X35) = s(bool,X36) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X37),s(sK0,sK6))) | s(bool,X36) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X38),s('type.semanticPrimitives.abort',sK7))) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK6))) | s(bool,X35) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X40),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f3921,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) | s(bool,X20) = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK6))) | s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X24),s('type.semanticPrimitives.abort',sK7))) | s(bool,X20) = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK6)))) ) | $spl50),
  inference(superposition,[],[f2702,f2702])).
fof(f4264,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK6))) | s(bool,X7) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X6))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3918,f1349])).
fof(f3918,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK6))) | s(bool,X7) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X9),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X5)) | p(s(bool,X6))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f4262,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1))) ) | ($spl16 | $spl50)),
  inference(subsumption_resolution,[],[f3917,f1349])).
fof(f3917,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl50),
  inference(superposition,[],[f2702,f81])).
fof(f3916,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6)))) ) | $spl58),
  inference(cnf_transformation,[],[f3916_D])).
fof(f3916_D,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | ~p(s(bool,X9)) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3909,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,X14) | s(bool,X14) = s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK6))) | ~p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X12),s('type.semanticPrimitives.abort',sK7))))) ) | ($spl16 | $spl54)),
  inference(superposition,[],[f3627,f84])).
fof(f3627,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X6),s('type.semanticPrimitives.abort',sK7))) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6)))) ) | ($spl16 | $spl54)),
  inference(resolution,[],[f3614,f1349])).
fof(f3908,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK6))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X8),s('type.semanticPrimitives.abort',sK7)))) | ~p(s(bool,X9))) ) | ($spl16 | $spl54)),
  inference(superposition,[],[f3627,f84])).
fof(f3915,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X6) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X5))) ) | ($spl16 | $spl54)),
  inference(subsumption_resolution,[],[f3907,f2794])).
fof(f3907,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | s(bool,X6) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X5))) ) | ($spl16 | $spl54)),
  inference(superposition,[],[f3627,f81])).
fof(f3913,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7))))) ) | ($spl16 | $spl54)),
  inference(subsumption_resolution,[],[f3906,f2794])).
fof(f3906,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X2) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7))))) ) | ($spl16 | $spl54)),
  inference(superposition,[],[f3627,f81])).
fof(f1314,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7)))) ) | $spl34),
  inference(cnf_transformation,[],[f1314_D])).
fof(f1314_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3873,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X5))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3823,f537])).
fof(f537,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X1),s('type.semanticPrimitives.abort',sK7)))) ) | $spl22),
  inference(resolution,[],[f487,f84])).
fof(f3823,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f3601,f81])).
fof(f3601,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) != s(bool,happ(s(fun(sK0,bool),X13),s(sK0,sK6))) | s(bool,X14) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X15),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f3518,f2059])).
fof(f2059,plain,(
  ( ! [X12,X13] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X13),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK6)))) ) | ($spl16 | $spl22)),
  inference(resolution,[],[f537,f477])).
fof(f3518,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7)))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f3516])).
fof(f3516,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X0),s('type.semanticPrimitives.abort',sK7)))) ) | $spl22),
  inference(equality_factoring,[],[f2191])).
fof(f2191,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X1) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X3),s('type.semanticPrimitives.abort',sK7)))) ) | $spl22),
  inference(resolution,[],[f2056,f537])).
fof(f2056,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X6)) | s(bool,X4) = s(bool,X6) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7)))) ) | $spl22),
  inference(resolution,[],[f537,f81])).
fof(f3872,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X3),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3822,f537])).
fof(f3822,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X3),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f3601,f81])).
fof(f3614,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7)))) ) | $spl54),
  inference(cnf_transformation,[],[f3614_D])).
fof(f3614_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3615,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X6),s('type.semanticPrimitives.abort',sK7))) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X5))) ) | $spl22),
  inference(subsumption_resolution,[],[f3562,f537])).
fof(f3562,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X6),s('type.semanticPrimitives.abort',sK7))) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X7),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl22),
  inference(superposition,[],[f3518,f81])).
fof(f3613,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X3),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1))) ) | $spl22),
  inference(subsumption_resolution,[],[f3561,f537])).
fof(f3561,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X3),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl22),
  inference(superposition,[],[f3518,f81])).
fof(f2048,plain,(
  ( ! [X94,X93,X89] : (s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | $spl48),
  inference(cnf_transformation,[],[f2048_D])).
fof(f2048_D,plain,(
  ( ! [X94,X93,X89] : (s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2702,plain,(
  ( ! [X66,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7)))) ) | $spl50),
  inference(cnf_transformation,[],[f2702_D])).
fof(f2702_D,plain,(
  ( ! [X66,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3558,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X88),s('type.semanticPrimitives.abort',sK7))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3492,f2056])).
fof(f3492,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (~p(s(bool,X89)) | s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X88),s('type.semanticPrimitives.abort',sK7))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f3046,f2191])).
fof(f3551,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X88),s('type.semanticPrimitives.abort',sK7))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3426,f2056])).
fof(f3426,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (~p(s(bool,X89)) | s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X88),s('type.semanticPrimitives.abort',sK7))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f3046,f2191])).
fof(f3542,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK6))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X67),s('type.semanticPrimitives.abort',sK7))) | s(bool,X63) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X68),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3327,f537])).
fof(f3327,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK6))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X67),s('type.semanticPrimitives.abort',sK7))) | s(bool,X63) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X68),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2794,f2191])).
fof(f3538,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X67),s('type.semanticPrimitives.abort',sK7))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X68),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(subsumption_resolution,[],[f3280,f537])).
fof(f3280,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X67),s('type.semanticPrimitives.abort',sK7))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X68),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl16 | $spl22)),
  inference(superposition,[],[f2794,f2191])).
fof(f2782,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,X5))) ) | $spl28),
  inference(subsumption_resolution,[],[f2729,f494])).
fof(f2729,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))) | s(bool,X4) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl28),
  inference(superposition,[],[f2681,f81])).
fof(f2681,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6)))) ) | $spl28),
  inference(duplicate_literal_removal,[],[f2679])).
fof(f2679,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6)))) ) | $spl28),
  inference(equality_factoring,[],[f1469])).
fof(f1469,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6)))) ) | $spl28),
  inference(resolution,[],[f1161,f494])).
fof(f1161,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) = s(bool,X2) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X1),s(sK0,sK6)))) ) | $spl28),
  inference(resolution,[],[f494,f81])).
fof(f2780,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1))) ) | $spl28),
  inference(subsumption_resolution,[],[f2728,f494])).
fof(f2728,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK6))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl28),
  inference(superposition,[],[f2681,f81])).
fof(f2725,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X90) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X91),s('type.semanticPrimitives.abort',sK7))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2655,f1161])).
fof(f2655,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (~p(s(bool,X89)) | s(bool,X90) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X91),s('type.semanticPrimitives.abort',sK7))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f2194,f1469])).
fof(f2194,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X17)) | s(bool,X14) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X16),s('type.semanticPrimitives.abort',sK7))) | s(bool,X15) = s(bool,X17) | s(bool,X14) = s(bool,X15)) ) | $spl22),
  inference(resolution,[],[f2056,f84])).
fof(f2717,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X90) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X91),s('type.semanticPrimitives.abort',sK7))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2589,f1161])).
fof(f2589,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (~p(s(bool,X89)) | s(bool,X90) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X91),s('type.semanticPrimitives.abort',sK7))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f2194,f1469])).
fof(f2707,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2490,f494])).
fof(f2490,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f2056,f1469])).
fof(f2701,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(subsumption_resolution,[],[f2443,f494])).
fof(f2443,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK6))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f2056,f1469])).
fof(f2047,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl28)),
  inference(subsumption_resolution,[],[f1968,f1106])).
fof(f1106,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) = s(bool,X2) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X1),s('type.semanticPrimitives.abort',sK7)))) ) | $spl18),
  inference(resolution,[],[f480,f81])).
fof(f1968,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (~p(s(bool,X89)) | s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK6))) | s(bool,X92) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK6))) | s(bool,X90) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,X93) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X94),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1471,f1439])).
fof(f1439,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X2),s('type.semanticPrimitives.abort',sK7))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(resolution,[],[f1106,f494])).
fof(f1471,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X11)) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),X10),s(sK0,sK6))) | s(bool,X9) = s(bool,X11) | s(bool,X8) = s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f1161,f84])).
fof(f2037,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,X83) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X84),s('type.semanticPrimitives.abort',sK7))) | s(bool,X85) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7))) | s(bool,X83) = s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(subsumption_resolution,[],[f1901,f1161])).
fof(f1901,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X82)) | s(bool,X83) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X84),s('type.semanticPrimitives.abort',sK7))) | s(bool,X85) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7))) | s(bool,X83) = s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,X86) = s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1441,f1439])).
fof(f1441,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,X11)) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X10),s('type.semanticPrimitives.abort',sK7))) | s(bool,X9) = s(bool,X11) | s(bool,X8) = s(bool,X9)) ) | $spl18),
  inference(resolution,[],[f1106,f84])).
fof(f2026,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6))) | s(bool,X57) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(subsumption_resolution,[],[f1802,f480])).
fof(f1802,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X57)) | s(bool,X56) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6))) | s(bool,X57) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7))) | s(bool,X56) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1161,f1439])).
fof(f2018,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s(bool,X50) = s(bool,X52) | s(bool,X52) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X53),s('type.semanticPrimitives.abort',sK7))) | s(bool,X50) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X54),s('type.semanticPrimitives.abort',sK7))) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(subsumption_resolution,[],[f1754,f494])).
fof(f1754,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (p(s(bool,X51)) | s(bool,X50) = s(bool,X52) | s(bool,X52) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X53),s('type.semanticPrimitives.abort',sK7))) | s(bool,X50) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X54),s('type.semanticPrimitives.abort',sK7))) | s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X55),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1106,f1439])).
fof(f1740,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,X44) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X43),s('type.semanticPrimitives.abort',sK7))) | s(bool,X44) = s(bool,X45) | s(bool,X45) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X46),s('type.semanticPrimitives.abort',sK7))) | s(bool,X47) = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK6))) | s(bool,X47) = s(bool,happ(s(fun(sK0,bool),X48),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1702,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(bool,X49) = s(bool,happ(s(fun(sK0,bool),X50),s(sK0,sK6))) | s(bool,X51) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X48),s('type.semanticPrimitives.abort',sK7))) | s(bool,X51) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X52),s('type.semanticPrimitives.abort',sK7))) | s(bool,X49) = s(bool,X53) | s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X54),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1700,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK6))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X38),s('type.semanticPrimitives.abort',sK7))) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7))) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1691,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s(bool,X51) = s(bool,happ(s(fun(sK0,bool),X50),s(sK0,sK6))) | s(bool,X51) = s(bool,X52) | s(bool,X52) = s(bool,happ(s(fun(sK0,bool),X53),s(sK0,sK6))) | s(bool,X54) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X49),s('type.semanticPrimitives.abort',sK7))) | s(bool,X54) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X55),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1656,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,X71) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X72),s('type.semanticPrimitives.abort',sK7))) | s(bool,X73) = s(bool,happ(s(fun(sK0,bool),X70),s(sK0,sK6))) | s(bool,X73) = s(bool,happ(s(fun(sK0,bool),X74),s(sK0,sK6))) | s(bool,X71) = s(bool,X75) | s(bool,X75) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X76),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1651,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X43),s('type.semanticPrimitives.abort',sK7))) | s(bool,X41) = s(bool,X44) | s(bool,X44) = s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK6))) | s(bool,X42) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X46),s('type.semanticPrimitives.abort',sK7))) | s(bool,X41) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1624,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X44),s('type.semanticPrimitives.abort',sK7))) | s(bool,X41) = s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK6))) | s(bool,X42) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X46),s('type.semanticPrimitives.abort',sK7))) | s(bool,X41) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1621,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X21),s('type.semanticPrimitives.abort',sK7))) | s(bool,X22) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X23),s('type.semanticPrimitives.abort',sK7))) | s(bool,X20) = s(bool,happ(s(fun(sK0,bool),X24),s(sK0,sK6))) | s(bool,X20) = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1596,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(bool,X35) = s(bool,X36) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X37),s('type.semanticPrimitives.abort',sK7))) | s(bool,X36) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK6))) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X39),s('type.semanticPrimitives.abort',sK7))) | s(bool,X35) = s(bool,happ(s(fun(sK0,bool),X40),s(sK0,sK6)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f1595,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK6))) | s(bool,X27) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X30),s('type.semanticPrimitives.abort',sK7))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK6))) | s(bool,X27) = s(bool,X32) | s(bool,X32) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X33),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl28)),
  inference(superposition,[],[f1439,f1439])).
fof(f477,plain,(
  ( ! [X59] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6))))) ) | $spl16),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  ( ! [X59] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1321,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1301,f480])).
fof(f1301,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X4),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6)))) | p(s(bool,X5))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f81,f1163])).
fof(f1163,plain,(
  ( ! [X6,X7] : (s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X6),s('type.semanticPrimitives.abort',sK7))) = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK6)))) ) | ($spl22 | $spl28)),
  inference(resolution,[],[f494,f487])).
fof(f1320,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X1),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6))))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1300,f480])).
fof(f1300,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X1),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK6))))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f81,f1163])).
fof(f1319,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6)))) | s(bool,X19) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X17),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1292,f480])).
fof(f1292,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK6)))) | s(bool,X19) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X17),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X19))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f81,f1163])).
fof(f1318,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK6)))) | s(bool,X16) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1291,f480])).
fof(f1291,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK6)))) | s(bool,X16) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X14),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X16))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f81,f1163])).
fof(f1315,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X8),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6))))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1283,f480])).
fof(f1283,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X8),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X6),s(sK0,sK6)))) | p(s(bool,X7))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1163,f81])).
fof(f1313,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))))) ) | ($spl18 | $spl22 | $spl28)),
  inference(subsumption_resolution,[],[f1282,f480])).
fof(f1282,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X5),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X4)) | p(s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK6))))) ) | ($spl22 | $spl28)),
  inference(superposition,[],[f1163,f81])).
fof(f1083,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | ($spl14 | $spl18)),
  inference(duplicate_literal_removal,[],[f944])).
fof(f944,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool)))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X0,bool))))) ) | ($spl14 | $spl18)),
  inference(superposition,[],[f870,f870])).
fof(f870,plain,(
  ( ! [X222,X223] : (s('type.semanticPrimitives.error_result'(bool),X222) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X223))) | s('type.semanticPrimitives.error_result'(bool),X222) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK11(X222,bool))))) ) | ($spl14 | $spl18)),
  inference(superposition,[],[f78,f617])).
fof(f617,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2)) ) | ($spl14 | $spl18)),
  inference(superposition,[],[f584,f584])).
fof(f584,plain,(
  ( ! [X66,X67] : (s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7)))) ) | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f480,f476])).
fof(f929,plain,(
  $false | ($spl0 | $spl9 | $spl12)),
  inference(subsumption_resolution,[],[f685,f93])).
fof(f93,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | $spl9),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f685,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | ($spl0 | $spl12)),
  inference(superposition,[],[f87,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun('type.semanticPrimitives.abort',X0),X1),s('type.semanticPrimitives.abort',sK7))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,X0),X2),s(fun('type.semanticPrimitives.abort',X0),X1)))) ) | $spl12),
  inference(superposition,[],[f77,f98])).
fof(f928,plain,(
  $false | ($spl2 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f923,f98])).
fof(f923,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f922])).
fof(f922,plain,(
  s(sK1,sK3) != s(sK1,sK3) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) | ($spl2 | $spl8)),
  inference(superposition,[],[f89,f94])).
fof(f707,plain,(
  $false | ($spl0 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f696,f91])).
fof(f91,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | $spl7),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f696,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | ($spl0 | $spl10)),
  inference(superposition,[],[f99,f87])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK6))) = s(X0,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,X0),X1),s(fun('type.semanticPrimitives.abort',X0),X2)))) ) | $spl10),
  inference(superposition,[],[f80,f96])).
fof(f706,plain,(
  $false | ($spl0 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f686,f91])).
fof(f686,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | ($spl0 | $spl10)),
  inference(superposition,[],[f87,f99])).
fof(f558,plain,(
  $false | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f493,f487])).
fof(f536,plain,(
  $false | ($spl1 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f473,f94])).
fof(f473,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | ($spl1 | $spl12)),
  inference(superposition,[],[f88,f312])).
fof(f494,plain,(
  ( ! [X80,X82] : (~p(s(bool,X82)) | s(bool,X82) = s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK6)))) ) | $spl28),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  ( ! [X80,X82] : (~p(s(bool,X82)) | s(bool,X82) = s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK6)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f495,plain,(
  ( ! [X85,X83,X84] : (s(bool,X85) = s(bool,happ(s(fun(sK0,bool),X83),s(sK0,sK6))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X84),s('type.semanticPrimitives.abort',sK7)))) | ~p(s(bool,X85))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f420,f99])).
fof(f420,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X84),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X85) = s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X83),s(fun('type.semanticPrimitives.abort',bool),X84))) | ~p(s(bool,X85))) ) | $spl12),
  inference(superposition,[],[f84,f312])).
fof(f492,plain,(
  ( ! [X80,X81,X82] : (s(bool,X82) = s(bool,happ(s(fun(sK0,bool),X80),s(sK0,sK6))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7)))) | ~p(s(bool,X82))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f419,f99])).
fof(f419,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X81),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X82) = s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X80),s(fun('type.semanticPrimitives.abort',bool),X81))) | ~p(s(bool,X82))) ) | $spl12),
  inference(superposition,[],[f84,f312])).
fof(f481,plain,(
  ( ! [X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6))))) ) | $spl20),
  inference(cnf_transformation,[],[f481_D])).
fof(f481_D,plain,(
  ( ! [X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f491,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,happ(s(fun(sK0,bool),X77),s(sK0,sK6)))) | s(bool,X79) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X78),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X79))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f418,f99])).
fof(f418,plain,(
  ( ! [X78,X79,X77] : (s(bool,X79) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X78),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X79)) | ~p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X77),s(fun('type.semanticPrimitives.abort',bool),X78))))) ) | $spl12),
  inference(superposition,[],[f84,f312])).
fof(f490,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(sK0,bool),X74),s(sK0,sK6)))) | s(bool,X76) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X75),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X76))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f417,f99])).
fof(f417,plain,(
  ( ! [X76,X74,X75] : (s(bool,X76) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X75),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X74),s(fun('type.semanticPrimitives.abort',bool),X75)))) | ~p(s(bool,X76))) ) | $spl12),
  inference(superposition,[],[f84,f312])).
fof(f488,plain,(
  ( ! [X70,X68] : (s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | p(s(bool,X70))) ) | $spl24),
  inference(cnf_transformation,[],[f488_D])).
fof(f488_D,plain,(
  ( ! [X70,X68] : (s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | p(s(bool,X70))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f489,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X72),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X73))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f416,f99])).
fof(f416,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X72),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X73) = s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X71),s(fun('type.semanticPrimitives.abort',bool),X72))) | p(s(bool,X73))) ) | $spl12),
  inference(superposition,[],[f81,f312])).
fof(f486,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) = s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6))) | p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) | p(s(bool,X70))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f415,f99])).
fof(f415,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7)))) | s(bool,X70) = s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X68),s(fun('type.semanticPrimitives.abort',bool),X69))) | p(s(bool,X70))) ) | $spl12),
  inference(superposition,[],[f81,f312])).
fof(f485,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6)))) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X67))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f414,f99])).
fof(f414,plain,(
  ( ! [X66,X67,X65] : (s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X65),s(fun('type.semanticPrimitives.abort',bool),X66)))) | p(s(bool,X67))) ) | $spl12),
  inference(superposition,[],[f81,f312])).
fof(f484,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK6)))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X63),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X64))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f413,f99])).
fof(f413,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X63),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X64)) | p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X62),s(fun('type.semanticPrimitives.abort',bool),X63))))) ) | $spl12),
  inference(superposition,[],[f81,f312])).
fof(f482,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,happ(s(fun(sK0,bool),X68),s(sK0,sK6)))) | s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X70))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f401,f99])).
fof(f401,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X69),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X68),s(fun('type.semanticPrimitives.abort',bool),X69)))) | ~p(s(bool,X70))) ) | $spl12),
  inference(superposition,[],[f312,f84])).
fof(f479,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,happ(s(fun(sK0,bool),X65),s(sK0,sK6)))) | s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X67))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f400,f99])).
fof(f400,plain,(
  ( ! [X66,X67,X65] : (s(bool,X67) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X66),s('type.semanticPrimitives.abort',sK7))) | ~p(s(bool,X67)) | ~p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X65),s(fun('type.semanticPrimitives.abort',bool),X66))))) ) | $spl12),
  inference(superposition,[],[f312,f84])).
fof(f478,plain,(
  ( ! [X64,X62,X63] : (p(s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK6)))) | s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X63),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X64))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f399,f99])).
fof(f399,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X63),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X64)) | p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X62),s(fun('type.semanticPrimitives.abort',bool),X63))))) ) | $spl12),
  inference(superposition,[],[f312,f81])).
fof(f475,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,happ(s(fun(sK0,bool),X59),s(sK0,sK6)))) | s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,X61))) ) | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f398,f99])).
fof(f398,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) = s(bool,happ(s(fun('type.semanticPrimitives.abort',bool),X60),s('type.semanticPrimitives.abort',sK7))) | p(s(bool,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,bool),X59),s(fun('type.semanticPrimitives.abort',bool),X60)))) | p(s(bool,X61))) ) | $spl12),
  inference(superposition,[],[f312,f81])).
fof(f311,plain,(
  $false | ($spl1 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f256,f92])).
fof(f256,plain,(
  s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | ($spl1 | $spl10)),
  inference(superposition,[],[f88,f99])).
fof(f87,plain,(
  s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | $spl0),
  inference(cnf_transformation,[],[f87_D])).
fof(f87_D,plain,(
  s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f96,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK6))) | $spl10),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK6))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f98,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) | $spl12),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f67,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK6))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  (s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK6))) & s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6)))) | (s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) & s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))))) & (s(sK1,sK3) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | (! [X7] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,X7))) | s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,X7)))) & ! [X6] : (s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | ? [X9] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X9))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X9)))) | ? [X8] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X8))))) & (s(X1,X3) != s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | (! [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) | s(X1,X3) != s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) & ! [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s(X1,X3) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6)))))))),
  inference(rectify,[],[f62])).
fof(f62,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | ? [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6))))) & (s(X1,X3) != s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | (! [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) | s(X1,X3) != s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) & ! [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s(X1,X3) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6)))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | (? [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6)))))) & (s(X1,X3) != s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) | (! [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) | s(X1,X3) != s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) & ! [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s(X1,X3) != s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6)))))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) <~> (? [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) <=> (? [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s(X1,X3) = s(X1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X0),X2),s(fun(X0,X1),X5),s(fun('type.semanticPrimitives.abort',X1),X4))) <=> (? [X7] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X7))) & s(X1,X3) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X7)))) | ? [X6] : (s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) & s(X1,X3) = s(X1,happ(s(fun('type.semanticPrimitives.abort',X1),X4),s('type.semanticPrimitives.abort',X6))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s(X7,X9) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X8),s(fun(X6,X7),X11),s(fun('type.semanticPrimitives.abort',X7),X10))) <=> (? [X12] : (s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X12))) & s(X7,X9) = s(X7,happ(s(fun(X6,X7),X11),s(X6,X12)))) | ? [X12] : (s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X12))) & s(X7,X9) = s(X7,happ(s(fun('type.semanticPrimitives.abort',X7),X10),s('type.semanticPrimitives.abort',X12))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s(X7,X9) = s(X7,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(X6),X8),s(fun(X6,X7),X11),s(fun('type.semanticPrimitives.abort',X7),X10))) <=> (? [X12] : (s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X12))) & s(X7,X9) = s(X7,happ(s(fun(X6,X7),X11),s(X6,X12)))) | ? [X12] : (s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X12))) & s(X7,X9) = s(X7,happ(s(fun('type.semanticPrimitives.abort',X7),X10),s('type.semanticPrimitives.abort',X12))))))),
  file('thm.closProps.eqs|split|7','thm.closProps.eqs|split|7')).
fof(f92,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | $spl6),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f68,plain,(
  s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK7))) | s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))),
  inference(cnf_transformation,[],[f64])).
fof(f94,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | $spl8),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f69,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | s('type.semanticPrimitives.error_result'(sK0),sK2) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,sK6))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))),
  inference(cnf_transformation,[],[f64])).
fof(f70,plain,(
  s(sK1,sK3) = s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',sK7))) | s(sK1,sK3) = s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,sK6))) | s(sK1,sK3) = s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))),
  inference(cnf_transformation,[],[f64])).
fof(f88,plain,(
  s(sK1,sK3) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) | $spl1),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  s(sK1,sK3) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f90,plain,(
  ( ! [X7] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,X7))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,X7)))) ) | $spl4),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  ( ! [X7] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,X7))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,X7)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f71,plain,(
  ( ! [X7] : (s(sK1,sK3) != s(sK1,happ(s(fun(sK0,sK1),sK5),s(sK0,X7))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rraise_1'(s(sK0,X7))) | s(sK1,sK3) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f89,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6)))) ) | $spl2),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f72,plain,(
  ( ! [X6] : (s(sK1,sK3) != s(sK1,happ(s(fun('type.semanticPrimitives.abort',sK1),sK4),s('type.semanticPrimitives.abort',X6))) | s('type.semanticPrimitives.error_result'(sK0),sK2) != s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s(sK1,sK3) != s(sK1,'const.semanticPrimitives.error_result_CASE_3'(s('type.semanticPrimitives.error_result'(sK0),sK2),s(fun(sK0,sK1),sK5),s(fun('type.semanticPrimitives.abort',sK1),sK4)))) )),
  inference(cnf_transformation,[],[f64])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 198
% Passive clauses: 7887
% Generated clauses: 22715
% Final active clauses: 91
% Final passive clauses: 2989
% Input formulas: 37
% Initial clauses: 20
% 
% Selected by SInE selection: 9
% SInE iterations: 1
% 
% Duplicate literals: 114
% Trivial inequalities: 43
% Fw subsumption resolutions: 849
% Fw demodulations: 4734
% Bw demodulations: 2
% 
% Simple tautologies: 131
% Equational tautologies: 31
% Forward subsumptions: 8870
% 
% Binary resolution: 218
% Factoring: 6
% Forward superposition: 7962
% Backward superposition: 7958
% Self superposition: 389
% Equality factoring: 118
% Equality resolution: 42
% 
% Split clauses: 188
% Split components: 382
% SAT solver clauses: 199
% SAT solver unit clauses: 1
% SAT solver binary clauses: 3
% 
% Sat splits: 188
% Sat splitting refutations: 10
% 
% Memory used [KB]: 23539
% Time elapsed: 1.017 s
% ------------------------------
% ------------------------------
% Success in time 12.79 s
