Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.okasaki_misc
% (11330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 2767
% Generated clauses: 13073
% Final active clauses: 125
% Final passive clauses: 2431
% Input formulas: 27
% Initial clauses: 67
% 
% Duplicate literals: 538
% Trivial inequalities: 280
% Fw subsumption resolutions: 1963
% Fw demodulations: 1216
% Bw demodulations: 7
% 
% Simple tautologies: 347
% Equational tautologies: 48
% Forward subsumptions: 5129
% 
% Binary resolution: 719
% Factoring: 2
% Forward superposition: 4655
% Backward superposition: 3346
% Self superposition: 125
% Equality factoring: 104
% Equality resolution: 35
% 
% Split clauses: 517
% Split components: 1055
% SAT solver clauses: 986
% SAT solver unit clauses: 39
% SAT solver binary clauses: 256
% 
% Sat splits: 517
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 986
% TWLsolver calls for satisfiability: 227
% 
% Memory used [KB]: 8571
% Time elapsed: 0.620 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.okasaki_misc
% (11357)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 1054
% Generated clauses: 11741
% Final active clauses: 91
% Final passive clauses: 765
% Input formulas: 27
% Initial clauses: 67
% 
% Duplicate literals: 687
% Trivial inequalities: 108
% Fw subsumption resolutions: 1068
% Fw demodulations: 1031
% Bw demodulations: 158
% Condensations: 1124
% Global subsumptions: 5
% 
% Simple tautologies: 167
% Equational tautologies: 46
% Forward subsumptions: 4848
% Bw demodulations to eq. taut.: 25
% 
% Binary resolution: 463
% Factoring: 1
% Forward superposition: 2365
% Backward superposition: 3271
% Self superposition: 120
% Equality factoring: 109
% Equality resolution: 17
% 
% Split clauses: 102
% Split components: 210
% SAT solver clauses: 1596
% SAT solver unit clauses: 62
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 102
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1575
% TWLsolver calls for satisfiability: 352
% 
% Memory used [KB]: 9466
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.okasaki_misc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.okasaki_misc
% SZS output start Proof for thm.okasaki_misc
fof(f12934,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f218,f228,f226,f2595,f2579,f2563,f2623,f2558,f2557,f2548,f2604,f2546,f2624,f2636,f2638,f2637,f2639,f2640,f2641,f2643,f2642,f2654,f2655,f2660,f2661,f2662,f2663,f2664,f2665,f2513,f2695,f2512,f2696,f2510,f2697,f2918,f2920,f3866,f3867,f6269,f6271,f6273,f6278,f6279,f6280,f6281,f6289,f6290,f6291,f6298,f6327,f6325,f7132,f7141,f7133,f7134,f7135,f7145,f7138,f7136,f7144,f7139,f7142,f8604,f8607,f8624,f8622,f10018,f10021,f10064,f10038,f10065,f10070,f10889,f10890,f2603,f11189,f11194,f11197,f11202,f11203,f11206,f11209,f11211,f11213,f11214,f11215,f11216,f11217,f11218,f11219,f11220,f11221,f11222,f11223,f11224,f11225,f11226,f11227,f11228,f11229,f11230,f11231,f11232,f11233,f11234,f11235,f11236,f11237,f11238,f11240,f11241,f11242,f11243,f11244,f11245,f11246,f11247,f11251,f11252,f11253,f11254,f11258,f11259,f11260,f11261,f11263,f11265,f11266,f11267,f11268,f11269,f11270,f11271,f11272,f11273,f11274,f11275,f11276,f11277,f11278,f11279,f11280,f11281,f12915])).
fof(f12915,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f12914,f112])).
fof(f112,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.okasaki_misc.BAG_EVERY_EL',aHL_FALSITY)).
fof(f12914,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f12913,f11372])).
fof(f11372,plain,(
  s(bool,'F_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f112,f225,f237])).
fof(f237,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f234,f112])).
fof(f234,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f153,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.okasaki_misc.BAG_EVERY_EL',aHL_BOOL_CASES)).
fof(f153,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.IMP_ANTISYM_AX')).
fof(f225,plain,(
  s(bool,'T_0') != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl1),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  s(bool,'T_0') != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f12913,plain,(
  p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f12912,f110])).
fof(f110,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  p(s(bool,'T_0'))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.TRUTH')).
fof(f12912,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | ($spl1 | $spl40)),
  inference(forward_demodulation,[],[f12873,f11308])).
fof(f11308,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl40),
  inference(unit_resulting_resolution,[],[f110,f6326,f156])).
fof(f156,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f70])).
fof(f6326,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl40),
  inference(cnf_transformation,[],[f6326_D])).
fof(f6326_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f12873,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | $spl1),
  inference(superposition,[],[f146,f12743])).
fof(f12743,plain,(
  s(sK0,sK2) = s(sK0,sK3('const.bag.EL_BAG_1'(s(sK0,sK2)),sK1,sK0)) | $spl1),
  inference(subsumption_resolution,[],[f12668,f110])).
fof(f12668,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK2) = s(sK0,sK3('const.bag.EL_BAG_1'(s(sK0,sK2)),sK1,sK0)) | $spl1),
  inference(superposition,[],[f1238,f11744])).
fof(f11744,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(sK0,sK3('const.bag.EL_BAG_1'(s(sK0,sK2)),sK1,sK0)),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl1),
  inference(subsumption_resolution,[],[f11656,f112])).
fof(f11656,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(sK0,sK3('const.bag.EL_BAG_1'(s(sK0,sK2)),sK1,sK0)),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl1),
  inference(superposition,[],[f532,f11372])).
fof(f532,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X2,bool),X4),s(fun(X2,'type.num.num'),X3)))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X2,sK3(X3,X4,X2)),s(fun(X2,'type.num.num'),X3)))) )),
  inference(subsumption_resolution,[],[f529,f112])).
fof(f529,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X2,bool),X4),s(fun(X2,'type.num.num'),X3)))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X2,sK3(X3,X4,X2)),s(fun(X2,'type.num.num'),X3)))) )),
  inference(superposition,[],[f145,f113])).
fof(f145,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) | ! [X4] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X4),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bag.BAG_IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,'type.num.num'),X2)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) | ! [X4] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X4),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2))))))),
  inference(rectify,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) | ! [X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2))))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) <=> ! [X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) <=> ! [X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) <=> ! [X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X3),s(fun(X0,'type.num.num'),X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X9] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X6,bool),X7),s(fun(X6,'type.num.num'),X9)))) <=> ! [X12] : (p(s(bool,'const.bag.BAG_IN_2'(s(X6,X12),s(fun(X6,'type.num.num'),X9)))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12))))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bag.BAG_EVERY')).
fof(f1238,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(subsumption_resolution,[],[f1237,f112])).
fof(f1237,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f1201,f197])).
fof(f197,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(unit_resulting_resolution,[],[f141,f120])).
fof(f120,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.EQ_CLAUSES|split|2')).
fof(f141,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X8] : ~p(s(bool,'const.bag.BAG_IN_2'(s(X6,X8),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0'))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bag.NOT_IN_EMPTY_BAG')).
fof(f1201,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) )),
  inference(superposition,[],[f150,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))) = s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))) = s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1] : s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))) = s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12] : s(fun(X6,'type.num.num'),'const.bag.EL_BAG_1'(s(X6,X12))) = s(fun(X6,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X6,X12),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0')))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bag.EL_BAG')).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1))))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) | (s(X0,X2) = s(X0,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1))))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9,X10,X11] : (p(s(bool,'const.bag.BAG_IN_2'(s(X6,X10),s(fun(X6,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X6,X11),s(fun(X6,'type.num.num'),X9)))))) <=> (s(X6,X10) = s(X6,X11) | p(s(bool,'const.bag.BAG_IN_2'(s(X6,X10),s(fun(X6,'type.num.num'),X9))))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bag.BAG_IN_BAG_INSERT')).
fof(f146,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f11281,plain,(
  $false | ($spl1 | $spl3 | $spl41)),
  inference(global_subsumption,[],[f225,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663,f228])).
fof(f4663,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0'))))))),
  inference(forward_demodulation,[],[f4613,f197])).
fof(f4613,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f197,f467])).
fof(f109,plain,(
  s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X2)))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X2)))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X6,bool),X7),s(fun(X6,'type.num.num'),'const.bag.EL_BAG_1'(s(X6,X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X6,bool),X7),s(fun(X6,'type.num.num'),'const.bag.EL_BAG_1'(s(X6,X8)))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.okasaki_misc.BAG_EVERY_EL')).
fof(f10752,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),X6)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f10718,f110])).
fof(f10718,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),X6)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1189,f202])).
fof(f10678,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X73),s(fun(bool,'type.num.num'),X74)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X70,bool),X71),s(X70,X72))),s(fun(bool,'type.num.num'),X74)))) | s(bool,'T_0') = s(bool,X73) | s(X70,X72) = s(X70,sK4(X72,X71,X70))) )),
  inference(superposition,[],[f1189,f400])).
fof(f10657,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X4),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'F_0') = s(bool,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),X5)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1189,f202])).
fof(f10581,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f10566,f110])).
fof(f10566,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1188,f202])).
fof(f10575,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X2)))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f10574])).
fof(f10574,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X26),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X27)))))) | s(bool,'T_0') = s(bool,X28) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X27)))) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(duplicate_literal_removal,[],[f10535])).
fof(f10535,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X26),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X27)))))) | s(bool,'T_0') = s(bool,X28) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X27)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1188,f202])).
fof(f10579,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f10578])).
fof(f10578,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f10529])).
fof(f10529,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X4)))))) | s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1188,f202])).
fof(f10075,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X6)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f10049,f110])).
fof(f10049,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X6)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1176,f202])).
fof(f10066,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X26),s(fun(bool,'type.num.num'),X28)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X28)))) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f10024,f190])).
fof(f10024,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X26),s(fun(bool,'type.num.num'),X28)))))) | s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X28)))) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1176,f202])).
fof(f9549,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f9534,f110])).
fof(f9534,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1175,f202])).
fof(f9541,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X26),s(fun(bool,'type.num.num'),X27)))))) | s(bool,'T_0') = s(bool,X28) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),X27)))) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(duplicate_literal_removal,[],[f9512])).
fof(f9512,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X26),s(fun(bool,'type.num.num'),X27)))))) | s(bool,'T_0') = s(bool,X28) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),X27)))) | s(bool,'T_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1175,f202])).
fof(f9547,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1))))) )),
  inference(condensation,[],[f9546])).
fof(f9546,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f9545])).
fof(f9545,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f9506])).
fof(f9506,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X4)))))) | s(bool,X3) = s(bool,X5) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1175,f202])).
fof(f9153,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X43),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X43,fun(bool,'type.num.num'))) | s(X40,X42) = s(X40,sK4(X42,X41,X40))) )),
  inference(superposition,[],[f399,f400])).
fof(f9144,plain,(
  ( ! [X26,X25] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X26),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25)))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0'))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X26,fun(bool,'type.num.num'))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f399,f113])).
fof(f9142,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X19)))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X20,fun(bool,'type.num.num'))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f399,f202])).
fof(f9137,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X7,fun(bool,'type.num.num'))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f399,f202])).
fof(f9136,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3,fun(bool,'type.num.num'))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f399,f202])).
fof(f9135,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(bool,'type.num.num'),bool),X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0')))))) | s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))) = s(fun(bool,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X1,fun(bool,'type.num.num'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f399,f113])).
fof(f8793,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X40,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X40),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X37,bool),X38),s(X37,X39)))))))) | s(X37,X39) = s(X37,sK4(X39,X38,X37))) )),
  inference(superposition,[],[f530,f400])).
fof(f8786,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X26,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X26),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f530,f113])).
fof(f8784,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X19),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X20,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X19)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f530,f202])).
fof(f8779,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X7,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f530,f202])).
fof(f8778,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f530,f202])).
fof(f8777,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X1,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0')))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f530,f113])).
fof(f8772,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5)))))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X6,bool))) )),
  inference(superposition,[],[f530,f202])).
fof(f8771,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) | s(bool,X4) = s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f530,f202])).
fof(f8620,plain,(
  ( ! [X8,X7,X9] : (s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,g(s(X7,sK5(X8,X9,bool,X7)))) | s(bool,'T_0') = s(bool,f(s(X7,sK5(X8,X9,bool,X7))))) )),
  inference(trivial_inequality_removal,[],[f8613])).
fof(f8613,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,g(s(X7,sK5(X8,X9,bool,X7)))) | s(bool,'T_0') = s(bool,f(s(X7,sK5(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f935,f113])).
fof(f8151,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X10),s(fun(bool,'type.num.num'),X9)))))) | s(bool,'T_0') = s(bool,X8) | s(bool,X11) = s(bool,'const.bag.BAG_IN_2'(s(bool,X8),s(fun(bool,'type.num.num'),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f8138,f110])).
fof(f8138,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X10),s(fun(bool,'type.num.num'),X9)))))) | s(bool,'T_0') = s(bool,X8) | s(bool,X11) = s(bool,'const.bag.BAG_IN_2'(s(bool,X8),s(fun(bool,'type.num.num'),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f797,f202])).
fof(f7974,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X7),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f7968,f110])).
fof(f7968,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X7),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f796,f202])).
fof(f7836,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f7823,f110])).
fof(f7823,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X8) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f750,f202])).
fof(f7835,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,sK4(X1,X0,bool))) )),
  inference(subsumption_resolution,[],[f7821,f110])).
fof(f7821,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,sK4(X1,X0,bool))) )),
  inference(superposition,[],[f750,f400])).
fof(f7654,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,X7) != s(bool,X8) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f7648,f110])).
fof(f7648,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,X8) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X8)))) | s(bool,'T_0') = s(bool,X7) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f749,f202])).
fof(f7653,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,sK4('F_0',X0,bool))) )),
  inference(subsumption_resolution,[],[f7646,f110])).
fof(f7646,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,sK4('F_0',X0,bool))) )),
  inference(superposition,[],[f749,f400])).
fof(f7547,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK3(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X0),s(fun(bool,'type.num.num'),X1)))) )),
  inference(subsumption_resolution,[],[f7541,f112])).
fof(f7541,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK3(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X0),s(fun(bool,'type.num.num'),X1)))) )),
  inference(superposition,[],[f613,f113])).
fof(f7540,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X1,bool))) )),
  inference(superposition,[],[f613,f138])).
fof(f7314,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X1)))) | s(bool,'T_0') = s(bool,sK3(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X0),s(fun(bool,'type.num.num'),X1)))) )),
  inference(subsumption_resolution,[],[f7307,f112])).
fof(f7307,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X1)))) | s(bool,'T_0') = s(bool,sK3(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X0),s(fun(bool,'type.num.num'),X1)))) )),
  inference(superposition,[],[f526,f113])).
fof(f7313,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | s(bool,'T_0') = s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X1,bool))) )),
  inference(forward_demodulation,[],[f7306,f138])).
fof(f7306,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')))))) | s(bool,'T_0') = s(bool,sK3('const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0')),X1,bool))) )),
  inference(superposition,[],[f526,f138])).
fof(f6546,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X1))))))) )),
  inference(condensation,[],[f6545])).
fof(f6545,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X23)))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) )),
  inference(trivial_inequality_removal,[],[f6535])).
fof(f6535,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X21) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X23)))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f453,f113])).
fof(f6547,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f6530])).
fof(f6530,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f453,f113])).
fof(f4877,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))),
  inference(forward_demodulation,[],[f4825,f197])).
fof(f4825,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f198,f1238])).
fof(f4875,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0'))))))),
  inference(forward_demodulation,[],[f4828,f197])).
fof(f4828,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f198,f1238])).
fof(f4827,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f190,f1238])).
fof(f4824,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f190,f1238])).
fof(f6244,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1))))))) )),
  inference(condensation,[],[f6243])).
fof(f6243,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1)))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f6242])).
fof(f6242,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X8),s(fun(bool,'type.num.num'),X9)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f6198,f202])).
fof(f6198,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X8),s(fun(bool,'type.num.num'),X9)))))) | s(bool,X7) != s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f452,f202])).
fof(f5986,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK4(X5,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5980,f110])).
fof(f5980,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK4(X5,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f423,f202])).
fof(f5985,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK4(X1,X0,bool)) | s(bool,'F_0') = s(bool,sK4('F_0',X0,bool))) )),
  inference(subsumption_resolution,[],[f5978,f110])).
fof(f5978,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK4(X1,X0,bool)) | s(bool,'F_0') = s(bool,sK4('F_0',X0,bool))) )),
  inference(superposition,[],[f423,f400])).
fof(f5826,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,'type.num.num'),X2))) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) )),
  inference(subsumption_resolution,[],[f5818,f112])).
fof(f5818,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,'type.num.num'),X2))) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) )),
  inference(superposition,[],[f532,f113])).
fof(f5825,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,sK3('const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X0)),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f5816,f138])).
fof(f5816,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,sK3('const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X0)),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))))) )),
  inference(superposition,[],[f532,f138])).
fof(f5662,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),X2)))))) )),
  inference(equality_resolution,[],[f485])).
fof(f5652,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,X43) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X43),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))),s(fun(bool,'type.num.num'),X44))))) | s(X40,X42) = s(X40,sK4(X42,X41,X40))) )),
  inference(superposition,[],[f485,f400])).
fof(f5648,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X31),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X32))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f485,f113])).
fof(f5646,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X24))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f485,f202])).
fof(f5642,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X9),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X7),s(fun(bool,'type.num.num'),X10))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f485,f202])).
fof(f5641,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),X5))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f485,f202])).
fof(f5640,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f485,f113])).
fof(f5630,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,X43) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X43),s(fun(bool,'type.num.num'),X44))))) | s(X40,X42) = s(X40,sK4(X42,X41,X40))) )),
  inference(superposition,[],[f485,f400])).
fof(f5626,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X31),s(fun(bool,'type.num.num'),X32))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f485,f113])).
fof(f5624,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X23),s(fun(bool,'type.num.num'),X24))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f485,f202])).
fof(f5620,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X9),s(fun(bool,'type.num.num'),X10))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f485,f202])).
fof(f5619,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X4),s(fun(bool,'type.num.num'),X5))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f485,f202])).
fof(f5618,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f485,f113])).
fof(f5664,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.BAG_INSERT_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),X2)))))) )),
  inference(forward_demodulation,[],[f5616,f138])).
fof(f5616,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.BAG_INSERT_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),X2)))))) )),
  inference(unit_resulting_resolution,[],[f138,f485])).
fof(f5046,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X131),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(bool,X131) = s(bool,happ(s(fun(X128,bool),X129),s(X128,X130))) | s(X128,X130) = s(X128,sK4(X130,X129,X128))) )),
  inference(superposition,[],[f1238,f400])).
fof(f5045,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X127)))))) | s(bool,X127) = s(bool,happ(s(fun(X124,bool),X125),s(X124,X126))) | s(X124,X126) = s(X124,sK4(X126,X125,X124))) )),
  inference(superposition,[],[f1238,f400])).
fof(f5044,plain,(
  ( ! [X123,X121,X122,X120] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X123),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(bool,X123) != s(bool,happ(s(fun(X120,bool),X121),s(X120,X122))) | s(X120,X122) = s(X120,sK4(X122,X121,X120))) )),
  inference(superposition,[],[f467,f400])).
fof(f5043,plain,(
  ( ! [X118,X116,X119,X117] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X119)))))) | s(bool,X119) != s(bool,happ(s(fun(X116,bool),X117),s(X116,X118))) | s(X116,X118) = s(X116,sK4(X118,X117,X116))) )),
  inference(superposition,[],[f467,f400])).
fof(f5040,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X108)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X105,bool),X106),s(X105,X107))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X109),s(fun(bool,'type.num.num'),X108)))))) | s(X105,X107) = s(X105,sK4(X107,X106,X105))) )),
  inference(superposition,[],[f152,f400])).
fof(f5039,plain,(
  ( ! [X103,X101,X102,X100,X104] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X103),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X104)))))) | s(bool,X103) != s(bool,happ(s(fun(X100,bool),X101),s(X100,X102))) | s(X100,X102) = s(X100,sK4(X102,X101,X100))) )),
  inference(superposition,[],[f151,f400])).
fof(f5038,plain,(
  ( ! [X99,X97,X95,X98,X96] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X98),s(fun(bool,'type.num.num'),X99)))))) | s(bool,X98) != s(bool,happ(s(fun(X95,bool),X96),s(X95,X97))) | s(X95,X97) = s(X95,sK4(X97,X96,X95))) )),
  inference(superposition,[],[f151,f400])).
fof(f5037,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X93),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X94)))))) | s(bool,X93) = s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X93),s(fun(bool,'type.num.num'),X94)))) | s(X90,X92) = s(X90,sK4(X92,X91,X90))) )),
  inference(superposition,[],[f150,f400])).
fof(f5036,plain,(
  ( ! [X88,X87,X85,X89,X86] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X88),s(fun(bool,'type.num.num'),X89)))))) | s(bool,X88) = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))),s(fun(bool,'type.num.num'),X89)))) | s(X85,X87) = s(X85,sK4(X87,X86,X85))) )),
  inference(superposition,[],[f150,f400])).
fof(f5035,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'T_0')))) | s(bool,X84) != s(bool,happ(s(fun(X80,bool),X81),s(X80,X82))) | p(s(bool,happ(s(fun(bool,bool),X83),s(bool,X84)))) | s(X80,X82) = s(X80,sK4(X82,X81,X80))) )),
  inference(superposition,[],[f149,f400])).
fof(f5034,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,happ(s(fun(bool,bool),X79),s(bool,'T_0')))) | s(bool,happ(s(fun(X76,bool),X77),s(X76,X78))) = s(bool,sK4(happ(s(fun(X76,bool),X77),s(X76,X78)),X79,bool)) | s(X76,X78) = s(X76,sK4(X78,X77,X76))) )),
  inference(superposition,[],[f147,f400])).
fof(f5074,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X62) != s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(bool,'T_0') = s(bool,X62) | s(X59,X61) = s(X59,sK4(X61,X60,X59))) )),
  inference(subsumption_resolution,[],[f5029,f110])).
fof(f5029,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | s(bool,X62) != s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(bool,'T_0') = s(bool,X62) | s(X59,X61) = s(X59,sK4(X61,X60,X59))) )),
  inference(superposition,[],[f204,f400])).
fof(f5073,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X44) != s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) | p(s(bool,X44)) | s(X41,X43) = s(X41,sK4(X43,X42,X41))) )),
  inference(subsumption_resolution,[],[f5026,f110])).
fof(f5026,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X44) != s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) | p(s(bool,X44)) | s(X41,X43) = s(X41,sK4(X43,X42,X41))) )),
  inference(superposition,[],[f187,f400])).
fof(f5072,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X40)) | s(bool,X40) = s(bool,happ(s(fun(X37,bool),X38),s(X37,X39))) | s(X37,X39) = s(X37,sK4(X39,X38,X37))) )),
  inference(subsumption_resolution,[],[f5025,f110])).
fof(f5025,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | s(bool,X40) = s(bool,happ(s(fun(X37,bool),X38),s(X37,X39))) | s(X37,X39) = s(X37,sK4(X39,X38,X37))) )),
  inference(superposition,[],[f156,f400])).
fof(f5071,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X10,bool),X11),s(fun(X10,'type.num.num'),X12)))) | s(X10,sK3(X12,X11,X10)) = s(X10,sK4(sK3(X12,X11,X10),X11,X10))) )),
  inference(subsumption_resolution,[],[f5014,f110])).
fof(f5014,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X10,bool),X11),s(fun(X10,'type.num.num'),X12)))) | s(X10,sK3(X12,X11,X10)) = s(X10,sK4(sK3(X12,X11,X10),X11,X10))) )),
  inference(superposition,[],[f146,f400])).
fof(f5070,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK4(X9,X8,X7)) = s(X7,sK4(sK4(X9,X8,X7),X8,X7))) )),
  inference(subsumption_resolution,[],[f5013,f110])).
fof(f5013,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK4(X9,X8,X7)) = s(X7,sK4(sK4(X9,X8,X7),X8,X7))) )),
  inference(superposition,[],[f148,f400])).
fof(f5069,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(subsumption_resolution,[],[f5011,f110])).
fof(f5011,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(superposition,[],[f149,f400])).
fof(f5067,plain,(
  ( ! [X39,X41,X40] : (s(fun(X39,'type.num.num'),'const.bag.EL_BAG_1'(s(X39,X40))) = s(fun(X39,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(X39,X40),s(fun(X39,'type.num.num'),'const.bag.EMPTY_BAG_0')),X41,fun(X39,'type.num.num'))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X39,'type.num.num'),bool),X41),s(fun(X39,'type.num.num'),'const.bag.EL_BAG_1'(s(X39,X40)))))) )),
  inference(forward_demodulation,[],[f5003,f138])).
fof(f5003,plain,(
  ( ! [X39,X41,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(fun(X39,'type.num.num'),bool),X41),s(fun(X39,'type.num.num'),'const.bag.EL_BAG_1'(s(X39,X40))))) | s(fun(X39,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X39,X40),s(fun(X39,'type.num.num'),'const.bag.EMPTY_BAG_0'))) = s(fun(X39,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(X39,X40),s(fun(X39,'type.num.num'),'const.bag.EMPTY_BAG_0')),X41,fun(X39,'type.num.num')))) )),
  inference(superposition,[],[f400,f138])).
fof(f5065,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(X35,X36),s(fun(X35,'type.num.num'),'const.bag.EMPTY_BAG_0')),X37,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'F_0')))) )),
  inference(forward_demodulation,[],[f5001,f197])).
fof(f5001,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'F_0'))) | s(bool,'const.bag.BAG_IN_2'(s(X35,X36),s(fun(X35,'type.num.num'),'const.bag.EMPTY_BAG_0'))) = s(bool,sK4('const.bag.BAG_IN_2'(s(X35,X36),s(fun(X35,'type.num.num'),'const.bag.EMPTY_BAG_0')),X37,bool))) )),
  inference(superposition,[],[f400,f197])).
fof(f4998,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0'))) | s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) = s(bool,sK4(happ(s(fun(X28,bool),X29),s(X28,X30)),X31,bool)) | s(X28,X30) = s(X28,sK4(X30,X29,X28))) )),
  inference(superposition,[],[f400,f400])).
fof(f4994,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X23),s(bool,X22))) | s(bool,'F_0') = s(bool,sK4('F_0',X23,bool)) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f400,f113])).
fof(f4992,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,'T_0') = s(bool,sK4('T_0',X17,bool)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f400,f202])).
fof(f4988,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,X6))) | s(bool,X5) = s(bool,sK4(X5,X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f400,f202])).
fof(f4987,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,X2) = s(bool,sK4(X2,X3,bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f400,f202])).
fof(f4986,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,X0) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f400,f113])).
fof(f4898,plain,(
  ( ! [X6,X4,X5,X3] : (s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(X3,X4),s(fun(X3,'type.num.num'),'const.bag.EL_BAG_1'(s(X3,X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f4869,f110])).
fof(f4869,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(X3,X4),s(fun(X3,'type.num.num'),'const.bag.EL_BAG_1'(s(X3,X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1238,f202])).
fof(f4856,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X22)))))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1238,f113])).
fof(f4849,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1238,f202])).
fof(f4848,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0')))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1238,f113])).
fof(f4838,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23)))))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1238,f113])).
fof(f4836,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X16),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17)))))) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1238,f202])).
fof(f4832,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7)))))) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1238,f202])).
fof(f4831,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3)))))) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1238,f202])).
fof(f4830,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1)))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1238,f113])).
fof(f4700,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X2)))))) )),
  inference(subsumption_resolution,[],[f4652,f112])).
fof(f4652,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X2)))))) )),
  inference(superposition,[],[f467,f113])).
fof(f4641,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X22)))))) | s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f467,f113])).
fof(f4639,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X17),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X16)))))) | s(bool,'T_0') != s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f467,f202])).
fof(f4635,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) | s(bool,X5) != s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f467,f202])).
fof(f4634,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | s(bool,X2) != s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f467,f202])).
fof(f4633,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'F_0')))))) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f467,f113])).
fof(f4624,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23)))))) | s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f467,f113])).
fof(f4622,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X16),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17)))))) | s(bool,'T_0') != s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f467,f202])).
fof(f4618,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7)))))) | s(bool,X5) != s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f467,f202])).
fof(f4617,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3)))))) | s(bool,X2) != s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f467,f202])).
fof(f4616,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1)))))) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f467,f113])).
fof(f4662,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.EL_BAG_1'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))))))))) )),
  inference(forward_demodulation,[],[f4614,f138])).
fof(f4614,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.EL_BAG_1'(s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f138,f467])).
fof(f749,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,X3) != s(bool,X5) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f149,f113])).
fof(f796,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X4)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),X4)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f152,f113])).
fof(f933,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,X5,bool))) )),
  inference(superposition,[],[f161,f113])).
fof(f935,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK5(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK5(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f161,f113])).
fof(f1175,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X4),s(fun(bool,'type.num.num'),X5)))))) | s(bool,X3) = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f150,f113])).
fof(f1188,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5)))))) | s(bool,X3) = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f150,f113])).
fof(f3869,plain,(
  ( ! [X36] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X36),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X36))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl3),
  inference(forward_demodulation,[],[f3868,f257])).
fof(f3868,plain,(
  ( ! [X36] : (s(bool,'const.bag.BAG_IN_2'(s(bool,X36),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X36))))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl3),
  inference(subsumption_resolution,[],[f3854,f112])).
fof(f3854,plain,(
  ( ! [X36] : (p(s(bool,'F_0')) | s(bool,'const.bag.BAG_IN_2'(s(bool,X36),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X36))))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl3),
  inference(superposition,[],[f3317,f257])).
fof(f3818,plain,(
  ( ! [X54,X52,X53] : (s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53))))) | s(bool,'F_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(subsumption_resolution,[],[f3797,f237])).
fof(f3797,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53))))) | s(bool,'F_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(superposition,[],[f3316,f1527])).
fof(f3817,plain,(
  ( ! [X48,X49] : (s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X49))))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(subsumption_resolution,[],[f3795,f110])).
fof(f3795,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X49))))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f3316,f3312])).
fof(f3802,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23))))) | s(bool,'F_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | s(bool,X22) = s(bool,X24)) ) | $spl3),
  inference(subsumption_resolution,[],[f3783,f110])).
fof(f3783,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23))))) | s(bool,'F_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | s(bool,X22) = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f3316,f1527])).
fof(f3798,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3))))) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f3776,f110])).
fof(f3776,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3))))) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f3316,f202])).
fof(f3773,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) != s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) | s(bool,X53) = s(bool,'const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53))))) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(subsumption_resolution,[],[f3748,f237])).
fof(f3748,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | s(bool,X53) != s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) | s(bool,X53) = s(bool,'const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53))))) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(superposition,[],[f2699,f1527])).
fof(f3772,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))) | s(bool,X49) = s(bool,'const.bag.BAG_IN_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X49))))) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(subsumption_resolution,[],[f3746,f110])).
fof(f3746,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | s(bool,X49) != s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))) | s(bool,X49) = s(bool,'const.bag.BAG_IN_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X49))))) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f2699,f3312])).
fof(f3771,plain,(
  ( ! [X47,X45,X46] : (s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X46))))))) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f3745,f110])).
fof(f3745,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X46))))))) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47)))))) ) | $spl3),
  inference(superposition,[],[f2699,f950])).
fof(f3694,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1)))))))))) ) | $spl3),
  inference(condensation,[],[f3693])).
fof(f3693,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | s(bool,'T_0') = s(bool,X53) | s(bool,X53) != s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))))))) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(subsumption_resolution,[],[f3672,f204])).
fof(f3672,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,X52)) | s(bool,'T_0') = s(bool,X53) | s(bool,X53) != s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))))))) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) ) | $spl3),
  inference(superposition,[],[f2635,f1527])).
fof(f3692,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X49) != s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))))))) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(subsumption_resolution,[],[f3670,f110])).
fof(f3670,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X49) | s(bool,X49) != s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X48),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X48))))))))) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f2635,f3312])).
fof(f3554,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X72),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X71),s(fun(bool,'type.num.num'),X73)))))) | s(bool,X72) != s(bool,'const.bag.BAG_IN_2'(s(bool,X71),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(superposition,[],[f151,f2625])).
fof(f3553,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X69),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X70)))))) | s(bool,X69) != s(bool,'const.bag.BAG_IN_2'(s(bool,X68),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X68))))) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f151,f3312])).
fof(f3530,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X71),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X72),s(fun(bool,'type.num.num'),X73)))))) | s(bool,X72) != s(bool,'const.bag.BAG_IN_2'(s(bool,X71),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X71))))) | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(superposition,[],[f151,f2625])).
fof(f3529,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X69),s(fun(bool,'type.num.num'),X70)))))) | s(bool,X69) != s(bool,'const.bag.BAG_IN_2'(s(bool,X68),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X68))))) | s(bool,'F_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f151,f3312])).
fof(f3474,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X53)))),X54,bool)) | s(bool,'F_0') = s(bool,X53)) ) | $spl3),
  inference(superposition,[],[f147,f3312])).
fof(f3494,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47))))))))) | s(bool,'F_0') = s(bool,X47)) ) | $spl3),
  inference(subsumption_resolution,[],[f3469,f110])).
fof(f3469,plain,(
  ( ! [X47] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47))))))))) | s(bool,'F_0') = s(bool,X47)) ) | $spl3),
  inference(superposition,[],[f2630,f3312])).
fof(f3490,plain,(
  ( ! [X24,X23] : (s(bool,X24) != s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f3455,f110])).
fof(f3455,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | s(bool,X24) != s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f204,f3312])).
fof(f3344,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2))))) | s(bool,X0) = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f2694,f1527])).
fof(f3309,plain,(
  ( ! [X52,X53] : (p(s(bool,happ(s(fun(bool,bool),X53),s(bool,X52)))) | s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52)))),X53,bool)) | s(bool,'F_0') = s(bool,X52)) ) | $spl3),
  inference(superposition,[],[f147,f2625])).
fof(f3318,plain,(
  ( ! [X50] : (s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))))))) | s(bool,'F_0') = s(bool,X50)) ) | $spl3),
  inference(subsumption_resolution,[],[f3307,f120])).
fof(f3307,plain,(
  ( ! [X50] : (~p(s(bool,X50)) | s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X50),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X50))))))))) | s(bool,'F_0') = s(bool,X50)) ) | $spl3),
  inference(superposition,[],[f2686,f2625])).
fof(f3317,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bag.BAG_IN_2'(s(bool,X13),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X13))))) | s(bool,'F_0') = s(bool,X13)) ) | $spl3),
  inference(subsumption_resolution,[],[f3286,f120])).
fof(f3286,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | s(bool,X14) != s(bool,'const.bag.BAG_IN_2'(s(bool,X13),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X13))))) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,X13)) ) | $spl3),
  inference(superposition,[],[f187,f2625])).
fof(f3316,plain,(
  ( ! [X12,X11] : (~p(s(bool,X12)) | s(bool,X12) = s(bool,'const.bag.BAG_IN_2'(s(bool,X11),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X11))))) | s(bool,'F_0') = s(bool,X11)) ) | $spl3),
  inference(subsumption_resolution,[],[f3285,f120])).
fof(f3285,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) = s(bool,'const.bag.BAG_IN_2'(s(bool,X11),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X11))))) | s(bool,'F_0') = s(bool,X11)) ) | $spl3),
  inference(superposition,[],[f156,f2625])).
fof(f3312,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2))))) | s(bool,'F_0') = s(bool,X2)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f3268])).
fof(f3268,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) ) | $spl3),
  inference(superposition,[],[f2625,f113])).
fof(f3159,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))))))) | ~p(s(bool,X0))) ) | $spl3),
  inference(condensation,[],[f3158])).
fof(f3158,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))))))) | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f3141,f204])).
fof(f3141,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))))))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(superposition,[],[f2630,f1527])).
fof(f3143,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15))))) | s(bool,X14) = s(bool,X15)) ) | $spl3),
  inference(subsumption_resolution,[],[f3129,f110])).
fof(f3129,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15))))) | s(bool,X14) = s(bool,X15)) ) | $spl3),
  inference(superposition,[],[f2630,f1527])).
fof(f3115,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X22) | s(bool,X20) = s(bool,'const.bag.BAG_IN_2'(s(bool,X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X20)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f3077,f190])).
fof(f3077,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X22) | s(bool,X20) = s(bool,'const.bag.BAG_IN_2'(s(bool,X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X20))))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl3),
  inference(superposition,[],[f2631,f1527])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f80])).
fof(f257,plain,(
  s(bool,'F_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f246,f113])).
fof(f388,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,X0) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f147,f113])).
fof(f389,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK4('F_0',X3,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f147,f113])).
fof(f396,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X4,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f390,f257])).
fof(f390,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X4,bool))) ) | $spl3),
  inference(superposition,[],[f147,f257])).
fof(f398,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(X6,X7),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0')),X8,bool)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f392,f197])).
fof(f392,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0')))) | s(bool,'const.bag.BAG_IN_2'(s(X6,X7),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0'))) = s(bool,sK4('const.bag.BAG_IN_2'(s(X6,X7),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0')),X8,bool))) )),
  inference(superposition,[],[f147,f197])).
fof(f423,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK4(X1,X2,bool))) )),
  inference(superposition,[],[f148,f113])).
fof(f425,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f396,f113])).
fof(f428,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f427,f112])).
fof(f427,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | $spl3),
  inference(superposition,[],[f396,f113])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f151,f113])).
fof(f453,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X4),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f151,f113])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f151,f113])).
fof(f460,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f151,f113])).
fof(f470,plain,(
  ( ! [X2] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2))))))) )),
  inference(forward_demodulation,[],[f450,f197])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2))))))) )),
  inference(unit_resulting_resolution,[],[f197,f151])).
fof(f526,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),X0)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X0)))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f145,f113])).
fof(f613,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X1),s(fun(bool,'type.num.num'),X0)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f146,f113])).
fof(f750,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'F_0') != s(bool,X8) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f149,f113])).
fof(f762,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,X0) != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(forward_demodulation,[],[f748,f257])).
fof(f748,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))))) | s(bool,X0) != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f149,f442])).
fof(f797,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),X7)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X8),s(fun(bool,'type.num.num'),X7)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f152,f113])).
fof(f812,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X1)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X2),s(fun(bool,'type.num.num'),X1)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(forward_demodulation,[],[f795,f257])).
fof(f795,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),X1)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X2),s(fun(bool,'type.num.num'),X1)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f152,f442])).
fof(f936,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f932,f257])).
fof(f932,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))) != s(X2,g(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) ) | $spl3),
  inference(superposition,[],[f161,f442])).
fof(f937,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) ) | $spl3),
  inference(forward_demodulation,[],[f934,f257])).
fof(f934,plain,(
  ( ! [X2,X0,X1] : (s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) ) | $spl3),
  inference(superposition,[],[f161,f442])).
fof(f1176,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X7),s(fun(bool,'type.num.num'),X8)))))) | s(bool,'F_0') = s(bool,X7) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X8)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f150,f113])).
fof(f1189,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X6),s(fun(bool,'type.num.num'),X8)))))) | s(bool,'F_0') = s(bool,X7) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),X8)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f150,f113])).
fof(f1207,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(forward_demodulation,[],[f1174,f257])).
fof(f1174,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f150,f442])).
fof(f1226,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(forward_demodulation,[],[f1187,f257])).
fof(f1187,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),X2)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f150,f442])).
fof(f2625,plain,(
  ( ! [X35] : (s(bool,X35) = s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35)) ) | $spl3),
  inference(subsumption_resolution,[],[f2126,f246])).
fof(f2126,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35) | s(bool,X35) = s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) ) | $spl3),
  inference(superposition,[],[f1527,f257])).
fof(f2698,plain,(
  ( ! [X92,X93] : (p(s(bool,happ(s(fun(bool,bool),X93),s(bool,X92)))) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X93,bool)) | s(bool,X92) = s(bool,'const.bag.BAG_IN_2'(s(bool,X92),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X92)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2308,f190])).
fof(f2308,plain,(
  ( ! [X92,X93] : (p(s(bool,happ(s(fun(bool,bool),X93),s(bool,X92)))) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X93,bool)) | s(bool,X92) = s(bool,'const.bag.BAG_IN_2'(s(bool,X92),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X92))))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl3),
  inference(superposition,[],[f396,f1527])).
fof(f2463,plain,(
  ( ! [X204,X202,X203,X201] : (s(bool,'T_0') != s(bool,f(s(X201,sK5(X202,X203,bool,X201)))) | s(fun(X201,bool),X202) = s(fun(X201,bool),X203) | s(bool,X204) = s(bool,'const.bag.BAG_IN_2'(s(bool,X204),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X204))))) | s(bool,X204) = s(bool,g(s(X201,sK5(X202,X203,bool,X201))))) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f2350,plain,(
  ( ! [X218,X216,X215,X217] : (s(bool,X218) != s(bool,f(s(X215,sK5(X216,X217,bool,X215)))) | s(fun(X215,bool),X216) = s(fun(X215,bool),X217) | s(bool,X218) = s(bool,'const.bag.BAG_IN_2'(s(bool,X218),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X218))))) | s(bool,'T_0') = s(bool,g(s(X215,sK5(X216,X217,bool,X215))))) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f2268,plain,(
  ( ! [X218,X216,X215,X217] : (s(bool,X218) != s(bool,f(s(X215,sK5(X216,X217,bool,X215)))) | s(fun(X215,bool),X216) = s(fun(X215,bool),X217) | s(bool,g(s(X215,sK5(X216,X217,bool,X215)))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(X215,sK5(X216,X217,bool,X215)))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(X215,sK5(X216,X217,bool,X215)))))))) | s(bool,'T_0') = s(bool,X218)) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f1624,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,'T_0') != s(bool,f(s(X187,sK5(X188,X189,bool,X187)))) | s(fun(X187,bool),X188) = s(fun(X187,bool),X189) | s(bool,X190) = s(bool,g(s(X187,sK5(X188,X189,bool,X187)))) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f161,f202])).
fof(f1443,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,X200) != s(bool,f(s(X197,sK5(X198,X199,bool,X197)))) | s(fun(X197,bool),X198) = s(fun(X197,bool),X199) | s(bool,'T_0') = s(bool,X200) | s(bool,'T_0') = s(bool,g(s(X197,sK5(X198,X199,bool,X197))))) )),
  inference(superposition,[],[f161,f202])).
fof(f144,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),X2)))) | ~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X4),s(fun(X0,'type.num.num'),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f2043,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X83,X86),s(fun(X83,'type.num.num'),X85)))) | p(s(bool,happ(s(fun(X83,bool),X84),s(X83,X86)))) | s(bool,X87) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X83,bool),X84),s(fun(X83,'type.num.num'),X85))) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f1596,f110])).
fof(f1596,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bag.BAG_IN_2'(s(X83,X86),s(fun(X83,'type.num.num'),X85)))) | p(s(bool,happ(s(fun(X83,bool),X84),s(X83,X86)))) | s(bool,X87) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X83,bool),X84),s(fun(X83,'type.num.num'),X85))) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f144,f202])).
fof(f2046,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X103,bool),X104),s(fun(X103,'type.num.num'),X105)))) | s(bool,X106) = s(bool,happ(s(fun(X103,bool),X104),s(X103,sK3(X105,X104,X103)))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f1601,f110])).
fof(f1601,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X103,bool),X104),s(fun(X103,'type.num.num'),X105)))) | s(bool,X106) = s(bool,happ(s(fun(X103,bool),X104),s(X103,sK3(X105,X104,X103)))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f146,f202])).
fof(f2716,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X97,X100),s(fun(X97,'type.num.num'),X99)))) | p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X100)))) | s(bool,X101) = s(bool,'const.bag.BAG_IN_2'(s(bool,X101),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X101))))) | s(bool,X101) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X97,bool),X98),s(fun(X97,'type.num.num'),X99)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2435,f110])).
fof(f2435,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bag.BAG_IN_2'(s(X97,X100),s(fun(X97,'type.num.num'),X99)))) | p(s(bool,happ(s(fun(X97,bool),X98),s(X97,X100)))) | s(bool,X101) = s(bool,'const.bag.BAG_IN_2'(s(bool,X101),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X101))))) | s(bool,X101) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X97,bool),X98),s(fun(X97,'type.num.num'),X99)))) ) | $spl3),
  inference(superposition,[],[f144,f1527])).
fof(f2719,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X117,bool),X118),s(fun(X117,'type.num.num'),X119)))) | s(bool,X120) = s(bool,'const.bag.BAG_IN_2'(s(bool,X120),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X120))))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK3(X119,X118,X117))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2440,f110])).
fof(f2440,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X117,bool),X118),s(fun(X117,'type.num.num'),X119)))) | s(bool,X120) = s(bool,'const.bag.BAG_IN_2'(s(bool,X120),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X120))))) | s(bool,X120) = s(bool,happ(s(fun(X117,bool),X118),s(X117,sK3(X119,X118,X117))))) ) | $spl3),
  inference(superposition,[],[f146,f1527])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,sK3('const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X0)),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f527,f138])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,sK3('const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')),X2,X0)),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))))) )),
  inference(superposition,[],[f145,f138])).
fof(f1444,plain,(
  ( ! [X202,X203,X201] : (~p(s(bool,happ(s(fun(bool,bool),X202),s(bool,X203)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X202),s(fun(bool,'type.num.num'),X201)))) | s(bool,'T_0') = s(bool,X203) | s(bool,'T_0') = s(bool,sK3(X201,X202,bool))) )),
  inference(superposition,[],[f146,f202])).
fof(f1445,plain,(
  ( ! [X206,X204,X205] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X206),s(fun(bool,'type.num.num'),X204)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X205),s(fun(bool,'type.num.num'),X204)))) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,sK3(X204,X205,bool))) )),
  inference(superposition,[],[f145,f202])).
fof(f1625,plain,(
  ( ! [X191,X192,X193] : (~p(s(bool,happ(s(fun(bool,bool),X192),s(bool,'T_0')))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X192),s(fun(bool,'type.num.num'),X191)))) | s(bool,X193) = s(bool,sK3(X191,X192,bool)) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f146,f202])).
fof(f1626,plain,(
  ( ! [X196,X194,X195] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X194)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X195),s(fun(bool,'type.num.num'),X194)))) | s(bool,X196) = s(bool,sK3(X194,X195,bool)) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f145,f202])).
fof(f2269,plain,(
  ( ! [X220,X219,X221] : (~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X221)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X220),s(fun(bool,'type.num.num'),X219)))) | s(bool,sK3(X219,X220,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK3(X219,X220,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK3(X219,X220,bool)))))) | s(bool,'T_0') = s(bool,X221)) ) | $spl3),
  inference(superposition,[],[f146,f1527])).
fof(f2270,plain,(
  ( ! [X222,X223,X224] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X224),s(fun(bool,'type.num.num'),X222)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X223),s(fun(bool,'type.num.num'),X222)))) | s(bool,sK3(X222,X223,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK3(X222,X223,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK3(X222,X223,bool)))))) | s(bool,'T_0') = s(bool,X224)) ) | $spl3),
  inference(superposition,[],[f145,f1527])).
fof(f2351,plain,(
  ( ! [X220,X219,X221] : (~p(s(bool,happ(s(fun(bool,bool),X220),s(bool,X221)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X220),s(fun(bool,'type.num.num'),X219)))) | s(bool,X221) = s(bool,'const.bag.BAG_IN_2'(s(bool,X221),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X221))))) | s(bool,'T_0') = s(bool,sK3(X219,X220,bool))) ) | $spl3),
  inference(superposition,[],[f146,f1527])).
fof(f2352,plain,(
  ( ! [X222,X223,X224] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X224),s(fun(bool,'type.num.num'),X222)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X223),s(fun(bool,'type.num.num'),X222)))) | s(bool,X224) = s(bool,'const.bag.BAG_IN_2'(s(bool,X224),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X224))))) | s(bool,'T_0') = s(bool,sK3(X222,X223,bool))) ) | $spl3),
  inference(superposition,[],[f145,f1527])).
fof(f2464,plain,(
  ( ! [X206,X207,X205] : (~p(s(bool,happ(s(fun(bool,bool),X206),s(bool,'T_0')))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X206),s(fun(bool,'type.num.num'),X205)))) | s(bool,X207) = s(bool,'const.bag.BAG_IN_2'(s(bool,X207),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X207))))) | s(bool,X207) = s(bool,sK3(X205,X206,bool))) ) | $spl3),
  inference(superposition,[],[f146,f1527])).
fof(f2465,plain,(
  ( ! [X210,X208,X209] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X208)))) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(bool,bool),X209),s(fun(bool,'type.num.num'),X208)))) | s(bool,X210) = s(bool,'const.bag.BAG_IN_2'(s(bool,X210),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X210))))) | s(bool,X210) = s(bool,sK3(X208,X209,bool))) ) | $spl3),
  inference(superposition,[],[f145,f1527])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X21,X22] : (! [X8] : (s(X6,X8) = s(X6,X22) => p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X8))))) <=> p(s(bool,happ(s(fun(X6,bool),X21),s(X6,X22)))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.UNWIND_FORALL_THM2')).
fof(f148,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f149,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f99])).
fof(f2044,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(X90,X92) != s(X90,X93) | p(s(bool,happ(s(fun(X90,bool),X91),s(X90,X93)))) | s(bool,X94) = s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1598,f110])).
fof(f1598,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,'T_0')) | s(X90,X92) != s(X90,X93) | p(s(bool,happ(s(fun(X90,bool),X91),s(X90,X93)))) | s(bool,X94) = s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f149,f202])).
fof(f2045,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,happ(s(fun(X99,bool),X100),s(X99,X101)))) | s(bool,X102) = s(bool,happ(s(fun(X99,bool),X100),s(X99,sK4(X101,X100,X99)))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f1600,f110])).
fof(f1600,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X99,bool),X100),s(X99,X101)))) | s(bool,X102) = s(bool,happ(s(fun(X99,bool),X100),s(X99,sK4(X101,X100,X99)))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f148,f202])).
fof(f2717,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(X104,X106) != s(X104,X107) | p(s(bool,happ(s(fun(X104,bool),X105),s(X104,X107)))) | s(bool,X108) = s(bool,'const.bag.BAG_IN_2'(s(bool,X108),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X108))))) | s(bool,X108) = s(bool,happ(s(fun(X104,bool),X105),s(X104,X106)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2437,f110])).
fof(f2437,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | s(X104,X106) != s(X104,X107) | p(s(bool,happ(s(fun(X104,bool),X105),s(X104,X107)))) | s(bool,X108) = s(bool,'const.bag.BAG_IN_2'(s(bool,X108),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X108))))) | s(bool,X108) = s(bool,happ(s(fun(X104,bool),X105),s(X104,X106)))) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2718,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,happ(s(fun(X113,bool),X114),s(X113,X115)))) | s(bool,X116) = s(bool,'const.bag.BAG_IN_2'(s(bool,X116),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X116))))) | s(bool,X116) = s(bool,happ(s(fun(X113,bool),X114),s(X113,sK4(X115,X114,X113))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2439,f110])).
fof(f2439,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X113,bool),X114),s(X113,X115)))) | s(bool,X116) = s(bool,'const.bag.BAG_IN_2'(s(bool,X116),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X116))))) | s(bool,X116) = s(bool,happ(s(fun(X113,bool),X114),s(X113,sK4(X115,X114,X113))))) ) | $spl3),
  inference(superposition,[],[f148,f1527])).
fof(f757,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun(bool,bool),X26),s(bool,'T_0')))) | s(bool,X27) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X25,bool)) | p(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f149,f402])).
fof(f1434,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,happ(s(fun(bool,bool),X164),s(bool,X163)))) | s(bool,X162) = s(bool,sK4(X162,X164,bool)) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f147,f202])).
fof(f1435,plain,(
  ( ! [X167,X165,X166,X168] : (~p(s(bool,happ(s(fun(bool,bool),X167),s(bool,X166)))) | s(bool,X165) != s(bool,X168) | p(s(bool,happ(s(fun(bool,bool),X167),s(bool,X168)))) | s(bool,'T_0') = s(bool,X166) | s(bool,'T_0') = s(bool,X165)) )),
  inference(superposition,[],[f149,f202])).
fof(f1446,plain,(
  ( ! [X208,X209,X207] : (~p(s(bool,happ(s(fun(bool,bool),X208),s(bool,X209)))) | p(s(bool,happ(s(fun(bool,bool),X208),s(bool,X207)))) | s(bool,'T_0') = s(bool,X209) | s(bool,'T_0') = s(bool,sK4(X207,X208,bool))) )),
  inference(superposition,[],[f148,f202])).
fof(f1525,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X8,bool)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(superposition,[],[f397,f202])).
fof(f1529,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X19,bool)) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(superposition,[],[f570,f202])).
fof(f1550,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X78)))) | s(bool,'T_0') = s(bool,sK4('T_0',X79,bool)) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f147,f202])).
fof(f1551,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,X81)))) | s(bool,'T_0') != s(bool,X83) | p(s(bool,happ(s(fun(bool,bool),X82),s(bool,X83)))) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f149,f202])).
fof(f1616,plain,(
  ( ! [X158,X156,X159,X157] : (~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0')))) | s(bool,X156) != s(bool,X158) | p(s(bool,happ(s(fun(bool,bool),X157),s(bool,X158)))) | s(bool,X156) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f149,f202])).
fof(f1627,plain,(
  ( ! [X198,X199,X197] : (~p(s(bool,happ(s(fun(bool,bool),X198),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X198),s(bool,X197)))) | s(bool,X199) = s(bool,sK4(X197,X198,bool)) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f148,f202])).
fof(f2182,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),X68),s(bool,X67)))) | s(bool,'const.bag.BAG_IN_2'(s(bool,X67),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X67))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,X67),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X67)))),X68,bool)) | s(bool,X67) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2183,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X70)))) | s(bool,X72) != s(bool,'const.bag.BAG_IN_2'(s(bool,X70),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X70))))) | p(s(bool,happ(s(fun(bool,bool),X71),s(bool,X72)))) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2259,plain,(
  ( ! [X182,X180,X181] : (p(s(bool,happ(s(fun(bool,bool),X182),s(bool,X181)))) | s(bool,X180) = s(bool,sK4(X180,X182,bool)) | s(bool,X180) = s(bool,'const.bag.BAG_IN_2'(s(bool,X180),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X180))))) | s(bool,'T_0') = s(bool,X181)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2260,plain,(
  ( ! [X185,X186,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X184)))) | s(bool,X183) != s(bool,X186) | p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X186)))) | s(bool,X183) = s(bool,'const.bag.BAG_IN_2'(s(bool,X183),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X183))))) | s(bool,'T_0') = s(bool,X184)) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2271,plain,(
  ( ! [X227,X225,X226] : (~p(s(bool,happ(s(fun(bool,bool),X226),s(bool,X227)))) | p(s(bool,happ(s(fun(bool,bool),X226),s(bool,X225)))) | s(bool,sK4(X225,X226,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(X225,X226,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK4(X225,X226,bool)))))) | s(bool,'T_0') = s(bool,X227)) ) | $spl3),
  inference(superposition,[],[f148,f1527])).
fof(f2341,plain,(
  ( ! [X182,X180,X181] : (p(s(bool,happ(s(fun(bool,bool),X182),s(bool,X181)))) | s(bool,X180) = s(bool,sK4(X180,X182,bool)) | s(bool,X181) = s(bool,'const.bag.BAG_IN_2'(s(bool,X181),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X181))))) | s(bool,'T_0') = s(bool,X180)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2342,plain,(
  ( ! [X185,X186,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X184)))) | s(bool,X183) != s(bool,X186) | p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X186)))) | s(bool,X184) = s(bool,'const.bag.BAG_IN_2'(s(bool,X184),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X184))))) | s(bool,'T_0') = s(bool,X183)) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2353,plain,(
  ( ! [X227,X225,X226] : (~p(s(bool,happ(s(fun(bool,bool),X226),s(bool,X227)))) | p(s(bool,happ(s(fun(bool,bool),X226),s(bool,X225)))) | s(bool,X227) = s(bool,'const.bag.BAG_IN_2'(s(bool,X227),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X227))))) | s(bool,'T_0') = s(bool,sK4(X225,X226,bool))) ) | $spl3),
  inference(superposition,[],[f148,f1527])).
fof(f2358,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X16,bool)) | s(bool,X17) = s(bool,'const.bag.BAG_IN_2'(s(bool,X17),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17))))) | s(bool,X15) = s(bool,X17)) ) | $spl3),
  inference(superposition,[],[f397,f1527])).
fof(f2362,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,happ(s(fun(bool,bool),X27),s(bool,X26)))) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X27,bool)) | s(bool,X28) = s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))) | s(bool,X26) = s(bool,X28)) ) | $spl3),
  inference(superposition,[],[f570,f1527])).
fof(f2385,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,happ(s(fun(bool,bool),X93),s(bool,X92)))) | s(bool,'T_0') = s(bool,sK4('T_0',X93,bool)) | s(bool,X94) = s(bool,'const.bag.BAG_IN_2'(s(bool,X94),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X94))))) | s(bool,X92) = s(bool,X94)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2386,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X95)))) | s(bool,'T_0') != s(bool,X97) | p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X97)))) | s(bool,X98) = s(bool,'const.bag.BAG_IN_2'(s(bool,X98),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X98))))) | s(bool,X95) = s(bool,X98)) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2455,plain,(
  ( ! [X173,X171,X172,X170] : (~p(s(bool,happ(s(fun(bool,bool),X171),s(bool,'T_0')))) | s(bool,X170) != s(bool,X172) | p(s(bool,happ(s(fun(bool,bool),X171),s(bool,X172)))) | s(bool,X173) = s(bool,'const.bag.BAG_IN_2'(s(bool,X173),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X173))))) | s(bool,X170) = s(bool,X173)) ) | $spl3),
  inference(superposition,[],[f149,f1527])).
fof(f2466,plain,(
  ( ! [X212,X213,X211] : (~p(s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X212),s(bool,X211)))) | s(bool,X213) = s(bool,'const.bag.BAG_IN_2'(s(bool,X213),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X213))))) | s(bool,X213) = s(bool,sK4(X211,X212,bool))) ) | $spl3),
  inference(superposition,[],[f148,f1527])).
fof(f397,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X5,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f391,f238])).
fof(f391,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) | s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X5,bool))) ) | $spl3),
  inference(superposition,[],[f147,f238])).
fof(f402,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f401,f112])).
fof(f401,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f397,f113])).
fof(f415,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) | s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X16,bool)) = s(bool,sK4(sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X16,bool),X17,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f147,f402])).
fof(f419,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f411,f110])).
fof(f411,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f156,f402])).
fof(f420,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X11,bool)) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f412,f110])).
fof(f412,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | s(bool,X12) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X11,bool)) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f187,f402])).
fof(f421,plain,(
  ( ! [X14,X13] : (s(bool,X14) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X13,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f413,f110])).
fof(f413,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,X14) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X13,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f204,f402])).
fof(f457,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X15),s(fun(bool,'type.num.num'),X16)))))) | s(bool,X15) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X14,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f151,f402])).
fof(f464,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X16)))))) | s(bool,X15) != s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X14,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f151,f402])).
fof(f570,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X10,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f553,f517])).
fof(f553,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X10,bool))) ) | $spl3),
  inference(superposition,[],[f147,f517])).
fof(f612,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X25)))),X26,bool)) | p(s(bool,happ(s(fun(bool,bool),X26),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f594,f495])).
fof(f594,plain,(
  ( ! [X26,X25] : (p(s(bool,happ(s(fun(bool,bool),X26),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X25))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X25)))),X26,bool))) ) | $spl3),
  inference(superposition,[],[f147,f495])).
fof(f804,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X26)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X25,bool)),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X27),s(fun(bool,'type.num.num'),X26)))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f152,f402])).
fof(f927,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))),X25,bool)) | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f905,f836])).
fof(f905,plain,(
  ( ! [X24,X25] : (p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))),X25,bool))) ) | $spl3),
  inference(superposition,[],[f147,f836])).
fof(f1092,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X50)))))),X51,bool)) | p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'T_0'))))) ) | $spl3),
  inference(forward_demodulation,[],[f1070,f950])).
fof(f1070,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),X51),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X50))))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X49),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X50)))))),X51,bool))) ) | $spl3),
  inference(superposition,[],[f147,f950])).
fof(f1185,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X33),s(fun(bool,'type.num.num'),X34)))))) | s(bool,X33) = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X32,bool)) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X32,bool)),s(fun(bool,'type.num.num'),X34)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f150,f402])).
fof(f1198,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X33),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X34)))))) | s(bool,X33) = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X32,bool)) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X33),s(fun(bool,'type.num.num'),X34)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f150,f402])).
fof(f1248,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f1247,f112])).
fof(f1247,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(forward_demodulation,[],[f1245,f257])).
fof(f1245,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X0,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f570,f442])).
fof(f1615,plain,(
  ( ! [X154,X155,X153] : (p(s(bool,happ(s(fun(bool,bool),X154),s(bool,'T_0')))) | s(bool,X153) = s(bool,sK4(X153,X154,bool)) | s(bool,X153) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f147,f202])).
fof(f2454,plain,(
  ( ! [X167,X169,X168] : (p(s(bool,happ(s(fun(bool,bool),X168),s(bool,'T_0')))) | s(bool,X167) = s(bool,sK4(X167,X168,bool)) | s(bool,X169) = s(bool,'const.bag.BAG_IN_2'(s(bool,X169),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X169))))) | s(bool,X167) = s(bool,X169)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2565,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool))))))) ) | $spl3),
  inference(condensation,[],[f2564])).
fof(f2564,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0')))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2111])).
fof(f2111,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X54,bool)))))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f1527,f402])).
fof(f2657,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))))))) ) | $spl3),
  inference(condensation,[],[f2656])).
fof(f2656,plain,(
  ( ! [X130,X129] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X129,bool)) | s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X130)) ) | $spl3),
  inference(subsumption_resolution,[],[f2241,f116])).
fof(f2241,plain,(
  ( ! [X130,X129] : (p(s(bool,X130)) | s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X129,bool)) | s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X130)) ) | $spl3),
  inference(superposition,[],[f570,f1527])).
fof(f2659,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))))))) ) | $spl3),
  inference(condensation,[],[f2658])).
fof(f2658,plain,(
  ( ! [X132,X131] : (s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X131,bool)) | s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X132)) ) | $spl3),
  inference(subsumption_resolution,[],[f2242,f116])).
fof(f2242,plain,(
  ( ! [X132,X131] : (p(s(bool,X132)) | s(bool,'T_0') = s(bool,sK4(happ(s(fun(sK0,bool),sK1),s(sK0,sK2)),X131,bool)) | s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X132)) ) | $spl3),
  inference(superposition,[],[f397,f1527])).
fof(f399,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,'type.num.num'),bool),X11),s(fun(X9,'type.num.num'),'const.bag.EL_BAG_1'(s(X9,X10)))))) | s(fun(X9,'type.num.num'),'const.bag.EL_BAG_1'(s(X9,X10))) = s(fun(X9,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(X9,X10),s(fun(X9,'type.num.num'),'const.bag.EMPTY_BAG_0')),X11,fun(X9,'type.num.num')))) )),
  inference(forward_demodulation,[],[f393,f138])).
fof(f393,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,'type.num.num'),bool),X11),s(fun(X9,'type.num.num'),'const.bag.EL_BAG_1'(s(X9,X10)))))) | s(fun(X9,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X9,X10),s(fun(X9,'type.num.num'),'const.bag.EMPTY_BAG_0'))) = s(fun(X9,'type.num.num'),sK4('const.bag.BAG_INSERT_2'(s(X9,X10),s(fun(X9,'type.num.num'),'const.bag.EMPTY_BAG_0')),X11,fun(X9,'type.num.num')))) )),
  inference(superposition,[],[f147,f138])).
fof(f770,plain,(
  ( ! [X0] : (s(sK0,sK2) != s(sK0,X0) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X0))))) ) | $spl3),
  inference(subsumption_resolution,[],[f759,f110])).
fof(f759,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(sK0,sK2) != s(sK0,X0) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X0))))) ) | $spl3),
  inference(superposition,[],[f149,f238])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),X1)))) | p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1))))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f2047,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(X119,X120) = s(X119,X121) | p(s(bool,'const.bag.BAG_IN_2'(s(X119,X120),s(fun(X119,'type.num.num'),X122)))) | s(bool,X123) = s(bool,'const.bag.BAG_IN_2'(s(X119,X120),s(fun(X119,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X119,X121),s(fun(X119,'type.num.num'),X122))))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(subsumption_resolution,[],[f1607,f110])).
fof(f1607,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,'T_0')) | s(X119,X120) = s(X119,X121) | p(s(bool,'const.bag.BAG_IN_2'(s(X119,X120),s(fun(X119,'type.num.num'),X122)))) | s(bool,X123) = s(bool,'const.bag.BAG_IN_2'(s(X119,X120),s(fun(X119,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X119,X121),s(fun(X119,'type.num.num'),X122))))) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f150,f202])).
fof(f2048,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,'const.bag.BAG_IN_2'(s(X129,X130),s(fun(X129,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X129,X132),s(fun(X129,'type.num.num'),X131)))))) | s(bool,X133) = s(bool,'const.bag.BAG_IN_2'(s(X129,X130),s(fun(X129,'type.num.num'),X131))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f1609,f110])).
fof(f1609,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(X129,X130),s(fun(X129,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X129,X132),s(fun(X129,'type.num.num'),X131)))))) | s(bool,X133) = s(bool,'const.bag.BAG_IN_2'(s(X129,X130),s(fun(X129,'type.num.num'),X131))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f152,f202])).
fof(f2720,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(X133,X134) = s(X133,X135) | p(s(bool,'const.bag.BAG_IN_2'(s(X133,X134),s(fun(X133,'type.num.num'),X136)))) | s(bool,X137) = s(bool,'const.bag.BAG_IN_2'(s(bool,X137),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X137))))) | s(bool,X137) = s(bool,'const.bag.BAG_IN_2'(s(X133,X134),s(fun(X133,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X133,X135),s(fun(X133,'type.num.num'),X136)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2446,f110])).
fof(f2446,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | s(X133,X134) = s(X133,X135) | p(s(bool,'const.bag.BAG_IN_2'(s(X133,X134),s(fun(X133,'type.num.num'),X136)))) | s(bool,X137) = s(bool,'const.bag.BAG_IN_2'(s(bool,X137),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X137))))) | s(bool,X137) = s(bool,'const.bag.BAG_IN_2'(s(X133,X134),s(fun(X133,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X133,X135),s(fun(X133,'type.num.num'),X136)))))) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2721,plain,(
  ( ! [X146,X144,X147,X145,X143] : (p(s(bool,'const.bag.BAG_IN_2'(s(X143,X144),s(fun(X143,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X143,X146),s(fun(X143,'type.num.num'),X145)))))) | s(bool,X147) = s(bool,'const.bag.BAG_IN_2'(s(bool,X147),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X147))))) | s(bool,X147) = s(bool,'const.bag.BAG_IN_2'(s(X143,X144),s(fun(X143,'type.num.num'),X145)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2448,f110])).
fof(f2448,plain,(
  ( ! [X146,X144,X147,X145,X143] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(X143,X144),s(fun(X143,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X143,X146),s(fun(X143,'type.num.num'),X145)))))) | s(bool,X147) = s(bool,'const.bag.BAG_IN_2'(s(bool,X147),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X147))))) | s(bool,X147) = s(bool,'const.bag.BAG_IN_2'(s(X143,X144),s(fun(X143,'type.num.num'),X145)))) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(superposition,[],[f151,f138])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(X0,X2),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X3),s(fun(X0,'type.num.num'),X1)))))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f101])).
fof(f1436,plain,(
  ( ! [X171,X169,X172,X170] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X170),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X171),s(fun(bool,'type.num.num'),X172)))))) | s(bool,X169) = s(bool,X171) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X169),s(fun(bool,'type.num.num'),X172)))) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f150,f202])).
fof(f1437,plain,(
  ( ! [X175,X173,X176,X174] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X175),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X174),s(fun(bool,'type.num.num'),X176)))))) | s(bool,X173) = s(bool,X175) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X175),s(fun(bool,'type.num.num'),X176)))) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f150,f202])).
fof(f1440,plain,(
  ( ! [X187,X185,X188,X186] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X186),s(fun(bool,'type.num.num'),X187)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X185),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X188),s(fun(bool,'type.num.num'),X187)))))) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f152,f202])).
fof(f1553,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X90),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X89),s(fun(bool,'type.num.num'),X91)))))) | s(bool,'T_0') = s(bool,X90) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X90),s(fun(bool,'type.num.num'),X91)))) | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f150,f202])).
fof(f1556,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X101),s(fun(bool,'type.num.num'),X102)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X103),s(fun(bool,'type.num.num'),X102)))))) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f152,f202])).
fof(f1617,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X161),s(fun(bool,'type.num.num'),X162)))))) | s(bool,X160) = s(bool,X161) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X160),s(fun(bool,'type.num.num'),X162)))) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f150,f202])).
fof(f1618,plain,(
  ( ! [X167,X165,X166,X164] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X165),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X166)))))) | s(bool,X164) = s(bool,X165) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X165),s(fun(bool,'type.num.num'),X166)))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f150,f202])).
fof(f2185,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X79),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X78),s(fun(bool,'type.num.num'),X80)))))) | s(bool,X79) = s(bool,'const.bag.BAG_IN_2'(s(bool,X78),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X78))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X79),s(fun(bool,'type.num.num'),X80)))) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2188,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X90),s(fun(bool,'type.num.num'),X91)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X90),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X90))))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X92),s(fun(bool,'type.num.num'),X91)))))) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f2261,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X188),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X189),s(fun(bool,'type.num.num'),X190)))))) | s(bool,X187) = s(bool,X189) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X187),s(fun(bool,'type.num.num'),X190)))) | s(bool,X187) = s(bool,'const.bag.BAG_IN_2'(s(bool,X187),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X187))))) | s(bool,'T_0') = s(bool,X188)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2262,plain,(
  ( ! [X191,X194,X192,X193] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X193),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X192),s(fun(bool,'type.num.num'),X194)))))) | s(bool,X191) = s(bool,X193) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X193),s(fun(bool,'type.num.num'),X194)))) | s(bool,X191) = s(bool,'const.bag.BAG_IN_2'(s(bool,X191),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X191))))) | s(bool,'T_0') = s(bool,X192)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2265,plain,(
  ( ! [X206,X204,X205,X203] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X204),s(fun(bool,'type.num.num'),X205)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X203),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X206),s(fun(bool,'type.num.num'),X205)))))) | s(bool,X203) = s(bool,'const.bag.BAG_IN_2'(s(bool,X203),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X203))))) | s(bool,'T_0') = s(bool,X204)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f2343,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X188),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X189),s(fun(bool,'type.num.num'),X190)))))) | s(bool,X187) = s(bool,X189) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X187),s(fun(bool,'type.num.num'),X190)))) | s(bool,X188) = s(bool,'const.bag.BAG_IN_2'(s(bool,X188),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X188))))) | s(bool,'T_0') = s(bool,X187)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2344,plain,(
  ( ! [X191,X194,X192,X193] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X193),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X192),s(fun(bool,'type.num.num'),X194)))))) | s(bool,X191) = s(bool,X193) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X193),s(fun(bool,'type.num.num'),X194)))) | s(bool,X192) = s(bool,'const.bag.BAG_IN_2'(s(bool,X192),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X192))))) | s(bool,'T_0') = s(bool,X191)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2347,plain,(
  ( ! [X206,X204,X205,X203] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X204),s(fun(bool,'type.num.num'),X205)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X203),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X206),s(fun(bool,'type.num.num'),X205)))))) | s(bool,X204) = s(bool,'const.bag.BAG_IN_2'(s(bool,X204),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X204))))) | s(bool,'T_0') = s(bool,X203)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f2388,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X104),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X103),s(fun(bool,'type.num.num'),X105)))))) | s(bool,'T_0') = s(bool,X104) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X104),s(fun(bool,'type.num.num'),X105)))) | s(bool,X106) = s(bool,'const.bag.BAG_IN_2'(s(bool,X106),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X106))))) | s(bool,X103) = s(bool,X106)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2391,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X115),s(fun(bool,'type.num.num'),X116)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X117),s(fun(bool,'type.num.num'),X116)))))) | s(bool,X118) = s(bool,'const.bag.BAG_IN_2'(s(bool,X118),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X118))))) | s(bool,X115) = s(bool,X118)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f2456,plain,(
  ( ! [X177,X175,X176,X174] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X175),s(fun(bool,'type.num.num'),X176)))))) | s(bool,X174) = s(bool,X175) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X174),s(fun(bool,'type.num.num'),X176)))) | s(bool,X177) = s(bool,'const.bag.BAG_IN_2'(s(bool,X177),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X177))))) | s(bool,X174) = s(bool,X177)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2457,plain,(
  ( ! [X180,X178,X181,X179] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X179),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X180)))))) | s(bool,X178) = s(bool,X179) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X179),s(fun(bool,'type.num.num'),X180)))) | s(bool,X181) = s(bool,'const.bag.BAG_IN_2'(s(bool,X181),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X181))))) | s(bool,X178) = s(bool,X181)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2184,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X74),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X75),s(fun(bool,'type.num.num'),X76)))))) | s(bool,X75) = s(bool,'const.bag.BAG_IN_2'(s(bool,X74),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X74))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X74),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X74))))),s(fun(bool,'type.num.num'),X76)))) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2186,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X82),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X83),s(fun(bool,'type.num.num'),X84)))))) | s(bool,X83) != s(bool,'const.bag.BAG_IN_2'(s(bool,X82),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X82))))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2187,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X87),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X86),s(fun(bool,'type.num.num'),X88)))))) | s(bool,X87) != s(bool,'const.bag.BAG_IN_2'(s(bool,X86),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X86))))) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2263,plain,(
  ( ! [X198,X196,X197,X195] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X196),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X197),s(fun(bool,'type.num.num'),X198)))))) | s(bool,X195) != s(bool,X197) | s(bool,X195) = s(bool,'const.bag.BAG_IN_2'(s(bool,X195),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X195))))) | s(bool,'T_0') = s(bool,X196)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2264,plain,(
  ( ! [X202,X200,X199,X201] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X201),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X200),s(fun(bool,'type.num.num'),X202)))))) | s(bool,X199) != s(bool,X201) | s(bool,X199) = s(bool,'const.bag.BAG_IN_2'(s(bool,X199),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X199))))) | s(bool,'T_0') = s(bool,X200)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2345,plain,(
  ( ! [X198,X196,X197,X195] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X196),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X197),s(fun(bool,'type.num.num'),X198)))))) | s(bool,X195) != s(bool,X197) | s(bool,X196) = s(bool,'const.bag.BAG_IN_2'(s(bool,X196),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X196))))) | s(bool,'T_0') = s(bool,X195)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2346,plain,(
  ( ! [X202,X200,X199,X201] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X201),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X200),s(fun(bool,'type.num.num'),X202)))))) | s(bool,X199) != s(bool,X201) | s(bool,X200) = s(bool,'const.bag.BAG_IN_2'(s(bool,X200),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X200))))) | s(bool,'T_0') = s(bool,X199)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2355,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X0) = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f110,f1527])).
fof(f2387,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X99),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X100),s(fun(bool,'type.num.num'),X101)))))) | s(bool,'T_0') = s(bool,X100) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X101)))) | s(bool,X102) = s(bool,'const.bag.BAG_IN_2'(s(bool,X102),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X102))))) | s(bool,X99) = s(bool,X102)) ) | $spl3),
  inference(superposition,[],[f150,f1527])).
fof(f2389,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X107),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X108),s(fun(bool,'type.num.num'),X109)))))) | s(bool,'T_0') != s(bool,X108) | s(bool,X110) = s(bool,'const.bag.BAG_IN_2'(s(bool,X110),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X110))))) | s(bool,X107) = s(bool,X110)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2390,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X112),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X111),s(fun(bool,'type.num.num'),X113)))))) | s(bool,'T_0') != s(bool,X112) | s(bool,X114) = s(bool,'const.bag.BAG_IN_2'(s(bool,X114),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X114))))) | s(bool,X111) = s(bool,X114)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2458,plain,(
  ( ! [X185,X182,X184,X183] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X183),s(fun(bool,'type.num.num'),X184)))))) | s(bool,X182) != s(bool,X183) | s(bool,X185) = s(bool,'const.bag.BAG_IN_2'(s(bool,X185),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X185))))) | s(bool,X182) = s(bool,X185)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2459,plain,(
  ( ! [X189,X187,X188,X186] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X187),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X188)))))) | s(bool,X186) != s(bool,X187) | s(bool,X189) = s(bool,'const.bag.BAG_IN_2'(s(bool,X189),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X189))))) | s(bool,X186) = s(bool,X189)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2460,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X191)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X190),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X192),s(fun(bool,'type.num.num'),X191)))))) | s(bool,X193) = s(bool,'const.bag.BAG_IN_2'(s(bool,X193),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X193))))) | s(bool,X190) = s(bool,X193)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f2586,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2585])).
fof(f2585,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(condensation,[],[f2086])).
fof(f2086,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,'const.bag.BAG_IN_2'(s(bool,X17),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X17),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(bool,X17),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X17))))))))) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f2630,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl3),
  inference(condensation,[],[f2629])).
fof(f2629,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X9),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X9))))) | s(bool,'T_0') = s(bool,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f2160,f204])).
fof(f2160,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X9),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X9))))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | $spl3),
  inference(superposition,[],[f116,f1527])).
fof(f2632,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X22)) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f2165,f237])).
fof(f2165,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f156,f1527])).
fof(f2633,plain,(
  ( ! [X26,X24,X25] : (s(bool,X25) != s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | p(s(bool,X25)) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl3),
  inference(subsumption_resolution,[],[f2166,f237])).
fof(f2166,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | s(bool,X25) != s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | p(s(bool,X25)) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) ) | $spl3),
  inference(superposition,[],[f187,f1527])).
fof(f2635,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1)))))) ) | $spl3),
  inference(condensation,[],[f2634])).
fof(f2634,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | s(bool,X40) != s(bool,'const.bag.BAG_IN_2'(s(bool,X39),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X39))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X41)) ) | $spl3),
  inference(subsumption_resolution,[],[f2171,f204])).
fof(f2171,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | s(bool,X40) != s(bool,'const.bag.BAG_IN_2'(s(bool,X39),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X39))))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl3),
  inference(superposition,[],[f204,f1527])).
fof(f2686,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s(bool,X5) = s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5)))))) ) | $spl3),
  inference(global_subsumption,[],[f109,f110,f112,f113,f116,f115,f120,f119,f138,f141,f146,f145,f144,f149,f148,f147,f152,f151,f150,f156,f153,f161,f187,f190,f198,f197,f202,f204,f218,f228,f237,f241,f238,f246,f258,f257,f305,f362,f388,f389,f396,f397,f398,f399,f400,f402,f419,f420,f421,f415,f423,f425,f428,f442,f443,f469,f470,f472,f452,f453,f457,f459,f460,f464,f467,f485,f495,f497,f517,f526,f530,f531,f532,f567,f570,f609,f612,f613,f762,f750,f757,f770,f783,f812,f797,f804,f821,f822,f836,f824,f926,f927,f936,f937,f950,f938,f1091,f1092,f1111,f1097,f1207,f1176,f1185,f1226,f1189,f1198,f1238,f1248,f1344,f1434,f1435,f1436,f1437,f1438,f1439,f1440,f1443,f1444,f1445,f1446,f1447,f1525,f1526,f1527,f1528,f1529,f1530,f1531,f1532,f1533,f1550,f1551,f1552,f1553,f1554,f1555,f1556,f1558,f2039,f2040,f2041,f2043,f2044,f2045,f2046,f2047,f2048,f1615,f1616,f1617,f1618,f1619,f1620,f1621,f1624,f1625,f1626,f1627,f1628,f2595,f2080,f2590,f2586,f2579,f2565,f2563,f2558,f2557,f2625,f2135,f2136,f2548,f2546,f2626,f2627,f2628,f2630,f2161,f2631,f2632,f2633,f2635,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2636,f2639,f2641,f2643,f2654,f2657,f2659,f2660,f2662,f2664,f2259,f2260,f2261,f2262,f2263,f2264,f2265,f2268,f2269,f2270,f2271,f2272,f2520])).
fof(f2699,plain,(
  ( ! [X97,X96] : (~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X96) = s(bool,'const.bag.BAG_IN_2'(s(bool,X96),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X96)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2310,f190])).
fof(f2310,plain,(
  ( ! [X97,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X97)) | s(bool,X96) = s(bool,'const.bag.BAG_IN_2'(s(bool,X96),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl3),
  inference(superposition,[],[f119,f1527])).
fof(f2712,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,'const.bag.BAG_IN_2'(s(bool,X18),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X18))))) | s(bool,X16) = s(bool,X18)) ) | $spl3),
  inference(subsumption_resolution,[],[f2401,f110])).
fof(f2401,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,'const.bag.BAG_IN_2'(s(bool,X18),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X18))))) | s(bool,X16) = s(bool,X18)) ) | $spl3),
  inference(superposition,[],[f156,f1527])).
fof(f2713,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X21) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,X19) = s(bool,X21)) ) | $spl3),
  inference(subsumption_resolution,[],[f2402,f110])).
fof(f2402,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X21) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,X19) = s(bool,X21)) ) | $spl3),
  inference(superposition,[],[f187,f1527])).
fof(f1438,plain,(
  ( ! [X177,X180,X178,X179] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X178),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X179),s(fun(bool,'type.num.num'),X180)))))) | s(bool,X177) != s(bool,X179) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X177)) )),
  inference(superposition,[],[f151,f202])).
fof(f1439,plain,(
  ( ! [X182,X184,X183,X181] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X183),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X182),s(fun(bool,'type.num.num'),X184)))))) | s(bool,X181) != s(bool,X183) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f151,f202])).
fof(f1552,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X85),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X86),s(fun(bool,'type.num.num'),X87)))))) | s(bool,'T_0') = s(bool,X86) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X87)))) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f150,f202])).
fof(f1554,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X93),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X94),s(fun(bool,'type.num.num'),X95)))))) | s(bool,'T_0') != s(bool,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f151,f202])).
fof(f1555,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X98),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X97),s(fun(bool,'type.num.num'),X99)))))) | s(bool,'T_0') != s(bool,X98) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f151,f202])).
fof(f1621,plain,(
  ( ! [X177,X178,X176,X179] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X177)))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X176),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X178),s(fun(bool,'type.num.num'),X177)))))) | s(bool,X176) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f152,f202])).
fof(f2628,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f2157,f237])).
fof(f2157,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,X2)) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f152,f1527])).
fof(f1620,plain,(
  ( ! [X175,X173,X174,X172] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X173),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X174)))))) | s(bool,X172) != s(bool,X173) | s(bool,X172) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f151,f202])).
fof(f1619,plain,(
  ( ! [X171,X169,X170,X168] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X169),s(fun(bool,'type.num.num'),X170)))))) | s(bool,X168) != s(bool,X169) | s(bool,X168) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) )),
  inference(superposition,[],[f151,f202])).
fof(f938,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X2))))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f822,f152])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f824,f152])).
fof(f2272,plain,(
  ( ! [X231,X229,X230,X228] : (s(X230,f(s(bool,X231))) != s(X230,g(s(bool,X231))) | s(fun(bool,X230),X228) = s(fun(bool,X230),X229) | s(bool,sK5(X228,X229,X230,bool)) = s(bool,'const.bag.BAG_IN_2'(s(bool,sK5(X228,X229,X230,bool)),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,sK5(X228,X229,X230,bool)))))) | s(bool,'T_0') = s(bool,X231)) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.BAG_INSERT_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),X2))))))) )),
  inference(forward_demodulation,[],[f451,f138])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(fun(X0,'type.num.num'),'type.num.num'),'const.bag.BAG_INSERT_2'(s(fun(X0,'type.num.num'),'const.bag.EL_BAG_1'(s(X0,X1))),s(fun(fun(X0,'type.num.num'),'type.num.num'),X2))))))) )),
  inference(unit_resulting_resolution,[],[f138,f151])).
fof(f2470,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(equality_factoring,[],[f1527])).
fof(f2490,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2468])).
fof(f2468,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,X0) | s(bool,X0) = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl3),
  inference(equality_factoring,[],[f1527])).
fof(f2467,plain,(
  ( ! [X216,X215,X217,X214] : (s(X216,f(s(bool,'T_0'))) != s(X216,g(s(bool,'T_0'))) | s(fun(bool,X216),X214) = s(fun(bool,X216),X215) | s(bool,X217) = s(bool,'const.bag.BAG_IN_2'(s(bool,X217),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X217))))) | s(bool,X217) = s(bool,sK5(X214,X215,X216,bool))) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f2714,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,'const.bag.BAG_IN_2'(s(bool,X32),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X32))))) | s(bool,X30) = s(bool,X32)) ) | $spl3),
  inference(subsumption_resolution,[],[f2405,f110])).
fof(f2405,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,'const.bag.BAG_IN_2'(s(bool,X32),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X32))))) | s(bool,X30) = s(bool,X32)) ) | $spl3),
  inference(superposition,[],[f204,f1527])).
fof(f2393,plain,(
  ( ! [X121,X122] : (s(bool,X121) != s(bool,'const.bag.BAG_IN_2'(s(bool,X121),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0'))) | s(bool,X122) = s(bool,'const.bag.BAG_IN_2'(s(bool,X122),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X122))))) | s(bool,X121) = s(bool,X122)) ) | $spl3),
  inference(superposition,[],[f198,f1527])).
fof(f2366,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) = s(bool,'const.bag.BAG_IN_2'(s(bool,X40),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X41),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X40),s(fun(bool,'type.num.num'),X42))))))) | s(bool,X43) = s(bool,'const.bag.BAG_IN_2'(s(bool,X43),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X43))))) | s(bool,X40) = s(bool,X43)) ) | $spl3),
  inference(superposition,[],[f950,f1527])).
fof(f2365,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X37) != s(bool,'const.bag.BAG_IN_2'(s(bool,X36),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X38),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X36))))))) | s(bool,'T_0') = s(bool,X37) | s(bool,X39) = s(bool,'const.bag.BAG_IN_2'(s(bool,X39),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X39))))) | s(bool,X36) = s(bool,X39)) ) | $spl3),
  inference(superposition,[],[f926,f1527])).
fof(f2364,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) = s(bool,'const.bag.BAG_IN_2'(s(bool,X33),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X34),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X33))))))) | s(bool,X35) = s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,X33) = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f836,f1527])).
fof(f2363,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) != s(bool,'const.bag.BAG_IN_2'(s(bool,X29),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X29),s(fun(bool,'type.num.num'),X31))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,'const.bag.BAG_IN_2'(s(bool,X32),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X32))))) | s(bool,X29) = s(bool,X32)) ) | $spl3),
  inference(superposition,[],[f609,f1527])).
fof(f2361,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,'const.bag.BAG_IN_2'(s(bool,X23),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X23))))) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,X23) = s(bool,X25)) ) | $spl3),
  inference(superposition,[],[f567,f1527])).
fof(f2360,plain,(
  ( ! [X21,X22] : (s(bool,X21) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X22))))) | s(bool,X21) = s(bool,X22)) ) | $spl3),
  inference(superposition,[],[f517,f1527])).
fof(f2359,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) = s(bool,'const.bag.BAG_IN_2'(s(bool,X18),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X18),s(fun(bool,'type.num.num'),X19))))) | s(bool,X20) = s(bool,'const.bag.BAG_IN_2'(s(bool,X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X20))))) | s(bool,X18) = s(bool,X20)) ) | $spl3),
  inference(superposition,[],[f495,f1527])).
fof(f2354,plain,(
  ( ! [X231,X229,X230,X228] : (s(X230,f(s(bool,X231))) != s(X230,g(s(bool,X231))) | s(fun(bool,X230),X228) = s(fun(bool,X230),X229) | s(bool,X231) = s(bool,'const.bag.BAG_IN_2'(s(bool,X231),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X231))))) | s(bool,'T_0') = s(bool,sK5(X228,X229,X230,bool))) ) | $spl3),
  inference(superposition,[],[f161,f1527])).
fof(f2694,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,X41) | s(bool,X41) = s(bool,'const.bag.BAG_IN_2'(s(bool,X41),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X41)))))) ) | $spl3),
  inference(global_subsumption,[],[f109,f110,f112,f113,f116,f115,f120,f119,f138,f141,f146,f145,f144,f149,f148,f147,f152,f151,f150,f156,f153,f161,f187,f190,f198,f197,f202,f204,f218,f228,f237,f241,f238,f246,f258,f257,f305,f362,f388,f389,f396,f397,f398,f399,f400,f402,f419,f420,f421,f415,f423,f425,f428,f442,f443,f469,f470,f472,f452,f453,f457,f459,f460,f464,f467,f485,f495,f497,f517,f526,f530,f531,f532,f567,f570,f609,f612,f613,f762,f750,f757,f770,f783,f812,f797,f804,f821,f822,f836,f824,f926,f927,f936,f937,f950,f938,f1091,f1092,f1111,f1097,f1207,f1176,f1185,f1226,f1189,f1198,f1238,f1248,f1344,f1434,f1435,f1436,f1437,f1438,f1439,f1440,f1443,f1444,f1445,f1446,f1447,f1525,f1526,f1527,f1528,f1529,f1530,f1531,f1532,f1533,f1550,f1551,f1552,f1553,f1554,f1555,f1556,f1558,f2039,f2040,f2041,f2043,f2044,f2045,f2046,f2047,f2048,f1615,f1616,f1617,f1618,f1619,f1620,f1621,f1624,f1625,f1626,f1627,f1628,f2595,f2080,f2590,f2586,f2579,f2565,f2563,f2558,f2557,f2625,f2135,f2136,f2548,f2546,f2626,f2627,f2628,f2630,f2161,f2631,f2632,f2633,f2635,f2182,f2183,f2184,f2185,f2186,f2187,f2188,f2636,f2639,f2641,f2643,f2654,f2657,f2659,f2660,f2662,f2664,f2259,f2260,f2261,f2262,f2263,f2264,f2265,f2268,f2269,f2270,f2271,f2272,f2520,f2686,f2693])).
fof(f2693,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.bag.BAG_IN_2'(s(bool,X41),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X41)))))) ) | $spl3),
  inference(forward_demodulation,[],[f2515,f238])).
fof(f2515,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.bag.BAG_IN_2'(s(bool,X41),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X41)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2288])).
fof(f2288,plain,(
  ( ! [X41,X40] : (s(bool,X41) != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.bag.BAG_IN_2'(s(bool,X41),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X41))))) | s(bool,'T_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f362,f1527])).
fof(f2520,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2275])).
fof(f2275,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f116,f1527])).
fof(f2631,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12)) ) | $spl3),
  inference(global_subsumption,[],[f2161,f109,f110,f112,f113,f116,f115,f120,f119,f138,f141,f146,f145,f144,f149,f148,f147,f152,f151,f150,f156,f153,f161,f187,f190,f198,f197,f202,f204,f218,f237,f305,f388,f389,f398,f399,f400,f423,f443,f469,f470,f452,f453,f459,f460,f467,f485,f526,f530,f531,f532,f613,f750,f783,f797,f1176,f1189,f1238,f1344,f1434,f1435,f1436,f1437,f1438,f1439,f1440,f1443,f1444,f1445,f1446,f1447,f1550,f1551,f1552,f1553,f1554,f1555,f1556,f1558,f2039,f2040,f2041,f2043,f2044,f2045,f2046,f2047,f2048,f1615,f1616,f1617,f1618,f1619,f1620,f1621,f1624,f1625,f1626,f1627,f1628])).
fof(f2161,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') != s(bool,X11) | ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X11),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X11)))))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl3),
  inference(superposition,[],[f119,f1527])).
fof(f2627,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.bag.BAG_IN_2'(s(bool,X29),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X29))))) | s(bool,X29) = s(bool,'const.bag.BAG_IN_2'(s(X27,X28),s(fun(X27,'type.num.num'),'const.bag.EMPTY_BAG_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f2154,f190])).
fof(f2154,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,'const.bag.BAG_IN_2'(s(bool,X29),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X29))))) | s(bool,X29) = s(bool,'const.bag.BAG_IN_2'(s(X27,X28),s(fun(X27,'type.num.num'),'const.bag.EMPTY_BAG_0')))) ) | $spl3),
  inference(superposition,[],[f1527,f197])).
fof(f2626,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,'const.bag.BAG_IN_2'(s(bool,X26),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X26))))) | s(bool,X26) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2153,f190])).
fof(f2153,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X26) = s(bool,'const.bag.BAG_IN_2'(s(bool,X26),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X26))))) | s(bool,X26) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) ) | $spl3),
  inference(superposition,[],[f1527,f257])).
fof(f2136,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6))))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(superposition,[],[f1527,f202])).
fof(f2135,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2))))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3))))) | s(bool,X0) = s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f2590,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2589])).
fof(f2589,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl3),
  inference(condensation,[],[f2082])).
fof(f2082,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) ) | $spl3),
  inference(superposition,[],[f1527,f202])).
fof(f2080,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1)))))) ) | $spl3),
  inference(superposition,[],[f1527,f113])).
fof(f1628,plain,(
  ( ! [X202,X200,X203,X201] : (s(X202,f(s(bool,'T_0'))) != s(X202,g(s(bool,'T_0'))) | s(fun(bool,X202),X200) = s(fun(bool,X202),X201) | s(bool,X203) = s(bool,sK5(X200,X201,X202,bool)) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f161,f202])).
fof(f2041,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1568,f110])).
fof(f1568,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f204,f202])).
fof(f2040,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1567,f110])).
fof(f1567,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f187,f202])).
fof(f2039,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1566,f110])).
fof(f1566,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f156,f202])).
fof(f1558,plain,(
  ( ! [X107,X108] : (s(bool,X107) != s(bool,'const.bag.BAG_IN_2'(s(bool,X107),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0'))) | s(bool,X107) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f198,f202])).
fof(f1533,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) = s(bool,'const.bag.BAG_IN_2'(s(bool,X32),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X33),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X32),s(fun(bool,'type.num.num'),X34))))))) | s(bool,X32) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl3),
  inference(superposition,[],[f950,f202])).
fof(f1532,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,'const.bag.BAG_IN_2'(s(bool,X28),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X30),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X28))))))) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl3),
  inference(superposition,[],[f926,f202])).
fof(f1531,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X26),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))))) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | $spl3),
  inference(superposition,[],[f836,f202])).
fof(f1530,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) != s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X21),s(fun(bool,'type.num.num'),X23))))) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f609,f202])).
fof(f1528,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) != s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15))))) | s(bool,'T_0') = s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | $spl3),
  inference(superposition,[],[f567,f202])).
fof(f1527,plain,(
  ( ! [X14,X13] : (s(bool,X13) = s(bool,'const.bag.BAG_IN_2'(s(bool,X13),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X13))))) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f517,f202])).
fof(f1526,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,'const.bag.BAG_IN_2'(s(bool,X10),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X10),s(fun(bool,'type.num.num'),X11))))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl3),
  inference(superposition,[],[f495,f202])).
fof(f1447,plain,(
  ( ! [X212,X210,X213,X211] : (s(X212,f(s(bool,X213))) != s(X212,g(s(bool,X213))) | s(fun(bool,X212),X210) = s(fun(bool,X212),X211) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,sK5(X210,X211,X212,bool))) )),
  inference(superposition,[],[f161,f202])).
fof(f1344,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f202,f202])).
fof(f1111,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f824,f116])).
fof(f1091,plain,(
  ( ! [X30,X31,X32] : (s(bool,X32) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X30),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X31))))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl3),
  inference(subsumption_resolution,[],[f1062,f110])).
fof(f1062,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X32) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X30),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X31))))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl3),
  inference(superposition,[],[f204,f950])).
fof(f950,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1)))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f822,f116])).
fof(f926,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X16),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl3),
  inference(subsumption_resolution,[],[f899,f110])).
fof(f899,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | s(bool,X17) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X16),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl3),
  inference(superposition,[],[f204,f836])).
fof(f824,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f821,f152])).
fof(f836,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f821,f116])).
fof(f822,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X4))))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f809,f110])).
fof(f809,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X4))))))))) ) | $spl3),
  inference(superposition,[],[f152,f495])).
fof(f821,plain,(
  ( ! [X3] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))) ) | $spl3),
  inference(subsumption_resolution,[],[f808,f110])).
fof(f808,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))) ) | $spl3),
  inference(superposition,[],[f152,f517])).
fof(f783,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(unit_resulting_resolution,[],[f531,f116])).
fof(f609,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X14))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl3),
  inference(subsumption_resolution,[],[f589,f110])).
fof(f589,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | s(bool,X15) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X14))))) | s(bool,'T_0') = s(bool,X15)) ) | $spl3),
  inference(superposition,[],[f204,f495])).
fof(f567,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f548,f110])).
fof(f548,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f204,f517])).
fof(f531,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) )),
  inference(subsumption_resolution,[],[f528,f112])).
fof(f528,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X0,bool),X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))))) )),
  inference(superposition,[],[f145,f197])).
fof(f517,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f497,f116])).
fof(f497,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl3),
  inference(superposition,[],[f472,f138])).
fof(f495,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f472,f116])).
fof(f485,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X2),s(fun(X0,'type.num.num'),X3)))))) )),
  inference(subsumption_resolution,[],[f468,f112])).
fof(f468,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X0,X2),s(fun(X0,'type.num.num'),X3)))))) )),
  inference(superposition,[],[f151,f113])).
fof(f472,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0))))))) ) | $spl3),
  inference(forward_demodulation,[],[f448,f238])).
fof(f448,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f238,f151])).
fof(f443,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.bag.BAG_IN_2'(s(X6,X7),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f440,f112])).
fof(f440,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'const.bag.BAG_IN_2'(s(X6,X7),s(fun(X6,'type.num.num'),'const.bag.EMPTY_BAG_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f237,f197])).
fof(f442,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f438,f112])).
fof(f438,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f237,f257])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(subsumption_resolution,[],[f395,f112])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f147,f113])).
fof(f362,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f358,f110])).
fof(f358,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f204,f238])).
fof(f305,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EMPTY_BAG_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f197,f113])).
fof(f258,plain,(
  ~p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | $spl3),
  inference(unit_resulting_resolution,[],[f246,f116])).
fof(f246,plain,(
  s(bool,'T_0') != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl3),
  inference(backward_demodulation,[],[f238,f109])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f228,f113])).
fof(f241,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl3),
  inference(unit_resulting_resolution,[],[f112,f228,f153])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f119,f113])).
fof(f202,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f113,f113])).
fof(f198,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) )),
  inference(unit_resulting_resolution,[],[f141,f115])).
fof(f190,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f112,f115])).
fof(f187,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f186])).
fof(f186,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f175])).
fof(f175,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8,X19,X16,X20] : ((s(bool,X8) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X16) = s(bool,X20))) => ((p(s(bool,X8)) => p(s(bool,X16))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.IMP_CONG')).
fof(f161,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.okasaki_misc.BAG_EVERY_EL',aHL_EXT)).
fof(f115,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.okasaki_misc.BAG_EVERY_EL','thm.bool.EQ_CLAUSES|split|1')).
fof(f116,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f11280,plain,(
  $false | ($spl3 | $spl41)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663,f228])).
fof(f11279,plain,(
  $false | ($spl3 | $spl8 | $spl41)),
  inference(global_subsumption,[],[f2623,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11278,plain,(
  $false | ($spl3 | $spl41 | $spl64)),
  inference(global_subsumption,[],[f10890,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11277,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11148,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11148,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl34),
  inference(trivial_inequality_removal,[],[f11141])).
fof(f11141,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) ) | $spl34),
  inference(superposition,[],[f3867,f113])).
fof(f11276,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f2640,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11275,plain,(
  $false | ($spl3 | $spl18 | $spl41)),
  inference(global_subsumption,[],[f2642,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11274,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f3867,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11273,plain,(
  $false | ($spl3 | $spl41)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109,f4663])).
fof(f11272,plain,(
  $false | ($spl3 | $spl22 | $spl41)),
  inference(global_subsumption,[],[f2661,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11271,plain,(
  $false | ($spl3 | $spl20 | $spl41)),
  inference(global_subsumption,[],[f2655,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11270,plain,(
  $false | ($spl3 | $spl26 | $spl41)),
  inference(global_subsumption,[],[f2665,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11269,plain,(
  $false | ($spl3 | $spl24 | $spl41)),
  inference(global_subsumption,[],[f2663,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11268,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11170,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11170,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(condensation,[],[f11169])).
fof(f11169,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3)))))) ) | $spl34),
  inference(subsumption_resolution,[],[f11143,f113])).
fof(f11143,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X3),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X3)))))) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11267,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11172,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11172,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl34),
  inference(condensation,[],[f11171])).
fof(f11171,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl34),
  inference(subsumption_resolution,[],[f11144,f113])).
fof(f11144,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,X6) | s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,X5),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11266,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f4705,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f4705,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl16),
  inference(unit_resulting_resolution,[],[f110,f2640])).
fof(f11265,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f11264,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11264,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(subsumption_resolution,[],[f4707,f110])).
fof(f4707,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(superposition,[],[f2640,f202])).
fof(f11263,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f11262,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11262,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl16),
  inference(subsumption_resolution,[],[f4710,f110])).
fof(f4710,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl16),
  inference(superposition,[],[f2640,f202])).
fof(f11261,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f4726,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f4726,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ~p(s(bool,X0))) ) | $spl16),
  inference(condensation,[],[f4725])).
fof(f4725,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl16),
  inference(subsumption_resolution,[],[f4712,f204])).
fof(f4712,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) ) | $spl16),
  inference(superposition,[],[f2640,f202])).
fof(f11260,plain,(
  $false | ($spl3 | $spl16 | $spl41)),
  inference(global_subsumption,[],[f4728,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f4728,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ~p(s(bool,X0))) ) | $spl16),
  inference(condensation,[],[f4727])).
fof(f4727,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X13)) ) | $spl16),
  inference(subsumption_resolution,[],[f4713,f204])).
fof(f4713,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl16),
  inference(superposition,[],[f2640,f202])).
fof(f11259,plain,(
  $false | ($spl3 | $spl41)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327,f109])).
fof(f11258,plain,(
  $false | ($spl3 | $spl41)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752,f6327])).
fof(f11254,plain,(
  $false | ($spl3 | $spl16 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f6587,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f6587,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl16 | $spl36)),
  inference(backward_demodulation,[],[f6572,f6578])).
fof(f6578,plain,(
  s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f6318,f2640])).
fof(f6318,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | $spl36),
  inference(forward_demodulation,[],[f6296,f6271])).
fof(f6296,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0))))))) | $spl36),
  inference(unit_resulting_resolution,[],[f6271,f467])).
fof(f6572,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl36),
  inference(unit_resulting_resolution,[],[f110,f6318,f156])).
fof(f11253,plain,(
  $false | ($spl3 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f6654,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f6654,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl36),
  inference(forward_demodulation,[],[f6601,f6572])).
fof(f6601,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))))))) | $spl36),
  inference(unit_resulting_resolution,[],[f6572,f467])).
fof(f11252,plain,(
  $false | ($spl3 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f6659,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f6659,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl36),
  inference(forward_demodulation,[],[f6597,f6572])).
fof(f6597,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl36),
  inference(unit_resulting_resolution,[],[f6572,f467])).
fof(f11251,plain,(
  $false | ($spl3 | $spl16 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f6670,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f6670,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl16 | $spl36)),
  inference(forward_demodulation,[],[f6669,f6572])).
fof(f6669,plain,(
  s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f6625,f110])).
fof(f6625,plain,(
  ~p(s(bool,'T_0')) | s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl16 | $spl36)),
  inference(superposition,[],[f2640,f6572])).
fof(f11247,plain,(
  $false | ($spl3 | $spl16 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f6991,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f6991,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl16 | $spl36)),
  inference(backward_demodulation,[],[f6972,f6978])).
fof(f6978,plain,(
  ( ! [X0] : (s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))))))))) ) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f6567,f2640])).
fof(f6567,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f6318,f152])).
fof(f6972,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.BAG_INSERT_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f110,f6567,f156])).
fof(f11246,plain,(
  $false | ($spl3 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f7211,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7211,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl42),
  inference(forward_demodulation,[],[f7153,f7138])).
fof(f7153,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0')))))))) | $spl42),
  inference(unit_resulting_resolution,[],[f7138,f467])).
fof(f11245,plain,(
  $false | ($spl3 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f7216,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7216,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl42),
  inference(forward_demodulation,[],[f7149,f7138])).
fof(f7149,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl42),
  inference(unit_resulting_resolution,[],[f7138,f467])).
fof(f11244,plain,(
  $false | ($spl3 | $spl16 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f7227,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7227,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl16 | $spl42)),
  inference(forward_demodulation,[],[f7226,f7138])).
fof(f7226,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0'))))))) | ($spl16 | $spl42)),
  inference(subsumption_resolution,[],[f7182,f110])).
fof(f7182,plain,(
  ~p(s(bool,'T_0')) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0'))))))) | ($spl16 | $spl42)),
  inference(superposition,[],[f2640,f7138])).
fof(f11243,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7266,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7266,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl51),
  inference(unit_resulting_resolution,[],[f7144,f1238])).
fof(f11242,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7267,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7267,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl51),
  inference(unit_resulting_resolution,[],[f110,f7144,f156])).
fof(f11241,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7270,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7270,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'F_0')))))))) | $spl51),
  inference(unit_resulting_resolution,[],[f7144,f1238])).
fof(f11240,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f11239,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11239,plain,(
  ( ! [X2] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,f(s(bool,'F_0'))),s(fun(bool,'type.num.num'),X2))))))) ) | $spl51),
  inference(forward_demodulation,[],[f7260,f197])).
fof(f7260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,f(s(bool,'F_0'))),s(fun(bool,'type.num.num'),X2))))))) ) | $spl51),
  inference(unit_resulting_resolution,[],[f197,f7144,f460])).
fof(f11238,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7298,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7298,plain,(
  ( ! [X2] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X2))))))) ) | $spl51),
  inference(forward_demodulation,[],[f7258,f197])).
fof(f7258,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0'))),s(fun(bool,'type.num.num'),X2))))))) ) | $spl51),
  inference(unit_resulting_resolution,[],[f197,f7144,f453])).
fof(f11237,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7256,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7256,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f112,f7144,f237])).
fof(f11236,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7301,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7301,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(forward_demodulation,[],[f7253,f197])).
fof(f7253,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) ) | $spl51),
  inference(unit_resulting_resolution,[],[f198,f7144,f202])).
fof(f11235,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7250,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7250,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f190,f7144,f202])).
fof(f11234,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7303,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7303,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(forward_demodulation,[],[f7248,f197])).
fof(f7248,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(X0,X1),s(fun(X0,'type.num.num'),'const.bag.EMPTY_BAG_0')))) ) | $spl51),
  inference(unit_resulting_resolution,[],[f198,f7144,f202])).
fof(f11233,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7245,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7245,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f190,f7144,f202])).
fof(f11232,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7244,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7244,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl51),
  inference(unit_resulting_resolution,[],[f7144,f116])).
fof(f11231,plain,(
  $false | ($spl3 | $spl41 | $spl51)),
  inference(global_subsumption,[],[f7243,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f7243,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f7144,f113])).
fof(f11230,plain,(
  $false | ($spl3 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f8919,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f8919,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl57),
  inference(forward_demodulation,[],[f8832,f8735])).
fof(f8735,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(unit_resulting_resolution,[],[f112,f8624,f237])).
fof(f8832,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0')))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f8735,f467])).
fof(f11229,plain,(
  $false | ($spl3 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f8924,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f8924,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl57),
  inference(forward_demodulation,[],[f8828,f8735])).
fof(f8828,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl57),
  inference(unit_resulting_resolution,[],[f8735,f467])).
fof(f11228,plain,(
  $false | ($spl3 | $spl16 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f8937,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f8937,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl16 | $spl57)),
  inference(forward_demodulation,[],[f8936,f8735])).
fof(f8936,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0'))))))) | ($spl16 | $spl57)),
  inference(subsumption_resolution,[],[f8872,f110])).
fof(f8872,plain,(
  ~p(s(bool,'T_0')) | s(bool,f(s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0'))))))) | ($spl16 | $spl57)),
  inference(superposition,[],[f2640,f8735])).
fof(f11227,plain,(
  $false | ($spl3 | $spl12 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f10953,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10953,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl36)),
  inference(forward_demodulation,[],[f10897,f6572])).
fof(f10897,plain,(
  s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl12 | $spl36)),
  inference(unit_resulting_resolution,[],[f6572,f2637])).
fof(f11226,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f10956,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10956,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl42)),
  inference(forward_demodulation,[],[f10894,f7138])).
fof(f10894,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0'))))))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f7138,f2637])).
fof(f11225,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f10957,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10957,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl57)),
  inference(forward_demodulation,[],[f10893,f8735])).
fof(f10893,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0'))))))) | ($spl12 | $spl57)),
  inference(unit_resulting_resolution,[],[f8735,f2637])).
fof(f11224,plain,(
  $false | ($spl3 | $spl12 | $spl41)),
  inference(global_subsumption,[],[f10949,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10949,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f10906])).
fof(f10906,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(superposition,[],[f2637,f202])).
fof(f11223,plain,(
  $false | ($spl3 | $spl12 | $spl41)),
  inference(global_subsumption,[],[f10948,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10948,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f10909])).
fof(f10909,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X7) = s(bool,'const.bag.BAG_IN_2'(s(bool,X7),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X7))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f2637,f202])).
fof(f11222,plain,(
  $false | ($spl3 | $spl12 | $spl41)),
  inference(global_subsumption,[],[f10947,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10947,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl12),
  inference(duplicate_literal_removal,[],[f10946])).
fof(f10946,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl12),
  inference(condensation,[],[f10945])).
fof(f10945,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f10912])).
fof(f10912,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X12) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl12),
  inference(superposition,[],[f2637,f202])).
fof(f11221,plain,(
  $false | ($spl3 | $spl12 | $spl41)),
  inference(global_subsumption,[],[f10944,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10944,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl12),
  inference(duplicate_literal_removal,[],[f10943])).
fof(f10943,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl12),
  inference(condensation,[],[f10942])).
fof(f10942,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f10913])).
fof(f10913,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X14) | s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl12),
  inference(superposition,[],[f2637,f202])).
fof(f11220,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f10959,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10959,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl57)),
  inference(forward_demodulation,[],[f10941,f8735])).
fof(f10941,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0'))))))) | ($spl12 | $spl57)),
  inference(trivial_inequality_removal,[],[f10918])).
fof(f10918,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,f(s(bool,'T_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,f(s(bool,'T_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,f(s(bool,'T_0'))))))) | ($spl12 | $spl57)),
  inference(superposition,[],[f2637,f8735])).
fof(f11219,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f10960,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10960,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl42)),
  inference(forward_demodulation,[],[f10940,f7138])).
fof(f10940,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0'))))))) | ($spl12 | $spl42)),
  inference(trivial_inequality_removal,[],[f10925])).
fof(f10925,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,g(s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,g(s(bool,'F_0'))))))) | ($spl12 | $spl42)),
  inference(superposition,[],[f2637,f7138])).
fof(f11218,plain,(
  $false | ($spl3 | $spl12 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f10964,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10964,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl12 | $spl36)),
  inference(forward_demodulation,[],[f10935,f6572])).
fof(f10935,plain,(
  s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl12 | $spl36)),
  inference(trivial_inequality_removal,[],[f10933])).
fof(f10933,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))))))) | ($spl12 | $spl36)),
  inference(superposition,[],[f2637,f6572])).
fof(f11217,plain,(
  $false | ($spl3 | $spl12 | $spl41)),
  inference(global_subsumption,[],[f10934,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f10934,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl12),
  inference(equality_resolution,[],[f2637])).
fof(f11216,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11110,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11110,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl34),
  inference(unit_resulting_resolution,[],[f190,f3867])).
fof(f11215,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11112,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11112,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11214,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11115,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11115,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11213,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11212,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11212,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X12),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X12))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl34),
  inference(subsumption_resolution,[],[f11118,f190])).
fof(f11118,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X12),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X12))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11211,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11210,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11210,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl34),
  inference(subsumption_resolution,[],[f11119,f190])).
fof(f11119,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl34),
  inference(superposition,[],[f3867,f202])).
fof(f11209,plain,(
  $false | ($spl3 | $spl34 | $spl41 | $spl57)),
  inference(global_subsumption,[],[f11208,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11208,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl57)),
  inference(subsumption_resolution,[],[f11207,f190])).
fof(f11207,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl57)),
  inference(forward_demodulation,[],[f11124,f8735])).
fof(f11124,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl34 | $spl57)),
  inference(superposition,[],[f3867,f8735])).
fof(f11206,plain,(
  $false | ($spl3 | $spl34 | $spl41 | $spl42)),
  inference(global_subsumption,[],[f11205,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11205,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl42)),
  inference(subsumption_resolution,[],[f11204,f190])).
fof(f11204,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl42)),
  inference(forward_demodulation,[],[f11131,f7138])).
fof(f11131,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl34 | $spl42)),
  inference(superposition,[],[f3867,f7138])).
fof(f11203,plain,(
  $false | ($spl3 | $spl34 | $spl41)),
  inference(global_subsumption,[],[f11133,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11133,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) | s(X26,X28) = s(X26,sK4(X28,X27,X26))) ) | $spl34),
  inference(superposition,[],[f3867,f400])).
fof(f11202,plain,(
  $false | ($spl3 | $spl34 | $spl36 | $spl41)),
  inference(global_subsumption,[],[f11201,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f4998,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4824,f4827,f228,f4875,f4877,f4663,f6327,f6547,f6546,f7132,f7133,f7134,f7135,f7136,f7313,f7314,f7540,f7547,f7653,f7654,f7835,f7836,f7974,f8151,f8604,f8607,f8620,f8771,f8772,f8777,f8778,f8779,f8784,f8786,f8793,f9135,f9136,f9137,f9142,f9144,f9153,f9547,f9541,f9549,f10018,f10021,f10066,f10038,f10075,f10579,f10575,f10581,f10657,f10678,f10752])).
fof(f11201,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl36)),
  inference(subsumption_resolution,[],[f11200,f190])).
fof(f11200,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | ($spl34 | $spl36)),
  inference(forward_demodulation,[],[f11139,f6572])).
fof(f11139,plain,(
  s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | ($spl34 | $spl36)),
  inference(superposition,[],[f3867,f6572])).
fof(f11197,plain,(
  $false | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11196,f112])).
fof(f11196,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11195,f6355])).
fof(f6355,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f112,f6288,f237])).
fof(f6288,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl0),
  inference(forward_demodulation,[],[f109,f226])).
fof(f11195,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11177,f110])).
fof(f11177,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ($spl0 | $spl36)),
  inference(superposition,[],[f4104,f6572])).
fof(f4104,plain,(
  ( ! [X0] : (~p(s(bool,'const.bag.BAG_IN_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X0))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4097,f110])).
fof(f4097,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bag.BAG_IN_2'(s(sK0,X0),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X0))))) ) | $spl0),
  inference(superposition,[],[f144,f226])).
fof(f11194,plain,(
  $false | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11193,f112])).
fof(f11193,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11192,f6355])).
fof(f11192,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11191,f6343])).
fof(f6343,plain,(
  s(sK0,sK2) = s(sK0,sK4(sK2,sK1,sK0)) | $spl0),
  inference(unit_resulting_resolution,[],[f6288,f400])).
fof(f11191,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11190,f110])).
fof(f11190,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11176,f6572])).
fof(f11176,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) | $spl0),
  inference(superposition,[],[f4104,f6343])).
fof(f11189,plain,(
  $false | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11188,f112])).
fof(f11188,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11187,f6355])).
fof(f11187,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11186,f6271])).
fof(f11186,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0))))) | ($spl0 | $spl36)),
  inference(subsumption_resolution,[],[f11185,f110])).
fof(f11185,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0))))) | ($spl0 | $spl36)),
  inference(forward_demodulation,[],[f11175,f6572])).
fof(f11175,plain,(
  ~p(s(bool,'const.bag.BAG_IN_2'(s(sK0,sK2),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0))))) | ($spl0 | $spl36)),
  inference(superposition,[],[f4104,f6271])).
fof(f2603,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) | $spl4),
  inference(cnf_transformation,[],[f2603_D])).
fof(f2603_D,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f10890,plain,(
  ( ! [X64,X63] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X64)))))) | s(bool,'F_0') = s(bool,X63) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X63))))),s(fun(bool,'type.num.num'),X64))))) ) | $spl64),
  inference(cnf_transformation,[],[f10890_D])).
fof(f10890_D,plain,(
  ( ! [X64,X63] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X64)))))) | s(bool,'F_0') = s(bool,X63) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X63))))),s(fun(bool,'type.num.num'),X64))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f10889,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X64)))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X63))))),s(fun(bool,'type.num.num'),X64)))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X63)) ) | $spl8),
  inference(subsumption_resolution,[],[f10824,f10882])).
fof(f10882,plain,(
  ( ! [X30,X31] : (s(bool,X31) != s(bool,'const.bag.BAG_IN_2'(s(bool,X30),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl8),
  inference(subsumption_resolution,[],[f10812,f120])).
fof(f10812,plain,(
  ( ! [X30,X31] : (~p(s(bool,X30)) | s(bool,X31) != s(bool,'const.bag.BAG_IN_2'(s(bool,X30),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X30))))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl8),
  inference(superposition,[],[f204,f2623])).
fof(f10824,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X64)))))) | s(bool,X65) = s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X63))))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(bool,X63),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X63))))),s(fun(bool,'type.num.num'),X64)))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X63)) ) | $spl8),
  inference(superposition,[],[f1188,f2623])).
fof(f10070,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(X67,X69) = s(X67,sK4(X69,X68,X67))) ) | $spl62),
  inference(cnf_transformation,[],[f10070_D])).
fof(f10070_D,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(X67,X69) = s(X67,sK4(X69,X68,X67))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f10065,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5))))) ) | $spl60),
  inference(cnf_transformation,[],[f10065_D])).
fof(f10065_D,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'T_0') = s(bool,X4) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10038,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X70),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X71)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X71)))) | s(bool,'T_0') = s(bool,X70) | s(X67,X69) = s(X67,sK4(X69,X68,X67))) )),
  inference(superposition,[],[f1176,f400])).
fof(f10064,plain,(
  ( ! [X6,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X6) | s(bool,X3) = s(bool,X6)) ) | $spl58),
  inference(cnf_transformation,[],[f10064_D])).
fof(f10064_D,plain,(
  ( ! [X6,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X6) | s(bool,X3) = s(bool,X6)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10021,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X16),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X17)))))) | s(bool,'F_0') = s(bool,X15) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X17)))) | s(bool,'T_0') = s(bool,X16) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1176,f202])).
fof(f10018,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bag.BAG_IN_2'(s(bool,X4),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X5)))))) | s(bool,'F_0') = s(bool,X3) | p(s(bool,'const.bag.BAG_IN_2'(s(bool,'F_0'),s(fun(bool,'type.num.num'),X5)))) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1176,f202])).
fof(f8622,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f8622_D])).
fof(f8622_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8624,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl57),
  inference(cnf_transformation,[],[f8624_D])).
fof(f8624_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f8607,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f935,f202])).
fof(f8604,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f935,f202])).
fof(f7142,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl48),
  inference(cnf_transformation,[],[f7142_D])).
fof(f7142_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7139,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl44),
  inference(cnf_transformation,[],[f7139_D])).
fof(f7139_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7144,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f7144_D])).
fof(f7144_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f7136,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f933,f202])).
fof(f7138,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl42),
  inference(cnf_transformation,[],[f7138_D])).
fof(f7138_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7145,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl52),
  inference(cnf_transformation,[],[f7145_D])).
fof(f7145_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7135,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f933,f202])).
fof(f7134,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f933,f202])).
fof(f7133,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f933,f202])).
fof(f7141,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl47),
  inference(cnf_transformation,[],[f7141_D])).
fof(f7141_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f7132,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f933,f113])).
fof(f6325,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) | $spl38),
  inference(cnf_transformation,[],[f6325_D])).
fof(f6325_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f6327,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl41),
  inference(cnf_transformation,[],[f6327_D])).
fof(f6327_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f6298,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK4(sK2,sK1,sK0))))) | $spl36),
  inference(superposition,[],[f148,f6271])).
fof(f6291,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877,f4827,f4875,f6288,f228])).
fof(f5453,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f4105,f202])).
fof(f5418,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f4105,f152])).
fof(f5423,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f110,f4105,f156])).
fof(f5424,plain,(
  ( ! [X0,X1] : (s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f4105,f4105,f156])).
fof(f5064,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X33)))),X34,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f5000,f4061])).
fof(f5000,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0'))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X33))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X33)))),X34,bool))) ) | $spl0),
  inference(superposition,[],[f400,f4061])).
fof(f5063,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X32,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f4999,f4110])).
fof(f4999,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0'))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X32,bool))) ) | $spl0),
  inference(superposition,[],[f400,f4110])).
fof(f5060,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X24,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f4995,f226])).
fof(f4995,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,'T_0'))) | s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X24,bool))) ) | $spl0),
  inference(superposition,[],[f400,f226])).
fof(f4814,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X30)))),X31,bool)) | p(s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f4773,f4061])).
fof(f4773,plain,(
  ( ! [X30,X31] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X30))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X30)))),X31,bool))) ) | $spl0),
  inference(superposition,[],[f147,f4061])).
fof(f4811,plain,(
  ( ! [X24,X23] : (s(bool,X24) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X23))))) | s(bool,'T_0') = s(bool,X24)) ) | $spl0),
  inference(subsumption_resolution,[],[f4768,f110])).
fof(f4768,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | s(bool,X24) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X23))))) | s(bool,'T_0') = s(bool,X24)) ) | $spl0),
  inference(superposition,[],[f204,f4061])).
fof(f4742,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f4061,f202])).
fof(f4198,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X9,bool)) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f4169,f4110])).
fof(f4169,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) = s(bool,sK4('const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))),X9,bool))) ) | $spl0),
  inference(superposition,[],[f147,f4110])).
fof(f4197,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f4166,f110])).
fof(f4166,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f204,f4110])).
fof(f4145,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(superposition,[],[f4110,f202])).
fof(f4116,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl0),
  inference(superposition,[],[f4074,f202])).
fof(f4105,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f4074,f152])).
fof(f4110,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f110,f4074,f156])).
fof(f4074,plain,(
  p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0')))))) | $spl0),
  inference(superposition,[],[f4042,f138])).
fof(f4067,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),X1)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f4042,f202])).
fof(f4056,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X0),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1))))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f4042,f152])).
fof(f4061,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f110,f4042,f156])).
fof(f4062,plain,(
  ( ! [X0,X1] : (s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X1)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f4042,f4042,f156])).
fof(f4051,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X8,bool)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f4028,f226])).
fof(f4028,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) | s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X8,bool))) ) | $spl0),
  inference(superposition,[],[f147,f226])).
fof(f4050,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f4025,f110])).
fof(f4025,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f204,f226])).
fof(f4042,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0))))))) ) | $spl0),
  inference(forward_demodulation,[],[f4010,f226])).
fof(f4010,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X0))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f226,f151])).
fof(f6290,plain,(
  $false | ($spl0 | $spl3 | $spl34)),
  inference(global_subsumption,[],[f3867,f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877,f4827,f4875,f6288])).
fof(f6289,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877,f4827,f4875,f6288])).
fof(f6281,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877,f4827,f4875])).
fof(f6280,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877,f4827])).
fof(f6279,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824,f4877])).
fof(f6278,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663,f4824])).
fof(f6273,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f110,f113,f116,f115,f138,f141,f156,f153,f161,f187,f190,f198,f202,f204,f218,f228,f237,f241,f238,f246,f258,f305,f362,f400,f442,f443,f472,f485,f495,f497,f517,f531,f567,f609,f783,f821,f822,f836,f824,f926,f950,f1091,f1111,f1344,f1447,f1526,f1527,f1528,f1530,f1531,f1532,f1533,f1558,f2039,f2040,f2041,f1628,f2595,f2080,f2590,f2579,f2563,f2558,f2557,f2135,f2136,f2548,f2546,f2626,f2627,f2161,f2631,f2636,f2639,f2641,f2643,f2654,f2660,f2662,f2664,f2520,f2693,f2694,f2513,f2512,f2510,f2354,f2359,f2360,f2361,f2363,f2364,f2365,f2366,f2393,f2714,f2467,f2490,f2470,f469,f2272,f1097,f938,f1619,f1620,f2628,f1621,f1555,f1554,f1552,f1439,f1438,f2713,f2712,f2699,f2686,f2635,f2633,f2632,f2630,f2586,f2460,f2459,f2458,f2390,f2389,f2387,f2355,f2346,f2345,f2264,f2263,f2187,f2186,f2184,f2457,f2456,f2391,f2388,f2347,f2344,f2343,f2265,f2262,f2261,f2188,f2185,f1618,f1617,f1556,f1553,f1440,f1437,f1436,f151,f467,f1238,f2721,f2720,f2048,f2047,f152,f150,f770,f399,f2659,f2657,f2565,f2454,f1615,f1248,f1198,f1185,f1092,f927,f804,f612,f570,f464,f457,f421,f420,f419,f415,f402,f397,f2466,f2455,f2386,f2385,f2362,f2358,f2353,f2342,f2341,f2271,f2260,f2259,f2183,f2182,f1627,f1616,f1551,f1550,f1529,f1525,f1446,f1435,f1434,f757,f2718,f2717,f2045,f2044,f149,f148,f147,f2465,f2464,f2352,f2351,f2270,f2269,f1626,f1625,f1445,f1444,f530,f2719,f2716,f2046,f2043,f532,f146,f145,f144,f1443,f1624,f2268,f2350,f2463,f2698,f2625,f1226,f1207,f1189,f1176,f937,f936,f812,f797,f762,f750,f613,f526,f470,f460,f459,f453,f452,f428,f425,f423,f398,f396,f389,f388,f257,f197,f120,f119,f112,f3115,f3143,f3159,f3312,f3316,f3317,f3318,f3309,f3344,f3490,f3494,f3474,f3529,f3530,f3553,f3554,f3692,f3694,f3771,f3772,f3773,f3798,f3802,f3817,f3818,f3866,f3869,f1188,f1175,f935,f933,f796,f749,f109,f226,f4042,f4050,f4051,f4062,f4061,f4056,f4067,f4074,f4104,f4110,f4105,f4116,f4145,f4197,f4198,f4662,f4616,f4617,f4618,f4622,f4624,f4633,f4634,f4635,f4639,f4641,f4700,f4742,f4811,f4814,f4830,f4831,f4832,f4836,f4838,f4848,f4849,f4856,f4898,f4986,f4987,f4988,f4992,f4994,f5060,f4998,f5063,f5064,f5065,f5067,f5069,f5070,f5071,f5072,f5073,f5074,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5043,f5044,f5045,f5046,f5424,f5423,f5418,f5453,f5664,f5618,f5619,f5620,f5624,f5626,f5630,f5640,f5641,f5642,f5646,f5648,f5652,f5662,f5825,f5826,f5985,f5986,f6244,f4663])).
fof(f6271,plain,(
  s(sK0,sK2) = s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0)) | $spl36),
  inference(cnf_transformation,[],[f6271_D])).
fof(f6271_D,plain,(
  s(sK0,sK2) = s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6269,plain,(
  ( ! [X0] : (s(sK0,sK2) = s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f6260,f116])).
fof(f6260,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(sK0,sK2) = s(sK0,sK4(sK4(sK2,sK1,sK0),sK1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f3949,f3909])).
fof(f3909,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f3879,f190])).
fof(f3879,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(superposition,[],[f227,f202])).
fof(f227,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl2),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3949,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK4(sK4(sK2,sK1,sK0),X0,sK0))) ) | $spl2),
  inference(forward_demodulation,[],[f3942,f3910])).
fof(f3910,plain,(
  s(sK0,sK2) = s(sK0,sK4(sK2,sK1,sK0)) | $spl2),
  inference(subsumption_resolution,[],[f3883,f112])).
fof(f3883,plain,(
  p(s(bool,'F_0')) | s(sK0,sK2) = s(sK0,sK4(sK2,sK1,sK0)) | $spl2),
  inference(superposition,[],[f147,f227])).
fof(f3942,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK2)))) | s(sK0,sK4(sK2,sK1,sK0)) = s(sK0,sK4(sK4(sK2,sK1,sK0),X0,sK0))) ) | $spl2),
  inference(superposition,[],[f147,f3910])).
fof(f3867,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35)) ) | $spl34),
  inference(cnf_transformation,[],[f3867_D])).
fof(f3867_D,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3866,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl3),
  inference(subsumption_resolution,[],[f3853,f116])).
fof(f3853,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | s(bool,'F_0') != s(bool,'const.bag.BAG_IN_2'(s(bool,X35),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X35))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl3),
  inference(superposition,[],[f3317,f113])).
fof(f2920,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2919,f2603])).
fof(f2919,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f2723,f115])).
fof(f2723,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f2603,f119])).
fof(f2918,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2722,f110])).
fof(f2722,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f2603,f112])).
fof(f2697,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,X52) = s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52)))))) ) | $spl32),
  inference(cnf_transformation,[],[f2697_D])).
fof(f2697_D,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,X52) = s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2510,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2293])).
fof(f2293,plain,(
  ( ! [X52,X53,X51] : (s(bool,X52) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,X53),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))))) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.bag.BAG_IN_2'(s(bool,X52),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X52))))) | s(bool,'T_0') = s(bool,X51)) ) | $spl3),
  inference(superposition,[],[f926,f1527])).
fof(f2696,plain,(
  ( ! [X47,X48] : (s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X48))))) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47)))))) ) | $spl30),
  inference(cnf_transformation,[],[f2696_D])).
fof(f2696_D,plain,(
  ( ! [X47,X48] : (s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X48))))) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2512,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X48))))) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2291])).
fof(f2291,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.BAG_INSERT_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),X48))))) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.bag.BAG_IN_2'(s(bool,X47),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X47))))) | s(bool,'T_0') = s(bool,X46)) ) | $spl3),
  inference(superposition,[],[f609,f1527])).
fof(f2695,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X45) = s(bool,'const.bag.BAG_IN_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X45)))))) ) | $spl28),
  inference(cnf_transformation,[],[f2695_D])).
fof(f2695_D,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,X45) = s(bool,'const.bag.BAG_IN_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X45)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2513,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,'const.bag.BAG_IN_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X45)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2290])).
fof(f2290,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,'const.bag.BAG_IN_2'(s(bool,'T_0'),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,'const.bag.BAG_IN_2'(s(bool,X45),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X45))))) | s(bool,'T_0') = s(bool,X44)) ) | $spl3),
  inference(superposition,[],[f567,f1527])).
fof(f2665,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X156,bool),X158),s(fun(X156,'type.num.num'),X157)))) | s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157)))))))) ) | $spl26),
  inference(cnf_transformation,[],[f2665_D])).
fof(f2665_D,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X156,bool),X158),s(fun(X156,'type.num.num'),X157)))) | s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157)))))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2664,plain,(
  ( ! [X158,X156,X159,X157] : (p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X156,bool),X158),s(fun(X156,'type.num.num'),X157)))) | s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))))))) | s(bool,'T_0') = s(bool,X159)) ) | $spl3),
  inference(subsumption_resolution,[],[f2249,f116])).
fof(f2249,plain,(
  ( ! [X158,X156,X159,X157] : (p(s(bool,X159)) | p(s(bool,'const.bag.BAG_EVERY_2'(s(fun(X156,bool),X158),s(fun(X156,'type.num.num'),X157)))) | s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X156,sK3(X157,X158,X156)),s(fun(X156,'type.num.num'),X157))))))) | s(bool,'T_0') = s(bool,X159)) ) | $spl3),
  inference(superposition,[],[f145,f1527])).
fof(f2663,plain,(
  ( ! [X146,X149,X147,X148] : (s(X146,X147) != s(X146,X148) | s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149)))))))))) ) | $spl24),
  inference(cnf_transformation,[],[f2663_D])).
fof(f2663_D,plain,(
  ( ! [X146,X149,X147,X148] : (s(X146,X147) != s(X146,X148) | s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149)))))))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2662,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(X146,X147) != s(X146,X148) | s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))))))) | s(bool,'T_0') = s(bool,X150)) ) | $spl3),
  inference(subsumption_resolution,[],[f2247,f116])).
fof(f2247,plain,(
  ( ! [X146,X149,X147,X150,X148] : (p(s(bool,X150)) | s(X146,X147) != s(X146,X148) | s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))) = s(bool,'const.bag.BAG_IN_2'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,'const.bag.BAG_IN_2'(s(X146,X147),s(fun(X146,'type.num.num'),'const.bag.BAG_INSERT_2'(s(X146,X148),s(fun(X146,'type.num.num'),X149))))))))) | s(bool,'T_0') = s(bool,X150)) ) | $spl3),
  inference(superposition,[],[f151,f1527])).
fof(f2661,plain,(
  ( ! [X133] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X133,bool)) | s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0')))))))) ) | $spl22),
  inference(cnf_transformation,[],[f2661_D])).
fof(f2661_D,plain,(
  ( ! [X133] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X133,bool)) | s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0')))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2660,plain,(
  ( ! [X134,X133] : (s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X133,bool)) | s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X134)) ) | $spl3),
  inference(subsumption_resolution,[],[f2243,f116])).
fof(f2243,plain,(
  ( ! [X134,X133] : (p(s(bool,X134)) | s(bool,'F_0') = s(bool,sK4('const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))),X133,bool)) | s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(bool,bool),X133),s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X134)) ) | $spl3),
  inference(superposition,[],[f396,f1527])).
fof(f2655,plain,(
  ( ! [X118,X119,X117] : (s(X117,X119) = s(X117,sK4(X119,X118,X117)) | s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))))))) ) | $spl20),
  inference(cnf_transformation,[],[f2655_D])).
fof(f2655_D,plain,(
  ( ! [X118,X119,X117] : (s(X117,X119) = s(X117,sK4(X119,X118,X117)) | s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119)))))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2654,plain,(
  ( ! [X118,X120,X119,X117] : (s(X117,X119) = s(X117,sK4(X119,X118,X117)) | s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))))))) | s(bool,'T_0') = s(bool,X120)) ) | $spl3),
  inference(subsumption_resolution,[],[f2238,f116])).
fof(f2238,plain,(
  ( ! [X118,X120,X119,X117] : (p(s(bool,X120)) | s(X117,X119) = s(X117,sK4(X119,X118,X117)) | s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))) = s(bool,'const.bag.BAG_IN_2'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,happ(s(fun(X117,bool),X118),s(X117,X119))))))) | s(bool,'T_0') = s(bool,X120)) ) | $spl3),
  inference(superposition,[],[f147,f1527])).
fof(f2642,plain,(
  ( ! [X12,X10] : (p(s(bool,X12)) | s(bool,X10) = s(bool,'const.bag.BAG_IN_2'(s(bool,X10),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X10))))) | s(bool,X10) = s(bool,X12)) ) | $spl18),
  inference(cnf_transformation,[],[f2642_D])).
fof(f2642_D,plain,(
  ( ! [X12,X10] : (p(s(bool,X12)) | s(bool,X10) = s(bool,'const.bag.BAG_IN_2'(s(bool,X10),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X10))))) | s(bool,X10) = s(bool,X12)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2643,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X15)) | s(bool,X13) = s(bool,X15) | s(bool,X13) = s(bool,'const.bag.BAG_IN_2'(s(bool,X13),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X13))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(subsumption_resolution,[],[f2197,f116])).
fof(f2197,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X13) = s(bool,X15) | s(bool,X13) = s(bool,'const.bag.BAG_IN_2'(s(bool,X13),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X13))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f153,f1527])).
fof(f2641,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,'const.bag.BAG_IN_2'(s(bool,X10),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(subsumption_resolution,[],[f2196,f116])).
fof(f2196,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,X10) = s(bool,'const.bag.BAG_IN_2'(s(bool,X10),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X10))))) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(superposition,[],[f153,f1527])).
fof(f2640,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) ) | $spl16),
  inference(cnf_transformation,[],[f2640_D])).
fof(f2640_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2639,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(subsumption_resolution,[],[f2194,f113])).
fof(f2194,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,X6) = s(bool,'const.bag.BAG_IN_2'(s(bool,X6),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X6))))) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(superposition,[],[f119,f1527])).
fof(f2637,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) ) | $spl12),
  inference(cnf_transformation,[],[f2637_D])).
fof(f2637_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2638,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl14),
  inference(cnf_transformation,[],[f2638_D])).
fof(f2638_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2636,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2192,f115])).
fof(f2192,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,'const.bag.BAG_IN_2'(s(bool,X2),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X2))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f115,f1527])).
fof(f2624,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15)))))) ) | $spl10),
  inference(cnf_transformation,[],[f2624_D])).
fof(f2624_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2546,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,X25) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2152])).
fof(f2152,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,X25) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,'T_0') = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f1527,f442])).
fof(f2604,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X27)))))) ) | $spl6),
  inference(cnf_transformation,[],[f2604_D])).
fof(f2604_D,plain,(
  ( ! [X27,X25] : (s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,X25) = s(bool,X27) | s(bool,X27) = s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X27)))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2548,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,'const.bag.BAG_IN_2'(s(bool,X12),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X12)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2149])).
fof(f2149,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,'const.bag.BAG_IN_2'(s(bool,X14),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X14))))) | s(bool,X12) = s(bool,X14) | s(bool,X12) = s(bool,'const.bag.BAG_IN_2'(s(bool,X12),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X12))))) | s(bool,'T_0') = s(bool,X13)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f2557,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'const.bag.BAG_IN_2'(s(bool,X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X20)))))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2119])).
fof(f2119,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,'const.bag.BAG_IN_2'(s(bool,X21),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X21))))) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'const.bag.BAG_IN_2'(s(bool,X20),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X20))))) | s(bool,'T_0') = s(bool,X19)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f2558,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2117])).
fof(f2117,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | s(bool,X15) = s(bool,'const.bag.BAG_IN_2'(s(bool,X15),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X15))))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f1527,f442])).
fof(f2623,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) ) | $spl8),
  inference(cnf_transformation,[],[f2623_D])).
fof(f2623_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2563,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2112])).
fof(f2112,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bag.BAG_IN_2'(s(bool,X1),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X1))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1527,f113])).
fof(f2579,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X22))))) | s(bool,X24) = s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | s(bool,X22) = s(bool,X24)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2097])).
fof(f2097,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,'const.bag.BAG_IN_2'(s(bool,X22),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X22))))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.bag.BAG_IN_2'(s(bool,X24),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X24))))) | s(bool,X22) = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f2595,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X27))))) | s(bool,X25) = s(bool,X27)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2066])).
fof(f2066,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,'const.bag.BAG_IN_2'(s(bool,X25),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X25))))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.bag.BAG_IN_2'(s(bool,X27),s(fun(bool,'type.num.num'),'const.bag.EL_BAG_1'(s(bool,X27))))) | s(bool,X25) = s(bool,X27)) ) | $spl3),
  inference(superposition,[],[f1527,f1527])).
fof(f226,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) | $spl0),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f228,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl3),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f218,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | s(bool,'T_0') = s(bool,'const.bag.BAG_EVERY_2'(s(fun(sK0,bool),sK1),s(fun(sK0,'type.num.num'),'const.bag.EL_BAG_1'(s(sK0,sK2)))))),
  inference(superposition,[],[f109,f113])).
% SZS output end Proof for thm.okasaki_misc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 195
% Passive clauses: 1063
% Generated clauses: 12971
% Final active clauses: 72
% Final passive clauses: 215
% Input formulas: 27
% Initial clauses: 67
% 
% Duplicate literals: 589
% Trivial inequalities: 129
% Fw subsumption resolutions: 831
% Fw demodulations: 2514
% Bw demodulations: 253
% Condensations: 689
% Global subsumptions: 78
% 
% Simple tautologies: 170
% Equational tautologies: 141
% Forward subsumptions: 6657
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 1119
% Factoring: 1
% Forward superposition: 3301
% Backward superposition: 2880
% Self superposition: 79
% Equality factoring: 50
% Equality resolution: 20
% 
% Split clauses: 33
% Split components: 71
% SAT solver clauses: 1288
% SAT solver unit clauses: 44
% SAT solver binary clauses: 222
% 
% Sat splits: 33
% Sat splitting refutations: 79
% 
% Memory used [KB]: 6652
% Time elapsed: 0.497 s
% ------------------------------
% ------------------------------
% Success in time 1.842 s
