Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (39009)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 123
% Passive clauses: 3421
% Generated clauses: 20165
% Final active clauses: 118
% Final passive clauses: 2955
% Input formulas: 37
% Initial clauses: 49
% 
% Duplicate literals: 453
% Trivial inequalities: 252
% Fw subsumption resolutions: 2525
% Fw demodulations: 2762
% Bw demodulations: 3
% 
% Simple tautologies: 200
% Equational tautologies: 62
% Forward subsumptions: 8599
% Fw demodulations to eq. taut.: 39
% 
% Binary resolution: 205
% Factoring: 1
% Forward superposition: 6527
% Backward superposition: 6799
% Self superposition: 323
% Equality factoring: 159
% Equality resolution: 42
% 
% Split clauses: 335
% Split components: 684
% SAT solver clauses: 609
% SAT solver unit clauses: 43
% SAT solver binary clauses: 242
% 
% Sat splits: 335
% 
% TWLsolver clauses: 601
% TWLsolver calls for satisfiability: 87
% 
% Memory used [KB]: 13304
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (39025)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 1446
% Generated clauses: 18781
% Final active clauses: 146
% Final passive clauses: 1158
% Input formulas: 37
% Initial clauses: 49
% 
% Duplicate literals: 310
% Trivial inequalities: 287
% Fw subsumption resolutions: 783
% Fw demodulations: 3013
% Bw demodulations: 128
% Condensations: 632
% 
% Simple tautologies: 155
% Equational tautologies: 29
% Forward subsumptions: 9981
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 106
% Factoring: 2
% Forward superposition: 5592
% Backward superposition: 6175
% Self superposition: 129
% Equality factoring: 39
% Equality resolution: 36
% 
% Split clauses: 30
% Split components: 60
% SAT solver clauses: 2114
% SAT solver unit clauses: 72
% SAT solver binary clauses: 37
% 
% Sat splits: 30
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2114
% TWLsolver calls for satisfiability: 277
% 
% Memory used [KB]: 19573
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f2655,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f108,f2585])).
fof(f2585,plain,(
  ( ! [X181] : (~p(s(bool,X181))) )),
  inference(forward_demodulation,[],[f2584,f138])).
fof(f138,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : ! [X9,X31,X32,X33,X34] : s(X8,X31) = s(X8,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X8),X9),s(X8,X31),s(fun('type.num.num',X8),X32),s(fun(bool,X8),X33),s(X8,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|0','thm.bvp_to_wordProps.tag_case_def|split|1')).
fof(f2584,plain,(
  ( ! [X182,X180,X184,X183,X181] : (~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',bool),X180),s(bool,X181),s(fun('type.num.num',bool),X182),s(fun(bool,bool),X183),s(bool,X184))))) )),
  inference(subsumption_resolution,[],[f2509,f2483])).
fof(f2483,plain,(
  ( ! [X6,X8,X7] : (s(X6,X8) = s(X6,happ(s(fun('type.num.num',X6),X7),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f1585,f138])).
fof(f1585,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(superposition,[],[f139,f107])).
fof(f107,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f90])).
fof(f90,plain,(
  ? [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|0','thm.bvp_to_wordProps.tag_distinct|split|0')).
fof(f139,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : ! [X6,X9,X31,X32,X33,X34] : s(X8,happ(s(fun('type.num.num',X8),X9),s('type.num.num',X6))) = s(X8,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X6))),s(fun('type.num.num',X8),X9),s(X8,X31),s(fun('type.num.num',X8),X32),s(fun(bool,X8),X33),s(X8,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|0','thm.bvp_to_wordProps.tag_case_def|split|0')).
fof(f2509,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X180),s('type.num.num',sK0))) | ~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'),s(fun('type.num.num',bool),X180),s(bool,X181),s(fun('type.num.num',bool),X182),s(fun(bool,bool),X183),s(bool,X184))))) )),
  inference(superposition,[],[f125,f1585])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X25] : (s(bool,'F_0') = s(bool,X25) <=> ~p(s(bool,X25)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f108,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_distinct|split|0',aHL_TRUTH)).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 72
% Passive clauses: 300
% Generated clauses: 2542
% Final active clauses: 70
% Final passive clauses: 195
% Input formulas: 37
% Initial clauses: 49
% 
% Duplicate literals: 4
% Trivial inequalities: 9
% Fw subsumption resolutions: 36
% Fw demodulations: 492
% Bw demodulations: 25
% Condensations: 4
% 
% Simple tautologies: 6
% Equational tautologies: 19
% Forward subsumptions: 1560
% Fw demodulations to eq. taut.: 76
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 173
% Factoring: 1
% Forward superposition: 899
% Backward superposition: 798
% Self superposition: 50
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 338
% SAT solver unit clauses: 73
% SAT solver binary clauses: 191
% 
% Sat splits: 16
% Sat splitting refutations: 4
% 
% Memory used [KB]: 6908
% Time elapsed: 0.127 s
% ------------------------------
% ------------------------------
% Success in time 1.459 s
