Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (133146)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 3218
% Generated clauses: 20925
% Final active clauses: 111
% Final passive clauses: 2974
% Input formulas: 21
% Initial clauses: 27
% 
% Duplicate literals: 764
% Trivial inequalities: 267
% Fw subsumption resolutions: 2276
% Fw demodulations: 3680
% Bw demodulations: 7
% 
% Simple tautologies: 235
% Equational tautologies: 71
% Forward subsumptions: 9699
% Fw demodulations to eq. taut.: 35
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 191
% Factoring: 2
% Forward superposition: 7573
% Backward superposition: 5744
% Self superposition: 261
% Equality factoring: 160
% Equality resolution: 28
% 
% Split clauses: 346
% Split components: 710
% SAT solver clauses: 790
% SAT solver unit clauses: 12
% SAT solver binary clauses: 58
% 
% Sat splits: 346
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 790
% TWLsolver calls for satisfiability: 74
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (133194)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 699
% Generated clauses: 15309
% Final active clauses: 89
% Final passive clauses: 573
% Input formulas: 21
% Initial clauses: 27
% 
% Duplicate literals: 892
% Trivial inequalities: 161
% Fw subsumption resolutions: 844
% Fw demodulations: 1449
% Bw demodulations: 24
% Condensations: 1880
% 
% Simple tautologies: 157
% Equational tautologies: 53
% Forward subsumptions: 7753
% Fw demodulations to eq. taut.: 13
% Bw demodulations to eq. taut.: 12
% 
% Binary resolution: 51
% Factoring: 1
% Forward superposition: 5020
% Backward superposition: 4040
% Self superposition: 232
% Equality factoring: 114
% Equality resolution: 21
% 
% Split clauses: 70
% Split components: 141
% SAT solver clauses: 1443
% SAT solver unit clauses: 38
% SAT solver binary clauses: 37
% 
% Sat splits: 70
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1443
% TWLsolver calls for satisfiability: 155
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f5794,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f66,f95,f93,f65,f94,f92,f97,f102,f109,f103,f111,f107,f105,f114,f112,f132,f136,f134,f133,f141,f1752,f1753,f1781,f1756,f1757,f1784,f1761,f1780,f1762,f1786,f1782,f1791,f1792,f1806,f3471,f3494,f3518,f3604,f3670,f3676,f5773,f5793])).
fof(f5793,plain,(
  $false | ($spl3 | $spl5 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5792,f114])).
fof(f5792,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f5755,f73])).
fof(f73,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.ind_type.FCONS|split|0')).
fof(f5755,plain,(
  ( ! [X70] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | ($spl3 | $spl5 | $spl6)),
  inference(superposition,[],[f73,f4202])).
fof(f4202,plain,(
  ( ! [X116,X115] : (s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X115),s('type.num.num',X116))) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X115),s('type.num.num',X116)))) ) | ($spl3 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f4166,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9] : ! [X8,X26,X27] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X26),s('type.num.num',X27))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X26))),s('type.num.num',X27)))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','ARITY_0_3_const.ind_type.FCONS')).
fof(f4166,plain,(
  ( ! [X116,X115] : (s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X115),s('type.num.num',X116))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun(fun('type.num.num',bool),fun('type.num.num',bool)),happ(s(fun(bool,fun(fun('type.num.num',bool),fun('type.num.num',bool))),'const.ind_type.FCONS_0'),s(bool,'T_0'))),s(fun('type.num.num',bool),X115))),s('type.num.num',X116)))) ) | ($spl3 | $spl5 | $spl6)),
  inference(superposition,[],[f78,f3777])).
fof(f3777,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(bool,X0),X1),s(bool,'T_0'))) = s(X0,happ(s(fun(bool,X0),X1),s(bool,'F_0')))) ) | ($spl3 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3765,f79])).
fof(f79,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun(bool,X9),X13),s(bool,X6))) = s(X9,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X6))),s(fun('type.num.num',X9),X10),s(X9,X11),s(fun('type.num.num',X9),X12),s(fun(bool,X9),X13),s(X9,X14)))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.bvp_to_wordProps.tag_case_def|split|3')).
fof(f3765,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(bool,X0),X1),s(bool,'F_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X1),s(X0,X5)))) ) | ($spl3 | $spl5 | $spl6)),
  inference(superposition,[],[f79,f3697])).
fof(f3697,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl3 | $spl5 | $spl6)),
  inference(backward_demodulation,[],[f3692,f109])).
fof(f3692,plain,(
  s(bool,'T_0') = s(bool,sK0) | ($spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f67,f3681,f83])).
fof(f83,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X33,X34] : ((p(s(bool,X33)) => p(s(bool,X34))) => ((p(s(bool,X34)) => p(s(bool,X33))) => s(bool,X33) = s(bool,X34)))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f3681,plain,(
  p(s(bool,sK0)) | ($spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f95,f106,f135])).
fof(f135,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f129,f68])).
fof(f68,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2',aHL_FALSITY)).
fof(f129,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f80,f72])).
fof(f72,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2',aHL_BOOL_CASES)).
fof(f80,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f57])).
fof(f106,plain,(
  s(bool,'T_0') != s(bool,sK1) | $spl5),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  s(bool,'T_0') != s(bool,sK1) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f67,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2',aHL_TRUTH)).
fof(f5773,plain,(
  $false | ($spl3 | $spl5 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5772,f114])).
fof(f5772,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f5727,f73])).
fof(f5727,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0')))) ) | ($spl3 | $spl5 | $spl6)),
  inference(superposition,[],[f4202,f73])).
fof(f3676,plain,(
  $false | ($spl9 | $spl10)),
  inference(subsumption_resolution,[],[f111,f112])).
fof(f3670,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3669,f112])).
fof(f3669,plain,(
  s(bool,'F_0') != s(bool,sK0) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f95,f112])).
fof(f3604,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl16)),
  inference(subsumption_resolution,[],[f1826,f3519])).
fof(f3519,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl3 | $spl4 | $spl16)),
  inference(subsumption_resolution,[],[f3496,f100])).
fof(f100,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f68,f72])).
fof(f3496,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,X1))) ) | ($spl0 | $spl3 | $spl4 | $spl16)),
  inference(backward_demodulation,[],[f3493,f134])).
fof(f3493,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f3452,f73])).
fof(f3452,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X69),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl3 | $spl4)),
  inference(superposition,[],[f73,f3261])).
fof(f3261,plain,(
  ( ! [X4,X5] : (s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X4),s('type.num.num',X5))) = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X4),s('type.num.num',X5)))) ) | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f3169,f78])).
fof(f3169,plain,(
  ( ! [X4,X5] : (s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X4),s('type.num.num',X5))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun(fun('type.num.num',bool),fun('type.num.num',bool)),happ(s(fun(bool,fun(fun('type.num.num',bool),fun('type.num.num',bool))),'const.ind_type.FCONS_0'),s(bool,'T_0'))),s(fun('type.num.num',bool),X4))),s('type.num.num',X5)))) ) | ($spl0 | $spl3 | $spl4)),
  inference(superposition,[],[f78,f2710])).
fof(f2710,plain,(
  ( ! [X8,X7] : (s(X7,happ(s(fun(bool,X7),X8),s(bool,'T_0'))) = s(X7,happ(s(fun(bool,X7),X8),s(bool,'F_0')))) ) | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f2579,f79])).
fof(f2579,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,happ(s(fun(bool,X7),X8),s(bool,'F_0'))) = s(X7,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X7),X9),s(X7,X10),s(fun('type.num.num',X7),X11),s(fun(bool,X7),X8),s(X7,X12)))) ) | ($spl0 | $spl3 | $spl4)),
  inference(superposition,[],[f79,f1805])).
fof(f1805,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl0 | $spl3 | $spl4)),
  inference(backward_demodulation,[],[f1798,f1793])).
fof(f1793,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f92,f107])).
fof(f1798,plain,(
  s(bool,'F_0') = s(bool,sK0) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f68,f1795,f135])).
fof(f1795,plain,(
  s(bool,'T_0') != s(bool,sK0) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f95,f107])).
fof(f1826,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl3 | $spl4)),
  inference(superposition,[],[f1795,f1798])).
fof(f3518,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f3495,f67])).
fof(f3495,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl4)),
  inference(backward_demodulation,[],[f3493,f68])).
fof(f3494,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl13)),
  inference(subsumption_resolution,[],[f3493,f114])).
fof(f3471,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl13)),
  inference(subsumption_resolution,[],[f3470,f114])).
fof(f3470,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f3425,f73])).
fof(f3425,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'T_0'),s(fun('type.num.num',bool),X7),s('type.num.num','const.num.|48|_0')))) ) | ($spl0 | $spl3 | $spl4)),
  inference(superposition,[],[f3261,f73])).
fof(f1806,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f1805,f1794])).
fof(f1794,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl0 | $spl4 | $spl7)),
  inference(backward_demodulation,[],[f1793,f108])).
fof(f108,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | $spl7),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1792,plain,(
  $false | ($spl3 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f107,f121])).
fof(f121,plain,(
  s(bool,'T_0') != s(bool,sK1) | ($spl3 | $spl9)),
  inference(backward_demodulation,[],[f116,f95])).
fof(f116,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl9),
  inference(unit_resulting_resolution,[],[f111,f72])).
fof(f1791,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1790,f1789])).
fof(f1789,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1788,f116])).
fof(f1788,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f92,f123])).
fof(f123,plain,(
  s(bool,'F_0') = s(bool,sK1) | ($spl3 | $spl9)),
  inference(unit_resulting_resolution,[],[f121,f72])).
fof(f1790,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl7 | $spl9)),
  inference(forward_demodulation,[],[f108,f116])).
fof(f1782,plain,(
  ( ! [X37,X38,X40] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X40))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),sK2(X38,'const.pair.FST_0','type.pair.prod'(bool,X37))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X40)) ) | $spl22),
  inference(cnf_transformation,[],[f1782_D])).
fof(f1782_D,plain,(
  ( ! [X37,X38,X40] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X40))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),sK2(X38,'const.pair.FST_0','type.pair.prod'(bool,X37))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1786,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f1786_D])).
fof(f1786_D,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1762,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,'F_0'))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X79))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),sK2(X76,'const.pair.FST_0','type.pair.prod'(bool,X75))))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X75),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f477,f1672])).
fof(f1672,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),sK2(X1,'const.pair.FST_0','type.pair.prod'(bool,X0))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1671,f68])).
fof(f1671,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),sK2(X1,'const.pair.FST_0','type.pair.prod'(bool,X0))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl9)),
  inference(forward_demodulation,[],[f1669,f123])).
fof(f1669,plain,(
  ( ! [X0,X1] : (p(s(bool,sK1)) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),sK2(X1,'const.pair.FST_0','type.pair.prod'(bool,X0))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl9)),
  inference(superposition,[],[f1664,f646])).
fof(f646,plain,(
  ( ! [X16] : (s(bool,sK1) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f642,f68])).
fof(f642,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,sK1) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl3 | $spl9)),
  inference(superposition,[],[f135,f123])).
fof(f1664,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),sK2(X1,'const.pair.FST_0','type.pair.prod'(bool,X0)))))) ) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1617,f67])).
fof(f1617,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),sK2(X1,'const.pair.FST_0','type.pair.prod'(bool,X0)))))) ) | ($spl3 | $spl9)),
  inference(superposition,[],[f75,f1584])).
fof(f1584,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4)))) ) | ($spl3 | $spl9)),
  inference(forward_demodulation,[],[f1532,f123])).
fof(f1532,plain,(
  ( ! [X4,X3] : (s(bool,sK1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4)))) ) | ($spl3 | $spl9)),
  inference(superposition,[],[f86,f646])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9,X15] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X15),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X15),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X15),X8)))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','ARITY_0_1_const.pair.FST')).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK2(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK2(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK2(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : ! [X31,X32] : (! [X28] : (s(X9,X28) = s(X9,X32) => p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X28))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X31)))))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.bool.JRH_INDUCT_UTIL')).
fof(f477,plain,(
  ( ! [X47,X50,X48,X49] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X49))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X50))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X47,bool),'const.pair.,_2'(s(X47,X48),s(bool,X49))))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f183,f89])).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9,X15] : ! [X28,X29] : s(X15,X29) = s(X15,'const.pair.SND_1'(s('type.pair.prod'(X9,X15),'const.pair.,_2'(s(X9,X28),s(X15,X29)))))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.pair.SND')).
fof(f183,plain,(
  ( ! [X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f159,f159])).
fof(f159,plain,(
  ( ! [X0] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f119,f72])).
fof(f119,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl6 | $spl9)),
  inference(backward_demodulation,[],[f116,f109])).
fof(f1780,plain,(
  ( ! [X35,X34] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),sK2(X35,'const.pair.FST_0','type.pair.prod'(bool,X34))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X34),bool),'const.pair.FST_0')))) ) | $spl18),
  inference(cnf_transformation,[],[f1780_D])).
fof(f1780_D,plain,(
  ( ! [X35,X34] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),sK2(X35,'const.pair.FST_0','type.pair.prod'(bool,X34))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X34),bool),'const.pair.FST_0')))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1761,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0'))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X71),sK2(X72,'const.pair.FST_0','type.pair.prod'(bool,X71))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X71),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f476,f1672])).
fof(f476,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X44,bool),'const.pair.,_2'(s(X44,X45),s(bool,X46))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X46)))) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f159,f89])).
fof(f1784,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52)))))) ) | $spl24),
  inference(cnf_transformation,[],[f1784_D])).
fof(f1784_D,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1757,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,'F_0'),s(X55,X56))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X57))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),sK2(X54,'const.pair.FST_0','type.pair.prod'(bool,X53))))))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f370,f1672])).
fof(f370,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X36))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X38))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,X36),s(X35,X37))))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f183,f88])).
fof(f88,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X15,X9] : ! [X28,X29] : s(X9,X28) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X15),'const.pair.,_2'(s(X9,X28),s(X15,X29)))))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.pair.FST')).
fof(f1756,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X51),'const.pair.,_2'(s(bool,'F_0'),s(X51,X52))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X49),sK2(X50,'const.pair.FST_0','type.pair.prod'(bool,X49))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X49),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f369,f1672])).
fof(f369,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X33),s(X32,X34))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X33)))) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f159,f88])).
fof(f1781,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1753,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X39),s('type.num.num','const.num.|48|_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X40))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),sK2(X38,'const.pair.FST_0','type.pair.prod'(bool,X37))))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X37),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f318,f1672])).
fof(f318,plain,(
  ( ! [X23,X21,X22] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X21))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X23))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f183,f73])).
fof(f1752,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'F_0'),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X34),sK2(X35,'const.pair.FST_0','type.pair.prod'(bool,X34))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X34),bool),'const.pair.FST_0')))) ) | ($spl3 | $spl6 | $spl9)),
  inference(superposition,[],[f317,f1672])).
fof(f317,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X19),s(fun('type.num.num',bool),X20),s('type.num.num','const.num.|48|_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X19)))) ) | ($spl6 | $spl9)),
  inference(superposition,[],[f159,f73])).
fof(f141,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f138,f67])).
fof(f138,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f133,f68])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f134,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f136,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f130,f100])).
fof(f130,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f80,f72])).
fof(f132,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f128,f100])).
fof(f128,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f80,f72])).
fof(f112,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f114,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl13),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f105,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f72])).
fof(f107,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl4),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  s(bool,'T_0') = s(bool,sK1) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f111,plain,(
  s(bool,'F_0') != s(bool,sK0) | $spl9),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  s(bool,'F_0') != s(bool,sK0) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f103,plain,(
  s(bool,'F_0') != s(bool,sK0) | s(bool,'T_0') = s(bool,sK1) | $spl3),
  inference(superposition,[],[f95,f72])).
fof(f109,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | $spl6),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f102,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | s(bool,'T_0') = s(bool,sK1) | $spl0),
  inference(superposition,[],[f92,f72])).
fof(f97,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f96])).
fof(f96,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f93,f94])).
fof(f92,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) | $spl0),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f94,plain,(
  s(bool,sK0) = s(bool,sK1) | $spl2),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  s(bool,sK0) = s(bool,sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f65,plain,(
  s(bool,sK0) = s(bool,sK1) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1)))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) | s(bool,sK0) = s(bool,sK1)) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) | s(bool,sK0) != s(bool,sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f59])).
fof(f59,plain,(
  ? [X0,X1] : ((s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1)) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) | s(bool,X0) != s(bool,X1)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ? [X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) <~> s(bool,X0) = s(bool,X1))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) <=> s(bool,X0) = s(bool,X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X6))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X7))) <=> s(bool,X6) = s(bool,X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X6))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X7))) <=> s(bool,X6) = s(bool,X7))),
  file('thm.bvp_to_wordProps.tag_|49||49||split|2','thm.bvp_to_wordProps.tag_|49||49||split|2')).
fof(f93,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) | $spl1),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f95,plain,(
  s(bool,sK0) != s(bool,sK1) | $spl3),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s(bool,sK0) != s(bool,sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f66,plain,(
  s(bool,sK0) != s(bool,sK1) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK0))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK1)))),
  inference(cnf_transformation,[],[f60])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 96
% Passive clauses: 514
% Generated clauses: 5794
% Final active clauses: 45
% Final passive clauses: 231
% Input formulas: 21
% Initial clauses: 27
% 
% Duplicate literals: 250
% Trivial inequalities: 9
% Fw subsumption resolutions: 108
% Fw demodulations: 642
% Bw demodulations: 93
% Condensations: 486
% Global subsumptions: 12
% 
% Simple tautologies: 41
% Equational tautologies: 67
% Forward subsumptions: 3581
% Fw demodulations to eq. taut.: 58
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 155
% Factoring: 1
% Forward superposition: 1676
% Backward superposition: 2013
% Self superposition: 138
% Equality factoring: 35
% Equality resolution: 2
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 567
% SAT solver unit clauses: 37
% SAT solver binary clauses: 148
% 
% Sat splits: 13
% Sat splitting refutations: 13
% 
% Memory used [KB]: 7803
% Time elapsed: 0.223 s
% ------------------------------
% ------------------------------
% Success in time 1.55 s
