Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSound
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSound
% SZS output start Proof for thm.typeSound
fof(f2189,plain,(
  $false),
  inference(subsumption_resolution,[],[f2188,f297])).
fof(f297,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.ast.TC_exn_0')),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.ast.TC_exn_0')),
  inference(flattening,[],[f68])).
fof(f68,axiom,(
  ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_exn_0')),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.ast.tctor_distinct|split|60')).
fof(f2188,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_exn_0')),
  inference(forward_demodulation,[],[f2183,f229])).
fof(f229,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',sK0)))),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f217])).
fof(f217,plain,(
  ? [X0] : s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0)))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ~! [X0] : s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ~! [X0] : ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X6)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.typeSound.typeSound_4_prove_ltb|split|7')).
fof(f2183,plain,(
  s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',sK0)))),
  inference(trivial_inequality_removal,[],[f2180])).
fof(f2180,plain,(
  s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.ast.TC_word64_0') | s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',sK0)))),
  inference(superposition,[],[f2042,f229])).
fof(f2042,plain,(
  ( ! [X7] : (s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X7))) | s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X7)))) )),
  inference(superposition,[],[f320,f1891])).
fof(f1891,plain,(
  ( ! [X0] : (s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0))) = s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1(X0)))) | s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0)))) )),
  inference(superposition,[],[f1874,f1779])).
fof(f1779,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1(X0)))) | s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0)))) )),
  inference(superposition,[],[f1763,f309])).
fof(f309,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1(X0))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : (s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK1(X0)))) | s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),sK2(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f176])).
fof(f176,plain,(
  ! [X0] : (? [X2] : s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X2))) | ? [X1] : s('type.semanticPrimitives.tid_or_exn',X0) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X16] : (? [X17] : s('type.semanticPrimitives.tid_or_exn',X16) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X17))) | ? [X17] : s('type.semanticPrimitives.tid_or_exn',X16) = s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X17))))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.semanticPrimitives.tid_or_exn_nchotomy')).
fof(f1763,plain,(
  ( ! [X39] : (s('type.ast.tctor','const.ast.TC_exn_0') = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X39)))))) )),
  inference(forward_demodulation,[],[f1701,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X21,X22] : ! [X23,X24] : s(X22,X23) = s(X22,happ(s(fun(X21,X22),'_dst_x0_1'(s(X22,X23))),s(X21,X24)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','_dst_x0_ax')).
fof(f1701,plain,(
  ( ! [X39] : (s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X39))))) = s('type.ast.tctor',happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'_dst_x0_1'(s('type.ast.tctor','const.ast.TC_exn_0'))),s('type.ast.id'('type.list.list'('type.string.char')),X39)))) )),
  inference(superposition,[],[f327,f306])).
fof(f306,plain,(
  ( ! [X0] : (s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0))) = s('type.ast.tctor','const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn',X0),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'_dst_x0_1'(s('type.ast.tctor','const.ast.TC_exn_0')))))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X0))) = s('type.ast.tctor','const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn',X0),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'_dst_x0_1'(s('type.ast.tctor','const.ast.TC_exn_0')))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X10] : s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn',X10))) = s('type.ast.tctor','const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn',X10),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'_dst_x0_1'(s('type.ast.tctor','const.ast.TC_exn_0')))))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.typeSystem.tid_exn_to_tc_def')).
fof(f327,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))),
  inference(flattening,[],[f206])).
fof(f206,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X11] : ! [X7,X14,X15] : s(X11,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X15),s('type.ast.id'('type.list.list'('type.string.char')),X7))) = s(X11,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeExn_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X14),s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X15)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.semanticPrimitives.tid_or_exn_case_def|split|1')).
fof(f1874,plain,(
  ( ! [X39] : (s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X39))) = s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X39)))))) )),
  inference(forward_demodulation,[],[f1808,f305])).
fof(f305,plain,(
  ( ! [X0] : (s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))) = s('type.ast.tctor',happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s('type.ast.id'('type.list.list'('type.string.char')),X0)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0))) = s('type.ast.tctor',happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s('type.ast.id'('type.list.list'('type.string.char')),X0)))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X18] : s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X18))) = s('type.ast.tctor',happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s('type.ast.id'('type.list.list'('type.string.char')),X18)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','ARITY_0_1_const.ast.TC_name')).
fof(f1808,plain,(
  ( ! [X39] : (s('type.ast.tctor','const.typeSystem.tid_exn_to_tc_1'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X39))))) = s('type.ast.tctor',happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.ast.tctor'),'const.ast.TC_name_0'),s('type.ast.id'('type.list.list'('type.string.char')),X39)))) )),
  inference(superposition,[],[f328,f306])).
fof(f328,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s('type.ast.id'('type.list.list'('type.string.char')),X1))) = s(X0,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X1))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X2),s(fun('type.ast.id'('type.list.list'('type.string.char')),X0),X3)))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X11] : ! [X7,X14,X15] : s(X11,happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X14),s('type.ast.id'('type.list.list'('type.string.char')),X7))) = s(X11,'const.semanticPrimitives.tid_or_exn_CASE_3'(s('type.semanticPrimitives.tid_or_exn','const.semanticPrimitives.TypeId_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7))),s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X14),s(fun('type.ast.id'('type.list.list'('type.string.char')),X11),X15)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.semanticPrimitives.tid_or_exn_case_def|split|0')).
fof(f320,plain,(
  ( ! [X0] : (s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : s('type.ast.tctor','const.ast.TC_word64_0') != s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ~s('type.ast.tctor','const.ast.TC_word64_0') = s('type.ast.tctor','const.ast.TC_name_1'(s('type.ast.id'('type.list.list'('type.string.char')),X7)))),
  file('thm.typeSound.typeSound_4_prove_ltb|split|7','thm.ast.tctor_distinct|split|5')).
% SZS output end Proof for thm.typeSound
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 135
% Passive clauses: 462
% Generated clauses: 1888
% Final active clauses: 127
% Final passive clauses: 288
% Input formulas: 100
% Initial clauses: 107
% 
% Duplicate literals: 163
% Trivial inequalities: 18
% Fw subsumption resolutions: 88
% Fw demodulations: 89
% Bw demodulations: 2
% 
% Simple tautologies: 63
% Equational tautologies: 43
% Forward subsumptions: 932
% 
% Binary resolution: 38
% Factoring: 1
% Forward superposition: 580
% Backward superposition: 595
% Self superposition: 105
% Equality factoring: 85
% Equality resolution: 1
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 55
% SAT solver unit clauses: 5
% SAT solver binary clauses: 31
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 55
% TWLsolver calls for satisfiability: 29
% 
% Memory used [KB]: 2046
% Time elapsed: 0.114 s
% ------------------------------
% ------------------------------
% Success in time 0.13 s
