Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f875,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f873,f793,f716,f220,f372,f219,f374,f368,f814,f222,f375,f370,f223,f373,f371,f369])).
fof(f369,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) | $spl1),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f371,plain,(
  p(s(bool,sK11)) | $spl2),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  p(s(bool,sK11)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f373,plain,(
  s(sK7,sK8) != s(sK7,sK10) | $spl5),
  inference(cnf_transformation,[],[f373_D])).
fof(f373_D,plain,(
  s(sK7,sK8) != s(sK7,sK10) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f223,plain,(
  s(sK7,sK8) != s(sK7,sK10) | p(s(bool,sK11)) | s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10)))),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  (s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) | (p(s(bool,sK11)) & s(sK7,sK8) = s(sK7,sK9)) | (~p(s(bool,sK11)) & s(sK7,sK8) = s(sK7,sK10))) & (s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) | ((~p(s(bool,sK11)) | s(sK7,sK8) != s(sK7,sK9)) & (p(s(bool,sK11)) | s(sK7,sK8) != s(sK7,sK10))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9,sK10,sK11])],[f151])).
fof(f151,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) | (p(s(bool,X4)) & s(X0,X1) = s(X0,X2)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X3))) & (s(X0,X1) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) | ((~p(s(bool,X4)) | s(X0,X1) != s(X0,X2)) & (p(s(bool,X4)) | s(X0,X1) != s(X0,X3)))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) | ((p(s(bool,X4)) & s(X0,X1) = s(X0,X2)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X3)))) & (s(X0,X1) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) | ((~p(s(bool,X4)) | s(X0,X1) != s(X0,X2)) & (p(s(bool,X4)) | s(X0,X1) != s(X0,X3)))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ? [X0,X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) <~> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X2)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X3))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ~! [X0,X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) <=> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X2)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X3))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X2),s(X0,X3))) <=> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X2)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X8),s(X6,X9))) <=> ((p(s(bool,X10)) & s(X6,X7) = s(X6,X8)) | (~p(s(bool,X10)) & s(X6,X7) = s(X6,X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X8),s(X6,X9))) <=> ((p(s(bool,X10)) & s(X6,X7) = s(X6,X8)) | (~p(s(bool,X10)) & s(X6,X7) = s(X6,X9))))),
  file('thm.closProps.closProps_68_prove_ltb','thm.closProps.closProps_68_prove_ltb')).
fof(f370,plain,(
  ~p(s(bool,sK11)) | $spl3),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  ~p(s(bool,sK11)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f375,plain,(
  s(sK7,sK8) != s(sK7,sK9) | $spl7),
  inference(cnf_transformation,[],[f375_D])).
fof(f375_D,plain,(
  s(sK7,sK8) != s(sK7,sK9) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f222,plain,(
  s(sK7,sK8) != s(sK7,sK9) | ~p(s(bool,sK11)) | s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10)))),
  inference(cnf_transformation,[],[f152])).
fof(f814,plain,(
  $false | ($spl0 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f808,f375])).
fof(f808,plain,(
  s(sK7,sK8) = s(sK7,sK9) | ($spl0 | $spl2)),
  inference(superposition,[],[f794,f298])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.closProps.closProps_68_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f794,plain,(
  s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,'T_0'),s(sK7,sK9),s(sK7,sK10))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f368,f728])).
fof(f728,plain,(
  s(bool,'T_0') = s(bool,sK11) | $spl2),
  inference(resolution,[],[f371,f241])).
fof(f241,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.closProps.closProps_68_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f368,plain,(
  s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) | $spl0),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f374,plain,(
  s(sK7,sK8) = s(sK7,sK9) | $spl6),
  inference(cnf_transformation,[],[f374_D])).
fof(f374_D,plain,(
  s(sK7,sK8) = s(sK7,sK9) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f219,plain,(
  ~p(s(bool,sK11)) | s(sK7,sK8) = s(sK7,sK9) | s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10)))),
  inference(cnf_transformation,[],[f152])).
fof(f372,plain,(
  s(sK7,sK8) = s(sK7,sK10) | $spl4),
  inference(cnf_transformation,[],[f372_D])).
fof(f372_D,plain,(
  s(sK7,sK8) = s(sK7,sK10) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f220,plain,(
  s(sK7,sK8) = s(sK7,sK10) | p(s(bool,sK11)) | s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK10)))),
  inference(cnf_transformation,[],[f152])).
fof(f716,plain,(
  $false | ($spl1 | $spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f713])).
fof(f713,plain,(
  s(sK7,sK8) != s(sK7,sK8) | ($spl1 | $spl3 | $spl4)),
  inference(superposition,[],[f586,f299])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X11),s(X6,X12)))),
  file('thm.closProps.closProps_68_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f586,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,'F_0'),s(sK7,sK9),s(sK7,sK8))) | ($spl1 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f585,f390])).
fof(f390,plain,(
  s(bool,'F_0') = s(bool,sK11) | $spl3),
  inference(resolution,[],[f245,f370])).
fof(f245,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.closProps.closProps_68_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f585,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK9),s(sK7,sK8))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f369,f372])).
fof(f793,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f792,f298])).
fof(f792,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,'T_0'),s(sK7,sK8),s(sK7,sK10))) | ($spl1 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f791,f728])).
fof(f791,plain,(
  s(sK7,sK8) != s(sK7,'const.bool.COND_3'(s(bool,sK11),s(sK7,sK8),s(sK7,sK10))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f369,f374])).
fof(f873,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f867,f373])).
fof(f867,plain,(
  s(sK7,sK8) = s(sK7,sK10) | ($spl0 | $spl3)),
  inference(superposition,[],[f863,f299])).
fof(f863,plain,(
  s(sK7,sK8) = s(sK7,'const.bool.COND_3'(s(bool,'F_0'),s(sK7,sK9),s(sK7,sK10))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f368,f816])).
fof(f816,plain,(
  s(bool,'F_0') = s(bool,sK11) | $spl3),
  inference(resolution,[],[f370,f245])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 92
% Passive clauses: 204
% Generated clauses: 655
% Final active clauses: 76
% Final passive clauses: 78
% Input formulas: 46
% Initial clauses: 146
% 
% Duplicate literals: 14
% Trivial inequalities: 3
% Fw subsumption resolutions: 39
% Fw demodulations: 59
% Bw demodulations: 3
% 
% Simple tautologies: 66
% Equational tautologies: 9
% Forward subsumptions: 241
% 
% Binary resolution: 176
% Factoring: 6
% Forward superposition: 56
% Backward superposition: 125
% Self superposition: 6
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 36
% Split components: 78
% SAT solver clauses: 56
% SAT solver unit clauses: 6
% SAT solver binary clauses: 31
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 36
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 51
% TWLsolver calls for satisfiability: 18
% 
% Memory used [KB]: 895
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 0.07 s
