Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.evalProps
% (120049)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 3400
% Generated clauses: 13454
% Final active clauses: 103
% Final passive clauses: 3110
% Input formulas: 22
% Initial clauses: 36
% 
% Duplicate literals: 731
% Trivial inequalities: 289
% Fw subsumption resolutions: 2047
% Fw demodulations: 502
% Bw demodulations: 1
% 
% Simple tautologies: 339
% Equational tautologies: 83
% Forward subsumptions: 4732
% Fw demodulations to eq. taut.: 29
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 142
% Factoring: 2
% Forward superposition: 4343
% Backward superposition: 4767
% Self superposition: 247
% Equality factoring: 183
% Equality resolution: 42
% 
% Split clauses: 433
% Split components: 896
% SAT solver clauses: 1272
% SAT solver unit clauses: 41
% SAT solver binary clauses: 401
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 433
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1269
% TWLsolver calls for satisfiability: 145
% 
% Memory used [KB]: 5756
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (120063)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 653
% Generated clauses: 9880
% Final active clauses: 84
% Final passive clauses: 556
% Input formulas: 22
% Initial clauses: 36
% 
% Duplicate literals: 557
% Trivial inequalities: 149
% Fw subsumption resolutions: 989
% Fw demodulations: 486
% Bw demodulations: 5
% Condensations: 1317
% Global subsumptions: 5
% 
% Simple tautologies: 171
% Equational tautologies: 44
% Forward subsumptions: 3909
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 64
% Factoring: 1
% Forward superposition: 2935
% Backward superposition: 2316
% Self superposition: 153
% Equality factoring: 102
% Equality resolution: 39
% 
% Split clauses: 173
% Split components: 364
% SAT solver clauses: 1542
% SAT solver unit clauses: 53
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 11
% SAT solver learnt literals: 1
% 
% Sat splits: 173
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1521
% TWLsolver calls for satisfiability: 221
% 
% Memory used [KB]: 8827
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f27503,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f75,f114,f112,f74,f113,f111,f510,f511,f512,f1535,f1568,f1525,f1595,f1593,f3155,f3167,f3156,f3157,f3158,f3159,f3168,f3160,f3170,f3172,f3161,f3171,f3174,f3162,f3175,f3165,f3164,f3330,f3363,f3361,f3483,f3484,f3485,f3486,f3487,f3497,f3488,f3499,f3501,f3489,f3500,f3503,f3490,f3504,f3494,f3607,f3610,f3629,f3612,f3496,f3630,f3911,f3912,f3913,f3915,f3916,f3923,f3925,f3926,f3927,f3931,f3933,f3934,f3935,f10513,f10669,f11463,f11462,f11548,f11549,f11859,f12064,f12108,f12112,f12111,f14908,f14915,f14974,f14975,f14878,f14976,f14983,f14985,f14986,f14987,f14988,f14989,f14992,f14994,f14996,f14997,f15001,f15005,f14984,f15061,f14990,f15064,f14995,f15072,f15073,f15075,f15076,f15078,f14993,f15431,f15432,f15433,f15435,f15436,f15447,f15452,f15489,f15500,f15505,f15510,f15522,f15526,f15527,f15528,f15530,f15543,f1567,f15545,f20607,f20608,f20787,f20788,f20862,f20863,f25031,f25032,f25055,f25056,f25057,f25058,f25063,f25065,f25067,f25069,f25071,f25072,f25073,f25211,f25064,f25212,f25066,f25213,f25068,f25214,f25070,f25216,f15062,f25074,f26825,f26829,f26827,f27466,f27467,f27501])).
fof(f27501,plain,(
  $false | ($spl0 | $spl3 | $spl114)),
  inference(subsumption_resolution,[],[f27486,f114])).
fof(f27486,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) | ($spl0 | $spl114)),
  inference(backward_demodulation,[],[f27485,f25841])).
fof(f25841,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | $spl0),
  inference(unit_resulting_resolution,[],[f111,f885])).
fof(f885,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) != s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X0,X2),X4),s('type.semanticPrimitives.error_result'(X0),X1))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f576,f95])).
fof(f95,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f39])).
fof(f39,plain,(
  ! [X0,X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X16] : (? [X10] : s('type.semanticPrimitives.error_result'(X6),X16) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X10))) | ? [X10] : s('type.semanticPrimitives.error_result'(X6),X16) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.semanticPrimitives.error_result_nchotomy')).
fof(f576,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X14))) != s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13)))))) )),
  inference(superposition,[],[f102,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : ! [X8,X9] : s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(X7,X6),X8),s(X7,X9))))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9)))))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.evalProps.map_error_result_def|split|0')).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X15,X10] : ~s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X15))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X10)))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.semanticPrimitives.error_result_distinct')).
fof(f27485,plain,(
  s('type.semanticPrimitives.abort',sK4) = s('type.semanticPrimitives.abort',sK6(sK3,sK1)) | ($spl0 | $spl114)),
  inference(backward_demodulation,[],[f26829,f26542])).
fof(f26542,plain,(
  s('type.semanticPrimitives.abort',sK6(sK3,sK1)) = s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) | $spl0),
  inference(unit_resulting_resolution,[],[f26198,f25884])).
fof(f25884,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0))) | s('type.semanticPrimitives.abort',X0) = s('type.semanticPrimitives.abort',sK6(sK3,sK1))) ) | $spl0),
  inference(superposition,[],[f98,f25841])).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2)) & (s('type.semanticPrimitives.abort',X1) != s('type.semanticPrimitives.abort',X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2)))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) <=> s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))) <=> s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',X2))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X10,X15] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X15))) <=> s('type.semanticPrimitives.abort',X10) = s('type.semanticPrimitives.abort',X15))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.semanticPrimitives.error_result_|49||49||split|1')).
fof(f26198,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)))) | $spl0),
  inference(forward_demodulation,[],[f26185,f596])).
fof(f596,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0))) )),
  inference(unit_resulting_resolution,[],[f530,f100])).
fof(f100,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X10,X15] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X10))) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X15))) <=> s(X6,X10) = s(X6,X15))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.semanticPrimitives.error_result_|49||49||split|0')).
fof(f530,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f102,f95])).
fof(f26185,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1))) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)))) | $spl0),
  inference(unit_resulting_resolution,[],[f596,f25888,f540])).
fof(f540,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(X9),X10) != s('type.semanticPrimitives.error_result'(X9),X11) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X9)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X9),X11))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X9)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,sK5(X10,X9)))))) | s('type.semanticPrimitives.error_result'(X9),X10) = s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X10,X9))))) )),
  inference(superposition,[],[f101,f95])).
fof(f101,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f72])).
fof(f25888,plain,(
  ( ! [X5] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK1)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rraise_1'(s(sK1,X5)))))) ) | $spl0),
  inference(superposition,[],[f433,f25841])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))) )),
  inference(unit_resulting_resolution,[],[f102,f100])).
fof(f27467,plain,(
  $false | ($spl0 | $spl113)),
  inference(subsumption_resolution,[],[f27425,f27151])).
fof(f27151,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1))))) ) | $spl113),
  inference(unit_resulting_resolution,[],[f26842,f98])).
fof(f26842,plain,(
  s('type.semanticPrimitives.abort',sK4) != s('type.semanticPrimitives.abort',sK6(sK3,sK1)) | $spl113),
  inference(unit_resulting_resolution,[],[f26827,f101])).
fof(f27425,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | $spl0),
  inference(superposition,[],[f111,f25887])).
fof(f25887,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(sK1,X3),X4),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl0),
  inference(superposition,[],[f107,f25841])).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) = s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X0,X1),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) = s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X0,X1),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) = s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X0,X1),X2),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X6] : ! [X8,X10] : s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10)))))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.evalProps.map_error_result_def|split|1')).
fof(f27466,plain,(
  $false | ($spl0 | $spl113)),
  inference(subsumption_resolution,[],[f27423,f27151])).
fof(f27423,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | $spl0),
  inference(superposition,[],[f25887,f111])).
fof(f26827,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | $spl113),
  inference(cnf_transformation,[],[f26827_D])).
fof(f26827_D,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f26829,plain,(
  s('type.semanticPrimitives.abort',sK4) = s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) | $spl114),
  inference(cnf_transformation,[],[f26829_D])).
fof(f26829_D,plain,(
  s('type.semanticPrimitives.abort',sK4) = s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f26825,plain,(
  s('type.semanticPrimitives.abort',sK4) = s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | $spl0),
  inference(forward_demodulation,[],[f26792,f1050])).
fof(f1050,plain,(
  s('type.semanticPrimitives.abort',sK4) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)),sK0)) | $spl0),
  inference(unit_resulting_resolution,[],[f1041,f98])).
fof(f1041,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)),sK0)))) | $spl0),
  inference(forward_demodulation,[],[f1022,f111])).
fof(f1022,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) = s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)),sK0)))) | $spl0),
  inference(unit_resulting_resolution,[],[f111,f539])).
fof(f539,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) | s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X7,X6))))) )),
  inference(superposition,[],[f102,f95])).
fof(f26792,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK6(sK3,sK1)))) | s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK1),sK3)),'type.semanticPrimitives.error_result'(sK1)),sK1)) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)),sK0)) | $spl0),
  inference(superposition,[],[f1092,f26542])).
fof(f1092,plain,(
  ( ! [X22] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X22))) | s('type.semanticPrimitives.abort',X22) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)),sK0))) ) | $spl0),
  inference(superposition,[],[f100,f1050])).
fof(f25074,plain,(
  ( ! [X263,X260,X262] : (s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,'T_0') = s(bool,X263) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X262)) ) | $spl110),
  inference(cnf_transformation,[],[f25074_D])).
fof(f25074_D,plain,(
  ( ! [X263,X260,X262] : (s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,'T_0') = s(bool,X263) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X262)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f15062,plain,(
  ( ! [X141] : (p(s(bool,X141))) ) | $spl80),
  inference(cnf_transformation,[],[f15062_D])).
fof(f15062_D,plain,(
  ( ! [X141] : (p(s(bool,X141))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f25216,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,X261)) | s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,X260) = s(bool,X262) | s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X263)) ) | $spl13),
  inference(subsumption_resolution,[],[f22441,f85])).
fof(f85,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.bool.EQ_CLAUSES|split|1')).
fof(f22441,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,X261)) | s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,X260) = s(bool,X262) | s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X263) | s(bool,'T_0') = s(bool,X261)) ) | $spl13),
  inference(superposition,[],[f14198,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f125,f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f83,f83])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.map_error_result_Rtype_error',aHL_BOOL_CASES)).
fof(f14198,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,g(s(bool,X32))) = s(bool,g(s(bool,g(s(bool,X32)))))) ) | $spl13),
  inference(subsumption_resolution,[],[f14180,f85])).
fof(f14180,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,g(s(bool,X32))) = s(bool,g(s(bool,g(s(bool,X32))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl13),
  inference(superposition,[],[f14021,f11907])).
fof(f11907,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f11833,f83])).
fof(f11833,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl13),
  inference(unit_resulting_resolution,[],[f76,f3163,f789])).
fof(f789,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f788,f596])).
fof(f788,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f766,f76])).
fof(f766,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X9)),bool)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f106,f691])).
fof(f691,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f596,f83])).
fof(f106,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.bool.IMP_ANTISYM_AX')).
fof(f3163,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | $spl13),
  inference(cnf_transformation,[],[f3163_D])).
fof(f3163_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f76,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.bool.TRUTH')).
fof(f14021,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,g(s(bool,X0)))) ) | $spl13),
  inference(factoring,[],[f12894])).
fof(f12894,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X10) = s(bool,g(s(bool,X9)))) ) | $spl13),
  inference(subsumption_resolution,[],[f12809,f85])).
fof(f12809,plain,(
  ( ! [X10,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X10) = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl13),
  inference(superposition,[],[f103,f11907])).
fof(f103,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f25070,plain,(
  ( ! [X255,X254,X252] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,'T_0') = s(bool,X255) | s(bool,X254) = s(bool,X255) | s(bool,X252) = s(bool,X254)) ) | $spl106),
  inference(cnf_transformation,[],[f25070_D])).
fof(f25070_D,plain,(
  ( ! [X255,X254,X252] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,'T_0') = s(bool,X255) | s(bool,X254) = s(bool,X255) | s(bool,X252) = s(bool,X254)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f25214,plain,(
  ( ! [X255,X253,X254,X252] : (p(s(bool,X253)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,X252) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl13),
  inference(subsumption_resolution,[],[f22439,f85])).
fof(f22439,plain,(
  ( ! [X255,X253,X254,X252] : (p(s(bool,X253)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,X252) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255) | s(bool,'T_0') = s(bool,X253)) ) | $spl13),
  inference(superposition,[],[f13913,f199])).
fof(f13913,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X32)))))))) ) | $spl13),
  inference(subsumption_resolution,[],[f13885,f85])).
fof(f13885,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X32))))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl13),
  inference(superposition,[],[f13357,f11907])).
fof(f13357,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X31)))))) ) | $spl13),
  inference(subsumption_resolution,[],[f13337,f85])).
fof(f13337,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X31))))) | s(bool,'T_0') = s(bool,X31)) ) | $spl13),
  inference(superposition,[],[f12893,f11907])).
fof(f12893,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') = s(bool,g(s(bool,X8)))) ) | $spl13),
  inference(subsumption_resolution,[],[f12808,f85])).
fof(f12808,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') = s(bool,g(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl13),
  inference(superposition,[],[f90,f11907])).
fof(f90,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.bool.EQ_CLAUSES|split|3')).
fof(f25068,plain,(
  ( ! [X251,X250,X248] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,'T_0') = s(bool,X251) | s(bool,X250) = s(bool,X251) | s(bool,X248) = s(bool,X250)) ) | $spl104),
  inference(cnf_transformation,[],[f25068_D])).
fof(f25068_D,plain,(
  ( ! [X251,X250,X248] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,'T_0') = s(bool,X251) | s(bool,X250) = s(bool,X251) | s(bool,X248) = s(bool,X250)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f25213,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X249)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,X248) = s(bool,X250) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251)) ) | $spl13),
  inference(subsumption_resolution,[],[f22438,f85])).
fof(f22438,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X249)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,X248) = s(bool,X250) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,'T_0') = s(bool,X249)) ) | $spl13),
  inference(superposition,[],[f13850,f199])).
fof(f13850,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X32)))))))) ) | $spl13),
  inference(subsumption_resolution,[],[f13823,f85])).
fof(f13823,plain,(
  ( ! [X32] : (p(s(bool,X32)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X32))))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl13),
  inference(superposition,[],[f13122,f11907])).
fof(f13122,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X31)))))) ) | $spl13),
  inference(subsumption_resolution,[],[f13104,f85])).
fof(f13104,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X31))))) | s(bool,'T_0') = s(bool,X31)) ) | $spl13),
  inference(superposition,[],[f12891,f11907])).
fof(f12891,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,g(s(bool,X5)))) ) | $spl13),
  inference(subsumption_resolution,[],[f12805,f85])).
fof(f12805,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl13),
  inference(superposition,[],[f85,f11907])).
fof(f25066,plain,(
  ( ! [X246,X244,X247] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,'T_0') = s(bool,X247) | s(bool,X246) = s(bool,X247) | s(bool,X244) = s(bool,X246)) ) | $spl102),
  inference(cnf_transformation,[],[f25066_D])).
fof(f25066_D,plain,(
  ( ! [X246,X244,X247] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,'T_0') = s(bool,X247) | s(bool,X246) = s(bool,X247) | s(bool,X244) = s(bool,X246)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f25212,plain,(
  ( ! [X246,X244,X247,X245] : (p(s(bool,X245)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,X244) = s(bool,X246) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247)) ) | $spl13),
  inference(subsumption_resolution,[],[f22437,f85])).
fof(f22437,plain,(
  ( ! [X246,X244,X247,X245] : (p(s(bool,X245)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,X244) = s(bool,X246) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247) | s(bool,'T_0') = s(bool,X245)) ) | $spl13),
  inference(superposition,[],[f13357,f199])).
fof(f25064,plain,(
  ( ! [X243,X242,X240] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,'T_0') = s(bool,X243) | s(bool,X242) = s(bool,X243) | s(bool,X240) = s(bool,X242)) ) | $spl100),
  inference(cnf_transformation,[],[f25064_D])).
fof(f25064_D,plain,(
  ( ! [X243,X242,X240] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,'T_0') = s(bool,X243) | s(bool,X242) = s(bool,X243) | s(bool,X240) = s(bool,X242)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f25211,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,X240) = s(bool,X242) | s(bool,X242) = s(bool,X243) | s(bool,'T_0') = s(bool,X243)) ) | $spl13),
  inference(subsumption_resolution,[],[f22436,f85])).
fof(f22436,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,X240) = s(bool,X242) | s(bool,X242) = s(bool,X243) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X241)) ) | $spl13),
  inference(superposition,[],[f13122,f199])).
fof(f25073,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,X261)) | s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,X260) = s(bool,X262) | s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X263)) ) | $spl13),
  inference(subsumption_resolution,[],[f21926,f85])).
fof(f21926,plain,(
  ( ! [X261,X263,X260,X262] : (p(s(bool,X261)) | s(bool,g(s(bool,X260))) = s(bool,g(s(bool,g(s(bool,X260))))) | s(bool,X260) = s(bool,X262) | s(bool,'T_0') = s(bool,X261) | s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X263)) ) | $spl13),
  inference(superposition,[],[f14198,f199])).
fof(f25072,plain,(
  ( ! [X259,X256,X258] : (s(bool,X256) = s(bool,g(s(bool,X256))) | s(bool,'T_0') = s(bool,X259) | s(bool,X258) = s(bool,X259) | s(bool,X256) = s(bool,X258)) ) | $spl108),
  inference(cnf_transformation,[],[f25072_D])).
fof(f25072_D,plain,(
  ( ! [X259,X256,X258] : (s(bool,X256) = s(bool,g(s(bool,X256))) | s(bool,'T_0') = s(bool,X259) | s(bool,X258) = s(bool,X259) | s(bool,X256) = s(bool,X258)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f25071,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,X257)) | s(bool,X256) = s(bool,g(s(bool,X256))) | s(bool,X256) = s(bool,X258) | s(bool,X258) = s(bool,X259) | s(bool,'T_0') = s(bool,X259)) ) | $spl13),
  inference(subsumption_resolution,[],[f21925,f85])).
fof(f21925,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,X257)) | s(bool,X256) = s(bool,g(s(bool,X256))) | s(bool,X256) = s(bool,X258) | s(bool,'T_0') = s(bool,X257) | s(bool,X258) = s(bool,X259) | s(bool,'T_0') = s(bool,X259)) ) | $spl13),
  inference(superposition,[],[f14021,f199])).
fof(f25069,plain,(
  ( ! [X255,X253,X254,X252] : (p(s(bool,X253)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,X252) = s(bool,X254) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl13),
  inference(subsumption_resolution,[],[f21924,f85])).
fof(f21924,plain,(
  ( ! [X255,X253,X254,X252] : (p(s(bool,X253)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X252))))))) | s(bool,X252) = s(bool,X254) | s(bool,'T_0') = s(bool,X253) | s(bool,X254) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | $spl13),
  inference(superposition,[],[f13913,f199])).
fof(f25067,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X249)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,X248) = s(bool,X250) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251)) ) | $spl13),
  inference(subsumption_resolution,[],[f21923,f85])).
fof(f21923,plain,(
  ( ! [X251,X249,X250,X248] : (p(s(bool,X249)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X248))))))) | s(bool,X248) = s(bool,X250) | s(bool,'T_0') = s(bool,X249) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251)) ) | $spl13),
  inference(superposition,[],[f13850,f199])).
fof(f25065,plain,(
  ( ! [X246,X244,X247,X245] : (p(s(bool,X245)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,X244) = s(bool,X246) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247)) ) | $spl13),
  inference(subsumption_resolution,[],[f21922,f85])).
fof(f21922,plain,(
  ( ! [X246,X244,X247,X245] : (p(s(bool,X245)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X244))))) | s(bool,X244) = s(bool,X246) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,X247) | s(bool,'T_0') = s(bool,X247)) ) | $spl13),
  inference(superposition,[],[f13357,f199])).
fof(f25063,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,X240) = s(bool,X242) | s(bool,X242) = s(bool,X243) | s(bool,'T_0') = s(bool,X243)) ) | $spl13),
  inference(subsumption_resolution,[],[f21921,f85])).
fof(f21921,plain,(
  ( ! [X243,X241,X242,X240] : (p(s(bool,X241)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X240))))) | s(bool,X240) = s(bool,X242) | s(bool,'T_0') = s(bool,X241) | s(bool,X242) = s(bool,X243) | s(bool,'T_0') = s(bool,X243)) ) | $spl13),
  inference(superposition,[],[f13122,f199])).
fof(f25058,plain,(
  ( ! [X229,X228,X226] : (s(bool,'F_0') = s(bool,g(s(bool,X226))) | s(bool,'T_0') = s(bool,X229) | s(bool,X228) = s(bool,X229) | s(bool,X226) = s(bool,X228)) ) | $spl98),
  inference(cnf_transformation,[],[f25058_D])).
fof(f25058_D,plain,(
  ( ! [X229,X228,X226] : (s(bool,'F_0') = s(bool,g(s(bool,X226))) | s(bool,'T_0') = s(bool,X229) | s(bool,X228) = s(bool,X229) | s(bool,X226) = s(bool,X228)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f25057,plain,(
  ( ! [X229,X227,X228,X226] : (p(s(bool,X227)) | s(bool,'F_0') = s(bool,g(s(bool,X226))) | s(bool,X226) = s(bool,X228) | s(bool,X228) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) ) | $spl13),
  inference(subsumption_resolution,[],[f21918,f85])).
fof(f21918,plain,(
  ( ! [X229,X227,X228,X226] : (p(s(bool,X227)) | s(bool,'F_0') = s(bool,g(s(bool,X226))) | s(bool,X226) = s(bool,X228) | s(bool,'T_0') = s(bool,X227) | s(bool,X228) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) ) | $spl13),
  inference(superposition,[],[f12893,f199])).
fof(f25056,plain,(
  ( ! [X218,X220,X221] : (s(bool,'T_0') != s(bool,g(s(bool,X218))) | s(bool,'T_0') = s(bool,X221) | s(bool,X220) = s(bool,X221) | s(bool,X218) = s(bool,X220)) ) | $spl96),
  inference(cnf_transformation,[],[f25056_D])).
fof(f25056_D,plain,(
  ( ! [X218,X220,X221] : (s(bool,'T_0') != s(bool,g(s(bool,X218))) | s(bool,'T_0') = s(bool,X221) | s(bool,X220) = s(bool,X221) | s(bool,X218) = s(bool,X220)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f25055,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | s(bool,'T_0') != s(bool,g(s(bool,X218))) | s(bool,X218) = s(bool,X220) | s(bool,X220) = s(bool,X221) | s(bool,'T_0') = s(bool,X221)) ) | $spl13),
  inference(subsumption_resolution,[],[f21916,f85])).
fof(f21916,plain,(
  ( ! [X218,X220,X219,X221] : (p(s(bool,X219)) | s(bool,'T_0') != s(bool,g(s(bool,X218))) | s(bool,X218) = s(bool,X220) | s(bool,'T_0') = s(bool,X219) | s(bool,X220) = s(bool,X221) | s(bool,'T_0') = s(bool,X221)) ) | $spl13),
  inference(superposition,[],[f12891,f199])).
fof(f25032,plain,(
  ( ! [X6,X4,X7] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X6)) ) | $spl94),
  inference(cnf_transformation,[],[f25032_D])).
fof(f25032_D,plain,(
  ( ! [X6,X4,X7] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X6)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f25031,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4) | s(bool,X4) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f21871,f85])).
fof(f21871,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f85,f199])).
fof(f20863,plain,(
  ( ! [X35] : (s(bool,g(s(bool,X35))) = s(bool,g(s(bool,g(s(bool,X35))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl92),
  inference(cnf_transformation,[],[f20863_D])).
fof(f20863_D,plain,(
  ( ! [X35] : (s(bool,g(s(bool,X35))) = s(bool,g(s(bool,g(s(bool,X35))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f20862,plain,(
  ( ! [X35,X36] : (s(bool,g(s(bool,X35))) = s(bool,g(s(bool,g(s(bool,X35))))) | p(s(bool,X36)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(forward_demodulation,[],[f20861,f596])).
fof(f20861,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))) = s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(subsumption_resolution,[],[f20826,f85])).
fof(f20826,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))) = s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,X36)) ) | $spl13),
  inference(superposition,[],[f14198,f613])).
fof(f613,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f530,f125])).
fof(f20788,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl90),
  inference(cnf_transformation,[],[f20788_D])).
fof(f20788_D,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20787,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | p(s(bool,X36)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(forward_demodulation,[],[f20786,f596])).
fof(f20786,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(subsumption_resolution,[],[f20743,f85])).
fof(f20743,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,X36)) ) | $spl13),
  inference(superposition,[],[f13913,f613])).
fof(f20608,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl88),
  inference(cnf_transformation,[],[f20608_D])).
fof(f20608_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f20607,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,X35))))))) | p(s(bool,X36)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(forward_demodulation,[],[f20606,f596])).
fof(f20606,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35)))) ) | $spl13),
  inference(subsumption_resolution,[],[f20563,f85])).
fof(f20563,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X35)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) | s(bool,'T_0') = s(bool,X36)) ) | $spl13),
  inference(superposition,[],[f13850,f613])).
fof(f15545,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) | $spl87),
  inference(cnf_transformation,[],[f15545_D])).
fof(f15545_D,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f1567,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,X32)) ) | $spl4),
  inference(cnf_transformation,[],[f1567_D])).
fof(f1567_D,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,X32)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f15543,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,g(s(bool,'F_0'))))) ) | $spl13),
  inference(condensation,[],[f15542])).
fof(f15542,plain,(
  ( ! [X52,X51] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X51)) ) | $spl13),
  inference(subsumption_resolution,[],[f12503,f13014])).
fof(f13014,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f13013])).
fof(f13013,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f12970])).
fof(f12970,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f423,f83])).
fof(f423,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f274,f76])).
fof(f274,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f127,f125])).
fof(f127,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f89,f83])).
fof(f89,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f66])).
fof(f12503,plain,(
  ( ! [X52,X51] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X51)) ) | $spl13),
  inference(superposition,[],[f789,f11856])).
fof(f11856,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(trivial_inequality_removal,[],[f11845])).
fof(f11845,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f3163,f125])).
fof(f15530,plain,(
  $false | $spl80),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964,f14965,f14881,f14974,f14878,f14986,f15009,f14729,f15011,f15012,f15013,f15014,f15015,f15016,f15017,f15018,f15019,f14749,f15020,f15022,f15024,f15025,f15026,f15027,f15028,f15039,f15052,f15079,f14826,f15080,f14833,f15082,f14837,f15083,f15084,f15085,f15086,f14844,f15062,f2576,f15456,f15465,f15466,f15496,f15499,f15509,f15521,f15529])).
fof(f15529,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl80),
  inference(subsumption_resolution,[],[f89,f15062])).
fof(f15521,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl80),
  inference(condensation,[],[f15520])).
fof(f15520,plain,(
  ( ! [X10,X9] : (s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) ) | $spl80),
  inference(subsumption_resolution,[],[f789,f15062])).
fof(f15509,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl80),
  inference(condensation,[],[f15508])).
fof(f15508,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | s(bool,'F_0') = s(bool,X7)) ) | $spl80),
  inference(subsumption_resolution,[],[f924,f15466])).
fof(f15499,plain,(
  ( ! [X57,X56] : (s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,X57) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool))) ) | $spl80),
  inference(forward_demodulation,[],[f15498,f15456])).
fof(f15498,plain,(
  ( ! [X57,X56] : (s(bool,X57) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool)) | s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl80),
  inference(subsumption_resolution,[],[f2159,f15466])).
fof(f15496,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61)))) ) | $spl80),
  inference(forward_demodulation,[],[f15495,f15456])).
fof(f15495,plain,(
  ( ! [X61,X60] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) | s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl80),
  inference(subsumption_resolution,[],[f2280,f15456])).
fof(f15466,plain,(
  ( ! [X146] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146)))) ) | $spl80),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964,f14965,f14881,f14974,f14878,f14986,f15009,f14729,f15011,f15012,f15013,f15014,f15015,f15016,f15017,f15018,f15019,f14749,f15020,f15022,f15024,f15025,f15026,f15027,f15028,f15039,f15052,f15079,f14826,f15080,f14833,f15082,f14837,f15083,f15084,f15085,f15086,f14844,f15062,f2576,f15456,f15465])).
fof(f15465,plain,(
  ( ! [X146,X149,X147,X148] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))) | s('type.semanticPrimitives.error_result'(X147),X148) = s('type.semanticPrimitives.error_result'(X147),'const.evalProps.map_error_result_2'(s(fun(bool,X147),X149),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),X148))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),'const.semanticPrimitives.Rraise_1'(s(X147,happ(s(fun(bool,X147),X149),s(bool,'T_0')))))))) ) | $spl80),
  inference(forward_demodulation,[],[f10091,f15456])).
fof(f15456,plain,(
  ( ! [X191] : (s(bool,'T_0') = s(bool,X191)) ) | $spl80),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964,f14965,f14881,f14974,f14878,f14986,f15009,f14729,f15011,f15012,f15013,f15014,f15015,f15016,f15017,f15018,f15019,f14749,f15020,f15022,f15024,f15025,f15026,f15027,f15028,f15039,f15052,f15079,f14826,f15080,f14833,f15082,f14837,f15083,f15084,f15085,f15086,f14844,f15062,f2576])).
fof(f2576,plain,(
  ( ! [X191,X189,X190,X188] : (s(X190,f(s(bool,X191))) != s(X190,g(s(bool,X191))) | s(fun(bool,X190),X188) = s(fun(bool,X190),X189) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X188,X189,X190,bool))),bool)) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f109,f689])).
fof(f14844,plain,(
  ( ! [X239,X240,X238] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X239))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X240))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X238)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X240)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X238))) | s(bool,'T_0') = s(bool,X239)) )),
  inference(superposition,[],[f619,f613])).
fof(f15086,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s('type.semanticPrimitives.error_result'(X233),X234) = s('type.semanticPrimitives.error_result'(X233),'const.evalProps.map_error_result_2'(s(fun(bool,X233),X235),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X231))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X233)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X233),X234))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X233)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X233),'const.semanticPrimitives.Rraise_1'(s(X233,happ(s(fun(bool,X233),X235),s(bool,X232))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X231))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(forward_demodulation,[],[f14842,f596])).
fof(f14842,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X233)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X233),X234))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X233)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X233),'const.semanticPrimitives.Rraise_1'(s(X233,happ(s(fun(bool,X233),X235),s(bool,X232))))))) | s('type.semanticPrimitives.error_result'(X233),X234) = s('type.semanticPrimitives.error_result'(X233),'const.evalProps.map_error_result_2'(s(fun(bool,X233),X235),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X231)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X231))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f586,f613])).
fof(f15085,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X228)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X228),X229))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X228)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X228),'const.evalProps.map_error_result_2'(s(fun(bool,X228),X230),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X226))))))) | s('type.semanticPrimitives.error_result'(X228),X229) != s('type.semanticPrimitives.error_result'(X228),'const.semanticPrimitives.Rraise_1'(s(X228,happ(s(fun(bool,X228),X230),s(bool,X227))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X226))) | s(bool,'T_0') = s(bool,X227)) )),
  inference(forward_demodulation,[],[f14841,f596])).
fof(f14841,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s('type.semanticPrimitives.error_result'(X228),X229) != s('type.semanticPrimitives.error_result'(X228),'const.semanticPrimitives.Rraise_1'(s(X228,happ(s(fun(bool,X228),X230),s(bool,X227))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X228)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X228),X229))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X228)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X228),'const.evalProps.map_error_result_2'(s(fun(bool,X228),X230),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X226)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X226))) | s(bool,'T_0') = s(bool,X227)) )),
  inference(superposition,[],[f584,f613])).
fof(f15084,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s('type.semanticPrimitives.error_result'(X223),X224) = s('type.semanticPrimitives.error_result'(X223),'const.semanticPrimitives.Rraise_1'(s(X223,happ(s(fun(bool,X223),X225),s(bool,X221))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),X224))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),'const.evalProps.map_error_result_2'(s(fun(bool,X223),X225),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X222))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X221))) | s(bool,'T_0') = s(bool,X222)) )),
  inference(forward_demodulation,[],[f14840,f596])).
fof(f14840,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),X224))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),'const.evalProps.map_error_result_2'(s(fun(bool,X223),X225),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X222))))))) | s('type.semanticPrimitives.error_result'(X223),X224) = s('type.semanticPrimitives.error_result'(X223),'const.semanticPrimitives.Rraise_1'(s(X223,happ(s(fun(bool,X223),X225),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X221)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X221))) | s(bool,'T_0') = s(bool,X222)) )),
  inference(superposition,[],[f580,f613])).
fof(f15083,plain,(
  ( ! [X218,X216,X220,X219,X217] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X218)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X218),X219))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X218)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X218),'const.semanticPrimitives.Rraise_1'(s(X218,happ(s(fun(bool,X218),X220),s(bool,X216))))))) | s('type.semanticPrimitives.error_result'(X218),X219) != s('type.semanticPrimitives.error_result'(X218),'const.evalProps.map_error_result_2'(s(fun(bool,X218),X220),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X217))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X216))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(forward_demodulation,[],[f14839,f596])).
fof(f14839,plain,(
  ( ! [X218,X216,X220,X219,X217] : (s('type.semanticPrimitives.error_result'(X218),X219) != s('type.semanticPrimitives.error_result'(X218),'const.evalProps.map_error_result_2'(s(fun(bool,X218),X220),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X217))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X218)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X218),X219))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X218)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X218),'const.semanticPrimitives.Rraise_1'(s(X218,happ(s(fun(bool,X218),X220),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X216)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X216))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(superposition,[],[f578,f613])).
fof(f14837,plain,(
  ( ! [X210,X208,X206,X209,X207] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X207))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X208,bool),X209),s('type.semanticPrimitives.error_result'(X208),'const.semanticPrimitives.Rraise_1'(s(X208,X210))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X206)),bool)) = s(bool,happ(s(fun(X208,bool),X209),s(X208,X210))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X206))) | s(bool,'T_0') = s(bool,X207)) )),
  inference(superposition,[],[f574,f613])).
fof(f15082,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s(X203,X204) = s(X203,happ(s(fun(bool,X203),X205),s(bool,X201))) | s('type.semanticPrimitives.error_result'(X203),'const.semanticPrimitives.Rraise_1'(s(X203,X204))) != s('type.semanticPrimitives.error_result'(X203),'const.evalProps.map_error_result_2'(s(fun(bool,X203),X205),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X202))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,X202)) )),
  inference(forward_demodulation,[],[f14836,f596])).
fof(f14836,plain,(
  ( ! [X204,X202,X205,X203,X201] : (s('type.semanticPrimitives.error_result'(X203),'const.semanticPrimitives.Rraise_1'(s(X203,X204))) != s('type.semanticPrimitives.error_result'(X203),'const.evalProps.map_error_result_2'(s(fun(bool,X203),X205),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X202))))) | s(X203,X204) = s(X203,happ(s(fun(bool,X203),X205),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X201)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X201))) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f574,f613])).
fof(f14833,plain,(
  ( ! [X189,X190,X188] : (s('type.semanticPrimitives.error_result'(bool),X190) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X189))) | s(bool,sK5(X190,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X188)),bool)) | s('type.semanticPrimitives.error_result'(bool),X190) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X190,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f537,f613])).
fof(f15080,plain,(
  ( ! [X173,X174,X172] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X172))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X174))) | s(bool,X173) != s(bool,X174) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X172))) | s(bool,'T_0') = s(bool,X173)) )),
  inference(forward_demodulation,[],[f14828,f596])).
fof(f14828,plain,(
  ( ! [X173,X174,X172] : (s(bool,X173) != s(bool,X174) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X174))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X172)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X172))) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f101,f613])).
fof(f14826,plain,(
  ( ! [X167,X166,X168] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X167))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X168))) | s(bool,X168) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X166)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X166))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f100,f613])).
fof(f15079,plain,(
  ( ! [X163,X161,X162,X160] : (s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,happ(s(fun(bool,X162),X163),s(bool,X161))))) = s('type.semanticPrimitives.error_result'(X162),'const.evalProps.map_error_result_2'(s(fun(bool,X162),X163),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X160))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X160))) | s(bool,'T_0') = s(bool,X161)) )),
  inference(subsumption_resolution,[],[f14824,f102])).
fof(f14824,plain,(
  ( ! [X163,X161,X162,X160] : (s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,happ(s(fun(bool,X162),X163),s(bool,X161))))) = s('type.semanticPrimitives.error_result'(X162),'const.evalProps.map_error_result_2'(s(fun(bool,X162),X163),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X160))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X160))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X160)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X160))) | s(bool,'T_0') = s(bool,X161)) )),
  inference(superposition,[],[f572,f613])).
fof(f15052,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(forward_demodulation,[],[f15051,f596])).
fof(f15051,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(condensation,[],[f15050])).
fof(f15050,plain,(
  ( ! [X85,X83,X84] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X85))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X83)),bool)))) | s(bool,'T_0') = s(bool,X85) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X83))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f14791,f83])).
fof(f14791,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,X84) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X85))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X83)),bool)))) | s(bool,'T_0') = s(bool,X85) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X83))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f480,f613])).
fof(f15039,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X0)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(forward_demodulation,[],[f15038,f596])).
fof(f15038,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(condensation,[],[f15037])).
fof(f15037,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X29)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29)))) )),
  inference(subsumption_resolution,[],[f14774,f85])).
fof(f14774,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X29)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f103,f613])).
fof(f15028,plain,(
  ( ! [X356,X358,X357,X359] : (s(X358,f(s(bool,X359))) != s(X358,g(s(bool,X359))) | s(fun(bool,X358),X356) = s(fun(bool,X358),X357) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X359))) | s(bool,'T_0') = s(bool,sK7(X356,X357,X358,bool))) )),
  inference(forward_demodulation,[],[f14758,f596])).
fof(f14758,plain,(
  ( ! [X356,X358,X357,X359] : (s(X358,f(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X359)),bool)))) != s(X358,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X359)),bool)))) | s(fun(bool,X358),X356) = s(fun(bool,X358),X357) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X359))) | s(bool,'T_0') = s(bool,sK7(X356,X357,X358,bool))) )),
  inference(superposition,[],[f109,f613])).
fof(f15027,plain,(
  ( ! [X349,X348] : (s('type.semanticPrimitives.error_result'(bool),X348) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X349))) | s('type.semanticPrimitives.error_result'(bool),X348) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X348,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X349))) | s(bool,'T_0') = s(bool,sK5(X348,bool))) )),
  inference(forward_demodulation,[],[f14755,f596])).
fof(f14755,plain,(
  ( ! [X349,X348] : (s('type.semanticPrimitives.error_result'(bool),X348) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X349)),bool)))) | s('type.semanticPrimitives.error_result'(bool),X348) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X348,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X349))) | s(bool,'T_0') = s(bool,sK5(X348,bool))) )),
  inference(superposition,[],[f95,f613])).
fof(f15026,plain,(
  ( ! [X344,X346,X345,X347] : (s('type.semanticPrimitives.error_result'(X346),'const.evalProps.map_error_result_2'(s(fun(bool,X346),X347),s('type.semanticPrimitives.error_result'(bool),X344))) = s('type.semanticPrimitives.error_result'(X346),'const.semanticPrimitives.Rraise_1'(s(X346,happ(s(fun(bool,X346),X347),s(bool,X345))))) | s('type.semanticPrimitives.error_result'(bool),X344) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X344,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X345))) | s(bool,'T_0') = s(bool,sK5(X344,bool))) )),
  inference(forward_demodulation,[],[f14754,f596])).
fof(f14754,plain,(
  ( ! [X344,X346,X345,X347] : (s('type.semanticPrimitives.error_result'(X346),'const.evalProps.map_error_result_2'(s(fun(bool,X346),X347),s('type.semanticPrimitives.error_result'(bool),X344))) = s('type.semanticPrimitives.error_result'(X346),'const.semanticPrimitives.Rraise_1'(s(X346,happ(s(fun(bool,X346),X347),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X345)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),X344) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X344,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X345))) | s(bool,'T_0') = s(bool,sK5(X344,bool))) )),
  inference(superposition,[],[f572,f613])).
fof(f15025,plain,(
  ( ! [X335,X337,X334,X336] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X337))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X334,bool),X335),s('type.semanticPrimitives.error_result'(X334),X336))) | s('type.semanticPrimitives.error_result'(X334),X336) = s('type.semanticPrimitives.error_result'(X334),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X336,X334)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X337))) | s(bool,'T_0') = s(bool,happ(s(fun(X334,bool),X335),s(X334,sK5(X336,X334))))) )),
  inference(forward_demodulation,[],[f14753,f596])).
fof(f14753,plain,(
  ( ! [X335,X337,X334,X336] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X337)),bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X334,bool),X335),s('type.semanticPrimitives.error_result'(X334),X336))) | s('type.semanticPrimitives.error_result'(X334),X336) = s('type.semanticPrimitives.error_result'(X334),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X336,X334)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X337))) | s(bool,'T_0') = s(bool,happ(s(fun(X334,bool),X335),s(X334,sK5(X336,X334))))) )),
  inference(superposition,[],[f572,f613])).
fof(f15024,plain,(
  ( ! [X325,X321,X323,X324,X320,X322] : (s('type.semanticPrimitives.error_result'(X324),'const.semanticPrimitives.Rraise_1'(s(X324,happ(s(fun(bool,X324),X325),s(bool,X323))))) = s('type.semanticPrimitives.error_result'(X324),'const.evalProps.map_error_result_2'(s(fun(bool,X324),X325),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X320,bool),X321),s(X320,X322))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X323))) | s(bool,'T_0') = s(bool,happ(s(fun(X320,bool),X321),s(X320,X322)))) )),
  inference(forward_demodulation,[],[f15023,f596])).
fof(f15023,plain,(
  ( ! [X325,X321,X323,X324,X320,X322] : (s('type.semanticPrimitives.error_result'(X324),'const.semanticPrimitives.Rraise_1'(s(X324,happ(s(fun(bool,X324),X325),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X323)),bool)))))) = s('type.semanticPrimitives.error_result'(X324),'const.evalProps.map_error_result_2'(s(fun(bool,X324),X325),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X320,bool),X321),s(X320,X322))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X323))) | s(bool,'T_0') = s(bool,happ(s(fun(X320,bool),X321),s(X320,X322)))) )),
  inference(forward_demodulation,[],[f14752,f108])).
fof(f14752,plain,(
  ( ! [X325,X321,X323,X324,X320,X322] : (s('type.semanticPrimitives.error_result'(X324),'const.semanticPrimitives.Rraise_1'(s(X324,happ(s(fun(bool,X324),X325),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X323)),bool)))))) = s('type.semanticPrimitives.error_result'(X324),'const.evalProps.map_error_result_2'(s(fun(bool,X324),X325),s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X320,bool),X321),s('type.semanticPrimitives.error_result'(X320),'const.semanticPrimitives.Rraise_1'(s(X320,X322))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X323))) | s(bool,'T_0') = s(bool,happ(s(fun(X320,bool),X321),s(X320,X322)))) )),
  inference(superposition,[],[f573,f613])).
fof(f15022,plain,(
  ( ! [X315,X316,X318,X317,X319] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X319))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X315,bool),X316),s(X315,X317))))))) | s('type.semanticPrimitives.error_result'(bool),X319) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X318))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X318))) | s(bool,'T_0') = s(bool,happ(s(fun(X315,bool),X316),s(X315,X317)))) )),
  inference(forward_demodulation,[],[f15021,f108])).
fof(f15021,plain,(
  ( ! [X315,X316,X318,X317,X319] : (s('type.semanticPrimitives.error_result'(bool),X319) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X318))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X319))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X315,bool),X316),s('type.semanticPrimitives.error_result'(X315),'const.semanticPrimitives.Rraise_1'(s(X315,X317))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X318))) | s(bool,'T_0') = s(bool,happ(s(fun(X315,bool),X316),s(X315,X317)))) )),
  inference(forward_demodulation,[],[f14751,f596])).
fof(f14751,plain,(
  ( ! [X315,X316,X318,X317,X319] : (s('type.semanticPrimitives.error_result'(bool),X319) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X318)),bool)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X319))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X315,bool),X316),s('type.semanticPrimitives.error_result'(X315),'const.semanticPrimitives.Rraise_1'(s(X315,X317))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X318))) | s(bool,'T_0') = s(bool,happ(s(fun(X315,bool),X316),s(X315,X317)))) )),
  inference(superposition,[],[f584,f613])).
fof(f15020,plain,(
  ( ! [X313,X310,X312,X314,X311] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X314))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X313))))) | s('type.semanticPrimitives.error_result'(bool),X314) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X310,bool),X311),s('type.semanticPrimitives.error_result'(X310),'const.semanticPrimitives.Rraise_1'(s(X310,X312))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X313))) | s(bool,'T_0') = s(bool,happ(s(fun(X310,bool),X311),s(X310,X312)))) )),
  inference(forward_demodulation,[],[f14750,f596])).
fof(f14750,plain,(
  ( ! [X313,X310,X312,X314,X311] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X314))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X313)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),X314) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X310,bool),X311),s('type.semanticPrimitives.error_result'(X310),'const.semanticPrimitives.Rraise_1'(s(X310,X312))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X313))) | s(bool,'T_0') = s(bool,happ(s(fun(X310,bool),X311),s(X310,X312)))) )),
  inference(superposition,[],[f586,f613])).
fof(f14749,plain,(
  ( ! [X308,X306,X305,X307] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X308)),bool)) != s(bool,f(s(X305,sK7(X306,X307,bool,X305)))) | s(fun(X305,bool),X306) = s(fun(X305,bool),X307) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X308))) | s(bool,'T_0') = s(bool,g(s(X305,sK7(X306,X307,bool,X305))))) )),
  inference(superposition,[],[f109,f613])).
fof(f15019,plain,(
  ( ! [X288,X287,X286] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X287))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X288))) | s(bool,X288) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X286)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X287))) | s(bool,'T_0') = s(bool,X286)) )),
  inference(forward_demodulation,[],[f14746,f596])).
fof(f14746,plain,(
  ( ! [X288,X287,X286] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X288))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X287)),bool)))) | s(bool,X288) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X286)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X287))) | s(bool,'T_0') = s(bool,X286)) )),
  inference(superposition,[],[f619,f613])).
fof(f15018,plain,(
  ( ! [X285,X283,X284] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X284))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X285))) | s(bool,X283) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X285)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X284))) | s(bool,'T_0') = s(bool,X283)) )),
  inference(forward_demodulation,[],[f14745,f596])).
fof(f14745,plain,(
  ( ! [X285,X283,X284] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X285))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X284)),bool)))) | s(bool,X283) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X285)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X284))) | s(bool,'T_0') = s(bool,X283)) )),
  inference(superposition,[],[f619,f613])).
fof(f15017,plain,(
  ( ! [X280,X277,X279,X276,X278] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X278)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X278),X279))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X278)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X278),'const.semanticPrimitives.Rraise_1'(s(X278,happ(s(fun(bool,X278),X280),s(bool,X277))))))) | s('type.semanticPrimitives.error_result'(X278),X279) = s('type.semanticPrimitives.error_result'(X278),'const.evalProps.map_error_result_2'(s(fun(bool,X278),X280),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X276))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X277))) | s(bool,'T_0') = s(bool,X276)) )),
  inference(forward_demodulation,[],[f14743,f596])).
fof(f14743,plain,(
  ( ! [X280,X277,X279,X276,X278] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X278)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X278),X279))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X278)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X278),'const.semanticPrimitives.Rraise_1'(s(X278,happ(s(fun(bool,X278),X280),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X277)),bool)))))))) | s('type.semanticPrimitives.error_result'(X278),X279) = s('type.semanticPrimitives.error_result'(X278),'const.evalProps.map_error_result_2'(s(fun(bool,X278),X280),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X276))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X277))) | s(bool,'T_0') = s(bool,X276)) )),
  inference(superposition,[],[f586,f613])).
fof(f15016,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s('type.semanticPrimitives.error_result'(X273),X274) != s('type.semanticPrimitives.error_result'(X273),'const.semanticPrimitives.Rraise_1'(s(X273,happ(s(fun(bool,X273),X275),s(bool,X272))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X273)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X273),X274))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X273)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X273),'const.evalProps.map_error_result_2'(s(fun(bool,X273),X275),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X271))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X272))) | s(bool,'T_0') = s(bool,X271)) )),
  inference(forward_demodulation,[],[f14742,f596])).
fof(f14742,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s('type.semanticPrimitives.error_result'(X273),X274) != s('type.semanticPrimitives.error_result'(X273),'const.semanticPrimitives.Rraise_1'(s(X273,happ(s(fun(bool,X273),X275),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X272)),bool)))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X273)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X273),X274))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X273)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X273),'const.evalProps.map_error_result_2'(s(fun(bool,X273),X275),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X271))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X272))) | s(bool,'T_0') = s(bool,X271)) )),
  inference(superposition,[],[f584,f613])).
fof(f15015,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X268),X269))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X268),'const.evalProps.map_error_result_2'(s(fun(bool,X268),X270),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X267))))))) | s('type.semanticPrimitives.error_result'(X268),X269) = s('type.semanticPrimitives.error_result'(X268),'const.semanticPrimitives.Rraise_1'(s(X268,happ(s(fun(bool,X268),X270),s(bool,X266))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X267))) | s(bool,'T_0') = s(bool,X266)) )),
  inference(forward_demodulation,[],[f14741,f596])).
fof(f14741,plain,(
  ( ! [X269,X267,X268,X270,X266] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X268),X269))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X268),'const.evalProps.map_error_result_2'(s(fun(bool,X268),X270),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X267)),bool)))))))) | s('type.semanticPrimitives.error_result'(X268),X269) = s('type.semanticPrimitives.error_result'(X268),'const.semanticPrimitives.Rraise_1'(s(X268,happ(s(fun(bool,X268),X270),s(bool,X266))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X267))) | s(bool,'T_0') = s(bool,X266)) )),
  inference(superposition,[],[f580,f613])).
fof(f15014,plain,(
  ( ! [X261,X263,X265,X262,X264] : (s('type.semanticPrimitives.error_result'(X263),X264) != s('type.semanticPrimitives.error_result'(X263),'const.evalProps.map_error_result_2'(s(fun(bool,X263),X265),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X262))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X263)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X263),X264))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X263)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X263),'const.semanticPrimitives.Rraise_1'(s(X263,happ(s(fun(bool,X263),X265),s(bool,X261))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X262))) | s(bool,'T_0') = s(bool,X261)) )),
  inference(forward_demodulation,[],[f14740,f596])).
fof(f14740,plain,(
  ( ! [X261,X263,X265,X262,X264] : (s('type.semanticPrimitives.error_result'(X263),X264) != s('type.semanticPrimitives.error_result'(X263),'const.evalProps.map_error_result_2'(s(fun(bool,X263),X265),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X262)),bool)))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X263)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X263),X264))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X263)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X263),'const.semanticPrimitives.Rraise_1'(s(X263,happ(s(fun(bool,X263),X265),s(bool,X261))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X262))) | s(bool,'T_0') = s(bool,X261)) )),
  inference(superposition,[],[f578,f613])).
fof(f15013,plain,(
  ( ! [X255,X253,X251,X254,X252] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X252))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X253,bool),X254),s('type.semanticPrimitives.error_result'(X253),'const.semanticPrimitives.Rraise_1'(s(X253,X255))))) | s(bool,X251) = s(bool,happ(s(fun(X253,bool),X254),s(X253,X255))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X252))) | s(bool,'T_0') = s(bool,X251)) )),
  inference(forward_demodulation,[],[f14738,f596])).
fof(f14738,plain,(
  ( ! [X255,X253,X251,X254,X252] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X252)),bool)))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X253,bool),X254),s('type.semanticPrimitives.error_result'(X253),'const.semanticPrimitives.Rraise_1'(s(X253,X255))))) | s(bool,X251) = s(bool,happ(s(fun(X253,bool),X254),s(X253,X255))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X252))) | s(bool,'T_0') = s(bool,X251)) )),
  inference(superposition,[],[f574,f613])).
fof(f15012,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s('type.semanticPrimitives.error_result'(X248),'const.semanticPrimitives.Rraise_1'(s(X248,X249))) != s('type.semanticPrimitives.error_result'(X248),'const.evalProps.map_error_result_2'(s(fun(bool,X248),X250),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X247))))) | s(X248,X249) = s(X248,happ(s(fun(bool,X248),X250),s(bool,X246))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X247))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(forward_demodulation,[],[f14737,f596])).
fof(f14737,plain,(
  ( ! [X249,X246,X250,X248,X247] : (s('type.semanticPrimitives.error_result'(X248),'const.semanticPrimitives.Rraise_1'(s(X248,X249))) != s('type.semanticPrimitives.error_result'(X248),'const.evalProps.map_error_result_2'(s(fun(bool,X248),X250),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X247)),bool)))))) | s(X248,X249) = s(X248,happ(s(fun(bool,X248),X250),s(bool,X246))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X247))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f574,f613])).
fof(f15011,plain,(
  ( ! [X235,X233,X234] : (s('type.semanticPrimitives.error_result'(bool),X235) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X234))) | s(bool,X233) = s(bool,sK5(X235,bool)) | s('type.semanticPrimitives.error_result'(bool),X235) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X235,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X234))) | s(bool,'T_0') = s(bool,X233)) )),
  inference(forward_demodulation,[],[f14734,f596])).
fof(f14734,plain,(
  ( ! [X235,X233,X234] : (s('type.semanticPrimitives.error_result'(bool),X235) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X234)),bool)))) | s(bool,X233) = s(bool,sK5(X235,bool)) | s('type.semanticPrimitives.error_result'(bool),X235) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X235,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X234))) | s(bool,'T_0') = s(bool,X233)) )),
  inference(superposition,[],[f537,f613])).
fof(f14729,plain,(
  ( ! [X218,X219,X217] : (s(bool,X219) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X218)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X217))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X219))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X218))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(superposition,[],[f101,f613])).
fof(f15009,plain,(
  ( ! [X212,X213,X211] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X212))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X213))) | s(bool,X211) = s(bool,X213) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X212))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(forward_demodulation,[],[f14727,f596])).
fof(f14727,plain,(
  ( ! [X212,X213,X211] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X213))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X212)),bool)))) | s(bool,X211) = s(bool,X213) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X212))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f100,f613])).
fof(f14881,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,X0) = s(bool,X1) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2)))) )),
  inference(condensation,[],[f14880])).
fof(f14880,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X47) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X46)),bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X46)))) )),
  inference(duplicate_literal_removal,[],[f14675])).
fof(f14675,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X47) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X46)),bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X46))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f423,f613])).
fof(f14965,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5)))) )),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f2576,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964])).
fof(f14964,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5)))) )),
  inference(forward_demodulation,[],[f14889,f596])).
fof(f14889,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X5)),bool))) | s(bool,'T_0') = s(bool,X4) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5)))) )),
  inference(duplicate_literal_removal,[],[f14661])).
fof(f14661,plain,(
  ( ! [X4,X5] : (~p(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X5)),bool))) | s(bool,'T_0') = s(bool,X4) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f86,f613])).
fof(f14934,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f14933])).
fof(f14933,plain,(
  ( ! [X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(subsumption_resolution,[],[f14609,f2823])).
fof(f14609,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(superposition,[],[f613,f689])).
fof(f14601,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool)) = s(bool,happ(s(fun(X77,bool),X78),s(X77,X79))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X80))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X77,bool),X78),s('type.semanticPrimitives.error_result'(X77),'const.semanticPrimitives.Rraise_1'(s(X77,X79)))),bool))) )),
  inference(superposition,[],[f613,f2283])).
fof(f14905,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(condensation,[],[f14904])).
fof(f14904,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X28)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool))) )),
  inference(duplicate_literal_removal,[],[f14586])).
fof(f14586,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X28)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f613,f689])).
fof(f14907,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(condensation,[],[f14906])).
fof(f14906,plain,(
  ( ! [X19,X17,X18] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X18)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(duplicate_literal_removal,[],[f14583])).
fof(f14583,plain,(
  ( ! [X19,X17,X18] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X18)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f613,f687])).
fof(f14248,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),X8) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X5,bool),X6),s('type.semanticPrimitives.error_result'(X5),X7))) | s('type.semanticPrimitives.error_result'(bool),X8) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X8,bool)) | s('type.semanticPrimitives.error_result'(X5),X7) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X7,X5))))) )),
  inference(superposition,[],[f1865,f572])).
fof(f14266,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s(X32,X34))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) != s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s(X32,X34))) | s(bool,'T_0') = s(bool,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34)))),'type.semanticPrimitives.error_result'(bool)),bool))) )),
  inference(forward_demodulation,[],[f14214,f2283])).
fof(f14214,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) != s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s(X32,X34))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),sK5('const.evalProps.map_error_result_2'(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34)))),'type.semanticPrimitives.error_result'(bool))) | s(bool,'T_0') = s(bool,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X32,'type.semanticPrimitives.error_result'(bool)),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34)))),'type.semanticPrimitives.error_result'(bool)),bool))) )),
  inference(superposition,[],[f1865,f2283])).
fof(f14265,plain,(
  ( ! [X30,X31] : (s('type.semanticPrimitives.error_result'(bool),X30) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X30) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X30)),'type.semanticPrimitives.error_result'(bool)),bool))) )),
  inference(forward_demodulation,[],[f14213,f596])).
fof(f14213,plain,(
  ( ! [X30,X31] : (s('type.semanticPrimitives.error_result'(bool),X30) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X31))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X30)),'type.semanticPrimitives.error_result'(bool))) | s(bool,'T_0') = s(bool,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X30)),'type.semanticPrimitives.error_result'(bool)),bool))) )),
  inference(superposition,[],[f1865,f596])).
fof(f14264,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X8,bool),X9),s(X8,happ(s(fun(X10,X8),X11),s(X10,X12))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X13))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X8,bool),X9),s(X8,happ(s(fun(X10,X8),X11),s(X10,X12))))))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X8,bool),X9),s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X10,X8),X11),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X12)))))),bool))) )),
  inference(forward_demodulation,[],[f14209,f573])).
fof(f14209,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X13))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X8,bool),X9),s(X8,happ(s(fun(X10,X8),X11),s(X10,X12))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X8,bool),X9),s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X10,X8),X11),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X12))))))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X8,bool),X9),s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X10,X8),X11),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X12)))))),bool))) )),
  inference(superposition,[],[f1865,f573])).
fof(f14263,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6)))),bool))) )),
  inference(forward_demodulation,[],[f14208,f108])).
fof(f14208,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X7))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6))))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6)))),bool))) )),
  inference(superposition,[],[f1865,f108])).
fof(f14002,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X32))) != s('type.semanticPrimitives.error_result'(X31),'const.evalProps.map_error_result_2'(s(fun(X33,X31),X34),s('type.semanticPrimitives.error_result'(X33),'const.evalProps.map_error_result_2'(s(fun(X25,X33),X35),s('type.semanticPrimitives.error_result'(X25),'const.evalProps.map_error_result_2'(s(fun(X26,X25),X27),s('type.semanticPrimitives.error_result'(X26),'const.evalProps.map_error_result_2'(s(fun(X28,X26),X29),s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,X30)))))))))))) )),
  inference(superposition,[],[f886,f573])).
fof(f14001,plain,(
  ( ! [X24,X23,X21,X19,X17,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X21))) != s('type.semanticPrimitives.error_result'(X20),'const.evalProps.map_error_result_2'(s(fun(X22,X20),X23),s('type.semanticPrimitives.error_result'(X22),'const.evalProps.map_error_result_2'(s(fun(X16,X22),X24),s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),X19))))))) | s('type.semanticPrimitives.error_result'(X17),X19) = s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X19,X17))))) )),
  inference(superposition,[],[f886,f572])).
fof(f13936,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X1,bool),X2),s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X3,X1),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,X5))))))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f886,f536])).
fof(f13720,plain,(
  ( ! [X532,X530,X533,X529,X531] : (s('type.semanticPrimitives.error_result'(X532),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X533))) != s('type.semanticPrimitives.error_result'(X532),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(X529,'type.semanticPrimitives.abort'),X530),s(X529,X531))))) | s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X533)),X532)) = s('type.semanticPrimitives.abort',sK5('const.evalProps.map_error_result_2'(s(fun(X529,'type.semanticPrimitives.abort'),X530),s('type.semanticPrimitives.error_result'(X529),'const.semanticPrimitives.Rraise_1'(s(X529,X531)))),'type.semanticPrimitives.abort'))) )),
  inference(superposition,[],[f800,f2283])).
fof(f13708,plain,(
  ( ! [X466,X469,X465,X467,X468] : (s('type.semanticPrimitives.error_result'(X468),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X469))) != s('type.semanticPrimitives.error_result'(X468),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',happ(s(fun(X465,'type.semanticPrimitives.abort'),X466),s(X465,X467))))) | s('type.semanticPrimitives.abort',X469) = s('type.semanticPrimitives.abort',sK5('const.evalProps.map_error_result_2'(s(fun(X465,'type.semanticPrimitives.abort'),X466),s('type.semanticPrimitives.error_result'(X465),'const.semanticPrimitives.Rraise_1'(s(X465,X467)))),'type.semanticPrimitives.abort'))) )),
  inference(superposition,[],[f98,f2283])).
fof(f13707,plain,(
  ( ! [X461,X463,X459,X460,X462,X458,X464] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X458)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X458),happ(s(fun(X459,'type.semanticPrimitives.error_result'(X458)),X460),s(X459,X461))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X458)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X458),'const.semanticPrimitives.Rraise_1'(s(X458,happ(s(fun(X462,X458),X463),s(X462,X464))))))) | s('type.semanticPrimitives.error_result'(X458),'const.evalProps.map_error_result_2'(s(fun(X462,X458),X463),s('type.semanticPrimitives.error_result'(X462),'const.semanticPrimitives.Rraise_1'(s(X462,X464))))) = s('type.semanticPrimitives.error_result'(X458),sK5('const.evalProps.map_error_result_2'(s(fun(X459,'type.semanticPrimitives.error_result'(X458)),X460),s('type.semanticPrimitives.error_result'(X459),'const.semanticPrimitives.Rraise_1'(s(X459,X461)))),'type.semanticPrimitives.error_result'(X458)))) )),
  inference(superposition,[],[f586,f2283])).
fof(f13705,plain,(
  ( ! [X445,X447,X449,X444,X446,X448,X450] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X444)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X444),happ(s(fun(X445,'type.semanticPrimitives.error_result'(X444)),X446),s(X445,X447))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X444)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X444),'const.evalProps.map_error_result_2'(s(fun(X448,X444),X449),s('type.semanticPrimitives.error_result'(X448),'const.semanticPrimitives.Rraise_1'(s(X448,X450))))))) | s('type.semanticPrimitives.error_result'(X444),'const.semanticPrimitives.Rraise_1'(s(X444,happ(s(fun(X448,X444),X449),s(X448,X450))))) = s('type.semanticPrimitives.error_result'(X444),sK5('const.evalProps.map_error_result_2'(s(fun(X445,'type.semanticPrimitives.error_result'(X444)),X446),s('type.semanticPrimitives.error_result'(X445),'const.semanticPrimitives.Rraise_1'(s(X445,X447)))),'type.semanticPrimitives.error_result'(X444)))) )),
  inference(superposition,[],[f580,f2283])).
fof(f13703,plain,(
  ( ! [X436,X432,X434,X433,X435] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X432)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X432),X436))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X432)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X432),happ(s(fun(X433,'type.semanticPrimitives.error_result'(X432)),X434),s(X433,X435))))) | s('type.semanticPrimitives.error_result'(X432),'const.semanticPrimitives.Rraise_1'(s(X432,sK5(X436,X432)))) = s('type.semanticPrimitives.error_result'(X432),sK5('const.evalProps.map_error_result_2'(s(fun(X433,'type.semanticPrimitives.error_result'(X432)),X434),s('type.semanticPrimitives.error_result'(X433),'const.semanticPrimitives.Rraise_1'(s(X433,X435)))),'type.semanticPrimitives.error_result'(X432))) | s('type.semanticPrimitives.error_result'(X432),X436) = s('type.semanticPrimitives.error_result'(X432),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X436,X432))))) )),
  inference(superposition,[],[f542,f2283])).
fof(f13788,plain,(
  ( ! [X428,X430,X429,X431,X427] : (s('type.semanticPrimitives.error_result'(X427),happ(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s(X428,X430))) = s('type.semanticPrimitives.error_result'(X427),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s('type.semanticPrimitives.error_result'(X428),'const.semanticPrimitives.Rraise_1'(s(X428,X430)))),'type.semanticPrimitives.error_result'(X427)),X427)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X427)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X427),X431))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X427)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X427),happ(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s(X428,X430))))) | s('type.semanticPrimitives.error_result'(X427),X431) = s('type.semanticPrimitives.error_result'(X427),'const.semanticPrimitives.Rraise_1'(s(X427,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s('type.semanticPrimitives.error_result'(X428),'const.semanticPrimitives.Rraise_1'(s(X428,X430)))),'type.semanticPrimitives.error_result'(X427)),X427))))) )),
  inference(forward_demodulation,[],[f13702,f2283])).
fof(f13702,plain,(
  ( ! [X428,X430,X429,X431,X427] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X427)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X427),X431))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X427)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X427),happ(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s(X428,X430))))) | s('type.semanticPrimitives.error_result'(X427),X431) = s('type.semanticPrimitives.error_result'(X427),'const.semanticPrimitives.Rraise_1'(s(X427,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s('type.semanticPrimitives.error_result'(X428),'const.semanticPrimitives.Rraise_1'(s(X428,X430)))),'type.semanticPrimitives.error_result'(X427)),X427)))) | s('type.semanticPrimitives.error_result'(X427),sK5('const.evalProps.map_error_result_2'(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s('type.semanticPrimitives.error_result'(X428),'const.semanticPrimitives.Rraise_1'(s(X428,X430)))),'type.semanticPrimitives.error_result'(X427))) = s('type.semanticPrimitives.error_result'(X427),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X428,'type.semanticPrimitives.error_result'(X427)),X429),s('type.semanticPrimitives.error_result'(X428),'const.semanticPrimitives.Rraise_1'(s(X428,X430)))),'type.semanticPrimitives.error_result'(X427)),X427))))) )),
  inference(superposition,[],[f542,f2283])).
fof(f13786,plain,(
  ( ! [X420,X418,X421,X417,X419] : (s('type.semanticPrimitives.error_result'(X417),happ(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s(X418,X420))) = s('type.semanticPrimitives.error_result'(X417),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s('type.semanticPrimitives.error_result'(X418),'const.semanticPrimitives.Rraise_1'(s(X418,X420)))),'type.semanticPrimitives.error_result'(X417)),X417)))) | s('type.semanticPrimitives.error_result'(X417),X421) != s('type.semanticPrimitives.error_result'(X417),happ(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s(X418,X420))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X417)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X417),X421))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X417)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X417),'const.semanticPrimitives.Rraise_1'(s(X417,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s('type.semanticPrimitives.error_result'(X418),'const.semanticPrimitives.Rraise_1'(s(X418,X420)))),'type.semanticPrimitives.error_result'(X417)),X417))))))) )),
  inference(forward_demodulation,[],[f13700,f2283])).
fof(f13700,plain,(
  ( ! [X420,X418,X421,X417,X419] : (s('type.semanticPrimitives.error_result'(X417),X421) != s('type.semanticPrimitives.error_result'(X417),happ(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s(X418,X420))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X417)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X417),X421))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X417)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X417),'const.semanticPrimitives.Rraise_1'(s(X417,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s('type.semanticPrimitives.error_result'(X418),'const.semanticPrimitives.Rraise_1'(s(X418,X420)))),'type.semanticPrimitives.error_result'(X417)),X417)))))) | s('type.semanticPrimitives.error_result'(X417),sK5('const.evalProps.map_error_result_2'(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s('type.semanticPrimitives.error_result'(X418),'const.semanticPrimitives.Rraise_1'(s(X418,X420)))),'type.semanticPrimitives.error_result'(X417))) = s('type.semanticPrimitives.error_result'(X417),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X418,'type.semanticPrimitives.error_result'(X417)),X419),s('type.semanticPrimitives.error_result'(X418),'const.semanticPrimitives.Rraise_1'(s(X418,X420)))),'type.semanticPrimitives.error_result'(X417)),X417))))) )),
  inference(superposition,[],[f540,f2283])).
fof(f13785,plain,(
  ( ! [X412,X414,X416,X413,X415] : (s('type.semanticPrimitives.error_result'(X412),happ(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s(X413,X415))) = s('type.semanticPrimitives.error_result'(X412),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s('type.semanticPrimitives.error_result'(X413),'const.semanticPrimitives.Rraise_1'(s(X413,X415)))),'type.semanticPrimitives.error_result'(X412)),X412)))) | s('type.semanticPrimitives.error_result'(X412),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X416))) != s('type.semanticPrimitives.error_result'(X412),happ(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s(X413,X415)))) )),
  inference(forward_demodulation,[],[f13699,f2283])).
fof(f13699,plain,(
  ( ! [X412,X414,X416,X413,X415] : (s('type.semanticPrimitives.error_result'(X412),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X416))) != s('type.semanticPrimitives.error_result'(X412),happ(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s(X413,X415))) | s('type.semanticPrimitives.error_result'(X412),sK5('const.evalProps.map_error_result_2'(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s('type.semanticPrimitives.error_result'(X413),'const.semanticPrimitives.Rraise_1'(s(X413,X415)))),'type.semanticPrimitives.error_result'(X412))) = s('type.semanticPrimitives.error_result'(X412),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X413,'type.semanticPrimitives.error_result'(X412)),X414),s('type.semanticPrimitives.error_result'(X413),'const.semanticPrimitives.Rraise_1'(s(X413,X415)))),'type.semanticPrimitives.error_result'(X412)),X412))))) )),
  inference(superposition,[],[f539,f2283])).
fof(f13784,plain,(
  ( ! [X408,X410,X407,X409,X411] : (s('type.semanticPrimitives.error_result'(X407),happ(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s(X408,X410))) = s('type.semanticPrimitives.error_result'(X407),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s('type.semanticPrimitives.error_result'(X408),'const.semanticPrimitives.Rraise_1'(s(X408,X410)))),'type.semanticPrimitives.error_result'(X407)),X407)))) | s('type.semanticPrimitives.error_result'(X407),'const.semanticPrimitives.Rraise_1'(s(X407,X411))) != s('type.semanticPrimitives.error_result'(X407),happ(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s(X408,X410))) | s(X407,X411) = s(X407,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s('type.semanticPrimitives.error_result'(X408),'const.semanticPrimitives.Rraise_1'(s(X408,X410)))),'type.semanticPrimitives.error_result'(X407)),X407))) )),
  inference(forward_demodulation,[],[f13698,f2283])).
fof(f13698,plain,(
  ( ! [X408,X410,X407,X409,X411] : (s('type.semanticPrimitives.error_result'(X407),'const.semanticPrimitives.Rraise_1'(s(X407,X411))) != s('type.semanticPrimitives.error_result'(X407),happ(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s(X408,X410))) | s(X407,X411) = s(X407,sK5(sK5('const.evalProps.map_error_result_2'(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s('type.semanticPrimitives.error_result'(X408),'const.semanticPrimitives.Rraise_1'(s(X408,X410)))),'type.semanticPrimitives.error_result'(X407)),X407)) | s('type.semanticPrimitives.error_result'(X407),sK5('const.evalProps.map_error_result_2'(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s('type.semanticPrimitives.error_result'(X408),'const.semanticPrimitives.Rraise_1'(s(X408,X410)))),'type.semanticPrimitives.error_result'(X407))) = s('type.semanticPrimitives.error_result'(X407),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.evalProps.map_error_result_2'(s(fun(X408,'type.semanticPrimitives.error_result'(X407)),X409),s('type.semanticPrimitives.error_result'(X408),'const.semanticPrimitives.Rraise_1'(s(X408,X410)))),'type.semanticPrimitives.error_result'(X407)),X407))))) )),
  inference(superposition,[],[f537,f2283])).
fof(f13697,plain,(
  ( ! [X405,X401,X403,X404,X406,X400,X402] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X400)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X400),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X404))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X400)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X400),happ(s(fun(X401,'type.semanticPrimitives.error_result'(X400)),X402),s(X401,X403))))) | s('type.semanticPrimitives.error_result'(X400),'const.evalProps.map_error_result_2'(s(fun(X405,X400),X406),s('type.semanticPrimitives.error_result'(X405),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X404))))) = s('type.semanticPrimitives.error_result'(X400),sK5('const.evalProps.map_error_result_2'(s(fun(X401,'type.semanticPrimitives.error_result'(X400)),X402),s('type.semanticPrimitives.error_result'(X401),'const.semanticPrimitives.Rraise_1'(s(X401,X403)))),'type.semanticPrimitives.error_result'(X400)))) )),
  inference(superposition,[],[f504,f2283])).
fof(f13653,plain,(
  ( ! [X208,X206,X209,X207] : (s(bool,X209) != s(bool,happ(s(fun(X206,bool),X207),s(X206,X208))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X209))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X206,bool),X207),s('type.semanticPrimitives.error_result'(X206),'const.semanticPrimitives.Rraise_1'(s(X206,X208)))),bool))) )),
  inference(superposition,[],[f475,f2283])).
fof(f13652,plain,(
  ( ! [X204,X202,X205,X203] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X202,bool),X203),s(X202,X204))))) | s(bool,X205) = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X202,bool),X203),s('type.semanticPrimitives.error_result'(X202),'const.semanticPrimitives.Rraise_1'(s(X202,X204)))),bool)) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f441,f2283])).
fof(f13651,plain,(
  ( ! [X198,X200,X199,X201] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X201))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X198,bool),X199),s(X198,X200))))) | s(bool,'F_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X198,bool),X199),s('type.semanticPrimitives.error_result'(X198),'const.semanticPrimitives.Rraise_1'(s(X198,X200)))),bool)) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f436,f2283])).
fof(f13650,plain,(
  ( ! [X196,X194,X197,X195] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X197))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X194,bool),X195),s(X194,X196))))) | s(bool,'F_0') = s(bool,X197) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X194,bool),X195),s('type.semanticPrimitives.error_result'(X194),'const.semanticPrimitives.Rraise_1'(s(X194,X196)))),bool))) )),
  inference(superposition,[],[f436,f2283])).
fof(f13649,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(bool,X192) != s(bool,happ(s(fun(X189,bool),X190),s(X189,X191))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X189,bool),X190),s('type.semanticPrimitives.error_result'(X189),'const.semanticPrimitives.Rraise_1'(s(X189,X191)))),bool)) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f423,f2283])).
fof(f13648,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s(bool,X187) != s(bool,happ(s(fun(X184,bool),X185),s(X184,X186))) | s(bool,'T_0') = s(bool,X187) | s(bool,X188) = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X184,bool),X185),s('type.semanticPrimitives.error_result'(X184),'const.semanticPrimitives.Rraise_1'(s(X184,X186)))),bool)) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f423,f2283])).
fof(f13631,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s('type.semanticPrimitives.error_result'(X118),'const.semanticPrimitives.Rraise_1'(s(X118,X122))) != s('type.semanticPrimitives.error_result'(X118),'const.semanticPrimitives.Rraise_1'(s(X118,happ(s(fun(X119,X118),X120),s(X119,X121))))) | s(X118,sK5('const.semanticPrimitives.Rraise_1'(s(X118,X122)),X118)) = s(X118,sK5('const.evalProps.map_error_result_2'(s(fun(X119,X118),X120),s('type.semanticPrimitives.error_result'(X119),'const.semanticPrimitives.Rraise_1'(s(X119,X121)))),X118))) )),
  inference(superposition,[],[f619,f2283])).
fof(f13624,plain,(
  ( ! [X78,X76,X74,X72,X77,X75,X73] : (s('type.semanticPrimitives.error_result'(X72),'const.semanticPrimitives.Rraise_1'(s(X72,happ(s(fun(X73,X72),X74),s(X73,X75))))) != s('type.semanticPrimitives.error_result'(X72),'const.evalProps.map_error_result_2'(s(fun(X76,X72),X77),s('type.semanticPrimitives.error_result'(X76),'const.semanticPrimitives.Rraise_1'(s(X76,X78))))) | s(X72,happ(s(fun(X76,X72),X77),s(X76,X78))) = s(X72,sK5('const.evalProps.map_error_result_2'(s(fun(X73,X72),X74),s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rraise_1'(s(X73,X75)))),X72))) )),
  inference(superposition,[],[f574,f2283])).
fof(f13620,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.semanticPrimitives.error_result'(X46),X50) != s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,happ(s(fun(X47,X46),X48),s(X47,X49))))) | s(X46,sK5(X50,X46)) = s(X46,sK5('const.evalProps.map_error_result_2'(s(fun(X47,X46),X48),s('type.semanticPrimitives.error_result'(X47),'const.semanticPrimitives.Rraise_1'(s(X47,X49)))),X46)) | s('type.semanticPrimitives.error_result'(X46),X50) = s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X50,X46))))) )),
  inference(superposition,[],[f537,f2283])).
fof(f13602,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X38) = s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool)) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X35,bool),X36),s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,X37)))),bool))) )),
  inference(superposition,[],[f2283,f689])).
fof(f13600,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(X27,bool),X28),s(X27,X29))) | s(bool,X30) = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X27,bool),X28),s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rraise_1'(s(X27,X29)))),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X30)),bool))) )),
  inference(superposition,[],[f2283,f689])).
fof(f13599,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X26)),bool)) = s(bool,happ(s(fun(X23,bool),X24),s(X23,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X23,bool),X24),s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rraise_1'(s(X23,X25)))),bool))) )),
  inference(superposition,[],[f2283,f687])).
fof(f13596,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X11,bool),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13)))),bool))) )),
  inference(superposition,[],[f2283,f125])).
fof(f13595,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,X9))) | s(bool,X10) = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X7,bool),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9)))),bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f2283,f125])).
fof(f13594,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,sK5('const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6)))),bool))) )),
  inference(superposition,[],[f2283,f83])).
fof(f13593,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(X34,happ(s(fun('type.semanticPrimitives.error_result'(X30),X34),X35),s('type.semanticPrimitives.error_result'(X30),'const.evalProps.map_error_result_2'(s(fun(X31,X30),X32),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,X33))))))) = s(X34,sK5('const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X30),X34),X35),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X30)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rraise_1'(s(X30,happ(s(fun(X31,X30),X32),s(X31,X33)))))))),X34))) )),
  inference(superposition,[],[f2283,f558])).
fof(f13592,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(X28,happ(s(fun('type.semanticPrimitives.error_result'(X24),X28),X29),s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rraise_1'(s(X24,happ(s(fun(X25,X24),X26),s(X25,X27))))))) = s(X28,sK5('const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X24),X28),X29),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X24)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X24),'const.evalProps.map_error_result_2'(s(fun(X25,X24),X26),s('type.semanticPrimitives.error_result'(X25),'const.semanticPrimitives.Rraise_1'(s(X25,X27)))))))),X28))) )),
  inference(superposition,[],[f2283,f558])).
fof(f13591,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X22,happ(s(fun(X16,X22),X23),s(X16,happ(s(fun(X17,X16),X18),s(X17,happ(s(fun(X19,X17),X20),s(X19,X21))))))) = s(X22,sK5('const.evalProps.map_error_result_2'(s(fun(X16,X22),X23),s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.evalProps.map_error_result_2'(s(fun(X19,X17),X20),s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rraise_1'(s(X19,X21)))))))),X22))) )),
  inference(superposition,[],[f2283,f573])).
fof(f13590,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s(X14,happ(s(fun(X10,X14),X15),s(X10,happ(s(fun(X11,X10),X12),s(X11,sK5(X13,X11)))))) = s(X14,sK5('const.evalProps.map_error_result_2'(s(fun(X10,X14),X15),s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),X13)))),X14)) | s('type.semanticPrimitives.error_result'(X11),X13) = s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X13,X11))))) )),
  inference(superposition,[],[f2283,f572])).
fof(f13589,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X8,happ(s(fun(X4,X8),X9),s(X4,happ(s(fun(X5,X4),X6),s(X5,X7))))) = s(X8,sK5('const.evalProps.map_error_result_2'(s(fun(X4,X8),X9),s('type.semanticPrimitives.error_result'(X4),'const.evalProps.map_error_result_2'(s(fun(X5,X4),X6),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X7)))))),X8))) )),
  inference(superposition,[],[f2283,f108])).
fof(f13588,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,sK5(X1,X0)))) = s(X2,sK5('const.evalProps.map_error_result_2'(s(fun(X0,X2),X3),s('type.semanticPrimitives.error_result'(X0),X1)),X2)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f2283,f95])).
fof(f13250,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X53))) != s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X54))) | s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X54)),X50)) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X51,X50),X52),s('type.semanticPrimitives.error_result'(X51),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X53)))),X50))) )),
  inference(superposition,[],[f800,f1042])).
fof(f13241,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) | s('type.semanticPrimitives.abort',X4) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))),X0))) )),
  inference(superposition,[],[f98,f1042])).
fof(f13238,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X0,bool)))) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(bool,X1),X2),s('type.semanticPrimitives.error_result'(bool),X0)),X1)))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f1042,f536])).
fof(f13226,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.abort',X0) = s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X1,X2),X3),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)))),X2))) )),
  inference(unit_resulting_resolution,[],[f1042,f98])).
fof(f13004,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f13003])).
fof(f13003,plain,(
  ( ! [X37,X35,X36] : (s(bool,X36) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(duplicate_literal_removal,[],[f12979])).
fof(f12979,plain,(
  ( ! [X37,X35,X36] : (s(bool,X36) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f423,f1327])).
fof(f13008,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f13007])).
fof(f13007,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X20)),bool)) | s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) )),
  inference(duplicate_literal_removal,[],[f12975])).
fof(f12975,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X20)),bool)) | s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f423,f687])).
fof(f12969,plain,(
  ( ! [X99,X97,X98] : (s(bool,X97) != s(bool,X98) | s(bool,'T_0') = s(bool,X98) | s(bool,X99) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X97)),bool)) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f423,f596])).
fof(f12424,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) | s('type.semanticPrimitives.abort',sK6(X0,bool)) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f800,f536])).
fof(f12423,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) != s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8))) | s('type.semanticPrimitives.abort',sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X6)),'type.semanticPrimitives.abort')) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X8)),X7))) )),
  inference(superposition,[],[f800,f596])).
fof(f12421,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0))) != s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) | s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)),X1)) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)),X2))) )),
  inference(superposition,[],[f800,f796])).
fof(f12420,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X0,bool))),bool)) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f800,f536])).
fof(f12213,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X19))) != s('type.semanticPrimitives.error_result'(X13),'const.evalProps.map_error_result_2'(s(fun(X14,X13),X15),s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X16,X14),X17),s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,X18))))))) | s(X13,sK5('const.semanticPrimitives.Rraise_1'(s(X13,X19)),X13)) = s(X13,happ(s(fun(X14,X13),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18)))))) )),
  inference(superposition,[],[f619,f573])).
fof(f12212,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,X12))) != s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X9,X8),X10),s('type.semanticPrimitives.error_result'(X9),X11))) | s(X8,sK5('const.semanticPrimitives.Rraise_1'(s(X8,X12)),X8)) = s(X8,happ(s(fun(X9,X8),X10),s(X9,sK5(X11,X9)))) | s('type.semanticPrimitives.error_result'(X9),X11) = s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X11,X9))))) )),
  inference(superposition,[],[f619,f572])).
fof(f12208,plain,(
  ( ! [X123,X124,X122] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X122))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X124))) | s('type.semanticPrimitives.abort',sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X124)),'type.semanticPrimitives.abort')) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X122)),X123))) )),
  inference(superposition,[],[f619,f796])).
fof(f12206,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),X120))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X116),'const.semanticPrimitives.Rraise_1'(s(X116,happ(s(fun(X117,X116),X118),s(X117,X119))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),X120)),'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X116),'const.evalProps.map_error_result_2'(s(fun(X117,X116),X118),s('type.semanticPrimitives.error_result'(X117),'const.semanticPrimitives.Rraise_1'(s(X117,X119)))))))) )),
  inference(superposition,[],[f619,f558])).
fof(f12205,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.evalProps.map_error_result_2'(s(fun(X112,X111),X113),s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rraise_1'(s(X112,X114))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115)),'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.semanticPrimitives.Rraise_1'(s(X111,happ(s(fun(X112,X111),X113),s(X112,X114)))))))) )),
  inference(superposition,[],[f619,f558])).
fof(f12204,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X104)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),X110))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X104)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),'const.evalProps.map_error_result_2'(s(fun(X105,X104),X106),s('type.semanticPrimitives.error_result'(X105),'const.evalProps.map_error_result_2'(s(fun(X107,X105),X108),s('type.semanticPrimitives.error_result'(X107),'const.semanticPrimitives.Rraise_1'(s(X107,X109))))))))) | s('type.semanticPrimitives.error_result'(X104),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),X110)),'type.semanticPrimitives.error_result'(X104))) = s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rraise_1'(s(X104,happ(s(fun(X105,X104),X106),s(X105,happ(s(fun(X107,X105),X108),s(X107,X109)))))))) )),
  inference(superposition,[],[f619,f573])).
fof(f12203,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X99)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),X103))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X99)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),'const.evalProps.map_error_result_2'(s(fun(X100,X99),X101),s('type.semanticPrimitives.error_result'(X100),X102))))) | s('type.semanticPrimitives.error_result'(X99),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),X103)),'type.semanticPrimitives.error_result'(X99))) = s('type.semanticPrimitives.error_result'(X99),'const.semanticPrimitives.Rraise_1'(s(X99,happ(s(fun(X100,X99),X101),s(X100,sK5(X102,X100)))))) | s('type.semanticPrimitives.error_result'(X100),X102) = s('type.semanticPrimitives.error_result'(X100),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X102,X100))))) )),
  inference(superposition,[],[f619,f572])).
fof(f12200,plain,(
  ( ! [X90,X89] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X89))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X90))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X89,bool)))) = s('type.semanticPrimitives.error_result'(bool),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X90)),'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'(bool),X89) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X89,bool))) )),
  inference(superposition,[],[f619,f536])).
fof(f12198,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X81)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),X87))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X81)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),'const.semanticPrimitives.Rraise_1'(s(X81,happ(s(fun(X82,X81),X83),s(X82,happ(s(fun(X84,X82),X85),s(X84,X86))))))))) | s('type.semanticPrimitives.error_result'(X81),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),X87)),'type.semanticPrimitives.error_result'(X81))) = s('type.semanticPrimitives.error_result'(X81),'const.evalProps.map_error_result_2'(s(fun(X82,X81),X83),s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(X84,X82),X85),s('type.semanticPrimitives.error_result'(X84),'const.semanticPrimitives.Rraise_1'(s(X84,X86)))))))) )),
  inference(superposition,[],[f619,f573])).
fof(f12192,plain,(
  ( ! [X61,X62,X60] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X62)),bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f619,f687])).
fof(f12191,plain,(
  ( ! [X59,X57,X58] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X58))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X57)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X58)),bool)) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f619,f689])).
fof(f12190,plain,(
  ( ! [X56,X55] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X56))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X55)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool)) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f619,f1335])).
fof(f12189,plain,(
  ( ! [X54,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X54))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X53)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X54)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X53)))) )),
  inference(superposition,[],[f619,f617])).
fof(f12180,plain,(
  ( ! [X37,X35,X36] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X36))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool)) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X37)),bool))) )),
  inference(superposition,[],[f619,f689])).
fof(f12176,plain,(
  ( ! [X28,X26,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,X26) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X28)),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f619,f689])).
fof(f12175,plain,(
  ( ! [X24,X23,X25] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) | s(bool,X23) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X25)),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X23)),bool)) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f619,f689])).
fof(f12174,plain,(
  ( ! [X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X21))) | s(bool,X20) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool)) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(superposition,[],[f619,f689])).
fof(f12170,plain,(
  ( ! [X10,X8,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X10))) | s(bool,X8) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X10)),bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f619,f125])).
fof(f12169,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X6))) | s(bool,X5) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X6)),bool)) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f619,f125])).
fof(f12167,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0)) = s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X2)),X0))) )),
  inference(superposition,[],[f619,f596])).
fof(f12157,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),X120))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X116),'const.semanticPrimitives.Rraise_1'(s(X116,happ(s(fun(X117,X116),X118),s(X117,X119))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),X120) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X116),'const.evalProps.map_error_result_2'(s(fun(X117,X116),X118),s('type.semanticPrimitives.error_result'(X117),'const.semanticPrimitives.Rraise_1'(s(X117,X119)))))))),'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X116))))) )),
  inference(superposition,[],[f619,f558])).
fof(f12156,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.evalProps.map_error_result_2'(s(fun(X112,X111),X113),s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rraise_1'(s(X112,X114))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.semanticPrimitives.Rraise_1'(s(X111,happ(s(fun(X112,X111),X113),s(X112,X114)))))))),'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))))) )),
  inference(superposition,[],[f619,f558])).
fof(f12155,plain,(
  ( ! [X109,X107,X105,X110,X108,X106,X104] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X104)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),X110))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X104)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),'const.evalProps.map_error_result_2'(s(fun(X105,X104),X106),s('type.semanticPrimitives.error_result'(X105),'const.evalProps.map_error_result_2'(s(fun(X107,X105),X108),s('type.semanticPrimitives.error_result'(X107),'const.semanticPrimitives.Rraise_1'(s(X107,X109))))))))) | s('type.semanticPrimitives.error_result'(X104),X110) = s('type.semanticPrimitives.error_result'(X104),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X104),'const.semanticPrimitives.Rraise_1'(s(X104,happ(s(fun(X105,X104),X106),s(X105,happ(s(fun(X107,X105),X108),s(X107,X109)))))))),'type.semanticPrimitives.error_result'(X104)))) )),
  inference(superposition,[],[f619,f573])).
fof(f12154,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X99)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),X103))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X99)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),'const.evalProps.map_error_result_2'(s(fun(X100,X99),X101),s('type.semanticPrimitives.error_result'(X100),X102))))) | s('type.semanticPrimitives.error_result'(X99),X103) = s('type.semanticPrimitives.error_result'(X99),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X99),'const.semanticPrimitives.Rraise_1'(s(X99,happ(s(fun(X100,X99),X101),s(X100,sK5(X102,X100))))))),'type.semanticPrimitives.error_result'(X99))) | s('type.semanticPrimitives.error_result'(X100),X102) = s('type.semanticPrimitives.error_result'(X100),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X102,X100))))) )),
  inference(superposition,[],[f619,f572])).
fof(f12152,plain,(
  ( ! [X92,X93,X91] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X91)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X91),X92))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X91)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X91),X93))) | s('type.semanticPrimitives.error_result'(X91),X93) = s('type.semanticPrimitives.error_result'(X91),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X91),'const.semanticPrimitives.Rraise_1'(s(X91,sK5(X92,X91))))),'type.semanticPrimitives.error_result'(X91))) | s('type.semanticPrimitives.error_result'(X91),X92) = s('type.semanticPrimitives.error_result'(X91),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X92,X91))))) )),
  inference(superposition,[],[f619,f95])).
fof(f12151,plain,(
  ( ! [X90,X89] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X89))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X90))) | s('type.semanticPrimitives.error_result'(bool),X90) = s('type.semanticPrimitives.error_result'(bool),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X89,bool))))),'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'(bool),X89) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X89,bool))) )),
  inference(superposition,[],[f619,f536])).
fof(f12149,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X81)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),X87))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X81)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),'const.semanticPrimitives.Rraise_1'(s(X81,happ(s(fun(X82,X81),X83),s(X82,happ(s(fun(X84,X82),X85),s(X84,X86))))))))) | s('type.semanticPrimitives.error_result'(X81),X87) = s('type.semanticPrimitives.error_result'(X81),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X81),'const.evalProps.map_error_result_2'(s(fun(X82,X81),X83),s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(X84,X82),X85),s('type.semanticPrimitives.error_result'(X84),'const.semanticPrimitives.Rraise_1'(s(X84,X86)))))))),'type.semanticPrimitives.error_result'(X81)))) )),
  inference(superposition,[],[f619,f573])).
fof(f12132,plain,(
  ( ! [X39,X38,X40] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X38))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X39))) | s(bool,X39) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f619,f125])).
fof(f12131,plain,(
  ( ! [X37,X35,X36] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X35))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X36))) | s(bool,X36) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X37)),bool))) )),
  inference(superposition,[],[f619,f689])).
fof(f12127,plain,(
  ( ! [X28,X26,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,X28) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X26)),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f619,f689])).
fof(f12125,plain,(
  ( ! [X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X21))) | s(bool,X21) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X20)),bool)) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(superposition,[],[f619,f689])).
fof(f11701,plain,(
  ( ! [X163,X162] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X162) = s(fun(bool,bool),X163) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X162,X163,bool,bool)))) | s(bool,'T_0') = s(bool,sK7(X162,X163,bool,bool))) )),
  inference(forward_demodulation,[],[f7901,f596])).
fof(f7901,plain,(
  ( ! [X163,X162] : (s(bool,'F_0') != s(bool,f(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) | s(fun(bool,bool),X162) = s(fun(bool,bool),X163) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X162,X163,bool,bool)))) | s(bool,'T_0') = s(bool,sK7(X162,X163,bool,bool))) )),
  inference(superposition,[],[f663,f1327])).
fof(f10091,plain,(
  ( ! [X146,X149,X147,X148] : (s('type.semanticPrimitives.error_result'(X147),X148) = s('type.semanticPrimitives.error_result'(X147),'const.evalProps.map_error_result_2'(s(fun(bool,X147),X149),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),X148))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),'const.semanticPrimitives.Rraise_1'(s(X147,happ(s(fun(bool,X147),X149),s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146)))) )),
  inference(forward_demodulation,[],[f10012,f596])).
fof(f10012,plain,(
  ( ! [X146,X149,X147,X148] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),X148))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),'const.semanticPrimitives.Rraise_1'(s(X147,happ(s(fun(bool,X147),X149),s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X147),X148) = s('type.semanticPrimitives.error_result'(X147),'const.evalProps.map_error_result_2'(s(fun(bool,X147),X149),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X146)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146)))) )),
  inference(superposition,[],[f586,f617])).
fof(f10090,plain,(
  ( ! [X144,X142,X145,X143] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X143)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X143),X144))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X143)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X143),'const.evalProps.map_error_result_2'(s(fun(bool,X143),X145),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142))))))) | s('type.semanticPrimitives.error_result'(X143),X144) != s('type.semanticPrimitives.error_result'(X143),'const.semanticPrimitives.Rraise_1'(s(X143,happ(s(fun(bool,X143),X145),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142)))) )),
  inference(forward_demodulation,[],[f10011,f596])).
fof(f10011,plain,(
  ( ! [X144,X142,X145,X143] : (s('type.semanticPrimitives.error_result'(X143),X144) != s('type.semanticPrimitives.error_result'(X143),'const.semanticPrimitives.Rraise_1'(s(X143,happ(s(fun(bool,X143),X145),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X143)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X143),X144))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X143)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X143),'const.evalProps.map_error_result_2'(s(fun(bool,X143),X145),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X142)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142)))) )),
  inference(superposition,[],[f584,f617])).
fof(f10089,plain,(
  ( ! [X140,X138,X141,X139] : (s('type.semanticPrimitives.error_result'(X139),X140) = s('type.semanticPrimitives.error_result'(X139),'const.semanticPrimitives.Rraise_1'(s(X139,happ(s(fun(bool,X139),X141),s(bool,X138))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X139)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X139),X140))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X139)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X139),'const.evalProps.map_error_result_2'(s(fun(bool,X139),X141),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X138)))) )),
  inference(forward_demodulation,[],[f10010,f596])).
fof(f10010,plain,(
  ( ! [X140,X138,X141,X139] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X139)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X139),X140))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X139)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X139),'const.evalProps.map_error_result_2'(s(fun(bool,X139),X141),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X139),X140) = s('type.semanticPrimitives.error_result'(X139),'const.semanticPrimitives.Rraise_1'(s(X139,happ(s(fun(bool,X139),X141),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X138)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X138)))) )),
  inference(superposition,[],[f580,f617])).
fof(f10088,plain,(
  ( ! [X134,X136,X135,X137] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X135)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X135),X136))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X135)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X135),'const.semanticPrimitives.Rraise_1'(s(X135,happ(s(fun(bool,X135),X137),s(bool,X134))))))) | s('type.semanticPrimitives.error_result'(X135),X136) != s('type.semanticPrimitives.error_result'(X135),'const.evalProps.map_error_result_2'(s(fun(bool,X135),X137),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X134)))) )),
  inference(forward_demodulation,[],[f10009,f596])).
fof(f10009,plain,(
  ( ! [X134,X136,X135,X137] : (s('type.semanticPrimitives.error_result'(X135),X136) != s('type.semanticPrimitives.error_result'(X135),'const.evalProps.map_error_result_2'(s(fun(bool,X135),X137),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X135)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X135),X136))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X135)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X135),'const.semanticPrimitives.Rraise_1'(s(X135,happ(s(fun(bool,X135),X137),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X134)),bool)))))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X134)))) )),
  inference(superposition,[],[f578,f617])).
fof(f10007,plain,(
  ( ! [X127,X128,X126,X129] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X127,bool),X128),s('type.semanticPrimitives.error_result'(X127),'const.semanticPrimitives.Rraise_1'(s(X127,X129))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X126)),bool)) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X126)))) )),
  inference(superposition,[],[f574,f617])).
fof(f10087,plain,(
  ( ! [X125,X123,X124,X122] : (s(X123,X124) = s(X123,happ(s(fun(bool,X123),X125),s(bool,X122))) | s('type.semanticPrimitives.error_result'(X123),'const.semanticPrimitives.Rraise_1'(s(X123,X124))) != s('type.semanticPrimitives.error_result'(X123),'const.evalProps.map_error_result_2'(s(fun(bool,X123),X125),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X122)))) )),
  inference(forward_demodulation,[],[f10006,f596])).
fof(f10006,plain,(
  ( ! [X125,X123,X124,X122] : (s('type.semanticPrimitives.error_result'(X123),'const.semanticPrimitives.Rraise_1'(s(X123,X124))) != s('type.semanticPrimitives.error_result'(X123),'const.evalProps.map_error_result_2'(s(fun(bool,X123),X125),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(X123,X124) = s(X123,happ(s(fun(bool,X123),X125),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X122)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X122)))) )),
  inference(superposition,[],[f574,f617])).
fof(f10003,plain,(
  ( ! [X112,X113] : (s('type.semanticPrimitives.error_result'(bool),X113) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,sK5(X113,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X112)),bool)) | s('type.semanticPrimitives.error_result'(bool),X113) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X113,bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X112)))) )),
  inference(superposition,[],[f537,f617])).
fof(f10085,plain,(
  ( ! [X101,X102] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X102))) | s(bool,'T_0') != s(bool,X102) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101)))) )),
  inference(forward_demodulation,[],[f9998,f596])).
fof(f9998,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,X102) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X102))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X101)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101)))) )),
  inference(superposition,[],[f101,f617])).
fof(f9996,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X98))) | s(bool,X98) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X97)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X97)))) )),
  inference(superposition,[],[f100,f617])).
fof(f10084,plain,(
  ( ! [X96] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X96))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X96)))) )),
  inference(subsumption_resolution,[],[f9995,f102])).
fof(f9995,plain,(
  ( ! [X96] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X96))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X96))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X96)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X96)))) )),
  inference(superposition,[],[f95,f617])).
fof(f10083,plain,(
  ( ! [X94,X95,X93] : (s('type.semanticPrimitives.error_result'(X94),'const.semanticPrimitives.Rraise_1'(s(X94,happ(s(fun(bool,X94),X95),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X94),'const.evalProps.map_error_result_2'(s(fun(bool,X94),X95),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93)))) )),
  inference(subsumption_resolution,[],[f9994,f102])).
fof(f9994,plain,(
  ( ! [X94,X95,X93] : (s('type.semanticPrimitives.error_result'(X94),'const.semanticPrimitives.Rraise_1'(s(X94,happ(s(fun(bool,X94),X95),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X94),'const.evalProps.map_error_result_2'(s(fun(bool,X94),X95),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X93)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X93)))) )),
  inference(superposition,[],[f572,f617])).
fof(f10081,plain,(
  ( ! [X56,X55] : (s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X55)),bool)) | s(bool,'F_0') = s(bool,X56) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55)))) )),
  inference(subsumption_resolution,[],[f9959,f76])).
fof(f9959,plain,(
  ( ! [X56,X55] : (~p(s(bool,'T_0')) | s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X55)),bool)) | s(bool,'F_0') = s(bool,X56) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55)))) )),
  inference(superposition,[],[f789,f617])).
fof(f10070,plain,(
  ( ! [X30,X31,X29] : (s(bool,X30) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X29)),bool)) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29)))) )),
  inference(subsumption_resolution,[],[f9947,f76])).
fof(f9947,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X29)),bool)) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29)))) )),
  inference(superposition,[],[f422,f617])).
fof(f10069,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X25)),bool)) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25)))) )),
  inference(subsumption_resolution,[],[f9945,f76])).
fof(f9945,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X25)),bool)) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25)))) )),
  inference(superposition,[],[f127,f617])).
fof(f10068,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | ~p(s(bool,X18)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X17)))) )),
  inference(forward_demodulation,[],[f10067,f596])).
fof(f10067,plain,(
  ( ! [X17,X18] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X17)))) )),
  inference(subsumption_resolution,[],[f9942,f76])).
fof(f9942,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X17)))) )),
  inference(superposition,[],[f106,f617])).
fof(f9546,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X16))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.evalProps.map_error_result_2'(s(fun(X13,X11),X14),s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X15))))))))) | s('type.semanticPrimitives.error_result'(X10),X16) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9545,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),X9))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),X8))))) | s('type.semanticPrimitives.error_result'(X5),X9) = s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,sK5(X8,X6)))))) | s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X8,X6))))) )),
  inference(superposition,[],[f586,f572])).
fof(f9536,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X32))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),X32) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X29,bool),X30),s('type.semanticPrimitives.error_result'(X29),'const.semanticPrimitives.Rraise_1'(s(X29,X31))))) | s(bool,X33) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X33)),bool))) )),
  inference(superposition,[],[f586,f689])).
fof(f9532,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X13))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X12))))) | s('type.semanticPrimitives.error_result'(bool),X13) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X9,bool),X10),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) )),
  inference(superposition,[],[f586,f125])).
fof(f9531,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X7))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),X7) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6))))) | s(bool,X8) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f586,f125])).
fof(f9530,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X3))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'(bool),X3) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X0,bool),X1),s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(superposition,[],[f586,f83])).
fof(f9526,plain,(
  ( ! [X222,X220,X225,X219,X223,X221,X224] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),X224))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),'const.semanticPrimitives.Rraise_1'(s(X223,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),X223),X225),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),'const.semanticPrimitives.Rraise_1'(s(X219,happ(s(fun(X220,X219),X221),s(X220,X222))))))))))))) | s('type.semanticPrimitives.error_result'(X223),X224) = s('type.semanticPrimitives.error_result'(X223),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),X223),X225),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),'const.evalProps.map_error_result_2'(s(fun(X220,X219),X221),s('type.semanticPrimitives.error_result'(X220),'const.semanticPrimitives.Rraise_1'(s(X220,X222)))))))))))) )),
  inference(superposition,[],[f586,f558])).
fof(f9525,plain,(
  ( ! [X212,X218,X216,X215,X213,X217,X214] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X216)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X216),X217))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X216)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X216),'const.semanticPrimitives.Rraise_1'(s(X216,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),X216),X218),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),'const.evalProps.map_error_result_2'(s(fun(X213,X212),X214),s('type.semanticPrimitives.error_result'(X213),'const.semanticPrimitives.Rraise_1'(s(X213,X215))))))))))))) | s('type.semanticPrimitives.error_result'(X216),X217) = s('type.semanticPrimitives.error_result'(X216),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),X216),X218),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),'const.semanticPrimitives.Rraise_1'(s(X212,happ(s(fun(X213,X212),X214),s(X213,X215)))))))))))) )),
  inference(superposition,[],[f586,f558])).
fof(f9524,plain,(
  ( ! [X210,X208,X206,X204,X211,X209,X207,X205,X203] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X209)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X209),X210))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X209)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X209),'const.semanticPrimitives.Rraise_1'(s(X209,happ(s(fun('type.semanticPrimitives.error_result'(X203),X209),X211),s('type.semanticPrimitives.error_result'(X203),'const.evalProps.map_error_result_2'(s(fun(X204,X203),X205),s('type.semanticPrimitives.error_result'(X204),'const.evalProps.map_error_result_2'(s(fun(X206,X204),X207),s('type.semanticPrimitives.error_result'(X206),'const.semanticPrimitives.Rraise_1'(s(X206,X208))))))))))))) | s('type.semanticPrimitives.error_result'(X209),X210) = s('type.semanticPrimitives.error_result'(X209),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X203),X209),X211),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X203)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X203),'const.semanticPrimitives.Rraise_1'(s(X203,happ(s(fun(X204,X203),X205),s(X204,happ(s(fun(X206,X204),X207),s(X206,X208)))))))))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9523,plain,(
  ( ! [X198,X196,X202,X200,X199,X197,X201] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X200)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X200),X201))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X200)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X200),'const.semanticPrimitives.Rraise_1'(s(X200,happ(s(fun('type.semanticPrimitives.error_result'(X196),X200),X202),s('type.semanticPrimitives.error_result'(X196),'const.evalProps.map_error_result_2'(s(fun(X197,X196),X198),s('type.semanticPrimitives.error_result'(X197),X199))))))))) | s('type.semanticPrimitives.error_result'(X200),X201) = s('type.semanticPrimitives.error_result'(X200),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X196),X200),X202),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X196)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X196),'const.semanticPrimitives.Rraise_1'(s(X196,happ(s(fun(X197,X196),X198),s(X197,sK5(X199,X197)))))))))) | s('type.semanticPrimitives.error_result'(X197),X199) = s('type.semanticPrimitives.error_result'(X197),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X199,X197))))) )),
  inference(superposition,[],[f586,f572])).
fof(f9522,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X193)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X193),X194))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X193)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X193),'const.semanticPrimitives.Rraise_1'(s(X193,happ(s(fun('type.semanticPrimitives.error_result'(X189),X193),X195),s('type.semanticPrimitives.error_result'(X189),'const.evalProps.map_error_result_2'(s(fun(X190,X189),X191),s('type.semanticPrimitives.error_result'(X190),'const.semanticPrimitives.Rraise_1'(s(X190,X192))))))))))) | s('type.semanticPrimitives.error_result'(X193),X194) = s('type.semanticPrimitives.error_result'(X193),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X189),X193),X195),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X189)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X189),'const.semanticPrimitives.Rraise_1'(s(X189,happ(s(fun(X190,X189),X191),s(X190,X192)))))))))) )),
  inference(superposition,[],[f586,f108])).
fof(f9521,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X186)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X186),X187))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X186)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X186),'const.semanticPrimitives.Rraise_1'(s(X186,happ(s(fun('type.semanticPrimitives.error_result'(X184),X186),X188),s('type.semanticPrimitives.error_result'(X184),X185))))))) | s('type.semanticPrimitives.error_result'(X186),X187) = s('type.semanticPrimitives.error_result'(X186),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X184),X186),X188),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X184)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X184),'const.semanticPrimitives.Rraise_1'(s(X184,sK5(X185,X184)))))))) | s('type.semanticPrimitives.error_result'(X184),X185) = s('type.semanticPrimitives.error_result'(X184),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X185,X184))))) )),
  inference(superposition,[],[f586,f95])).
fof(f9520,plain,(
  ( ! [X182,X180,X183,X181] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X181)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X181),X182))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X181)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X181),'const.semanticPrimitives.Rraise_1'(s(X181,happ(s(fun('type.semanticPrimitives.error_result'(bool),X181),X183),s('type.semanticPrimitives.error_result'(bool),X180))))))) | s('type.semanticPrimitives.error_result'(X181),X182) = s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X181),X183),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X180,bool)))))))) | s('type.semanticPrimitives.error_result'(bool),X180) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X180,bool))) )),
  inference(superposition,[],[f586,f536])).
fof(f9518,plain,(
  ( ! [X175,X173,X171,X169,X176,X174,X172,X170,X168] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X174)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X174),X175))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X174)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X174),'const.semanticPrimitives.Rraise_1'(s(X174,happ(s(fun('type.semanticPrimitives.error_result'(X168),X174),X176),s('type.semanticPrimitives.error_result'(X168),'const.semanticPrimitives.Rraise_1'(s(X168,happ(s(fun(X169,X168),X170),s(X169,happ(s(fun(X171,X169),X172),s(X171,X173))))))))))))) | s('type.semanticPrimitives.error_result'(X174),X175) = s('type.semanticPrimitives.error_result'(X174),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X168),X174),X176),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X168)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X168),'const.evalProps.map_error_result_2'(s(fun(X169,X168),X170),s('type.semanticPrimitives.error_result'(X169),'const.evalProps.map_error_result_2'(s(fun(X171,X169),X172),s('type.semanticPrimitives.error_result'(X171),'const.semanticPrimitives.Rraise_1'(s(X171,X173)))))))))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9517,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X165),X166))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X165),'const.semanticPrimitives.Rraise_1'(s(X165,happ(s(fun('type.semanticPrimitives.error_result'(X161),X165),X167),s('type.semanticPrimitives.error_result'(X161),'const.semanticPrimitives.Rraise_1'(s(X161,happ(s(fun(X162,X161),X163),s(X162,X164))))))))))) | s('type.semanticPrimitives.error_result'(X165),X166) = s('type.semanticPrimitives.error_result'(X165),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X161),X165),X167),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X161)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X161),'const.evalProps.map_error_result_2'(s(fun(X162,X161),X163),s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,X164)))))))))) )),
  inference(superposition,[],[f586,f108])).
fof(f9498,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),X86))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rraise_1'(s(X85,happ(s(fun(bool,X85),X87),s(bool,X84))))))) | s('type.semanticPrimitives.error_result'(X85),X86) = s('type.semanticPrimitives.error_result'(X85),'const.evalProps.map_error_result_2'(s(fun(bool,X85),X87),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f586,f83])).
fof(f9496,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(bool,X75),X77),s(bool,X74))))))) | s('type.semanticPrimitives.error_result'(X75),X76) = s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(bool,X75),X77),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f586,f125])).
fof(f9495,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),X71))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(bool,X70),X72),s(bool,X69))))))) | s('type.semanticPrimitives.error_result'(X70),X71) = s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(bool,X70),X72),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,X69) = s(bool,X73) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X73)),bool))) )),
  inference(superposition,[],[f586,f689])).
fof(f9487,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),X36))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(bool,X35),X37),s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X35),X36) = s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(bool,X35),X37),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34))))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f586,f689])).
fof(f9483,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X18),s(bool,X15))))))) | s('type.semanticPrimitives.error_result'(X16),X17) = s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X18),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X14))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f586,f125])).
fof(f9482,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X11))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X12),s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X10),X11) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X12),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f586,f125])).
fof(f9481,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(bool,X6),X8),s(bool,'F_0'))))))) | s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X8),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f586,f83])).
fof(f9478,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X33)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X33),'const.semanticPrimitives.Rraise_1'(s(X33,happ(s(fun(X39,X33),X40),s(X39,X41))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X33)),'const.evalProps.map_error_result_2'(s(fun(X34,'type.semanticPrimitives.error_result'(X33)),X35),s('type.semanticPrimitives.error_result'(X34),'const.evalProps.map_error_result_2'(s(fun(X36,X34),X37),s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,X38))))))) | s('type.semanticPrimitives.error_result'(X33),'const.evalProps.map_error_result_2'(s(fun(X39,X33),X40),s('type.semanticPrimitives.error_result'(X39),'const.semanticPrimitives.Rraise_1'(s(X39,X41))))) = s('type.semanticPrimitives.error_result'(X33),happ(s(fun(X34,'type.semanticPrimitives.error_result'(X33)),X35),s(X34,happ(s(fun(X36,X34),X37),s(X36,X38)))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9477,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.evalProps.map_error_result_2'(s(fun(X27,'type.semanticPrimitives.error_result'(X26)),X28),s('type.semanticPrimitives.error_result'(X27),X29))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.semanticPrimitives.Rraise_1'(s(X26,happ(s(fun(X30,X26),X31),s(X30,X32))))))) | s('type.semanticPrimitives.error_result'(X26),happ(s(fun(X27,'type.semanticPrimitives.error_result'(X26)),X28),s(X27,sK5(X29,X27)))) = s('type.semanticPrimitives.error_result'(X26),'const.evalProps.map_error_result_2'(s(fun(X30,X26),X31),s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rraise_1'(s(X30,X32))))) | s('type.semanticPrimitives.error_result'(X27),X29) = s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X29,X27))))) )),
  inference(superposition,[],[f586,f572])).
fof(f9472,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(X77,X75),X78),s(X77,X79))))))) | s('type.semanticPrimitives.error_result'(X75),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76)),'type.semanticPrimitives.error_result'(X75))) = s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(X77,X75),X78),s('type.semanticPrimitives.error_result'(X77),'const.semanticPrimitives.Rraise_1'(s(X77,X79)))))) )),
  inference(superposition,[],[f586,f596])).
fof(f9471,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),happ(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s(X69,X71))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,happ(s(fun(X66,X65),X67),s(X66,X68))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.evalProps.map_error_result_2'(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X71))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.evalProps.map_error_result_2'(s(fun(X66,X65),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68)))))))) )),
  inference(superposition,[],[f586,f558])).
fof(f9470,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),happ(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s(X62,X64))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(X59,X58),X60),s('type.semanticPrimitives.error_result'(X59),'const.semanticPrimitives.Rraise_1'(s(X59,X61))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.evalProps.map_error_result_2'(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rraise_1'(s(X62,X64))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,happ(s(fun(X59,X58),X60),s(X59,X61)))))))) )),
  inference(superposition,[],[f586,f558])).
fof(f9469,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X55,X49),X56),s(X55,X57))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X50,X49),X51),s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X52,X50),X53),s('type.semanticPrimitives.error_result'(X52),'const.semanticPrimitives.Rraise_1'(s(X52,X54))))))))) | s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X55,X49),X56),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,X57))))) = s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X50,X49),X51),s(X50,happ(s(fun(X52,X50),X53),s(X52,X54)))))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9468,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),X45))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X46,X42),X47),s(X46,X48))))))) | s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X46,X42),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,X48))))) = s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,sK5(X45,X43)))))) | s('type.semanticPrimitives.error_result'(X43),X45) = s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X45,X43))))) )),
  inference(superposition,[],[f586,f572])).
fof(f9464,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X20,X14),X21),s(X20,X22))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))))) | s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X20,X14),X21),s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,X22))))) = s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))))) )),
  inference(superposition,[],[f586,f573])).
fof(f9463,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X8,X7),X9),s(X8,X10))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))))) | s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X8,X7),X9),s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,X10))))) = s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X11,X7),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13)))))) )),
  inference(superposition,[],[f586,f108])).
fof(f9270,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s('type.semanticPrimitives.error_result'(X10),X16) != s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.evalProps.map_error_result_2'(s(fun(X13,X11),X14),s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X15))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X16))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))))))))) )),
  inference(superposition,[],[f584,f573])).
fof(f9269,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'(X5),X9) != s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),X8))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),X9))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,sK5(X8,X6)))))))) | s('type.semanticPrimitives.error_result'(X6),X8) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X8,X6))))) )),
  inference(superposition,[],[f584,f572])).
fof(f9312,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X43))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X39,bool),X40),s(X39,X41))))))) | s('type.semanticPrimitives.error_result'(bool),X43) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X42)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) )),
  inference(forward_demodulation,[],[f9262,f108])).
fof(f9262,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.semanticPrimitives.error_result'(bool),X43) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X42))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X43))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X39,bool),X40),s('type.semanticPrimitives.error_result'(X39),'const.semanticPrimitives.Rraise_1'(s(X39,X41))))))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X42)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) )),
  inference(superposition,[],[f584,f689])).
fof(f9309,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X32))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))))))) | s('type.semanticPrimitives.error_result'(bool),X32) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,X33) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X33)),bool))) )),
  inference(forward_demodulation,[],[f9260,f108])).
fof(f9260,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.semanticPrimitives.error_result'(bool),X32) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X32))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X29,bool),X30),s('type.semanticPrimitives.error_result'(X29),'const.semanticPrimitives.Rraise_1'(s(X29,X31))))))) | s(bool,X33) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X33)),bool))) )),
  inference(superposition,[],[f584,f689])).
fof(f9304,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X13))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))))))) | s('type.semanticPrimitives.error_result'(bool),X13) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X12))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) )),
  inference(forward_demodulation,[],[f9256,f108])).
fof(f9256,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),X13) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X12))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X13))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X9,bool),X10),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) )),
  inference(superposition,[],[f584,f125])).
fof(f9303,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X7))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))))) | s('type.semanticPrimitives.error_result'(bool),X7) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,X8) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(forward_demodulation,[],[f9255,f108])).
fof(f9255,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),X7) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X7))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X4,bool),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6))))))) | s(bool,X8) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f584,f125])).
fof(f9250,plain,(
  ( ! [X222,X220,X225,X219,X223,X221,X224] : (s('type.semanticPrimitives.error_result'(X223),X224) != s('type.semanticPrimitives.error_result'(X223),'const.semanticPrimitives.Rraise_1'(s(X223,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),X223),X225),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),'const.semanticPrimitives.Rraise_1'(s(X219,happ(s(fun(X220,X219),X221),s(X220,X222))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),X224))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X223)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X223),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),X223),X225),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),'const.evalProps.map_error_result_2'(s(fun(X220,X219),X221),s('type.semanticPrimitives.error_result'(X220),'const.semanticPrimitives.Rraise_1'(s(X220,X222)))))))))))))) )),
  inference(superposition,[],[f584,f558])).
fof(f9249,plain,(
  ( ! [X212,X218,X216,X215,X213,X217,X214] : (s('type.semanticPrimitives.error_result'(X216),X217) != s('type.semanticPrimitives.error_result'(X216),'const.semanticPrimitives.Rraise_1'(s(X216,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),X216),X218),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),'const.evalProps.map_error_result_2'(s(fun(X213,X212),X214),s('type.semanticPrimitives.error_result'(X213),'const.semanticPrimitives.Rraise_1'(s(X213,X215))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X216)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X216),X217))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X216)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X216),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),X216),X218),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),'const.semanticPrimitives.Rraise_1'(s(X212,happ(s(fun(X213,X212),X214),s(X213,X215)))))))))))))) )),
  inference(superposition,[],[f584,f558])).
fof(f9248,plain,(
  ( ! [X210,X208,X206,X204,X211,X209,X207,X205,X203] : (s('type.semanticPrimitives.error_result'(X209),X210) != s('type.semanticPrimitives.error_result'(X209),'const.semanticPrimitives.Rraise_1'(s(X209,happ(s(fun('type.semanticPrimitives.error_result'(X203),X209),X211),s('type.semanticPrimitives.error_result'(X203),'const.evalProps.map_error_result_2'(s(fun(X204,X203),X205),s('type.semanticPrimitives.error_result'(X204),'const.evalProps.map_error_result_2'(s(fun(X206,X204),X207),s('type.semanticPrimitives.error_result'(X206),'const.semanticPrimitives.Rraise_1'(s(X206,X208))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X209)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X209),X210))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X209)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X209),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X203),X209),X211),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X203)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X203),'const.semanticPrimitives.Rraise_1'(s(X203,happ(s(fun(X204,X203),X205),s(X204,happ(s(fun(X206,X204),X207),s(X206,X208)))))))))))))) )),
  inference(superposition,[],[f584,f573])).
fof(f9247,plain,(
  ( ! [X198,X196,X202,X200,X199,X197,X201] : (s('type.semanticPrimitives.error_result'(X200),X201) != s('type.semanticPrimitives.error_result'(X200),'const.semanticPrimitives.Rraise_1'(s(X200,happ(s(fun('type.semanticPrimitives.error_result'(X196),X200),X202),s('type.semanticPrimitives.error_result'(X196),'const.evalProps.map_error_result_2'(s(fun(X197,X196),X198),s('type.semanticPrimitives.error_result'(X197),X199))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X200)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X200),X201))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X200)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X200),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X196),X200),X202),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X196)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X196),'const.semanticPrimitives.Rraise_1'(s(X196,happ(s(fun(X197,X196),X198),s(X197,sK5(X199,X197)))))))))))) | s('type.semanticPrimitives.error_result'(X197),X199) = s('type.semanticPrimitives.error_result'(X197),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X199,X197))))) )),
  inference(superposition,[],[f584,f572])).
fof(f9246,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (s('type.semanticPrimitives.error_result'(X193),X194) != s('type.semanticPrimitives.error_result'(X193),'const.semanticPrimitives.Rraise_1'(s(X193,happ(s(fun('type.semanticPrimitives.error_result'(X189),X193),X195),s('type.semanticPrimitives.error_result'(X189),'const.evalProps.map_error_result_2'(s(fun(X190,X189),X191),s('type.semanticPrimitives.error_result'(X190),'const.semanticPrimitives.Rraise_1'(s(X190,X192))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X193)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X193),X194))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X193)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X193),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X189),X193),X195),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X189)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X189),'const.semanticPrimitives.Rraise_1'(s(X189,happ(s(fun(X190,X189),X191),s(X190,X192)))))))))))) )),
  inference(superposition,[],[f584,f108])).
fof(f9245,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s('type.semanticPrimitives.error_result'(X186),X187) != s('type.semanticPrimitives.error_result'(X186),'const.semanticPrimitives.Rraise_1'(s(X186,happ(s(fun('type.semanticPrimitives.error_result'(X184),X186),X188),s('type.semanticPrimitives.error_result'(X184),X185))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X186)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X186),X187))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X186)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X186),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X184),X186),X188),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X184)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X184),'const.semanticPrimitives.Rraise_1'(s(X184,sK5(X185,X184)))))))))) | s('type.semanticPrimitives.error_result'(X184),X185) = s('type.semanticPrimitives.error_result'(X184),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X185,X184))))) )),
  inference(superposition,[],[f584,f95])).
fof(f9244,plain,(
  ( ! [X182,X180,X183,X181] : (s('type.semanticPrimitives.error_result'(X181),X182) != s('type.semanticPrimitives.error_result'(X181),'const.semanticPrimitives.Rraise_1'(s(X181,happ(s(fun('type.semanticPrimitives.error_result'(bool),X181),X183),s('type.semanticPrimitives.error_result'(bool),X180))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X181)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X181),X182))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X181)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X181),X183),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X180,bool)))))))))) | s('type.semanticPrimitives.error_result'(bool),X180) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X180,bool))) )),
  inference(superposition,[],[f584,f536])).
fof(f9242,plain,(
  ( ! [X175,X173,X171,X169,X176,X174,X172,X170,X168] : (s('type.semanticPrimitives.error_result'(X174),X175) != s('type.semanticPrimitives.error_result'(X174),'const.semanticPrimitives.Rraise_1'(s(X174,happ(s(fun('type.semanticPrimitives.error_result'(X168),X174),X176),s('type.semanticPrimitives.error_result'(X168),'const.semanticPrimitives.Rraise_1'(s(X168,happ(s(fun(X169,X168),X170),s(X169,happ(s(fun(X171,X169),X172),s(X171,X173))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X174)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X174),X175))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X174)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X174),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X168),X174),X176),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X168)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X168),'const.evalProps.map_error_result_2'(s(fun(X169,X168),X170),s('type.semanticPrimitives.error_result'(X169),'const.evalProps.map_error_result_2'(s(fun(X171,X169),X172),s('type.semanticPrimitives.error_result'(X171),'const.semanticPrimitives.Rraise_1'(s(X171,X173)))))))))))))) )),
  inference(superposition,[],[f584,f573])).
fof(f9241,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162] : (s('type.semanticPrimitives.error_result'(X165),X166) != s('type.semanticPrimitives.error_result'(X165),'const.semanticPrimitives.Rraise_1'(s(X165,happ(s(fun('type.semanticPrimitives.error_result'(X161),X165),X167),s('type.semanticPrimitives.error_result'(X161),'const.semanticPrimitives.Rraise_1'(s(X161,happ(s(fun(X162,X161),X163),s(X162,X164))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X165),X166))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X165),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X161),X165),X167),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X161)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X161),'const.evalProps.map_error_result_2'(s(fun(X162,X161),X163),s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,X164)))))))))))) )),
  inference(superposition,[],[f584,f108])).
fof(f9222,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.semanticPrimitives.error_result'(X85),X86) != s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rraise_1'(s(X85,happ(s(fun(bool,X85),X87),s(bool,X84))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),X86))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),'const.evalProps.map_error_result_2'(s(fun(bool,X85),X87),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f584,f83])).
fof(f9220,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s('type.semanticPrimitives.error_result'(X75),X76) != s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(bool,X75),X77),s(bool,X74))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(bool,X75),X77),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f584,f125])).
fof(f9219,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.semanticPrimitives.error_result'(X70),X71) != s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(bool,X70),X72),s(bool,X69))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),X71))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(bool,X70),X72),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s(bool,X69) = s(bool,X73) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X73)),bool))) )),
  inference(superposition,[],[f584,f689])).
fof(f9211,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'(X35),X36) != s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(bool,X35),X37),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),X36))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(bool,X35),X37),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34))))))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f584,f689])).
fof(f9207,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),X17) != s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X18),s(bool,X15))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X18),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X14))))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f584,f125])).
fof(f9206,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),X11) != s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X12),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X11))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X12),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))))))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f584,f125])).
fof(f9205,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(bool,X6),X8),s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X8),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f584,f83])).
fof(f9201,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),happ(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s(X69,X71))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,happ(s(fun(X66,X65),X67),s(X66,X68))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.evalProps.map_error_result_2'(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X71))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.evalProps.map_error_result_2'(s(fun(X66,X65),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68)))))))))) )),
  inference(superposition,[],[f584,f558])).
fof(f9200,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),happ(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s(X62,X64))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(X59,X58),X60),s('type.semanticPrimitives.error_result'(X59),'const.semanticPrimitives.Rraise_1'(s(X59,X61))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.evalProps.map_error_result_2'(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rraise_1'(s(X62,X64))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,happ(s(fun(X59,X58),X60),s(X59,X61)))))))))) )),
  inference(superposition,[],[f584,f558])).
fof(f9199,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X55,X49),X56),s(X55,X57))))) != s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X50,X49),X51),s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X52,X50),X53),s('type.semanticPrimitives.error_result'(X52),'const.semanticPrimitives.Rraise_1'(s(X52,X54))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X55,X49),X56),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,X57))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X50,X49),X51),s(X50,happ(s(fun(X52,X50),X53),s(X52,X54)))))))))) )),
  inference(superposition,[],[f584,f573])).
fof(f9198,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),X45))) != s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X46,X42),X47),s(X46,X48))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X46,X42),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,X48))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,sK5(X45,X43)))))))) | s('type.semanticPrimitives.error_result'(X43),X45) = s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X45,X43))))) )),
  inference(superposition,[],[f584,f572])).
fof(f9194,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X20,X14),X21),s(X20,X22))))) != s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X20,X14),X21),s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,X22))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))))))) )),
  inference(superposition,[],[f584,f573])).
fof(f9193,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X8,X7),X9),s(X8,X10))))) != s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X8,X7),X9),s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,X10))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X11,X7),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13)))))))) )),
  inference(superposition,[],[f584,f108])).
fof(f8785,plain,(
  ( ! [X26,X24,X23,X21,X19,X27,X25,X22,X20] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X25)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X25),X26))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X25)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X25),'const.evalProps.map_error_result_2'(s(fun(X19,X25),X27),s('type.semanticPrimitives.error_result'(X19),'const.evalProps.map_error_result_2'(s(fun(X20,X19),X21),s('type.semanticPrimitives.error_result'(X20),'const.evalProps.map_error_result_2'(s(fun(X22,X20),X23),s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rraise_1'(s(X22,X24))))))))))) | s('type.semanticPrimitives.error_result'(X25),X26) = s('type.semanticPrimitives.error_result'(X25),'const.semanticPrimitives.Rraise_1'(s(X25,happ(s(fun(X19,X25),X27),s(X19,happ(s(fun(X20,X19),X21),s(X20,happ(s(fun(X22,X20),X23),s(X22,X24)))))))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8784,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X12,X16),X18),s('type.semanticPrimitives.error_result'(X12),'const.evalProps.map_error_result_2'(s(fun(X13,X12),X14),s('type.semanticPrimitives.error_result'(X13),X15))))))) | s('type.semanticPrimitives.error_result'(X16),X17) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X12,X16),X18),s(X12,happ(s(fun(X13,X12),X14),s(X13,sK5(X15,X13)))))))) | s('type.semanticPrimitives.error_result'(X13),X15) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X15,X13))))) )),
  inference(superposition,[],[f580,f572])).
fof(f8778,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X221] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),X220))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X219)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X219),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),X219),X221),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X215),'const.semanticPrimitives.Rraise_1'(s(X215,happ(s(fun(X216,X215),X217),s(X216,X218))))))))))))) | s('type.semanticPrimitives.error_result'(X219),X220) = s('type.semanticPrimitives.error_result'(X219),'const.semanticPrimitives.Rraise_1'(s(X219,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),X219),X221),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X215),'const.evalProps.map_error_result_2'(s(fun(X216,X215),X217),s('type.semanticPrimitives.error_result'(X216),'const.semanticPrimitives.Rraise_1'(s(X216,X218)))))))))))) )),
  inference(superposition,[],[f580,f558])).
fof(f8777,plain,(
  ( ! [X212,X210,X208,X213,X211,X209,X214] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),X213))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X212)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X212),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),X212),X214),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X208),'const.evalProps.map_error_result_2'(s(fun(X209,X208),X210),s('type.semanticPrimitives.error_result'(X209),'const.semanticPrimitives.Rraise_1'(s(X209,X211))))))))))))) | s('type.semanticPrimitives.error_result'(X212),X213) = s('type.semanticPrimitives.error_result'(X212),'const.semanticPrimitives.Rraise_1'(s(X212,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),X212),X214),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X208),'const.semanticPrimitives.Rraise_1'(s(X208,happ(s(fun(X209,X208),X210),s(X209,X211)))))))))))) )),
  inference(superposition,[],[f580,f558])).
fof(f8776,plain,(
  ( ! [X206,X204,X202,X200,X199,X207,X205,X203,X201] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X205)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X205),X206))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X205)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X205),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X199),X205),X207),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X199)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X199),'const.evalProps.map_error_result_2'(s(fun(X200,X199),X201),s('type.semanticPrimitives.error_result'(X200),'const.evalProps.map_error_result_2'(s(fun(X202,X200),X203),s('type.semanticPrimitives.error_result'(X202),'const.semanticPrimitives.Rraise_1'(s(X202,X204))))))))))))) | s('type.semanticPrimitives.error_result'(X205),X206) = s('type.semanticPrimitives.error_result'(X205),'const.semanticPrimitives.Rraise_1'(s(X205,happ(s(fun('type.semanticPrimitives.error_result'(X199),X205),X207),s('type.semanticPrimitives.error_result'(X199),'const.semanticPrimitives.Rraise_1'(s(X199,happ(s(fun(X200,X199),X201),s(X200,happ(s(fun(X202,X200),X203),s(X202,X204)))))))))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8775,plain,(
  ( ! [X198,X196,X194,X192,X197,X195,X193] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X196)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X196),X197))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X196)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X196),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X192),X196),X198),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X192)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X192),'const.evalProps.map_error_result_2'(s(fun(X193,X192),X194),s('type.semanticPrimitives.error_result'(X193),X195))))))))) | s('type.semanticPrimitives.error_result'(X196),X197) = s('type.semanticPrimitives.error_result'(X196),'const.semanticPrimitives.Rraise_1'(s(X196,happ(s(fun('type.semanticPrimitives.error_result'(X192),X196),X198),s('type.semanticPrimitives.error_result'(X192),'const.semanticPrimitives.Rraise_1'(s(X192,happ(s(fun(X193,X192),X194),s(X193,sK5(X195,X193)))))))))) | s('type.semanticPrimitives.error_result'(X193),X195) = s('type.semanticPrimitives.error_result'(X193),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X195,X193))))) )),
  inference(superposition,[],[f580,f572])).
fof(f8774,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X189)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X189),X190))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X189)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X189),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X185),X189),X191),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X185)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X185),'const.evalProps.map_error_result_2'(s(fun(X186,X185),X187),s('type.semanticPrimitives.error_result'(X186),'const.semanticPrimitives.Rraise_1'(s(X186,X188))))))))))) | s('type.semanticPrimitives.error_result'(X189),X190) = s('type.semanticPrimitives.error_result'(X189),'const.semanticPrimitives.Rraise_1'(s(X189,happ(s(fun('type.semanticPrimitives.error_result'(X185),X189),X191),s('type.semanticPrimitives.error_result'(X185),'const.semanticPrimitives.Rraise_1'(s(X185,happ(s(fun(X186,X185),X187),s(X186,X188)))))))))) )),
  inference(superposition,[],[f580,f108])).
fof(f8773,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X182)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X182),X183))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X182)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X182),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X180),X182),X184),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),X181))))))) | s('type.semanticPrimitives.error_result'(X182),X183) = s('type.semanticPrimitives.error_result'(X182),'const.semanticPrimitives.Rraise_1'(s(X182,happ(s(fun('type.semanticPrimitives.error_result'(X180),X182),X184),s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rraise_1'(s(X180,sK5(X181,X180)))))))) | s('type.semanticPrimitives.error_result'(X180),X181) = s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X181,X180))))) )),
  inference(superposition,[],[f580,f95])).
fof(f8772,plain,(
  ( ! [X177,X178,X176,X179] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X177)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X177),X178))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X177)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X177),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X177),X179),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X176))))))) | s('type.semanticPrimitives.error_result'(X177),X178) = s('type.semanticPrimitives.error_result'(X177),'const.semanticPrimitives.Rraise_1'(s(X177,happ(s(fun('type.semanticPrimitives.error_result'(bool),X177),X179),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X176,bool)))))))) | s('type.semanticPrimitives.error_result'(bool),X176) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X176,bool))) )),
  inference(superposition,[],[f580,f536])).
fof(f8770,plain,(
  ( ! [X167,X165,X171,X169,X166,X164,X172,X170,X168] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X170)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X170),X171))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X170)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X170),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X164),X170),X172),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X164)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X164),'const.semanticPrimitives.Rraise_1'(s(X164,happ(s(fun(X165,X164),X166),s(X165,happ(s(fun(X167,X165),X168),s(X167,X169))))))))))))) | s('type.semanticPrimitives.error_result'(X170),X171) = s('type.semanticPrimitives.error_result'(X170),'const.semanticPrimitives.Rraise_1'(s(X170,happ(s(fun('type.semanticPrimitives.error_result'(X164),X170),X172),s('type.semanticPrimitives.error_result'(X164),'const.evalProps.map_error_result_2'(s(fun(X165,X164),X166),s('type.semanticPrimitives.error_result'(X165),'const.evalProps.map_error_result_2'(s(fun(X167,X165),X168),s('type.semanticPrimitives.error_result'(X167),'const.semanticPrimitives.Rraise_1'(s(X167,X169)))))))))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8769,plain,(
  ( ! [X158,X163,X161,X159,X157,X162,X160] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X161)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X161),X162))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X161)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X161),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X157),X161),X163),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X157)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X157),'const.semanticPrimitives.Rraise_1'(s(X157,happ(s(fun(X158,X157),X159),s(X158,X160))))))))))) | s('type.semanticPrimitives.error_result'(X161),X162) = s('type.semanticPrimitives.error_result'(X161),'const.semanticPrimitives.Rraise_1'(s(X161,happ(s(fun('type.semanticPrimitives.error_result'(X157),X161),X163),s('type.semanticPrimitives.error_result'(X157),'const.evalProps.map_error_result_2'(s(fun(X158,X157),X159),s('type.semanticPrimitives.error_result'(X158),'const.semanticPrimitives.Rraise_1'(s(X158,X160)))))))))) )),
  inference(superposition,[],[f580,f108])).
fof(f8751,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),X86))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),'const.evalProps.map_error_result_2'(s(fun(bool,X85),X87),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X84))))))) | s('type.semanticPrimitives.error_result'(X85),X86) = s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rraise_1'(s(X85,happ(s(fun(bool,X85),X87),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f580,f83])).
fof(f8749,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(bool,X75),X77),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X74))))))) | s('type.semanticPrimitives.error_result'(X75),X76) = s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(bool,X75),X77),s(bool,'T_0'))))) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f580,f125])).
fof(f8748,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),X71))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(bool,X70),X72),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X69))))))) | s('type.semanticPrimitives.error_result'(X70),X71) = s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(bool,X70),X72),s(bool,'T_0'))))) | s(bool,X69) = s(bool,X73) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X73)),bool))) )),
  inference(superposition,[],[f580,f689])).
fof(f8740,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),X36))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(bool,X35),X37),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X35),X36) = s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(bool,X35),X37),s(bool,X34))))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f580,f689])).
fof(f8736,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X18),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))))))) | s('type.semanticPrimitives.error_result'(X16),X17) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X18),s(bool,X14))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f580,f125])).
fof(f8735,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X11))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X12),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s('type.semanticPrimitives.error_result'(X10),X11) = s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X12),s(bool,X9))))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f580,f125])).
fof(f8734,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X8),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))))) | s('type.semanticPrimitives.error_result'(X6),X7) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(bool,X6),X8),s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f580,f83])).
fof(f8731,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X33)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X33),'const.evalProps.map_error_result_2'(s(fun(X39,X33),X40),s('type.semanticPrimitives.error_result'(X39),'const.semanticPrimitives.Rraise_1'(s(X39,X41))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X33)),'const.evalProps.map_error_result_2'(s(fun(X34,'type.semanticPrimitives.error_result'(X33)),X35),s('type.semanticPrimitives.error_result'(X34),'const.evalProps.map_error_result_2'(s(fun(X36,X34),X37),s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,X38))))))) | s('type.semanticPrimitives.error_result'(X33),'const.semanticPrimitives.Rraise_1'(s(X33,happ(s(fun(X39,X33),X40),s(X39,X41))))) = s('type.semanticPrimitives.error_result'(X33),happ(s(fun(X34,'type.semanticPrimitives.error_result'(X33)),X35),s(X34,happ(s(fun(X36,X34),X37),s(X36,X38)))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8730,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.evalProps.map_error_result_2'(s(fun(X27,'type.semanticPrimitives.error_result'(X26)),X28),s('type.semanticPrimitives.error_result'(X27),X29))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.evalProps.map_error_result_2'(s(fun(X30,X26),X31),s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rraise_1'(s(X30,X32))))))) | s('type.semanticPrimitives.error_result'(X26),'const.semanticPrimitives.Rraise_1'(s(X26,happ(s(fun(X30,X26),X31),s(X30,X32))))) = s('type.semanticPrimitives.error_result'(X26),happ(s(fun(X27,'type.semanticPrimitives.error_result'(X26)),X28),s(X27,sK5(X29,X27)))) | s('type.semanticPrimitives.error_result'(X27),X29) = s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X29,X27))))) )),
  inference(superposition,[],[f580,f572])).
fof(f8725,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(X77,X75),X78),s('type.semanticPrimitives.error_result'(X77),'const.semanticPrimitives.Rraise_1'(s(X77,X79))))))) | s('type.semanticPrimitives.error_result'(X75),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76)),'type.semanticPrimitives.error_result'(X75))) = s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(X77,X75),X78),s(X77,X79)))))) )),
  inference(superposition,[],[f580,f596])).
fof(f8724,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.evalProps.map_error_result_2'(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X71))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,happ(s(fun(X66,X65),X67),s(X66,X68))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),happ(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s(X69,X71))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.evalProps.map_error_result_2'(s(fun(X66,X65),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68)))))))) )),
  inference(superposition,[],[f580,f558])).
fof(f8723,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.evalProps.map_error_result_2'(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rraise_1'(s(X62,X64))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(X59,X58),X60),s('type.semanticPrimitives.error_result'(X59),'const.semanticPrimitives.Rraise_1'(s(X59,X61))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),happ(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s(X62,X64))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,happ(s(fun(X59,X58),X60),s(X59,X61)))))))) )),
  inference(superposition,[],[f580,f558])).
fof(f8722,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X55,X49),X56),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,X57))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X50,X49),X51),s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X52,X50),X53),s('type.semanticPrimitives.error_result'(X52),'const.semanticPrimitives.Rraise_1'(s(X52,X54))))))))) | s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X55,X49),X56),s(X55,X57))))) = s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X50,X49),X51),s(X50,happ(s(fun(X52,X50),X53),s(X52,X54)))))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8721,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),X45))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X46,X42),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,X48))))))) | s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X46,X42),X47),s(X46,X48))))) = s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,sK5(X45,X43)))))) | s('type.semanticPrimitives.error_result'(X43),X45) = s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X45,X43))))) )),
  inference(superposition,[],[f580,f572])).
fof(f8720,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(X36,X35),X37),s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,X38))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(X39,X35),X40),s('type.semanticPrimitives.error_result'(X39),'const.semanticPrimitives.Rraise_1'(s(X39,X41))))))) | s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(X36,X35),X37),s(X36,X38))))) = s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(X39,X35),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f580,f108])).
fof(f8717,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X20,X14),X21),s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,X22))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))))) | s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X20,X14),X21),s(X20,X22))))) = s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))))) )),
  inference(superposition,[],[f580,f573])).
fof(f8716,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X8,X7),X9),s(X8,X10))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X11,X7),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13))))))) | s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))) = s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X8,X7),X9),s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,X10)))))) )),
  inference(superposition,[],[f580,f108])).
fof(f8380,plain,(
  ( ! [X26,X24,X23,X21,X19,X27,X25,X22,X20] : (s('type.semanticPrimitives.error_result'(X25),X26) != s('type.semanticPrimitives.error_result'(X25),'const.evalProps.map_error_result_2'(s(fun(X19,X25),X27),s('type.semanticPrimitives.error_result'(X19),'const.evalProps.map_error_result_2'(s(fun(X20,X19),X21),s('type.semanticPrimitives.error_result'(X20),'const.evalProps.map_error_result_2'(s(fun(X22,X20),X23),s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rraise_1'(s(X22,X24))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X25)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X25),X26))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X25)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X25),'const.semanticPrimitives.Rraise_1'(s(X25,happ(s(fun(X19,X25),X27),s(X19,happ(s(fun(X20,X19),X21),s(X20,happ(s(fun(X22,X20),X23),s(X22,X24)))))))))))) )),
  inference(superposition,[],[f578,f573])).
fof(f8379,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),X17) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X12,X16),X18),s('type.semanticPrimitives.error_result'(X12),'const.evalProps.map_error_result_2'(s(fun(X13,X12),X14),s('type.semanticPrimitives.error_result'(X13),X15))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X12,X16),X18),s(X12,happ(s(fun(X13,X12),X14),s(X13,sK5(X15,X13)))))))))) | s('type.semanticPrimitives.error_result'(X13),X15) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X15,X13))))) )),
  inference(superposition,[],[f578,f572])).
fof(f8373,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (s('type.semanticPrimitives.error_result'(X215),X216) != s('type.semanticPrimitives.error_result'(X215),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X211)),X215),X217),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X211))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X211)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X211),'const.semanticPrimitives.Rraise_1'(s(X211,happ(s(fun(X212,X211),X213),s(X212,X214))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X215),X216))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X215)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X215),'const.semanticPrimitives.Rraise_1'(s(X215,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X211)),X215),X217),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X211)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X211),'const.evalProps.map_error_result_2'(s(fun(X212,X211),X213),s('type.semanticPrimitives.error_result'(X212),'const.semanticPrimitives.Rraise_1'(s(X212,X214)))))))))))))) )),
  inference(superposition,[],[f578,f558])).
fof(f8372,plain,(
  ( ! [X210,X208,X206,X204,X209,X207,X205] : (s('type.semanticPrimitives.error_result'(X208),X209) != s('type.semanticPrimitives.error_result'(X208),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X204)),X208),X210),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X204))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X204)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X204),'const.evalProps.map_error_result_2'(s(fun(X205,X204),X206),s('type.semanticPrimitives.error_result'(X205),'const.semanticPrimitives.Rraise_1'(s(X205,X207))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X208),X209))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X208)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X208),'const.semanticPrimitives.Rraise_1'(s(X208,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X204)),X208),X210),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X204)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X204),'const.semanticPrimitives.Rraise_1'(s(X204,happ(s(fun(X205,X204),X206),s(X205,X207)))))))))))))) )),
  inference(superposition,[],[f578,f558])).
fof(f8371,plain,(
  ( ! [X198,X196,X202,X200,X199,X197,X195,X203,X201] : (s('type.semanticPrimitives.error_result'(X201),X202) != s('type.semanticPrimitives.error_result'(X201),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X195),X201),X203),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X195)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X195),'const.evalProps.map_error_result_2'(s(fun(X196,X195),X197),s('type.semanticPrimitives.error_result'(X196),'const.evalProps.map_error_result_2'(s(fun(X198,X196),X199),s('type.semanticPrimitives.error_result'(X198),'const.semanticPrimitives.Rraise_1'(s(X198,X200))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X201)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X201),X202))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X201)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X201),'const.semanticPrimitives.Rraise_1'(s(X201,happ(s(fun('type.semanticPrimitives.error_result'(X195),X201),X203),s('type.semanticPrimitives.error_result'(X195),'const.semanticPrimitives.Rraise_1'(s(X195,happ(s(fun(X196,X195),X197),s(X196,happ(s(fun(X198,X196),X199),s(X198,X200)))))))))))))) )),
  inference(superposition,[],[f578,f573])).
fof(f8370,plain,(
  ( ! [X191,X189,X194,X192,X190,X188,X193] : (s('type.semanticPrimitives.error_result'(X192),X193) != s('type.semanticPrimitives.error_result'(X192),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X188),X192),X194),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X188)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X188),'const.evalProps.map_error_result_2'(s(fun(X189,X188),X190),s('type.semanticPrimitives.error_result'(X189),X191))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X192)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X192),X193))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X192)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X192),'const.semanticPrimitives.Rraise_1'(s(X192,happ(s(fun('type.semanticPrimitives.error_result'(X188),X192),X194),s('type.semanticPrimitives.error_result'(X188),'const.semanticPrimitives.Rraise_1'(s(X188,happ(s(fun(X189,X188),X190),s(X189,sK5(X191,X189)))))))))))) | s('type.semanticPrimitives.error_result'(X189),X191) = s('type.semanticPrimitives.error_result'(X189),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X191,X189))))) )),
  inference(superposition,[],[f578,f572])).
fof(f8369,plain,(
  ( ! [X187,X185,X182,X186,X184,X183,X181] : (s('type.semanticPrimitives.error_result'(X185),X186) != s('type.semanticPrimitives.error_result'(X185),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X181),X185),X187),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X181)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun(X182,X181),X183),s('type.semanticPrimitives.error_result'(X182),'const.semanticPrimitives.Rraise_1'(s(X182,X184))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X185)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X185),X186))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X185)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X185),'const.semanticPrimitives.Rraise_1'(s(X185,happ(s(fun('type.semanticPrimitives.error_result'(X181),X185),X187),s('type.semanticPrimitives.error_result'(X181),'const.semanticPrimitives.Rraise_1'(s(X181,happ(s(fun(X182,X181),X183),s(X182,X184)))))))))))) )),
  inference(superposition,[],[f578,f108])).
fof(f8368,plain,(
  ( ! [X177,X180,X178,X176,X179] : (s('type.semanticPrimitives.error_result'(X178),X179) != s('type.semanticPrimitives.error_result'(X178),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X176),X178),X180),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X176)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X176),X177))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X178)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X178),X179))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X178)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X178),'const.semanticPrimitives.Rraise_1'(s(X178,happ(s(fun('type.semanticPrimitives.error_result'(X176),X178),X180),s('type.semanticPrimitives.error_result'(X176),'const.semanticPrimitives.Rraise_1'(s(X176,sK5(X177,X176)))))))))) | s('type.semanticPrimitives.error_result'(X176),X177) = s('type.semanticPrimitives.error_result'(X176),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X177,X176))))) )),
  inference(superposition,[],[f578,f95])).
fof(f8367,plain,(
  ( ! [X175,X173,X174,X172] : (s('type.semanticPrimitives.error_result'(X173),X174) != s('type.semanticPrimitives.error_result'(X173),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X173),X175),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X172))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X173)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X173),X174))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X173)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X173),'const.semanticPrimitives.Rraise_1'(s(X173,happ(s(fun('type.semanticPrimitives.error_result'(bool),X173),X175),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X172,bool)))))))))) | s('type.semanticPrimitives.error_result'(bool),X172) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X172,bool))) )),
  inference(superposition,[],[f578,f536])).
fof(f8365,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162,X160,X168] : (s('type.semanticPrimitives.error_result'(X166),X167) != s('type.semanticPrimitives.error_result'(X166),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X160),X166),X168),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X160)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X160),'const.semanticPrimitives.Rraise_1'(s(X160,happ(s(fun(X161,X160),X162),s(X161,happ(s(fun(X163,X161),X164),s(X163,X165))))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),X167))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun('type.semanticPrimitives.error_result'(X160),X166),X168),s('type.semanticPrimitives.error_result'(X160),'const.evalProps.map_error_result_2'(s(fun(X161,X160),X162),s('type.semanticPrimitives.error_result'(X161),'const.evalProps.map_error_result_2'(s(fun(X163,X161),X164),s('type.semanticPrimitives.error_result'(X163),'const.semanticPrimitives.Rraise_1'(s(X163,X165)))))))))))))) )),
  inference(superposition,[],[f578,f573])).
fof(f8364,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X153] : (s('type.semanticPrimitives.error_result'(X157),X158) != s('type.semanticPrimitives.error_result'(X157),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X153),X157),X159),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X153)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X153),'const.semanticPrimitives.Rraise_1'(s(X153,happ(s(fun(X154,X153),X155),s(X154,X156))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X157)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X157),X158))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X157)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X157),'const.semanticPrimitives.Rraise_1'(s(X157,happ(s(fun('type.semanticPrimitives.error_result'(X153),X157),X159),s('type.semanticPrimitives.error_result'(X153),'const.evalProps.map_error_result_2'(s(fun(X154,X153),X155),s('type.semanticPrimitives.error_result'(X154),'const.semanticPrimitives.Rraise_1'(s(X154,X156)))))))))))) )),
  inference(superposition,[],[f578,f108])).
fof(f8347,plain,(
  ( ! [X87,X85,X86,X84] : (s('type.semanticPrimitives.error_result'(X85),X86) != s('type.semanticPrimitives.error_result'(X85),'const.evalProps.map_error_result_2'(s(fun(bool,X85),X87),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X84))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),X86))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X85)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X85),'const.semanticPrimitives.Rraise_1'(s(X85,happ(s(fun(bool,X85),X87),s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f578,f83])).
fof(f8345,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s('type.semanticPrimitives.error_result'(X75),X76) != s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(bool,X75),X77),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X74))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),X76))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(bool,X75),X77),s(bool,'T_0'))))))) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f578,f125])).
fof(f8344,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s('type.semanticPrimitives.error_result'(X70),X71) != s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(bool,X70),X72),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X69))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),X71))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(bool,X70),X72),s(bool,'T_0'))))))) | s(bool,X69) = s(bool,X73) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X73)),bool))) )),
  inference(superposition,[],[f578,f689])).
fof(f8336,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'(X35),X36) != s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(bool,X35),X37),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),X36))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(bool,X35),X37),s(bool,X34))))))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f578,f689])).
fof(f8332,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),X17) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X18),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X17))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X18),s(bool,X14))))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f578,f125])).
fof(f8331,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),X11) != s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X12),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X11))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X12),s(bool,X9))))))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f578,f125])).
fof(f8330,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(X6),X7) != s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X8),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),X7))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X6)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(bool,X6),X8),s(bool,X5))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f578,f83])).
fof(f8326,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.evalProps.map_error_result_2'(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X71))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,happ(s(fun(X66,X65),X67),s(X66,X68))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),happ(s(fun(X69,'type.semanticPrimitives.error_result'(X65)),X70),s(X69,X71))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.evalProps.map_error_result_2'(s(fun(X66,X65),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68)))))))))) )),
  inference(superposition,[],[f578,f558])).
fof(f8325,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.evalProps.map_error_result_2'(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rraise_1'(s(X62,X64))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(X59,X58),X60),s('type.semanticPrimitives.error_result'(X59),'const.semanticPrimitives.Rraise_1'(s(X59,X61))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),happ(s(fun(X62,'type.semanticPrimitives.error_result'(X58)),X63),s(X62,X64))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,happ(s(fun(X59,X58),X60),s(X59,X61)))))))))) )),
  inference(superposition,[],[f578,f558])).
fof(f8324,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51,X49] : (s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X55,X49),X56),s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,X57))))) != s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X50,X49),X51),s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X52,X50),X53),s('type.semanticPrimitives.error_result'(X52),'const.semanticPrimitives.Rraise_1'(s(X52,X54))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X55,X49),X56),s(X55,X57))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X50,X49),X51),s(X50,happ(s(fun(X52,X50),X53),s(X52,X54)))))))))) )),
  inference(superposition,[],[f578,f573])).
fof(f8323,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),X45))) != s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X46,X42),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,X48))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X46,X42),X47),s(X46,X48))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,sK5(X45,X43)))))))) | s('type.semanticPrimitives.error_result'(X43),X45) = s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X45,X43))))) )),
  inference(superposition,[],[f578,f572])).
fof(f8322,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(X36,X35),X37),s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,X38))))) != s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(X39,X35),X40),s('type.semanticPrimitives.error_result'(X39),'const.semanticPrimitives.Rraise_1'(s(X39,X41))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(X36,X35),X37),s(X36,X38))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(X39,X35),X40),s(X39,X41)))))))) )),
  inference(superposition,[],[f578,f108])).
fof(f8319,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X20,X14),X21),s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,X22))))) != s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X20,X14),X21),s(X20,X22))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))))))) )),
  inference(superposition,[],[f578,f573])).
fof(f8318,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X8,X7),X9),s(X8,X10))))) != s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X11,X7),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X7)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X8,X7),X9),s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,X10)))))))) )),
  inference(superposition,[],[f578,f108])).
fof(f7727,plain,(
  ( ! [X134,X132,X130,X138,X136,X135,X133,X131,X137] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X130)),'const.evalProps.map_error_result_2'(s(fun(X136,'type.semanticPrimitives.error_result'(X130)),X137),s('type.semanticPrimitives.error_result'(X136),'const.semanticPrimitives.Rraise_1'(s(X136,X138))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X130)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X130),'const.semanticPrimitives.Rraise_1'(s(X130,happ(s(fun(X131,X130),X132),s(X131,happ(s(fun(X133,X131),X134),s(X133,X135))))))))) | s('type.semanticPrimitives.error_result'(X130),happ(s(fun(X136,'type.semanticPrimitives.error_result'(X130)),X137),s(X136,X138))) = s('type.semanticPrimitives.error_result'(X130),'const.evalProps.map_error_result_2'(s(fun(X131,X130),X132),s('type.semanticPrimitives.error_result'(X131),'const.evalProps.map_error_result_2'(s(fun(X133,X131),X134),s('type.semanticPrimitives.error_result'(X133),'const.semanticPrimitives.Rraise_1'(s(X133,X135)))))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7726,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X129] : (s('type.semanticPrimitives.error_result'(X127),'const.semanticPrimitives.Rraise_1'(s(X127,X128))) != s('type.semanticPrimitives.error_result'(X127),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X121),X127),X129),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X121)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X121),'const.semanticPrimitives.Rraise_1'(s(X121,happ(s(fun(X122,X121),X123),s(X122,happ(s(fun(X124,X122),X125),s(X124,X126))))))))))) | s(X127,X128) = s(X127,happ(s(fun('type.semanticPrimitives.error_result'(X121),X127),X129),s('type.semanticPrimitives.error_result'(X121),'const.evalProps.map_error_result_2'(s(fun(X122,X121),X123),s('type.semanticPrimitives.error_result'(X122),'const.evalProps.map_error_result_2'(s(fun(X124,X122),X125),s('type.semanticPrimitives.error_result'(X124),'const.semanticPrimitives.Rraise_1'(s(X124,X126)))))))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7724,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),X112) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X106),'const.semanticPrimitives.Rraise_1'(s(X106,happ(s(fun(X107,X106),X108),s(X107,happ(s(fun(X109,X107),X110),s(X109,X111))))))))) | s('type.semanticPrimitives.error_result'(X106),sK5(X112,'type.semanticPrimitives.error_result'(X106))) = s('type.semanticPrimitives.error_result'(X106),'const.evalProps.map_error_result_2'(s(fun(X107,X106),X108),s('type.semanticPrimitives.error_result'(X107),'const.evalProps.map_error_result_2'(s(fun(X109,X107),X110),s('type.semanticPrimitives.error_result'(X109),'const.semanticPrimitives.Rraise_1'(s(X109,X111))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),X112) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X112,'type.semanticPrimitives.error_result'(X106)))))) )),
  inference(superposition,[],[f537,f573])).
fof(f7719,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s('type.semanticPrimitives.error_result'(X70),X76) != s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(X71,X70),X72),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),X76))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(X71,X70),X72),s('type.semanticPrimitives.error_result'(X71),'const.evalProps.map_error_result_2'(s(fun(X73,X71),X74),s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rraise_1'(s(X73,X75)))))))))) )),
  inference(superposition,[],[f101,f573])).
fof(f7717,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X56)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X56),X62))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X56)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,happ(s(fun(X57,X56),X58),s(X57,happ(s(fun(X59,X57),X60),s(X59,X61))))))))) | s('type.semanticPrimitives.error_result'(X56),X62) = s('type.semanticPrimitives.error_result'(X56),'const.evalProps.map_error_result_2'(s(fun(X57,X56),X58),s('type.semanticPrimitives.error_result'(X57),'const.evalProps.map_error_result_2'(s(fun(X59,X57),X60),s('type.semanticPrimitives.error_result'(X59),'const.semanticPrimitives.Rraise_1'(s(X59,X61)))))))) )),
  inference(superposition,[],[f100,f573])).
fof(f7716,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),X55))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X49)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,happ(s(fun(X50,X49),X51),s(X50,happ(s(fun(X52,X50),X53),s(X52,X54))))))))) | s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,sK5(X55,X49)))) = s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X50,X49),X51),s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X52,X50),X53),s('type.semanticPrimitives.error_result'(X52),'const.semanticPrimitives.Rraise_1'(s(X52,X54))))))) | s('type.semanticPrimitives.error_result'(X49),X55) = s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X55,X49))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7766,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),X48))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,happ(s(fun(X45,X43),X46),s(X45,X47))))))))) | s('type.semanticPrimitives.error_result'(X42),X48) = s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,sK5('const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),'const.evalProps.map_error_result_2'(s(fun(X45,X43),X46),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rraise_1'(s(X45,X47)))))),X42))))) )),
  inference(subsumption_resolution,[],[f7715,f886])).
fof(f7715,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),X48))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X42)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,happ(s(fun(X43,X42),X44),s(X43,happ(s(fun(X45,X43),X46),s(X45,X47))))))))) | s('type.semanticPrimitives.error_result'(X42),X48) = s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rraise_1'(s(X42,sK5('const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),'const.evalProps.map_error_result_2'(s(fun(X45,X43),X46),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rraise_1'(s(X45,X47)))))),X42)))) | s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),'const.evalProps.map_error_result_2'(s(fun(X45,X43),X46),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rraise_1'(s(X45,X47))))))) = s('type.semanticPrimitives.error_result'(X42),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X43,X42),X44),s('type.semanticPrimitives.error_result'(X43),'const.evalProps.map_error_result_2'(s(fun(X45,X43),X46),s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rraise_1'(s(X45,X47)))))),X42))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7714,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (s('type.semanticPrimitives.error_result'(X35),X41) != s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,happ(s(fun(X36,X35),X37),s(X36,happ(s(fun(X38,X36),X39),s(X38,X40))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,sK5(X41,X35)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X35)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(X36,X35),X37),s('type.semanticPrimitives.error_result'(X36),'const.evalProps.map_error_result_2'(s(fun(X38,X36),X39),s('type.semanticPrimitives.error_result'(X38),'const.semanticPrimitives.Rraise_1'(s(X38,X40))))))))) | s('type.semanticPrimitives.error_result'(X35),X41) = s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X41,X35))))) )),
  inference(superposition,[],[f540,f573])).
fof(f7765,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.semanticPrimitives.error_result'(X28),X34) != s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,happ(s(fun(X29,X28),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X28)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X28),X34))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X28)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,sK5('const.evalProps.map_error_result_2'(s(fun(X29,X28),X30),s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X31,X29),X32),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,X33)))))),X28))))))) )),
  inference(subsumption_resolution,[],[f7713,f886])).
fof(f7713,plain,(
  ( ! [X30,X28,X33,X31,X29,X34,X32] : (s('type.semanticPrimitives.error_result'(X28),X34) != s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,happ(s(fun(X29,X28),X30),s(X29,happ(s(fun(X31,X29),X32),s(X31,X33))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X28)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X28),X34))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X28)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,sK5('const.evalProps.map_error_result_2'(s(fun(X29,X28),X30),s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X31,X29),X32),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,X33)))))),X28)))))) | s('type.semanticPrimitives.error_result'(X28),'const.evalProps.map_error_result_2'(s(fun(X29,X28),X30),s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X31,X29),X32),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,X33))))))) = s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X29,X28),X30),s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X31,X29),X32),s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,X33)))))),X28))))) )),
  inference(superposition,[],[f540,f573])).
fof(f7764,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,X20))) != s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))) | s(X14,X20) = s(X14,sK5('const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))),X14))) )),
  inference(subsumption_resolution,[],[f7711,f886])).
fof(f7711,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,X20))) != s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun(X15,X14),X16),s(X15,happ(s(fun(X17,X15),X18),s(X17,X19))))))) | s(X14,X20) = s(X14,sK5('const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))),X14)) | s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))))) = s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X17,X15),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19)))))),X14))))) )),
  inference(superposition,[],[f537,f573])).
fof(f7706,plain,(
  ( ! [X269,X271,X273,X275,X268,X270,X276,X272,X274] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.evalProps.map_error_result_2'(s(fun(X274,'type.semanticPrimitives.error_result'(X268)),X275),s('type.semanticPrimitives.error_result'(X274),'const.semanticPrimitives.Rraise_1'(s(X274,X276))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X268)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X268),'const.evalProps.map_error_result_2'(s(fun(X269,X268),X270),s('type.semanticPrimitives.error_result'(X269),'const.evalProps.map_error_result_2'(s(fun(X271,X269),X272),s('type.semanticPrimitives.error_result'(X271),'const.semanticPrimitives.Rraise_1'(s(X271,X273))))))))) | s('type.semanticPrimitives.error_result'(X268),happ(s(fun(X274,'type.semanticPrimitives.error_result'(X268)),X275),s(X274,X276))) = s('type.semanticPrimitives.error_result'(X268),'const.semanticPrimitives.Rraise_1'(s(X268,happ(s(fun(X269,X268),X270),s(X269,happ(s(fun(X271,X269),X272),s(X271,X273)))))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7705,plain,(
  ( ! [X261,X263,X259,X265,X267,X260,X262,X264,X266] : (s('type.semanticPrimitives.error_result'(X265),'const.semanticPrimitives.Rraise_1'(s(X265,X266))) != s('type.semanticPrimitives.error_result'(X265),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X259),X265),X267),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X259)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X259),'const.evalProps.map_error_result_2'(s(fun(X260,X259),X261),s('type.semanticPrimitives.error_result'(X260),'const.evalProps.map_error_result_2'(s(fun(X262,X260),X263),s('type.semanticPrimitives.error_result'(X262),'const.semanticPrimitives.Rraise_1'(s(X262,X264))))))))))) | s(X265,X266) = s(X265,happ(s(fun('type.semanticPrimitives.error_result'(X259),X265),X267),s('type.semanticPrimitives.error_result'(X259),'const.semanticPrimitives.Rraise_1'(s(X259,happ(s(fun(X260,X259),X261),s(X260,happ(s(fun(X262,X260),X263),s(X262,X264)))))))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7703,plain,(
  ( ! [X249,X246,X244,X250,X248,X247,X245] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X244)),X250) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X244)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X244),'const.evalProps.map_error_result_2'(s(fun(X245,X244),X246),s('type.semanticPrimitives.error_result'(X245),'const.evalProps.map_error_result_2'(s(fun(X247,X245),X248),s('type.semanticPrimitives.error_result'(X247),'const.semanticPrimitives.Rraise_1'(s(X247,X249))))))))) | s('type.semanticPrimitives.error_result'(X244),sK5(X250,'type.semanticPrimitives.error_result'(X244))) = s('type.semanticPrimitives.error_result'(X244),'const.semanticPrimitives.Rraise_1'(s(X244,happ(s(fun(X245,X244),X246),s(X245,happ(s(fun(X247,X245),X248),s(X247,X249))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X244)),X250) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X244)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X250,'type.semanticPrimitives.error_result'(X244)))))) )),
  inference(superposition,[],[f537,f573])).
fof(f7695,plain,(
  ( ! [X191,X189,X187,X192,X190,X188,X193] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X187)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X187),X193))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X187)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X187),'const.evalProps.map_error_result_2'(s(fun(X188,X187),X189),s('type.semanticPrimitives.error_result'(X188),'const.evalProps.map_error_result_2'(s(fun(X190,X188),X191),s('type.semanticPrimitives.error_result'(X190),'const.semanticPrimitives.Rraise_1'(s(X190,X192))))))))) | s('type.semanticPrimitives.error_result'(X187),'const.semanticPrimitives.Rraise_1'(s(X187,sK5(X193,X187)))) = s('type.semanticPrimitives.error_result'(X187),'const.semanticPrimitives.Rraise_1'(s(X187,happ(s(fun(X188,X187),X189),s(X188,happ(s(fun(X190,X188),X191),s(X190,X192))))))) | s('type.semanticPrimitives.error_result'(X187),X193) = s('type.semanticPrimitives.error_result'(X187),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X193,X187))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7763,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (s('type.semanticPrimitives.error_result'(X180),X186) = s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rraise_1'(s(X180,happ(s(fun(X181,X180),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),X186))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),'const.evalProps.map_error_result_2'(s(fun(X181,X180),X182),s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun(X183,X181),X184),s('type.semanticPrimitives.error_result'(X183),'const.semanticPrimitives.Rraise_1'(s(X183,X185)))))))))) )),
  inference(forward_demodulation,[],[f7762,f596])).
fof(f7762,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),X186))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),'const.evalProps.map_error_result_2'(s(fun(X181,X180),X182),s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun(X183,X181),X184),s('type.semanticPrimitives.error_result'(X183),'const.semanticPrimitives.Rraise_1'(s(X183,X185))))))))) | s('type.semanticPrimitives.error_result'(X180),X186) = s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rraise_1'(s(X180,sK5('const.semanticPrimitives.Rraise_1'(s(X180,happ(s(fun(X181,X180),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185)))))),X180))))) )),
  inference(subsumption_resolution,[],[f7694,f102])).
fof(f7694,plain,(
  ( ! [X185,X182,X180,X186,X184,X183,X181] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),X186))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X180)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X180),'const.evalProps.map_error_result_2'(s(fun(X181,X180),X182),s('type.semanticPrimitives.error_result'(X181),'const.evalProps.map_error_result_2'(s(fun(X183,X181),X184),s('type.semanticPrimitives.error_result'(X183),'const.semanticPrimitives.Rraise_1'(s(X183,X185))))))))) | s('type.semanticPrimitives.error_result'(X180),X186) = s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rraise_1'(s(X180,sK5('const.semanticPrimitives.Rraise_1'(s(X180,happ(s(fun(X181,X180),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185)))))),X180)))) | s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rraise_1'(s(X180,happ(s(fun(X181,X180),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))))) = s('type.semanticPrimitives.error_result'(X180),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X180,happ(s(fun(X181,X180),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185)))))),X180))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7693,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X179] : (s('type.semanticPrimitives.error_result'(X173),X179) != s('type.semanticPrimitives.error_result'(X173),'const.evalProps.map_error_result_2'(s(fun(X174,X173),X175),s('type.semanticPrimitives.error_result'(X174),'const.evalProps.map_error_result_2'(s(fun(X176,X174),X177),s('type.semanticPrimitives.error_result'(X176),'const.semanticPrimitives.Rraise_1'(s(X176,X178))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X173)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X173),'const.semanticPrimitives.Rraise_1'(s(X173,sK5(X179,X173)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X173)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X173),'const.semanticPrimitives.Rraise_1'(s(X173,happ(s(fun(X174,X173),X175),s(X174,happ(s(fun(X176,X174),X177),s(X176,X178))))))))) | s('type.semanticPrimitives.error_result'(X173),X179) = s('type.semanticPrimitives.error_result'(X173),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X179,X173))))) )),
  inference(superposition,[],[f540,f573])).
fof(f7761,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),X172))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun(X167,X166),X168),s(X167,happ(s(fun(X169,X167),X170),s(X169,X171))))))))) | s('type.semanticPrimitives.error_result'(X166),X172) != s('type.semanticPrimitives.error_result'(X166),'const.evalProps.map_error_result_2'(s(fun(X167,X166),X168),s('type.semanticPrimitives.error_result'(X167),'const.evalProps.map_error_result_2'(s(fun(X169,X167),X170),s('type.semanticPrimitives.error_result'(X169),'const.semanticPrimitives.Rraise_1'(s(X169,X171)))))))) )),
  inference(forward_demodulation,[],[f7760,f596])).
fof(f7760,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s('type.semanticPrimitives.error_result'(X166),X172) != s('type.semanticPrimitives.error_result'(X166),'const.evalProps.map_error_result_2'(s(fun(X167,X166),X168),s('type.semanticPrimitives.error_result'(X167),'const.evalProps.map_error_result_2'(s(fun(X169,X167),X170),s('type.semanticPrimitives.error_result'(X169),'const.semanticPrimitives.Rraise_1'(s(X169,X171))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),X172))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rraise_1'(s(X166,sK5('const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun(X167,X166),X168),s(X167,happ(s(fun(X169,X167),X170),s(X169,X171)))))),X166))))))) )),
  inference(subsumption_resolution,[],[f7692,f102])).
fof(f7692,plain,(
  ( ! [X167,X171,X169,X166,X172,X170,X168] : (s('type.semanticPrimitives.error_result'(X166),X172) != s('type.semanticPrimitives.error_result'(X166),'const.evalProps.map_error_result_2'(s(fun(X167,X166),X168),s('type.semanticPrimitives.error_result'(X167),'const.evalProps.map_error_result_2'(s(fun(X169,X167),X170),s('type.semanticPrimitives.error_result'(X169),'const.semanticPrimitives.Rraise_1'(s(X169,X171))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),X172))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X166)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rraise_1'(s(X166,sK5('const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun(X167,X166),X168),s(X167,happ(s(fun(X169,X167),X170),s(X169,X171)))))),X166)))))) | s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun(X167,X166),X168),s(X167,happ(s(fun(X169,X167),X170),s(X169,X171))))))) = s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X166,happ(s(fun(X167,X166),X168),s(X167,happ(s(fun(X169,X167),X170),s(X169,X171)))))),X166))))) )),
  inference(superposition,[],[f540,f573])).
fof(f7759,plain,(
  ( ! [X158,X156,X154,X152,X157,X155,X153] : (s('type.semanticPrimitives.error_result'(X152),'const.semanticPrimitives.Rraise_1'(s(X152,X158))) != s('type.semanticPrimitives.error_result'(X152),'const.evalProps.map_error_result_2'(s(fun(X153,X152),X154),s('type.semanticPrimitives.error_result'(X153),'const.evalProps.map_error_result_2'(s(fun(X155,X153),X156),s('type.semanticPrimitives.error_result'(X155),'const.semanticPrimitives.Rraise_1'(s(X155,X157))))))) | s(X152,X158) = s(X152,sK5('const.semanticPrimitives.Rraise_1'(s(X152,happ(s(fun(X153,X152),X154),s(X153,happ(s(fun(X155,X153),X156),s(X155,X157)))))),X152))) )),
  inference(subsumption_resolution,[],[f7690,f102])).
fof(f7690,plain,(
  ( ! [X158,X156,X154,X152,X157,X155,X153] : (s('type.semanticPrimitives.error_result'(X152),'const.semanticPrimitives.Rraise_1'(s(X152,X158))) != s('type.semanticPrimitives.error_result'(X152),'const.evalProps.map_error_result_2'(s(fun(X153,X152),X154),s('type.semanticPrimitives.error_result'(X153),'const.evalProps.map_error_result_2'(s(fun(X155,X153),X156),s('type.semanticPrimitives.error_result'(X155),'const.semanticPrimitives.Rraise_1'(s(X155,X157))))))) | s(X152,X158) = s(X152,sK5('const.semanticPrimitives.Rraise_1'(s(X152,happ(s(fun(X153,X152),X154),s(X153,happ(s(fun(X155,X153),X156),s(X155,X157)))))),X152)) | s('type.semanticPrimitives.error_result'(X152),'const.semanticPrimitives.Rraise_1'(s(X152,happ(s(fun(X153,X152),X154),s(X153,happ(s(fun(X155,X153),X156),s(X155,X157))))))) = s('type.semanticPrimitives.error_result'(X152),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X152,happ(s(fun(X153,X152),X154),s(X153,happ(s(fun(X155,X153),X156),s(X155,X157)))))),X152))))) )),
  inference(superposition,[],[f537,f573])).
fof(f7688,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122,X120,X119] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X119)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X119),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X125))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X119)),'const.evalProps.map_error_result_2'(s(fun(X120,'type.semanticPrimitives.error_result'(X119)),X121),s('type.semanticPrimitives.error_result'(X120),'const.evalProps.map_error_result_2'(s(fun(X122,X120),X123),s('type.semanticPrimitives.error_result'(X122),'const.semanticPrimitives.Rraise_1'(s(X122,X124))))))) | s('type.semanticPrimitives.error_result'(X119),'const.evalProps.map_error_result_2'(s(fun(X126,X119),X127),s('type.semanticPrimitives.error_result'(X126),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X125))))) = s('type.semanticPrimitives.error_result'(X119),happ(s(fun(X120,'type.semanticPrimitives.error_result'(X119)),X121),s(X120,happ(s(fun(X122,X120),X123),s(X122,X124)))))) )),
  inference(superposition,[],[f504,f573])).
fof(f7687,plain,(
  ( ! [X118,X116,X114,X112,X117,X115,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X112)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X112),X118))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X112)),'const.evalProps.map_error_result_2'(s(fun(X113,'type.semanticPrimitives.error_result'(X112)),X114),s('type.semanticPrimitives.error_result'(X113),'const.evalProps.map_error_result_2'(s(fun(X115,X113),X116),s('type.semanticPrimitives.error_result'(X115),'const.semanticPrimitives.Rraise_1'(s(X115,X117))))))) | s('type.semanticPrimitives.error_result'(X112),X118) = s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rraise_1'(s(X112,sK5(happ(s(fun(X113,'type.semanticPrimitives.error_result'(X112)),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))),X112)))) | s('type.semanticPrimitives.error_result'(X112),happ(s(fun(X113,'type.semanticPrimitives.error_result'(X112)),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117))))) = s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(happ(s(fun(X113,'type.semanticPrimitives.error_result'(X112)),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))),X112))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7686,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X105)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X105),X111))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X105)),'const.evalProps.map_error_result_2'(s(fun(X106,'type.semanticPrimitives.error_result'(X105)),X107),s('type.semanticPrimitives.error_result'(X106),'const.evalProps.map_error_result_2'(s(fun(X108,X106),X109),s('type.semanticPrimitives.error_result'(X108),'const.semanticPrimitives.Rraise_1'(s(X108,X110))))))) | s('type.semanticPrimitives.error_result'(X105),'const.semanticPrimitives.Rraise_1'(s(X105,sK5(X111,X105)))) = s('type.semanticPrimitives.error_result'(X105),happ(s(fun(X106,'type.semanticPrimitives.error_result'(X105)),X107),s(X106,happ(s(fun(X108,X106),X109),s(X108,X110))))) | s('type.semanticPrimitives.error_result'(X105),X111) = s('type.semanticPrimitives.error_result'(X105),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X111,X105))))) )),
  inference(superposition,[],[f542,f573])).
fof(f7685,plain,(
  ( ! [X88,X87,X85,X83,X81,X89,X86,X84,X82] : (s('type.semanticPrimitives.error_result'(X87),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X88))) != s('type.semanticPrimitives.error_result'(X87),'const.evalProps.map_error_result_2'(s(fun(X81,X87),X89),s('type.semanticPrimitives.error_result'(X81),'const.evalProps.map_error_result_2'(s(fun(X82,X81),X83),s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(X84,X82),X85),s('type.semanticPrimitives.error_result'(X84),'const.semanticPrimitives.Rraise_1'(s(X84,X86)))))))))) )),
  inference(superposition,[],[f576,f573])).
fof(f7684,plain,(
  ( ! [X80,X78,X76,X74,X72,X79,X77,X75,X73] : (s('type.semanticPrimitives.error_result'(X72),'const.evalProps.map_error_result_2'(s(fun(X78,X72),X79),s('type.semanticPrimitives.error_result'(X78),'const.semanticPrimitives.Rraise_1'(s(X78,X80))))) != s('type.semanticPrimitives.error_result'(X72),'const.evalProps.map_error_result_2'(s(fun(X73,X72),X74),s('type.semanticPrimitives.error_result'(X73),'const.evalProps.map_error_result_2'(s(fun(X75,X73),X76),s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,X77))))))) | s(X72,happ(s(fun(X78,X72),X79),s(X78,X80))) = s(X72,happ(s(fun(X73,X72),X74),s(X73,happ(s(fun(X75,X73),X76),s(X75,X77)))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7683,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65,X63] : (s('type.semanticPrimitives.error_result'(X69),'const.semanticPrimitives.Rraise_1'(s(X69,X70))) != s('type.semanticPrimitives.error_result'(X69),'const.evalProps.map_error_result_2'(s(fun(X63,X69),X71),s('type.semanticPrimitives.error_result'(X63),'const.evalProps.map_error_result_2'(s(fun(X64,X63),X65),s('type.semanticPrimitives.error_result'(X64),'const.evalProps.map_error_result_2'(s(fun(X66,X64),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68))))))))) | s(X69,X70) = s(X69,happ(s(fun(X63,X69),X71),s(X63,happ(s(fun(X64,X63),X65),s(X64,happ(s(fun(X66,X64),X67),s(X66,X68)))))))) )),
  inference(superposition,[],[f574,f573])).
fof(f7682,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,happ(s(fun(X55,X61),X62),s(X55,happ(s(fun(X56,X55),X57),s(X56,happ(s(fun(X58,X56),X59),s(X58,X60))))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X61),'const.evalProps.map_error_result_2'(s(fun(X55,X61),X62),s('type.semanticPrimitives.error_result'(X55),'const.evalProps.map_error_result_2'(s(fun(X56,X55),X57),s('type.semanticPrimitives.error_result'(X56),'const.evalProps.map_error_result_2'(s(fun(X58,X56),X59),s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,X60)))))))))))) )),
  inference(superposition,[],[f558,f573])).
fof(f7681,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s('type.semanticPrimitives.error_result'(X48),X54) != s('type.semanticPrimitives.error_result'(X48),'const.evalProps.map_error_result_2'(s(fun(X49,X48),X50),s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X51,X49),X52),s('type.semanticPrimitives.error_result'(X51),'const.semanticPrimitives.Rraise_1'(s(X51,X53))))))) | s(X48,sK5(X54,X48)) = s(X48,happ(s(fun(X49,X48),X50),s(X49,happ(s(fun(X51,X49),X52),s(X51,X53))))) | s('type.semanticPrimitives.error_result'(X48),X54) = s('type.semanticPrimitives.error_result'(X48),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X54,X48))))) )),
  inference(superposition,[],[f537,f573])).
fof(f7680,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X41)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X47))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X41)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X41),'const.evalProps.map_error_result_2'(s(fun(X42,X41),X43),s('type.semanticPrimitives.error_result'(X42),'const.evalProps.map_error_result_2'(s(fun(X44,X42),X45),s('type.semanticPrimitives.error_result'(X44),'const.semanticPrimitives.Rraise_1'(s(X44,X46)))))))))) )),
  inference(superposition,[],[f433,f573])).
fof(f7674,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X3,X1),X4),s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,X5)))))))))) )),
  inference(superposition,[],[f558,f573])).
fof(f7672,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s('type.semanticPrimitives.error_result'(X44),'const.semanticPrimitives.Rraise_1'(s(X44,happ(s(fun(X45,X44),X46),s(X45,happ(s(fun('type.semanticPrimitives.error_result'(X40),X45),X47),s('type.semanticPrimitives.error_result'(X40),'const.evalProps.map_error_result_2'(s(fun(X41,X40),X42),s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rraise_1'(s(X41,X43))))))))))) = s('type.semanticPrimitives.error_result'(X44),'const.evalProps.map_error_result_2'(s(fun(X45,X44),X46),s('type.semanticPrimitives.error_result'(X45),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X40),X45),X47),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X40)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X40),'const.semanticPrimitives.Rraise_1'(s(X40,happ(s(fun(X41,X40),X42),s(X41,X43)))))))))))) )),
  inference(superposition,[],[f573,f558])).
fof(f7671,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,happ(s(fun(X37,X36),X38),s(X37,happ(s(fun('type.semanticPrimitives.error_result'(X32),X37),X39),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,happ(s(fun(X33,X32),X34),s(X33,X35))))))))))) = s('type.semanticPrimitives.error_result'(X36),'const.evalProps.map_error_result_2'(s(fun(X37,X36),X38),s('type.semanticPrimitives.error_result'(X37),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X32),X37),X39),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X32)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X32),'const.evalProps.map_error_result_2'(s(fun(X33,X32),X34),s('type.semanticPrimitives.error_result'(X33),'const.semanticPrimitives.Rraise_1'(s(X33,X35)))))))))))) )),
  inference(superposition,[],[f573,f558])).
fof(f7670,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X22] : (s('type.semanticPrimitives.error_result'(X28),'const.semanticPrimitives.Rraise_1'(s(X28,happ(s(fun(X29,X28),X30),s(X29,happ(s(fun(X22,X29),X31),s(X22,happ(s(fun(X23,X22),X24),s(X23,happ(s(fun(X25,X23),X26),s(X25,X27))))))))))) = s('type.semanticPrimitives.error_result'(X28),'const.evalProps.map_error_result_2'(s(fun(X29,X28),X30),s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X22,X29),X31),s('type.semanticPrimitives.error_result'(X22),'const.evalProps.map_error_result_2'(s(fun(X23,X22),X24),s('type.semanticPrimitives.error_result'(X23),'const.evalProps.map_error_result_2'(s(fun(X25,X23),X26),s('type.semanticPrimitives.error_result'(X25),'const.semanticPrimitives.Rraise_1'(s(X25,X27)))))))))))) )),
  inference(superposition,[],[f573,f573])).
fof(f7669,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X18),'const.evalProps.map_error_result_2'(s(fun(X19,X18),X20),s('type.semanticPrimitives.error_result'(X19),'const.evalProps.map_error_result_2'(s(fun(X14,X19),X21),s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun(X15,X14),X16),s('type.semanticPrimitives.error_result'(X15),X17))))))) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rraise_1'(s(X18,happ(s(fun(X19,X18),X20),s(X19,happ(s(fun(X14,X19),X21),s(X14,happ(s(fun(X15,X14),X16),s(X15,sK5(X17,X15)))))))))) | s('type.semanticPrimitives.error_result'(X15),X17) = s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X17,X15))))) )),
  inference(superposition,[],[f573,f572])).
fof(f7668,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(X11,X10),X12),s(X11,happ(s(fun(X6,X11),X13),s(X6,happ(s(fun(X7,X6),X8),s(X7,X9))))))))) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.evalProps.map_error_result_2'(s(fun(X6,X11),X13),s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,X9)))))))))) )),
  inference(superposition,[],[f573,f108])).
fof(f7619,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rraise_1'(s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(X9,X7),X10),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11)))))))) )),
  inference(superposition,[],[f573,f108])).
fof(f7746,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.semanticPrimitives.error_result'(X38),'const.semanticPrimitives.Rraise_1'(s(X38,happ(s(fun(bool,X38),X39),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X38),'const.evalProps.map_error_result_2'(s(fun(bool,X38),X39),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))))))) | s(bool,X40) = s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X40)),bool))) )),
  inference(forward_demodulation,[],[f7612,f108])).
fof(f7612,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.semanticPrimitives.error_result'(X38),'const.semanticPrimitives.Rraise_1'(s(X38,happ(s(fun(bool,X38),X39),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X38),'const.evalProps.map_error_result_2'(s(fun(bool,X38),X39),s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X35,bool),X36),s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,X37))))))) | s(bool,X40) = s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X40)),bool))) )),
  inference(superposition,[],[f573,f689])).
fof(f7409,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X18),X22))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.evalProps.map_error_result_2'(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s('type.semanticPrimitives.error_result'(X19),X21))) | s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rraise_1'(s(X18,sK5(X22,X18)))) = s('type.semanticPrimitives.error_result'(X18),happ(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s(X19,sK5(X21,X19)))) | s('type.semanticPrimitives.error_result'(X18),X22) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X22,X18)))) | s('type.semanticPrimitives.error_result'(X19),X21) = s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X21,X19))))) )),
  inference(superposition,[],[f542,f572])).
fof(f7404,plain,(
  ( ! [X39,X37,X38] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X39))) | s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,sK5(X39,X37)))) = s('type.semanticPrimitives.error_result'(X37),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37))) | s('type.semanticPrimitives.error_result'(X37),X39) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X39,X37))))) )),
  inference(superposition,[],[f542,f596])).
fof(f7402,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),X35))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,happ(s(fun(X32,X31),X33),s(X32,X34))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),sK5(X35,'type.semanticPrimitives.error_result'(X31))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),'const.evalProps.map_error_result_2'(s(fun(X32,X31),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),X35) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X35,'type.semanticPrimitives.error_result'(X31)))))) )),
  inference(superposition,[],[f542,f558])).
fof(f7401,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),X30))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.evalProps.map_error_result_2'(s(fun(X27,X26),X28),s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rraise_1'(s(X27,X29))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),sK5(X30,'type.semanticPrimitives.error_result'(X26))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.semanticPrimitives.Rraise_1'(s(X26,happ(s(fun(X27,X26),X28),s(X27,X29))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),X30) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X30,'type.semanticPrimitives.error_result'(X26)))))) )),
  inference(superposition,[],[f542,f558])).
fof(f7400,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),X25))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),'const.evalProps.map_error_result_2'(s(fun(X22,X21),X23),s('type.semanticPrimitives.error_result'(X22),X24))))) | s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,sK5(X25,X21)))) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,happ(s(fun(X22,X21),X23),s(X22,sK5(X24,X22)))))) | s('type.semanticPrimitives.error_result'(X21),X25) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X25,X21)))) | s('type.semanticPrimitives.error_result'(X22),X24) = s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X24,X22))))) )),
  inference(superposition,[],[f542,f572])).
fof(f7399,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),X20))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))))) | s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,sK5(X20,X16)))) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19))))) | s('type.semanticPrimitives.error_result'(X16),X20) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X20,X16))))) )),
  inference(superposition,[],[f542,f108])).
fof(f7398,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),X14))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),X15))) | s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13)))) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(X15,X13)))) | s('type.semanticPrimitives.error_result'(X13),X15) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X15,X13)))) | s('type.semanticPrimitives.error_result'(X13),X14) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X14,X13))))) )),
  inference(superposition,[],[f542,f95])).
fof(f7432,plain,(
  ( ! [X12,X11] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X11))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X12))) | s('type.semanticPrimitives.error_result'(bool),X12) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X12,bool)))) | s('type.semanticPrimitives.error_result'(bool),X11) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X11,bool))) )),
  inference(subsumption_resolution,[],[f7397,f102])).
fof(f7397,plain,(
  ( ! [X12,X11] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X11))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X12))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X11,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5(X12,bool)))) | s('type.semanticPrimitives.error_result'(bool),X12) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X12,bool)))) | s('type.semanticPrimitives.error_result'(bool),X11) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X11,bool))) )),
  inference(superposition,[],[f542,f536])).
fof(f7395,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),X9))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))))))) | s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,sK5(X9,X5)))) = s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))) | s('type.semanticPrimitives.error_result'(X5),X9) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X9,X5))))) )),
  inference(superposition,[],[f542,f108])).
fof(f7393,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X18),X22))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.evalProps.map_error_result_2'(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s('type.semanticPrimitives.error_result'(X19),X21))) | s('type.semanticPrimitives.error_result'(X18),X22) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rraise_1'(s(X18,sK5(happ(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s(X19,sK5(X21,X19))),X18)))) | s('type.semanticPrimitives.error_result'(X18),happ(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s(X19,sK5(X21,X19)))) = s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(happ(s(fun(X19,'type.semanticPrimitives.error_result'(X18)),X20),s(X19,sK5(X21,X19))),X18)))) | s('type.semanticPrimitives.error_result'(X19),X21) = s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X21,X19))))) )),
  inference(superposition,[],[f542,f572])).
fof(f7392,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.evalProps.map_error_result_2'(s(fun(X14,'type.semanticPrimitives.error_result'(X13)),X15),s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,X16))))) | s('type.semanticPrimitives.error_result'(X13),X17) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(happ(s(fun(X14,'type.semanticPrimitives.error_result'(X13)),X15),s(X14,X16)),X13)))) | s('type.semanticPrimitives.error_result'(X13),happ(s(fun(X14,'type.semanticPrimitives.error_result'(X13)),X15),s(X14,X16))) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(happ(s(fun(X14,'type.semanticPrimitives.error_result'(X13)),X15),s(X14,X16)),X13))))) )),
  inference(superposition,[],[f542,f108])).
fof(f7391,plain,(
  ( ! [X12,X10,X11] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),X11) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X12))) | s('type.semanticPrimitives.error_result'(X10),X12) = s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,sK5(sK5(X11,'type.semanticPrimitives.error_result'(X10)),X10)))) | s('type.semanticPrimitives.error_result'(X10),sK5(X11,'type.semanticPrimitives.error_result'(X10))) = s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5(X11,'type.semanticPrimitives.error_result'(X10)),X10)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),X11) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X11,'type.semanticPrimitives.error_result'(X10)))))) )),
  inference(superposition,[],[f542,f95])).
fof(f7427,plain,(
  ( ! [X39,X37,X38] : (s('type.semanticPrimitives.error_result'(X37),X38) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X39))) | s('type.semanticPrimitives.error_result'(X37),X39) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37))))) )),
  inference(forward_demodulation,[],[f7388,f596])).
fof(f7388,plain,(
  ( ! [X39,X37,X38] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X39))) | s('type.semanticPrimitives.error_result'(X37),X39) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37)))) | s('type.semanticPrimitives.error_result'(X37),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37))) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37))))) )),
  inference(superposition,[],[f542,f596])).
fof(f6745,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),X1))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5(X1,X0)))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(equality_resolution,[],[f540])).
fof(f6742,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),X35) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,happ(s(fun(X32,X31),X33),s(X32,X34))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),sK5(X35,'type.semanticPrimitives.error_result'(X31))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),'const.evalProps.map_error_result_2'(s(fun(X32,X31),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),X35) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X35,'type.semanticPrimitives.error_result'(X31)))))) )),
  inference(superposition,[],[f540,f558])).
fof(f6741,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),X30) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.evalProps.map_error_result_2'(s(fun(X27,X26),X28),s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rraise_1'(s(X27,X29))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),sK5(X30,'type.semanticPrimitives.error_result'(X26))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X26),'const.semanticPrimitives.Rraise_1'(s(X26,happ(s(fun(X27,X26),X28),s(X27,X29))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),X30) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X26)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X30,'type.semanticPrimitives.error_result'(X26)))))) )),
  inference(superposition,[],[f540,f558])).
fof(f6740,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.semanticPrimitives.error_result'(X21),X25) != s('type.semanticPrimitives.error_result'(X21),'const.evalProps.map_error_result_2'(s(fun(X22,X21),X23),s('type.semanticPrimitives.error_result'(X22),X24))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,sK5(X25,X21)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,happ(s(fun(X22,X21),X23),s(X22,sK5(X24,X22)))))))) | s('type.semanticPrimitives.error_result'(X21),X25) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X25,X21)))) | s('type.semanticPrimitives.error_result'(X22),X24) = s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X24,X22))))) )),
  inference(superposition,[],[f540,f572])).
fof(f6739,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),X20) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,sK5(X20,X16)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19))))))) | s('type.semanticPrimitives.error_result'(X16),X20) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X20,X16))))) )),
  inference(superposition,[],[f540,f108])).
fof(f6738,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'(X13),X14) != s('type.semanticPrimitives.error_result'(X13),X15) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(X15,X13)))))) | s('type.semanticPrimitives.error_result'(X13),X15) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X15,X13)))) | s('type.semanticPrimitives.error_result'(X13),X14) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X14,X13))))) )),
  inference(superposition,[],[f540,f95])).
fof(f6735,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'(X5),X9) != s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,sK5(X9,X5)))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X5)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))))) | s('type.semanticPrimitives.error_result'(X5),X9) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X9,X5))))) )),
  inference(superposition,[],[f540,f108])).
fof(f6769,plain,(
  ( ! [X39,X37,X38] : (s('type.semanticPrimitives.error_result'(X37),X38) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37)))) | s('type.semanticPrimitives.error_result'(X37),X38) != s('type.semanticPrimitives.error_result'(X37),X39) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X39))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37))))))) )),
  inference(forward_demodulation,[],[f6733,f596])).
fof(f6733,plain,(
  ( ! [X39,X37,X38] : (s('type.semanticPrimitives.error_result'(X37),X38) != s('type.semanticPrimitives.error_result'(X37),X39) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X39))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X37)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rraise_1'(s(X37,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37)))))) | s('type.semanticPrimitives.error_result'(X37),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37))) = s('type.semanticPrimitives.error_result'(X37),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X37),X38)),'type.semanticPrimitives.error_result'(X37)),X37))))) )),
  inference(superposition,[],[f540,f596])).
fof(f6589,plain,(
  ( ! [X167,X165,X171,X169,X166,X170,X168] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.evalProps.map_error_result_2'(s(fun(X169,'type.semanticPrimitives.error_result'(X165)),X170),s('type.semanticPrimitives.error_result'(X169),'const.semanticPrimitives.Rraise_1'(s(X169,X171))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X165)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X165),'const.evalProps.map_error_result_2'(s(fun(X166,X165),X167),s('type.semanticPrimitives.error_result'(X166),X168))))) | s('type.semanticPrimitives.error_result'(X165),happ(s(fun(X169,'type.semanticPrimitives.error_result'(X165)),X170),s(X169,X171))) = s('type.semanticPrimitives.error_result'(X165),'const.semanticPrimitives.Rraise_1'(s(X165,happ(s(fun(X166,X165),X167),s(X166,sK5(X168,X166)))))) | s('type.semanticPrimitives.error_result'(X166),X168) = s('type.semanticPrimitives.error_result'(X166),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X168,X166))))) )),
  inference(superposition,[],[f574,f572])).
fof(f6588,plain,(
  ( ! [X158,X163,X161,X159,X164,X162,X160] : (s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,X163))) != s('type.semanticPrimitives.error_result'(X162),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X158),X162),X164),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X158)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X158),'const.evalProps.map_error_result_2'(s(fun(X159,X158),X160),s('type.semanticPrimitives.error_result'(X159),X161))))))) | s(X162,X163) = s(X162,happ(s(fun('type.semanticPrimitives.error_result'(X158),X162),X164),s('type.semanticPrimitives.error_result'(X158),'const.semanticPrimitives.Rraise_1'(s(X158,happ(s(fun(X159,X158),X160),s(X159,sK5(X161,X159)))))))) | s('type.semanticPrimitives.error_result'(X159),X161) = s('type.semanticPrimitives.error_result'(X159),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X161,X159))))) )),
  inference(superposition,[],[f574,f572])).
fof(f6586,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),X151) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X147),'const.evalProps.map_error_result_2'(s(fun(X148,X147),X149),s('type.semanticPrimitives.error_result'(X148),X150))))) | s('type.semanticPrimitives.error_result'(X147),sK5(X151,'type.semanticPrimitives.error_result'(X147))) = s('type.semanticPrimitives.error_result'(X147),'const.semanticPrimitives.Rraise_1'(s(X147,happ(s(fun(X148,X147),X149),s(X148,sK5(X150,X148)))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),X151) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X147)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X151,'type.semanticPrimitives.error_result'(X147))))) | s('type.semanticPrimitives.error_result'(X148),X150) = s('type.semanticPrimitives.error_result'(X148),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X150,X148))))) )),
  inference(superposition,[],[f537,f572])).
fof(f6581,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s('type.semanticPrimitives.error_result'(X121),X125) != s('type.semanticPrimitives.error_result'(X121),'const.evalProps.map_error_result_2'(s(fun(X122,X121),X123),s('type.semanticPrimitives.error_result'(X122),X124))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X121)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X121),X125))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X121)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X121),'const.semanticPrimitives.Rraise_1'(s(X121,happ(s(fun(X122,X121),X123),s(X122,sK5(X124,X122)))))))) | s('type.semanticPrimitives.error_result'(X122),X124) = s('type.semanticPrimitives.error_result'(X122),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X124,X122))))) )),
  inference(superposition,[],[f101,f572])).
fof(f6579,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),X115))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.evalProps.map_error_result_2'(s(fun(X112,X111),X113),s('type.semanticPrimitives.error_result'(X112),X114))))) | s('type.semanticPrimitives.error_result'(X111),X115) = s('type.semanticPrimitives.error_result'(X111),'const.semanticPrimitives.Rraise_1'(s(X111,happ(s(fun(X112,X111),X113),s(X112,sK5(X114,X112)))))) | s('type.semanticPrimitives.error_result'(X112),X114) = s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X114,X112))))) )),
  inference(superposition,[],[f100,f572])).
fof(f6606,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,X105))) != s('type.semanticPrimitives.error_result'(X101),'const.evalProps.map_error_result_2'(s(fun(X102,X101),X103),s('type.semanticPrimitives.error_result'(X102),X104))) | s(X101,X105) = s(X101,sK5('const.semanticPrimitives.Rraise_1'(s(X101,happ(s(fun(X102,X101),X103),s(X102,sK5(X104,X102))))),X101)) | s('type.semanticPrimitives.error_result'(X102),X104) = s('type.semanticPrimitives.error_result'(X102),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X104,X102))))) )),
  inference(subsumption_resolution,[],[f6577,f102])).
fof(f6577,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,X105))) != s('type.semanticPrimitives.error_result'(X101),'const.evalProps.map_error_result_2'(s(fun(X102,X101),X103),s('type.semanticPrimitives.error_result'(X102),X104))) | s(X101,X105) = s(X101,sK5('const.semanticPrimitives.Rraise_1'(s(X101,happ(s(fun(X102,X101),X103),s(X102,sK5(X104,X102))))),X101)) | s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,happ(s(fun(X102,X101),X103),s(X102,sK5(X104,X102)))))) = s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X101,happ(s(fun(X102,X101),X103),s(X102,sK5(X104,X102))))),X101)))) | s('type.semanticPrimitives.error_result'(X102),X104) = s('type.semanticPrimitives.error_result'(X102),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X104,X102))))) )),
  inference(superposition,[],[f537,f572])).
fof(f6572,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X78)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X78),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X82))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X78)),'const.evalProps.map_error_result_2'(s(fun(X79,'type.semanticPrimitives.error_result'(X78)),X80),s('type.semanticPrimitives.error_result'(X79),X81))) | s('type.semanticPrimitives.error_result'(X78),happ(s(fun(X79,'type.semanticPrimitives.error_result'(X78)),X80),s(X79,sK5(X81,X79)))) = s('type.semanticPrimitives.error_result'(X78),'const.evalProps.map_error_result_2'(s(fun(X83,X78),X84),s('type.semanticPrimitives.error_result'(X83),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X82))))) | s('type.semanticPrimitives.error_result'(X79),X81) = s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X81,X79))))) )),
  inference(superposition,[],[f504,f572])).
fof(f6571,plain,(
  ( ! [X76,X74,X77,X75] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X77))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X74,bool),X75),s('type.semanticPrimitives.error_result'(X74),X76))) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,happ(s(fun(X74,bool),X75),s(X74,sK5(X76,X74)))) | s('type.semanticPrimitives.error_result'(X74),X76) = s('type.semanticPrimitives.error_result'(X74),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X76,X74))))) )),
  inference(superposition,[],[f436,f572])).
fof(f6570,plain,(
  ( ! [X70,X72,X71,X73] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X73))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X70,bool),X71),s('type.semanticPrimitives.error_result'(X70),X72))) | s(bool,'F_0') = s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X72,X70)))) | s(bool,'T_0') = s(bool,X73) | s('type.semanticPrimitives.error_result'(X70),X72) = s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X72,X70))))) )),
  inference(superposition,[],[f436,f572])).
fof(f6569,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X66,bool),X67),s('type.semanticPrimitives.error_result'(X66),X68))) | s(bool,X69) = s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X68,X66)))) | s(bool,'T_0') = s(bool,X69) | s('type.semanticPrimitives.error_result'(X66),X68) = s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X68,X66))))) )),
  inference(superposition,[],[f441,f572])).
fof(f6568,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X64))) != s('type.semanticPrimitives.error_result'(X63),'const.evalProps.map_error_result_2'(s(fun(X59,X63),X65),s('type.semanticPrimitives.error_result'(X59),'const.evalProps.map_error_result_2'(s(fun(X60,X59),X61),s('type.semanticPrimitives.error_result'(X60),X62))))) | s('type.semanticPrimitives.error_result'(X60),X62) = s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X62,X60))))) )),
  inference(superposition,[],[f576,f572])).
fof(f6567,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (s('type.semanticPrimitives.error_result'(X52),'const.evalProps.map_error_result_2'(s(fun(X53,X52),X54),s('type.semanticPrimitives.error_result'(X53),X55))) != s('type.semanticPrimitives.error_result'(X52),'const.evalProps.map_error_result_2'(s(fun(X56,X52),X57),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X58))))) | s(X52,happ(s(fun(X56,X52),X57),s(X56,X58))) = s(X52,happ(s(fun(X53,X52),X54),s(X53,sK5(X55,X53)))) | s('type.semanticPrimitives.error_result'(X53),X55) = s('type.semanticPrimitives.error_result'(X53),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X55,X53))))) )),
  inference(superposition,[],[f574,f572])).
fof(f6566,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s('type.semanticPrimitives.error_result'(X49),'const.semanticPrimitives.Rraise_1'(s(X49,X50))) != s('type.semanticPrimitives.error_result'(X49),'const.evalProps.map_error_result_2'(s(fun(X45,X49),X51),s('type.semanticPrimitives.error_result'(X45),'const.evalProps.map_error_result_2'(s(fun(X46,X45),X47),s('type.semanticPrimitives.error_result'(X46),X48))))) | s(X49,X50) = s(X49,happ(s(fun(X45,X49),X51),s(X45,happ(s(fun(X46,X45),X47),s(X46,sK5(X48,X46)))))) | s('type.semanticPrimitives.error_result'(X46),X48) = s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X48,X46))))) )),
  inference(superposition,[],[f574,f572])).
fof(f6565,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X43)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X43),'const.evalProps.map_error_result_2'(s(fun(X39,X43),X44),s('type.semanticPrimitives.error_result'(X39),'const.evalProps.map_error_result_2'(s(fun(X40,X39),X41),s('type.semanticPrimitives.error_result'(X40),X42))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X43)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X43),'const.semanticPrimitives.Rraise_1'(s(X43,happ(s(fun(X39,X43),X44),s(X39,happ(s(fun(X40,X39),X41),s(X40,sK5(X42,X40)))))))))) | s('type.semanticPrimitives.error_result'(X40),X42) = s('type.semanticPrimitives.error_result'(X40),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X42,X40))))) )),
  inference(superposition,[],[f558,f572])).
fof(f6564,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'(X34),X38) != s('type.semanticPrimitives.error_result'(X34),'const.evalProps.map_error_result_2'(s(fun(X35,X34),X36),s('type.semanticPrimitives.error_result'(X35),X37))) | s(X34,sK5(X38,X34)) = s(X34,happ(s(fun(X35,X34),X36),s(X35,sK5(X37,X35)))) | s('type.semanticPrimitives.error_result'(X34),X38) = s('type.semanticPrimitives.error_result'(X34),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X38,X34)))) | s('type.semanticPrimitives.error_result'(X35),X37) = s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X37,X35))))) )),
  inference(superposition,[],[f537,f572])).
fof(f6563,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X29)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X29),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X33))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X29)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X29),'const.evalProps.map_error_result_2'(s(fun(X30,X29),X31),s('type.semanticPrimitives.error_result'(X30),X32))))) | s('type.semanticPrimitives.error_result'(X30),X32) = s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X32,X30))))) )),
  inference(superposition,[],[f433,f572])).
fof(f6562,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.semanticPrimitives.error_result'(X27),'const.evalProps.map_error_result_2'(s(fun(X23,X27),X28),s('type.semanticPrimitives.error_result'(X23),'const.evalProps.map_error_result_2'(s(fun(X24,X23),X25),s('type.semanticPrimitives.error_result'(X24),X26))))) = s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rraise_1'(s(X27,happ(s(fun(X23,X27),X28),s(X23,happ(s(fun(X24,X23),X25),s(X24,sK5(X26,X24)))))))) | s('type.semanticPrimitives.error_result'(X24),X26) = s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X26,X24))))) )),
  inference(superposition,[],[f108,f572])).
fof(f6557,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,sK5(X3,X1)))))))) | s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X3,X1))))) )),
  inference(superposition,[],[f558,f572])).
fof(f6556,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(X4),'const.evalProps.map_error_result_2'(s(fun(X5,X4),X6),s('type.semanticPrimitives.error_result'(X5),X7))) = s('type.semanticPrimitives.error_result'(X4),'const.evalProps.map_error_result_2'(s(fun(X5,X4),X6),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,sK5(X7,X5)))))) | s('type.semanticPrimitives.error_result'(X5),X7) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X7,X5))))) )),
  inference(superposition,[],[f572,f108])).
fof(f6553,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X31,bool),X32),s('type.semanticPrimitives.error_result'(X31),X33))) | s('type.semanticPrimitives.error_result'(X31),X33) = s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X33,X31)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool)) | s(bool,'T_0') = s(bool,happ(s(fun(X31,bool),X32),s(X31,sK5(X33,X31))))) )),
  inference(superposition,[],[f572,f689])).
fof(f6552,plain,(
  ( ! [X30,X28,X29,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X30))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X27,bool),X28),s('type.semanticPrimitives.error_result'(X27),X29))) | s('type.semanticPrimitives.error_result'(X27),X29) = s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X29,X27)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,happ(s(fun(X27,bool),X28),s(X27,sK5(X29,X27))))),bool)) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f572,f689])).
fof(f6551,plain,(
  ( ! [X26,X24,X23,X25] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X23,bool),X24),s('type.semanticPrimitives.error_result'(X23),X25))) | s('type.semanticPrimitives.error_result'(X23),X25) = s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X25,X23)))) | s(bool,X26) = s(bool,happ(s(fun(X23,bool),X24),s(X23,sK5(X25,X23)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X26)),bool))) )),
  inference(superposition,[],[f572,f689])).
fof(f6547,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X10))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X7,bool),X8),s('type.semanticPrimitives.error_result'(X7),X9))) | s('type.semanticPrimitives.error_result'(X7),X9) = s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X9,X7)))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK5(X9,X7))))) )),
  inference(superposition,[],[f572,f125])).
fof(f6546,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X3,bool),X4),s('type.semanticPrimitives.error_result'(X3),X5))) | s('type.semanticPrimitives.error_result'(X3),X5) = s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X5,X3)))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK5(X5,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f572,f125])).
fof(f6545,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X0,bool),X1),s('type.semanticPrimitives.error_result'(X0),X2))) | s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X2,X0)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X0))))) )),
  inference(superposition,[],[f572,f83])).
fof(f6539,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.semanticPrimitives.error_result'(X36),'const.evalProps.map_error_result_2'(s(fun(bool,X36),X37),s('type.semanticPrimitives.error_result'(bool),X35))) = s('type.semanticPrimitives.error_result'(X36),'const.semanticPrimitives.Rraise_1'(s(X36,happ(s(fun(bool,X36),X37),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),X35) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X35,bool)))) | s(bool,X38) = s(bool,sK5(X35,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f572,f689])).
fof(f6535,plain,(
  ( ! [X21,X19,X22,X20] : (s('type.semanticPrimitives.error_result'(X21),'const.evalProps.map_error_result_2'(s(fun(bool,X21),X22),s('type.semanticPrimitives.error_result'(bool),X19))) = s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,happ(s(fun(bool,X21),X22),s(bool,X20))))) | s('type.semanticPrimitives.error_result'(bool),X19) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X19,bool)))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,sK5(X19,bool))) )),
  inference(superposition,[],[f572,f125])).
fof(f6534,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X17),s('type.semanticPrimitives.error_result'(bool),X15))) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X17),s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),X15) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X15,bool)))) | s(bool,X18) = s(bool,sK5(X15,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f572,f125])).
fof(f6533,plain,(
  ( ! [X14,X12,X13] : (s('type.semanticPrimitives.error_result'(X13),'const.evalProps.map_error_result_2'(s(fun(bool,X13),X14),s('type.semanticPrimitives.error_result'(bool),X12))) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,happ(s(fun(bool,X13),X14),s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'(bool),X12) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X12,bool)))) | s(bool,'T_0') = s(bool,sK5(X12,bool))) )),
  inference(superposition,[],[f572,f83])).
fof(f6595,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X11),s(bool,X9))))) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X11),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f6532,f102])).
fof(f6532,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(bool,X10),X11),s(bool,X9))))) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X11),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool)))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f572,f687])).
fof(f6594,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(bool,X5),X6),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(bool,X5),X6),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))))) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f6531,f102])).
fof(f6531,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(bool,X5),X6),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(bool,X5),X6),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)))) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f572,f689])).
fof(f6322,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X0))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X1))) | s('type.semanticPrimitives.error_result'(bool),X1) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X2,bool),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X0,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f504,f536])).
fof(f6314,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X50),X51))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X52))))) | s('type.semanticPrimitives.error_result'(X50),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X50),X51)),'type.semanticPrimitives.error_result'(X50))) = s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X53,X50),X54),s('type.semanticPrimitives.error_result'(X53),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X52)))))) )),
  inference(superposition,[],[f504,f596])).
fof(f6308,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X18))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X17,bool)))) = s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X19,bool),X20),s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X18))))) | s('type.semanticPrimitives.error_result'(bool),X17) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X17,bool))) )),
  inference(superposition,[],[f504,f536])).
fof(f6305,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))) | s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X5,X0),X6),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))) )),
  inference(superposition,[],[f504,f107])).
fof(f6154,plain,(
  ( ! [X118,X116,X114,X119,X117,X115,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113))),'const.evalProps.map_error_result_2'(s(fun(X117,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113))),X118),s('type.semanticPrimitives.error_result'(X117),'const.semanticPrimitives.Rraise_1'(s(X117,X119))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X113),'const.semanticPrimitives.Rraise_1'(s(X113,happ(s(fun(X114,X113),X115),s(X114,X116))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113)),happ(s(fun(X117,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113))),X118),s(X117,X119))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X113)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X113),'const.evalProps.map_error_result_2'(s(fun(X114,X113),X115),s('type.semanticPrimitives.error_result'(X114),'const.semanticPrimitives.Rraise_1'(s(X114,X116)))))))) )),
  inference(superposition,[],[f574,f558])).
fof(f6153,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106] : (s('type.semanticPrimitives.error_result'(X110),'const.semanticPrimitives.Rraise_1'(s(X110,X111))) != s('type.semanticPrimitives.error_result'(X110),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),X110),X112),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X106),'const.semanticPrimitives.Rraise_1'(s(X106,happ(s(fun(X107,X106),X108),s(X107,X109))))))))))) | s(X110,X111) = s(X110,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),X110),X112),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X106),'const.evalProps.map_error_result_2'(s(fun(X107,X106),X108),s('type.semanticPrimitives.error_result'(X107),'const.semanticPrimitives.Rraise_1'(s(X107,X109)))))))))) )),
  inference(superposition,[],[f574,f558])).
fof(f6152,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101))),X105) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X101),'const.semanticPrimitives.Rraise_1'(s(X101,happ(s(fun(X102,X101),X103),s(X102,X104))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101)),sK5(X105,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101)))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X101),'const.evalProps.map_error_result_2'(s(fun(X102,X101),X103),s('type.semanticPrimitives.error_result'(X102),'const.semanticPrimitives.Rraise_1'(s(X102,X104))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101))),X105) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101))),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X105,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X101))))))) )),
  inference(superposition,[],[f537,f558])).
fof(f6147,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),X79) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.semanticPrimitives.Rraise_1'(s(X75,happ(s(fun(X76,X75),X77),s(X76,X78))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),X79))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X75)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X75),'const.evalProps.map_error_result_2'(s(fun(X76,X75),X77),s('type.semanticPrimitives.error_result'(X76),'const.semanticPrimitives.Rraise_1'(s(X76,X78)))))))))) )),
  inference(superposition,[],[f101,f558])).
fof(f6145,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),X69))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.semanticPrimitives.Rraise_1'(s(X65,happ(s(fun(X66,X65),X67),s(X66,X68))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),X69) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X65)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X65),'const.evalProps.map_error_result_2'(s(fun(X66,X65),X67),s('type.semanticPrimitives.error_result'(X66),'const.semanticPrimitives.Rraise_1'(s(X66,X68)))))))) )),
  inference(superposition,[],[f100,f558])).
fof(f6166,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),X59))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,happ(s(fun(X56,X55),X57),s(X56,X58))))))) | s('type.semanticPrimitives.error_result'(X55),X59) = s('type.semanticPrimitives.error_result'(X55),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.evalProps.map_error_result_2'(s(fun(X56,X55),X57),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X58)))))),'type.semanticPrimitives.error_result'(X55)))) )),
  inference(subsumption_resolution,[],[f6143,f102])).
fof(f6143,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),X59))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.semanticPrimitives.Rraise_1'(s(X55,happ(s(fun(X56,X55),X57),s(X56,X58))))))) | s('type.semanticPrimitives.error_result'(X55),X59) = s('type.semanticPrimitives.error_result'(X55),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.evalProps.map_error_result_2'(s(fun(X56,X55),X57),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X58)))))),'type.semanticPrimitives.error_result'(X55))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.evalProps.map_error_result_2'(s(fun(X56,X55),X57),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X58))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X55)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X55),'const.evalProps.map_error_result_2'(s(fun(X56,X55),X57),s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,X58)))))),'type.semanticPrimitives.error_result'(X55)))))) )),
  inference(superposition,[],[f537,f558])).
fof(f6137,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rraise_1'(s(X22,happ(s(fun('type.semanticPrimitives.error_result'(X18),X22),X23),s('type.semanticPrimitives.error_result'(X18),'const.evalProps.map_error_result_2'(s(fun(X19,X18),X20),s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rraise_1'(s(X19,X21))))))))) = s('type.semanticPrimitives.error_result'(X22),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X18),X22),X23),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X18)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X18),'const.semanticPrimitives.Rraise_1'(s(X18,happ(s(fun(X19,X18),X20),s(X19,X21)))))))))) )),
  inference(superposition,[],[f108,f558])).
fof(f6130,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118))),'const.evalProps.map_error_result_2'(s(fun(X122,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118))),X123),s('type.semanticPrimitives.error_result'(X122),'const.semanticPrimitives.Rraise_1'(s(X122,X124))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X118),'const.evalProps.map_error_result_2'(s(fun(X119,X118),X120),s('type.semanticPrimitives.error_result'(X119),'const.semanticPrimitives.Rraise_1'(s(X119,X121))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118)),happ(s(fun(X122,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118))),X123),s(X122,X124))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X118),'const.semanticPrimitives.Rraise_1'(s(X118,happ(s(fun(X119,X118),X120),s(X119,X121)))))))) )),
  inference(superposition,[],[f574,f558])).
fof(f6129,plain,(
  ( ! [X111,X116,X114,X112,X117,X115,X113] : (s('type.semanticPrimitives.error_result'(X115),'const.semanticPrimitives.Rraise_1'(s(X115,X116))) != s('type.semanticPrimitives.error_result'(X115),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115),X117),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.evalProps.map_error_result_2'(s(fun(X112,X111),X113),s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rraise_1'(s(X112,X114))))))))))) | s(X115,X116) = s(X115,happ(s(fun('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),X115),X117),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X111)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X111),'const.semanticPrimitives.Rraise_1'(s(X111,happ(s(fun(X112,X111),X113),s(X112,X114)))))))))) )),
  inference(superposition,[],[f574,f558])).
fof(f6128,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))),X110) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X106),'const.evalProps.map_error_result_2'(s(fun(X107,X106),X108),s('type.semanticPrimitives.error_result'(X107),'const.semanticPrimitives.Rraise_1'(s(X107,X109))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),sK5(X110,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X106),'const.semanticPrimitives.Rraise_1'(s(X106,happ(s(fun(X107,X106),X108),s(X107,X109))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))),X110) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X110,'type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X106))))))) )),
  inference(superposition,[],[f537,f558])).
fof(f6123,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80)),X84) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X80),'const.evalProps.map_error_result_2'(s(fun(X81,X80),X82),s('type.semanticPrimitives.error_result'(X81),'const.semanticPrimitives.Rraise_1'(s(X81,X83))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80)),X84))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X80)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X80),'const.semanticPrimitives.Rraise_1'(s(X80,happ(s(fun(X81,X80),X82),s(X81,X83)))))))))) )),
  inference(superposition,[],[f101,f558])).
fof(f6121,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),X74))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.evalProps.map_error_result_2'(s(fun(X71,X70),X72),s('type.semanticPrimitives.error_result'(X71),'const.semanticPrimitives.Rraise_1'(s(X71,X73))))))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),X74) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X70)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X70),'const.semanticPrimitives.Rraise_1'(s(X70,happ(s(fun(X71,X70),X72),s(X71,X73)))))))) )),
  inference(superposition,[],[f100,f558])).
fof(f6165,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),X64))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.evalProps.map_error_result_2'(s(fun(X61,X60),X62),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X63))))))) | s('type.semanticPrimitives.error_result'(X60),X64) = s('type.semanticPrimitives.error_result'(X60),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,happ(s(fun(X61,X60),X62),s(X61,X63)))))),'type.semanticPrimitives.error_result'(X60)))) )),
  inference(subsumption_resolution,[],[f6119,f102])).
fof(f6119,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),X64))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.evalProps.map_error_result_2'(s(fun(X61,X60),X62),s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,X63))))))) | s('type.semanticPrimitives.error_result'(X60),X64) = s('type.semanticPrimitives.error_result'(X60),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,happ(s(fun(X61,X60),X62),s(X61,X63)))))),'type.semanticPrimitives.error_result'(X60))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,happ(s(fun(X61,X60),X62),s(X61,X63))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X60)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X60),'const.semanticPrimitives.Rraise_1'(s(X60,happ(s(fun(X61,X60),X62),s(X61,X63)))))),'type.semanticPrimitives.error_result'(X60)))))) )),
  inference(superposition,[],[f537,f558])).
fof(f6113,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.semanticPrimitives.error_result'(X27),'const.semanticPrimitives.Rraise_1'(s(X27,happ(s(fun('type.semanticPrimitives.error_result'(X23),X27),X28),s('type.semanticPrimitives.error_result'(X23),'const.semanticPrimitives.Rraise_1'(s(X23,happ(s(fun(X24,X23),X25),s(X24,X26))))))))) = s('type.semanticPrimitives.error_result'(X27),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X23),X27),X28),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X23)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X23),'const.evalProps.map_error_result_2'(s(fun(X24,X23),X25),s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rraise_1'(s(X24,X26)))))))))) )),
  inference(superposition,[],[f108,f558])).
fof(f6108,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X20)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X20),'const.semanticPrimitives.Rraise_1'(s(X20,happ(s(fun('type.semanticPrimitives.error_result'(X16),X20),X21),s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X20)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X20),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X16),X20),X21),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X16)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19)))))))))))) )),
  inference(superposition,[],[f558,f558])).
fof(f6107,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,happ(s(fun('type.semanticPrimitives.error_result'(X10),X14),X15),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,happ(s(fun(X11,X10),X12),s(X11,X13))))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X14)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X14),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X10),X14),X15),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rraise_1'(s(X11,X13)))))))))))) )),
  inference(superposition,[],[f558,f558])).
fof(f6106,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X8)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,happ(s(fun(X4,X8),X9),s(X4,happ(s(fun(X5,X4),X6),s(X5,X7))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X8)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X4,X8),X9),s('type.semanticPrimitives.error_result'(X4),'const.evalProps.map_error_result_2'(s(fun(X5,X4),X6),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X7)))))))))) )),
  inference(superposition,[],[f558,f108])).
fof(f6105,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X2)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X0,X2),X3),s('type.semanticPrimitives.error_result'(X0),X1))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X2)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,happ(s(fun(X0,X2),X3),s(X0,sK5(X1,X0)))))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f558,f95])).
fof(f6062,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))))),'type.semanticPrimitives.error_result'(X0)))) )),
  inference(unit_resulting_resolution,[],[f102,f558,f537])).
fof(f6065,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))),'type.semanticPrimitives.error_result'(X0)))) )),
  inference(unit_resulting_resolution,[],[f102,f558,f537])).
fof(f6066,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))))))) )),
  inference(unit_resulting_resolution,[],[f558,f101])).
fof(f115,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f79,f76])).
fof(f79,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f11])).
fof(f11,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.bool.NOT_CLAUSES|split|1')).
fof(f434,plain,(
  s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f116,f100])).
fof(f436,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f100,f83])).
fof(f441,plain,(
  ( ! [X15,X16] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f100,f83])).
fof(f457,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f434,f100])).
fof(f475,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f101,f83])).
fof(f480,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,X16) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f101,f83])).
fof(f536,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(bool),X8) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X8,bool)))) | s('type.semanticPrimitives.error_result'(bool),X8) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X8,bool))) )),
  inference(superposition,[],[f95,f83])).
fof(f617,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8)))) )),
  inference(superposition,[],[f530,f83])).
fof(f658,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK7(X16,X17,X18,bool))) )),
  inference(superposition,[],[f109,f83])).
fof(f663,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK7(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK7(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f109,f83])).
fof(f773,plain,(
  ( ! [X24,X23] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) | s(bool,X24) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X23)),bool)) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f100,f691])).
fof(f790,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool)) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f769,f76])).
fof(f769,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | s(bool,X18) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool)) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(superposition,[],[f127,f691])).
fof(f791,plain,(
  ( ! [X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f772,f102])).
fof(f772,plain,(
  ( ! [X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool)))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f95,f691])).
fof(f792,plain,(
  ( ! [X28,X27] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X27))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f775,f596])).
fof(f775,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') != s(bool,X28) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X28))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f101,f691])).
fof(f860,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool)) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f852,f76])).
fof(f852,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool)) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f789,f691])).
fof(f894,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f791,f125])).
fof(f900,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f791,f791])).
fof(f906,plain,(
  ( ! [X11] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11))) | s(bool,'F_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X11)),bool))) )),
  inference(superposition,[],[f791,f596])).
fof(f907,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK5(X2,bool)) | s('type.semanticPrimitives.error_result'(bool),X2) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X2,bool))))) )),
  inference(superposition,[],[f791,f95])).
fof(f909,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,happ(s(fun(bool,X6),X7),s(bool,'T_0'))))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X7),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f108,f791])).
fof(f913,plain,(
  ( ! [X14,X15] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X15))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X14))))) | s('type.semanticPrimitives.error_result'(bool),X15) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f100,f791])).
fof(f915,plain,(
  ( ! [X19,X18] : (s('type.semanticPrimitives.error_result'(bool),X19) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X18))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X19))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f101,f791])).
fof(f923,plain,(
  ( ! [X6,X4,X5] : (s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(bool,X5),X6),s(bool,X4))))) = s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(bool,X5),X6),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f108,f791])).
fof(f924,plain,(
  ( ! [X8,X7] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X8))) | s(bool,X7) = s(bool,X8) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f100,f791])).
fof(f928,plain,(
  ( ! [X17,X18] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X18))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),X18) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X17))) | s(bool,'F_0') = s(bool,X17)) )),
  inference(superposition,[],[f100,f791])).
fof(f930,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.error_result'(bool),X22) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X22))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X21))))) | s(bool,'F_0') = s(bool,X21)) )),
  inference(superposition,[],[f101,f791])).
fof(f1006,plain,(
  ( ! [X30,X31,X32] : (s(bool,X31) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X30)),bool)) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f989,f76])).
fof(f989,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X31) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X30)),bool)) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f422,f691])).
fof(f1011,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0')))))))),
  inference(unit_resulting_resolution,[],[f457,f100])).
fof(f1115,plain,(
  ( ! [X23,X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X23))) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(superposition,[],[f436,f596])).
fof(f1118,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,sK5(X4,bool)) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X4,bool))))) )),
  inference(superposition,[],[f436,f95])).
fof(f1128,plain,(
  ( ! [X24,X25] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X25))) | s(bool,'F_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X24)),bool)) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f436,f596])).
fof(f1131,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s(bool,'F_0') = s(bool,sK5(X4,bool)) | s(bool,'T_0') = s(bool,X5) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X4,bool))))) )),
  inference(superposition,[],[f436,f95])).
fof(f1166,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X20))) | s(bool,X21) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X20)),bool)) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f441,f596])).
fof(f1168,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.error_result'(bool),X3) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,X4) = s(bool,sK5(X3,bool)) | s(bool,'T_0') = s(bool,X4) | s('type.semanticPrimitives.error_result'(bool),X3) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X3,bool))))) )),
  inference(superposition,[],[f441,f95])).
fof(f1327,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1326])).
fof(f1326,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(duplicate_literal_removal,[],[f1202])).
fof(f1202,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f687,f83])).
fof(f1335,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1334])).
fof(f1334,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) )),
  inference(duplicate_literal_removal,[],[f1193])).
fof(f1193,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f687,f83])).
fof(f1390,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X29) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X28)),bool))) )),
  inference(superposition,[],[f475,f596])).
fof(f1410,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1)))) )),
  inference(condensation,[],[f1409])).
fof(f1409,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f1379])).
fof(f1379,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f475,f687])).
fof(f1413,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f1412])).
fof(f1412,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f1411])).
fof(f1411,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1391,f687])).
fof(f1391,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f475,f687])).
fof(f1429,plain,(
  ( ! [X19] : (s(bool,X19) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool))) )),
  inference(superposition,[],[f1327,f596])).
fof(f1476,plain,(
  ( ! [X52,X51] : (s(bool,X52) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f101,f1327])).
fof(f1516,plain,(
  ( ! [X50,X49] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X49))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) | s(bool,X50) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f100,f1327])).
fof(f1773,plain,(
  ( ! [X54,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X54))) | s(bool,X54) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X53)),bool)) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f100,f1335])).
fof(f1863,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) | s('type.semanticPrimitives.abort',X1) = s('type.semanticPrimitives.abort',sK6(X0,bool)) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) )),
  inference(superposition,[],[f98,f536])).
fof(f1865,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X4,bool))) )),
  inference(superposition,[],[f102,f536])).
fof(f1866,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X6,bool)))) = s('type.semanticPrimitives.error_result'(X7),'const.evalProps.map_error_result_2'(s(fun(bool,X7),X8),s('type.semanticPrimitives.error_result'(bool),X6))) | s('type.semanticPrimitives.error_result'(bool),X6) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X6,bool))) )),
  inference(superposition,[],[f107,f536])).
fof(f1867,plain,(
  ( ! [X10,X9] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X9))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X10))))) | s('type.semanticPrimitives.error_result'(bool),X9) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X9,bool))) )),
  inference(superposition,[],[f433,f536])).
fof(f1868,plain,(
  ( ! [X12,X11] : (s('type.semanticPrimitives.error_result'(bool),X11) != s('type.semanticPrimitives.error_result'(bool),X12) | s('type.semanticPrimitives.error_result'(bool),X12) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X12,bool)))) | s('type.semanticPrimitives.error_result'(bool),X11) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X11,bool))) )),
  inference(superposition,[],[f539,f536])).
fof(f1869,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.semanticPrimitives.error_result'(bool),X13) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X14,bool),X15),s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,X16))))) | s('type.semanticPrimitives.error_result'(bool),X13) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X13,bool))) )),
  inference(superposition,[],[f576,f536])).
fof(f1872,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X20))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X21))) | s('type.semanticPrimitives.error_result'(bool),X21) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X20,bool)))) | s('type.semanticPrimitives.error_result'(bool),X20) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X20,bool))) )),
  inference(superposition,[],[f100,f536])).
fof(f1874,plain,(
  ( ! [X24,X25] : (s('type.semanticPrimitives.error_result'(bool),X24) != s('type.semanticPrimitives.error_result'(bool),X25) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X25))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X24,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X24) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X24,bool))) )),
  inference(superposition,[],[f101,f536])).
fof(f1915,plain,(
  ( ! [X19,X20,X18] : (s('type.semanticPrimitives.error_result'(X19),'const.semanticPrimitives.Rraise_1'(s(X19,happ(s(fun(bool,X19),X20),s(bool,'F_0'))))) = s('type.semanticPrimitives.error_result'(X19),'const.evalProps.map_error_result_2'(s(fun(bool,X19),X20),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X18))))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f108,f1413])).
fof(f1920,plain,(
  ( ! [X30,X29] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X30))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X29))))) | s('type.semanticPrimitives.error_result'(bool),X30) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f100,f1413])).
fof(f1922,plain,(
  ( ! [X33,X34] : (s('type.semanticPrimitives.error_result'(bool),X34) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X34))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f101,f1413])).
fof(f1938,plain,(
  ( ! [X17,X15,X16] : (s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(bool,X16),X17),s(bool,X15))))) = s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X17),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f108,f1413])).
fof(f1944,plain,(
  ( ! [X30,X31] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X31))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'(bool),X31) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f100,f1413])).
fof(f1946,plain,(
  ( ! [X35,X34] : (s('type.semanticPrimitives.error_result'(bool),X35) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X35))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X34))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f101,f1413])).
fof(f2068,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool))) )),
  inference(superposition,[],[f101,f1429])).
fof(f2116,plain,(
  ( ! [X61,X60] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) | s(bool,X61) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool))) )),
  inference(superposition,[],[f100,f1429])).
fof(f2159,plain,(
  ( ! [X57,X56] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))) | s(bool,X57) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool)) | s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f100,f1429])).
fof(f2239,plain,(
  ( ! [X57,X56] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool))) )),
  inference(forward_demodulation,[],[f2066,f596])).
fof(f2066,plain,(
  ( ! [X57,X56] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool))) )),
  inference(superposition,[],[f100,f1429])).
fof(f2241,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') != s(bool,f(s(X79,sK7(X80,X81,bool,X79)))) | s(fun(X79,bool),X80) = s(fun(X79,bool),X81) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(X79,sK7(X80,X81,bool,X79))))),bool))) )),
  inference(forward_demodulation,[],[f2075,f596])).
fof(f2075,plain,(
  ( ! [X80,X81,X79] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) != s(bool,f(s(X79,sK7(X80,X81,bool,X79)))) | s(fun(X79,bool),X80) = s(fun(X79,bool),X81) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(X79,sK7(X80,X81,bool,X79))))),bool))) )),
  inference(superposition,[],[f109,f1429])).
fof(f2242,plain,(
  ( ! [X85] : (s('type.semanticPrimitives.error_result'(bool),X85) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X85) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X85,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK5(X85,bool))),bool))) )),
  inference(forward_demodulation,[],[f2076,f596])).
fof(f2076,plain,(
  ( ! [X85] : (s('type.semanticPrimitives.error_result'(bool),X85) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) | s('type.semanticPrimitives.error_result'(bool),X85) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X85,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK5(X85,bool))),bool))) )),
  inference(superposition,[],[f95,f1429])).
fof(f2245,plain,(
  ( ! [X88,X87,X89] : (s(X89,f(s(bool,'F_0'))) != s(X89,g(s(bool,'F_0'))) | s(fun(bool,X89),X87) = s(fun(bool,X89),X88) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X87,X88,X89,bool))),bool))) )),
  inference(forward_demodulation,[],[f2078,f596])).
fof(f2078,plain,(
  ( ! [X88,X87,X89] : (s(X89,f(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) != s(X89,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) | s(fun(bool,X89),X87) = s(fun(bool,X89),X88) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X87,X88,X89,bool))),bool))) )),
  inference(superposition,[],[f109,f1429])).
fof(f2273,plain,(
  ( ! [X24,X23] : (s(bool,X24) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X23)),bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(subsumption_resolution,[],[f2140,f76])).
fof(f2140,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | s(bool,X24) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X23)),bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f127,f1429])).
fof(f2274,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X27) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(subsumption_resolution,[],[f2142,f76])).
fof(f2142,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X27)),bool)) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X27) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f422,f1429])).
fof(f2278,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X47)),bool)) | s(bool,'F_0') = s(bool,X48) | s(bool,X47) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(subsumption_resolution,[],[f2152,f76])).
fof(f2152,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | s(bool,X48) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X47)),bool)) | s(bool,'F_0') = s(bool,X48) | s(bool,X47) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f789,f1429])).
fof(f2279,plain,(
  ( ! [X55] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55))) | s(bool,X55) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(subsumption_resolution,[],[f2158,f102])).
fof(f2158,plain,(
  ( ! [X55] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X55)),bool)))) | s(bool,X55) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f95,f1429])).
fof(f2280,plain,(
  ( ! [X61,X60] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) | s(bool,'T_0') != s(bool,X61) | s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(forward_demodulation,[],[f2161,f596])).
fof(f2161,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') != s(bool,X61) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X60)),bool)))) | s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f101,f1429])).
fof(f2309,plain,(
  ( ! [X35,X34] : (s('type.semanticPrimitives.error_result'(bool),X35) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,sK5(X35,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X34)),bool)) | s('type.semanticPrimitives.error_result'(bool),X35) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X35,bool)))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f537,f1335])).
fof(f2317,plain,(
  ( ! [X57,X56] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X57) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X56))) | s('type.semanticPrimitives.error_result'(bool),sK5(X57,'type.semanticPrimitives.error_result'(bool))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X56,bool)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X57) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X57,'type.semanticPrimitives.error_result'(bool))))) | s('type.semanticPrimitives.error_result'(bool),X56) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X56,bool))) )),
  inference(superposition,[],[f537,f536])).
fof(f2320,plain,(
  ( ! [X66,X67] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X67) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X66))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),sK5(X67,'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X67) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X67,'type.semanticPrimitives.error_result'(bool))))) | s(bool,'F_0') = s(bool,X66)) )),
  inference(superposition,[],[f537,f791])).
fof(f2321,plain,(
  ( ! [X68,X69] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X69) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X68))) = s('type.semanticPrimitives.error_result'(bool),sK5(X69,'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X69) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X69,'type.semanticPrimitives.error_result'(bool))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f537,f1413])).
fof(f2322,plain,(
  ( ! [X70,X71] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X71) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X70))) = s('type.semanticPrimitives.error_result'(bool),sK5(X71,'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X71) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X71,'type.semanticPrimitives.error_result'(bool))))) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f537,f791])).
fof(f2323,plain,(
  ( ! [X72,X73] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X73) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X72))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),sK5(X73,'type.semanticPrimitives.error_result'(bool))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),X73) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X73,'type.semanticPrimitives.error_result'(bool))))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f537,f1413])).
fof(f2331,plain,(
  ( ! [X12,X13] : (s('type.semanticPrimitives.error_result'(bool),X13) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,X12) = s(bool,sK5(X13,bool)) | s('type.semanticPrimitives.error_result'(bool),X13) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X13,bool)))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f537,f791])).
fof(f2338,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X21))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s(bool,X22) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,'F_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2290,f102])).
fof(f2290,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X21))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X22))) | s(bool,X22) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)))) | s(bool,'F_0') = s(bool,X21)) )),
  inference(superposition,[],[f537,f791])).
fof(f2342,plain,(
  ( ! [X23,X22] : (s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,X22) = s(bool,sK5(X23,bool)) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X23,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(forward_demodulation,[],[f2304,f596])).
fof(f2304,plain,(
  ( ! [X23,X22] : (s('type.semanticPrimitives.error_result'(bool),X23) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) | s(bool,X22) = s(bool,sK5(X23,bool)) | s('type.semanticPrimitives.error_result'(bool),X23) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X23,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X22)),bool))) )),
  inference(superposition,[],[f537,f1429])).
fof(f2343,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,X32) | s('type.semanticPrimitives.error_result'(bool),X33) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,sK5(X33,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X32)),bool)) | s('type.semanticPrimitives.error_result'(bool),X33) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X33,bool))))) )),
  inference(forward_demodulation,[],[f2308,f596])).
fof(f2308,plain,(
  ( ! [X33,X32] : (s('type.semanticPrimitives.error_result'(bool),X33) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,sK5(X33,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X32)),bool)) | s('type.semanticPrimitives.error_result'(bool),X33) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X33,bool)))) | s(bool,X32) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f537,f1429])).
fof(f2344,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,sK5(X42,bool)) | s('type.semanticPrimitives.error_result'(bool),X42) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X41))) | s('type.semanticPrimitives.error_result'(bool),X42) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X42,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X41)),bool))) )),
  inference(forward_demodulation,[],[f2312,f596])).
fof(f2312,plain,(
  ( ! [X41,X42] : (s('type.semanticPrimitives.error_result'(bool),X42) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X41))) | s(bool,sK5(X42,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),X42) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X42,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X41)),bool))) )),
  inference(superposition,[],[f537,f1429])).
fof(f2700,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X13))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X14))) | s(bool,'F_0') = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X15)),bool))) )),
  inference(superposition,[],[f791,f689])).
fof(f3113,plain,(
  ( ! [X78,X79,X77] : (s(bool,X78) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X77)),bool)) | s(bool,'F_0') = s(bool,X78) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2677,f76])).
fof(f2677,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,'T_0')) | s(bool,X78) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X77)),bool)) | s(bool,'F_0') = s(bool,X78) | s(bool,X77) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f789,f689])).
fof(f3137,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X63) = s(bool,X65) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X65)),bool))) )),
  inference(subsumption_resolution,[],[f2756,f76])).
fof(f2756,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X63) = s(bool,X65) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X65)),bool))) )),
  inference(superposition,[],[f789,f689])).
fof(f3366,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X6,bool),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))) | s(bool,X5) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f574,f83])).
fof(f3379,plain,(
  ( ! [X85,X83,X84,X82] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X83,bool),X84),s('type.semanticPrimitives.error_result'(X83),'const.semanticPrimitives.Rraise_1'(s(X83,X85))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X82)),bool)) = s(bool,happ(s(fun(X83,bool),X84),s(X83,X85))) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f574,f1335])).
fof(f3386,plain,(
  ( ! [X127,X125,X128,X126] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X125))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.evalProps.map_error_result_2'(s(fun(X126,'type.semanticPrimitives.error_result'(bool)),X127),s('type.semanticPrimitives.error_result'(X126),'const.semanticPrimitives.Rraise_1'(s(X126,X128))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X125,bool)))) = s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X126,'type.semanticPrimitives.error_result'(bool)),X127),s(X126,X128))) | s('type.semanticPrimitives.error_result'(bool),X125) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X125,bool))) )),
  inference(superposition,[],[f574,f536])).
fof(f3389,plain,(
  ( ! [X146,X147,X145,X148] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.evalProps.map_error_result_2'(s(fun(X146,'type.semanticPrimitives.error_result'(bool)),X147),s('type.semanticPrimitives.error_result'(X146),'const.semanticPrimitives.Rraise_1'(s(X146,X148))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X145))) = s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X146,'type.semanticPrimitives.error_result'(bool)),X147),s(X146,X148))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f574,f1413])).
fof(f3390,plain,(
  ( ! [X152,X151,X149,X150] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.evalProps.map_error_result_2'(s(fun(X150,'type.semanticPrimitives.error_result'(bool)),X151),s('type.semanticPrimitives.error_result'(X150),'const.semanticPrimitives.Rraise_1'(s(X150,X152))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X149))) = s('type.semanticPrimitives.error_result'(bool),happ(s(fun(X150,'type.semanticPrimitives.error_result'(bool)),X151),s(X150,X152))) | s(bool,'F_0') = s(bool,X149)) )),
  inference(superposition,[],[f574,f791])).
fof(f3397,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X21,bool),X22),s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,X23))))) | s(bool,X20) = s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f574,f791])).
fof(f3400,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X7))) != s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(bool,X6),X8),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s(X6,X7) = s(X6,happ(s(fun(bool,X6),X8),s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f574,f83])).
fof(f3414,plain,(
  ( ! [X74,X72,X75,X73] : (s('type.semanticPrimitives.error_result'(X73),'const.semanticPrimitives.Rraise_1'(s(X73,X74))) != s('type.semanticPrimitives.error_result'(X73),'const.evalProps.map_error_result_2'(s(fun(bool,X73),X75),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X72))))) | s(X73,X74) = s(X73,happ(s(fun(bool,X73),X75),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f574,f83])).
fof(f3426,plain,(
  ( ! [X130,X128,X131,X129] : (s('type.semanticPrimitives.error_result'(X129),'const.semanticPrimitives.Rraise_1'(s(X129,X130))) != s('type.semanticPrimitives.error_result'(X129),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X129),X131),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),X128))))) | s(X129,X130) = s(X129,happ(s(fun('type.semanticPrimitives.error_result'(bool),X129),X131),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X128,bool)))))) | s('type.semanticPrimitives.error_result'(bool),X128) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X128,bool))) )),
  inference(superposition,[],[f574,f536])).
fof(f3429,plain,(
  ( ! [X146,X144,X147,X145] : (s('type.semanticPrimitives.error_result'(X145),'const.semanticPrimitives.Rraise_1'(s(X145,X146))) != s('type.semanticPrimitives.error_result'(X145),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X145),X147),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X144))))))) | s(X145,X146) = s(X145,happ(s(fun('type.semanticPrimitives.error_result'(bool),X145),X147),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X144)) )),
  inference(superposition,[],[f574,f791])).
fof(f3430,plain,(
  ( ! [X151,X149,X150,X148] : (s('type.semanticPrimitives.error_result'(X149),'const.semanticPrimitives.Rraise_1'(s(X149,X150))) != s('type.semanticPrimitives.error_result'(X149),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X149),X151),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))))) | s(X149,X150) = s(X149,happ(s(fun('type.semanticPrimitives.error_result'(bool),X149),X151),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X148))))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f574,f1413])).
fof(f3431,plain,(
  ( ! [X154,X152,X155,X153] : (s('type.semanticPrimitives.error_result'(X153),'const.semanticPrimitives.Rraise_1'(s(X153,X154))) != s('type.semanticPrimitives.error_result'(X153),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X153),X155),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | s(X153,X154) = s(X153,happ(s(fun('type.semanticPrimitives.error_result'(bool),X153),X155),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152))))) | s(bool,'F_0') = s(bool,X152)) )),
  inference(superposition,[],[f574,f791])).
fof(f3432,plain,(
  ( ! [X158,X156,X159,X157] : (s('type.semanticPrimitives.error_result'(X157),'const.semanticPrimitives.Rraise_1'(s(X157,X158))) != s('type.semanticPrimitives.error_result'(X157),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(bool),X157),X159),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X156))))))) | s(X157,X158) = s(X157,happ(s(fun('type.semanticPrimitives.error_result'(bool),X157),X159),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f574,f1413])).
fof(f3438,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X14))) != s('type.semanticPrimitives.error_result'(X13),'const.evalProps.map_error_result_2'(s(fun(bool,X13),X15),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X12))))) | s(X13,X14) = s(X13,happ(s(fun(bool,X13),X15),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f574,f791])).
fof(f3445,plain,(
  ( ! [X54,X56,X55,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X54,bool),X55),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X56))))) | s(bool,X53) = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X53)),bool))) )),
  inference(forward_demodulation,[],[f3376,f596])).
fof(f3376,plain,(
  ( ! [X54,X56,X55,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X54,bool),X55),s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rraise_1'(s(X54,X56))))) | s(bool,X53) = s(bool,happ(s(fun(X54,bool),X55),s(X54,X56))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X53)),bool))) )),
  inference(superposition,[],[f574,f1429])).
fof(f3446,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'F_0') = s(bool,X78) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X79,bool),X80),s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rraise_1'(s(X79,X81))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X78)),bool)) = s(bool,happ(s(fun(X79,bool),X80),s(X79,X81)))) )),
  inference(forward_demodulation,[],[f3378,f596])).
fof(f3378,plain,(
  ( ! [X80,X78,X81,X79] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X79,bool),X80),s('type.semanticPrimitives.error_result'(X79),'const.semanticPrimitives.Rraise_1'(s(X79,X81))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X78)),bool)) = s(bool,happ(s(fun(X79,bool),X80),s(X79,X81))) | s(bool,X78) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f574,f1429])).
fof(f3459,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,X81) | s(X82,X83) = s(X82,happ(s(fun(bool,X82),X84),s(bool,X81))) | s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rraise_1'(s(X82,X83))) != s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(bool,X82),X84),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0')))))) )),
  inference(forward_demodulation,[],[f3458,f596])).
fof(f3458,plain,(
  ( ! [X83,X81,X84,X82] : (s(X82,X83) = s(X82,happ(s(fun(bool,X82),X84),s(bool,X81))) | s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rraise_1'(s(X82,X83))) != s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(bool,X82),X84),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(bool,X81) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(forward_demodulation,[],[f3416,f596])).
fof(f3416,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.semanticPrimitives.error_result'(X82),'const.semanticPrimitives.Rraise_1'(s(X82,X83))) != s('type.semanticPrimitives.error_result'(X82),'const.evalProps.map_error_result_2'(s(fun(bool,X82),X84),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(X82,X83) = s(X82,happ(s(fun(bool,X82),X84),s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X81)),bool)))) | s(bool,X81) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(superposition,[],[f574,f1429])).
fof(f659,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f109,f125])).
fof(f661,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK7(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK7(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f109,f125])).
fof(f1249,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X109)),bool)) != s(bool,f(s(X106,sK7(X107,X108,bool,X106)))) | s(fun(X106,bool),X107) = s(fun(X106,bool),X108) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,g(s(X106,sK7(X107,X108,bool,X106))))) )),
  inference(superposition,[],[f109,f687])).
fof(f2573,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X179) != s(bool,f(s(X176,sK7(X177,X178,bool,X176)))) | s(fun(X176,bool),X177) = s(fun(X176,bool),X178) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(X176,sK7(X177,X178,bool,X176))))),bool)) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f109,f689])).
fof(f2645,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X179) != s(bool,f(s(X176,sK7(X177,X178,bool,X176)))) | s(fun(X176,bool),X177) = s(fun(X176,bool),X178) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X179)),bool)) | s(bool,'T_0') = s(bool,g(s(X176,sK7(X177,X178,bool,X176))))) )),
  inference(superposition,[],[f109,f689])).
fof(f2797,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'T_0') != s(bool,f(s(X163,sK7(X164,X165,bool,X163)))) | s(fun(X163,bool),X164) = s(fun(X163,bool),X165) | s(bool,X166) = s(bool,g(s(X163,sK7(X164,X165,bool,X163)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X166)),bool))) )),
  inference(superposition,[],[f109,f689])).
fof(f3628,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK7(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f3617])).
fof(f3617,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK7(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f663,f83])).
fof(f3638,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,f(s(X0,sK7(X1,X2,bool,X0))))),bool))) )),
  inference(condensation,[],[f3637])).
fof(f3637,plain,(
  ( ! [X30,X28,X29,X27] : (s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,'T_0') = s(bool,g(s(X27,sK7(X28,X29,bool,X27)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,f(s(X27,sK7(X28,X29,bool,X27))))),bool)) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f3624,f83])).
fof(f3624,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'F_0') != s(bool,X30) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,'T_0') = s(bool,g(s(X27,sK7(X28,X29,bool,X27)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,f(s(X27,sK7(X28,X29,bool,X27))))),bool)) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f663,f689])).
fof(f3367,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X10,bool),X11),s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X12))))) | s(bool,X9) = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f574,f125])).
fof(f3368,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X16,bool),X17),s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,X18))))) | s(bool,X14) = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f574,f125])).
fof(f3372,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X35,bool),X36),s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,X37))))) | s(bool,X34) = s(bool,happ(s(fun(X35,bool),X36),s(X35,X37))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f574,f689])).
fof(f3373,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X40))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X41,bool),X42),s('type.semanticPrimitives.error_result'(X41),'const.semanticPrimitives.Rraise_1'(s(X41,X43))))) | s(bool,X39) = s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X39)),bool)) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f574,f689])).
fof(f3374,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X45))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X46,bool),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rraise_1'(s(X46,X48))))) | s(bool,X44) = s(bool,happ(s(fun(X46,bool),X47),s(X46,X48))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X45)),bool)) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f574,f689])).
fof(f3377,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X74,bool),X75),s('type.semanticPrimitives.error_result'(X74),'const.semanticPrimitives.Rraise_1'(s(X74,X76))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X73)),bool)) = s(bool,happ(s(fun(X74,bool),X75),s(X74,X76))) | s(bool,X73) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f574,f689])).
fof(f3380,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X87))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X88,bool),X89),s('type.semanticPrimitives.error_result'(X88),'const.semanticPrimitives.Rraise_1'(s(X88,X90))))) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X86)),bool)) = s(bool,happ(s(fun(X88,bool),X89),s(X88,X90))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f574,f687])).
fof(f532,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X0,bool)))) | s(bool,X1) = s(bool,sK5(X0,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f95,f125])).
fof(f534,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X4,bool)))) | s(bool,'T_0') = s(bool,sK5(X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f95,f125])).
fof(f2298,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.error_result'(bool),X6) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,X5) = s(bool,sK5(X6,bool)) | s('type.semanticPrimitives.error_result'(bool),X6) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X6,bool)))) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f537,f125])).
fof(f2299,plain,(
  ( ! [X10,X8,X9] : (s('type.semanticPrimitives.error_result'(bool),X10) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))) | s(bool,X8) = s(bool,sK5(X10,bool)) | s('type.semanticPrimitives.error_result'(bool),X10) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X10,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f537,f125])).
fof(f2305,plain,(
  ( ! [X26,X24,X25] : (s('type.semanticPrimitives.error_result'(bool),X25) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X24))) | s(bool,'T_0') = s(bool,sK5(X25,bool)) | s('type.semanticPrimitives.error_result'(bool),X25) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X25,bool)))) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f537,f125])).
fof(f2310,plain,(
  ( ! [X37,X38,X36] : (s('type.semanticPrimitives.error_result'(bool),X38) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X37))) | s(bool,sK5(X38,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X36)),bool)) | s('type.semanticPrimitives.error_result'(bool),X38) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X38,bool)))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f537,f687])).
fof(f2570,plain,(
  ( ! [X163,X164,X162] : (s('type.semanticPrimitives.error_result'(bool),X164) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X163))) | s(bool,X162) = s(bool,sK5(X164,bool)) | s('type.semanticPrimitives.error_result'(bool),X164) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X164,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X162)),bool)) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f537,f689])).
fof(f2574,plain,(
  ( ! [X185,X184] : (s('type.semanticPrimitives.error_result'(bool),X184) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X185))) | s('type.semanticPrimitives.error_result'(bool),X184) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X184,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK5(X184,bool))),bool)) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f95,f689])).
fof(f2642,plain,(
  ( ! [X163,X164,X162] : (s('type.semanticPrimitives.error_result'(bool),X164) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X163))) | s(bool,X162) = s(bool,sK5(X164,bool)) | s('type.semanticPrimitives.error_result'(bool),X164) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X164,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X163)),bool)) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f537,f689])).
fof(f2646,plain,(
  ( ! [X185,X184] : (s('type.semanticPrimitives.error_result'(bool),X184) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X185))) | s('type.semanticPrimitives.error_result'(bool),X184) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X184,bool)))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X185)),bool)) | s(bool,'T_0') = s(bool,sK5(X184,bool))) )),
  inference(superposition,[],[f95,f689])).
fof(f2694,plain,(
  ( ! [X123,X121,X122] : (s('type.semanticPrimitives.error_result'(bool),X122) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,sK5(X122,bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X121)),bool)) | s('type.semanticPrimitives.error_result'(bool),X122) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X122,bool)))) | s(bool,X121) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f537,f689])).
fof(f2731,plain,(
  ( ! [X109,X107,X108] : (s('type.semanticPrimitives.error_result'(bool),X108) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X107))) | s(bool,'T_0') = s(bool,sK5(X108,bool)) | s('type.semanticPrimitives.error_result'(bool),X108) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X108,bool)))) | s(bool,X107) = s(bool,X109) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X109)),bool))) )),
  inference(superposition,[],[f537,f689])).
fof(f2794,plain,(
  ( ! [X152,X151,X150] : (s('type.semanticPrimitives.error_result'(bool),X151) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,X150) = s(bool,sK5(X151,bool)) | s('type.semanticPrimitives.error_result'(bool),X151) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X151,bool)))) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X152)),bool))) )),
  inference(superposition,[],[f537,f689])).
fof(f2798,plain,(
  ( ! [X171,X170] : (s('type.semanticPrimitives.error_result'(bool),X170) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'(bool),X170) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X170,bool)))) | s(bool,X171) = s(bool,sK5(X170,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X171)),bool))) )),
  inference(superposition,[],[f95,f689])).
fof(f1341,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f1340])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f1188])).
fof(f1188,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f687,f687])).
fof(f3437,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X10))) != s('type.semanticPrimitives.error_result'(X9),'const.evalProps.map_error_result_2'(s(fun(X5,X9),X11),s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))))) | s(X9,X10) = s(X9,happ(s(fun(X5,X9),X11),s(X5,happ(s(fun(X6,X5),X7),s(X6,X8)))))) )),
  inference(superposition,[],[f574,f108])).
fof(f3436,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X3))) != s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X0,X2),X4),s('type.semanticPrimitives.error_result'(X0),X1))) | s(X2,X3) = s(X2,happ(s(fun(X0,X2),X4),s(X0,sK5(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f574,f95])).
fof(f3428,plain,(
  ( ! [X142,X140,X138,X143,X141,X139,X137] : (s('type.semanticPrimitives.error_result'(X141),'const.semanticPrimitives.Rraise_1'(s(X141,X142))) != s('type.semanticPrimitives.error_result'(X141),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X137),X141),X143),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X137)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X137),'const.evalProps.map_error_result_2'(s(fun(X138,X137),X139),s('type.semanticPrimitives.error_result'(X138),'const.semanticPrimitives.Rraise_1'(s(X138,X140))))))))) | s(X141,X142) = s(X141,happ(s(fun('type.semanticPrimitives.error_result'(X137),X141),X143),s('type.semanticPrimitives.error_result'(X137),'const.semanticPrimitives.Rraise_1'(s(X137,happ(s(fun(X138,X137),X139),s(X138,X140)))))))) )),
  inference(superposition,[],[f574,f108])).
fof(f3427,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s('type.semanticPrimitives.error_result'(X134),'const.semanticPrimitives.Rraise_1'(s(X134,X135))) != s('type.semanticPrimitives.error_result'(X134),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X132),X134),X136),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X132)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X132),X133))))) | s(X134,X135) = s(X134,happ(s(fun('type.semanticPrimitives.error_result'(X132),X134),X136),s('type.semanticPrimitives.error_result'(X132),'const.semanticPrimitives.Rraise_1'(s(X132,sK5(X133,X132)))))) | s('type.semanticPrimitives.error_result'(X132),X133) = s('type.semanticPrimitives.error_result'(X132),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X133,X132))))) )),
  inference(superposition,[],[f574,f95])).
fof(f3424,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119] : (s('type.semanticPrimitives.error_result'(X122),'const.semanticPrimitives.Rraise_1'(s(X122,X123))) != s('type.semanticPrimitives.error_result'(X122),'const.evalProps.map_error_result_2'(s(fun('type.semanticPrimitives.error_result'(X118),X122),X124),s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X118)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X118),'const.semanticPrimitives.Rraise_1'(s(X118,happ(s(fun(X119,X118),X120),s(X119,X121))))))))) | s(X122,X123) = s(X122,happ(s(fun('type.semanticPrimitives.error_result'(X118),X122),X124),s('type.semanticPrimitives.error_result'(X118),'const.evalProps.map_error_result_2'(s(fun(X119,X118),X120),s('type.semanticPrimitives.error_result'(X119),'const.semanticPrimitives.Rraise_1'(s(X119,X121)))))))) )),
  inference(superposition,[],[f574,f108])).
fof(f3412,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s('type.semanticPrimitives.error_result'(X63),'const.semanticPrimitives.Rraise_1'(s(X63,X64))) != s('type.semanticPrimitives.error_result'(X63),'const.evalProps.map_error_result_2'(s(fun(bool,X63),X65),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))))) | s(X63,X64) = s(X63,happ(s(fun(bool,X63),X65),s(bool,'T_0'))) | s(bool,X62) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f574,f125])).
fof(f3411,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,X59))) != s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(bool,X58),X60),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))))) | s(X58,X59) = s(X58,happ(s(fun(bool,X58),X60),s(bool,'T_0'))) | s(bool,X57) = s(bool,X61) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X61)),bool))) )),
  inference(superposition,[],[f574,f689])).
fof(f3406,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.error_result'(X35),'const.semanticPrimitives.Rraise_1'(s(X35,X36))) != s('type.semanticPrimitives.error_result'(X35),'const.evalProps.map_error_result_2'(s(fun(bool,X35),X37),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(X35,X36) = s(X35,happ(s(fun(bool,X35),X37),s(bool,X34))) | s(bool,X34) = s(bool,X38) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X38)),bool))) )),
  inference(superposition,[],[f574,f689])).
fof(f3402,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,X17))) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(bool,X16),X18),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))))) | s(X16,X17) = s(X16,happ(s(fun(bool,X16),X18),s(bool,X14))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f574,f125])).
fof(f3401,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rraise_1'(s(X10,X11))) != s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(bool,X10),X12),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(X10,X11) = s(X10,happ(s(fun(bool,X10),X12),s(bool,X9))) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f574,f125])).
fof(f3395,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))) != s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X9,X5),X10),s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rraise_1'(s(X9,X11))))) | s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))) = s(X5,happ(s(fun(X9,X5),X10),s(X9,X11)))) )),
  inference(superposition,[],[f574,f108])).
fof(f3392,plain,(
  ( ! [X163,X161,X164,X162,X160] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X160))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.evalProps.map_error_result_2'(s(fun(X162,'type.semanticPrimitives.abort'),X163),s('type.semanticPrimitives.error_result'(X162),'const.semanticPrimitives.Rraise_1'(s(X162,X164))))) | s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X160)),X161)) = s('type.semanticPrimitives.abort',happ(s(fun(X162,'type.semanticPrimitives.abort'),X163),s(X162,X164)))) )),
  inference(superposition,[],[f574,f796])).
fof(f3388,plain,(
  ( ! [X134,X140,X138,X136,X135,X139,X137] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X134)),'const.evalProps.map_error_result_2'(s(fun(X138,'type.semanticPrimitives.error_result'(X134)),X139),s('type.semanticPrimitives.error_result'(X138),'const.semanticPrimitives.Rraise_1'(s(X138,X140))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X134)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X134),'const.evalProps.map_error_result_2'(s(fun(X135,X134),X136),s('type.semanticPrimitives.error_result'(X135),'const.semanticPrimitives.Rraise_1'(s(X135,X137))))))) | s('type.semanticPrimitives.error_result'(X134),happ(s(fun(X138,'type.semanticPrimitives.error_result'(X134)),X139),s(X138,X140))) = s('type.semanticPrimitives.error_result'(X134),'const.semanticPrimitives.Rraise_1'(s(X134,happ(s(fun(X135,X134),X136),s(X135,X137)))))) )),
  inference(superposition,[],[f574,f108])).
fof(f3387,plain,(
  ( ! [X132,X130,X133,X131,X129] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X129)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X129),X130))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X129)),'const.evalProps.map_error_result_2'(s(fun(X131,'type.semanticPrimitives.error_result'(X129)),X132),s('type.semanticPrimitives.error_result'(X131),'const.semanticPrimitives.Rraise_1'(s(X131,X133))))) | s('type.semanticPrimitives.error_result'(X129),'const.semanticPrimitives.Rraise_1'(s(X129,sK5(X130,X129)))) = s('type.semanticPrimitives.error_result'(X129),happ(s(fun(X131,'type.semanticPrimitives.error_result'(X129)),X132),s(X131,X133))) | s('type.semanticPrimitives.error_result'(X129),X130) = s('type.semanticPrimitives.error_result'(X129),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X130,X129))))) )),
  inference(superposition,[],[f574,f95])).
fof(f3384,plain,(
  ( ! [X121,X118,X116,X120,X119,X117,X115] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X115)),'const.evalProps.map_error_result_2'(s(fun(X119,'type.semanticPrimitives.error_result'(X115)),X120),s('type.semanticPrimitives.error_result'(X119),'const.semanticPrimitives.Rraise_1'(s(X119,X121))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X115)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X115),'const.semanticPrimitives.Rraise_1'(s(X115,happ(s(fun(X116,X115),X117),s(X116,X118))))))) | s('type.semanticPrimitives.error_result'(X115),happ(s(fun(X119,'type.semanticPrimitives.error_result'(X115)),X120),s(X119,X121))) = s('type.semanticPrimitives.error_result'(X115),'const.evalProps.map_error_result_2'(s(fun(X116,X115),X117),s('type.semanticPrimitives.error_result'(X116),'const.semanticPrimitives.Rraise_1'(s(X116,X118)))))) )),
  inference(superposition,[],[f574,f108])).
fof(f3383,plain,(
  ( ! [X111,X109,X114,X112,X110,X108,X113] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X108)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X108),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X111))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X108)),'const.evalProps.map_error_result_2'(s(fun(X112,'type.semanticPrimitives.error_result'(X108)),X113),s('type.semanticPrimitives.error_result'(X112),'const.semanticPrimitives.Rraise_1'(s(X112,X114))))) | s('type.semanticPrimitives.error_result'(X108),happ(s(fun(X112,'type.semanticPrimitives.error_result'(X108)),X113),s(X112,X114))) = s('type.semanticPrimitives.error_result'(X108),'const.evalProps.map_error_result_2'(s(fun(X109,X108),X110),s('type.semanticPrimitives.error_result'(X109),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X111)))))) )),
  inference(superposition,[],[f574,f107])).
fof(f3365,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) != s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X2,X0),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4))))) | s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0)) = s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))) )),
  inference(superposition,[],[f574,f596])).
fof(f3202,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool))) )),
  inference(superposition,[],[f2823,f689])).
fof(f3235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,X0) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool))) )),
  inference(condensation,[],[f3234])).
fof(f3234,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool))) )),
  inference(duplicate_literal_removal,[],[f3188])).
fof(f3188,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f2823,f689])).
fof(f2811,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f689])).
fof(f2823,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f2801])).
fof(f2801,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f689])).
fof(f2800,plain,(
  ( ! [X177,X175,X176,X174] : (s(X176,f(s(bool,'T_0'))) != s(X176,g(s(bool,'T_0'))) | s(fun(bool,X176),X174) = s(fun(bool,X176),X175) | s(bool,X177) = s(bool,sK7(X174,X175,X176,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X177)),bool))) )),
  inference(superposition,[],[f109,f689])).
fof(f2789,plain,(
  ( ! [X134,X136,X135] : (s(bool,'T_0') != s(bool,X135) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X134))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X135))) | s(bool,X134) = s(bool,X136) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X136)),bool))) )),
  inference(superposition,[],[f101,f689])).
fof(f2787,plain,(
  ( ! [X130,X128,X129] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X129))) | s(bool,X128) = s(bool,X129) | s(bool,X128) = s(bool,X130) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X130)),bool))) )),
  inference(superposition,[],[f100,f689])).
fof(f2726,plain,(
  ( ! [X92,X93,X91] : (s(bool,X91) != s(bool,X92) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X92))) | s(bool,X91) = s(bool,X93) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X93)),bool))) )),
  inference(superposition,[],[f101,f689])).
fof(f3131,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool))) )),
  inference(subsumption_resolution,[],[f2716,f434])).
fof(f2716,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) != s(bool,X63) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X64)),bool))) )),
  inference(superposition,[],[f475,f689])).
fof(f2701,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X20)),bool)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool))) )),
  inference(superposition,[],[f2204,f689])).
fof(f2687,plain,(
  ( ! [X101,X99,X100] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))) | s(bool,X100) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X99)),bool)) | s(bool,X99) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f100,f689])).
fof(f3117,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X97))) | s(bool,X97) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2686,f102])).
fof(f2686,plain,(
  ( ! [X97,X98] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X97))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X97))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X97)),bool)))) | s(bool,X97) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f95,f689])).
fof(f3109,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,X48) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X47)),bool)) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2667,f76])).
fof(f2667,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | s(bool,X48) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X47)),bool)) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f422,f689])).
fof(f3108,plain,(
  ( ! [X43,X41,X42] : (s(bool,X42) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X41)),bool)) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2665,f76])).
fof(f2665,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | s(bool,X42) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X41)),bool)) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f127,f689])).
fof(f2648,plain,(
  ( ! [X191,X189,X190,X188] : (s(X190,f(s(bool,X191))) != s(X190,g(s(bool,X191))) | s(fun(bool,X190),X188) = s(fun(bool,X190),X189) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X191)),bool)) | s(bool,'T_0') = s(bool,sK7(X188,X189,X190,bool))) )),
  inference(superposition,[],[f109,f689])).
fof(f3095,plain,(
  ( ! [X116,X114,X115] : (s(bool,X114) != s(bool,X115) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X115))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X116))) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool))) )),
  inference(subsumption_resolution,[],[f2622,f116])).
fof(f2622,plain,(
  ( ! [X116,X114,X115] : (s(bool,X114) != s(bool,X115) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X115))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X116))) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool)) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f480,f689])).
fof(f3092,plain,(
  ( ! [X111,X112,X110] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X110))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X111))) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X110)),bool))) )),
  inference(subsumption_resolution,[],[f2620,f116])).
fof(f2620,plain,(
  ( ! [X111,X112,X110] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X110))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X111))) | s(bool,X111) = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X110)),bool)) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f441,f689])).
fof(f2565,plain,(
  ( ! [X146,X147,X148] : (s(bool,X147) != s(bool,X148) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X146))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X148))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X146)),bool)) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f101,f689])).
fof(f2563,plain,(
  ( ! [X142,X140,X141] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X141))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X142))) | s(bool,X140) = s(bool,X142) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X140)),bool)) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f100,f689])).
fof(f2482,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X6)),bool)) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f689,f125])).
fof(f2481,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X2)),bool)) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X3)),bool))) )),
  inference(superposition,[],[f689,f689])).
fof(f2350,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2204,f125])).
fof(f2328,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.semanticPrimitives.error_result'(X3),X7) != s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X4,X3),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6))))) | s(X3,sK5(X7,X3)) = s(X3,happ(s(fun(X4,X3),X5),s(X4,X6))) | s('type.semanticPrimitives.error_result'(X3),X7) = s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X7,X3))))) )),
  inference(superposition,[],[f537,f108])).
fof(f2327,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) != s('type.semanticPrimitives.error_result'(X0),X2) | s(X0,sK5(X1,X0)) = s(X0,sK5(X2,X0)) | s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X2,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f537,f95])).
fof(f2325,plain,(
  ( ! [X76,X77,X75] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),X77) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X75))) | s('type.semanticPrimitives.abort',sK5(X77,'type.semanticPrimitives.abort')) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X75)),X76)) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),X77) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X77,'type.semanticPrimitives.abort'))))) )),
  inference(superposition,[],[f537,f796])).
fof(f2319,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),X65) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X61),'const.evalProps.map_error_result_2'(s(fun(X62,X61),X63),s('type.semanticPrimitives.error_result'(X62),'const.semanticPrimitives.Rraise_1'(s(X62,X64))))))) | s('type.semanticPrimitives.error_result'(X61),sK5(X65,'type.semanticPrimitives.error_result'(X61))) = s('type.semanticPrimitives.error_result'(X61),'const.semanticPrimitives.Rraise_1'(s(X61,happ(s(fun(X62,X61),X63),s(X62,X64))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),X65) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X61)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X65,'type.semanticPrimitives.error_result'(X61)))))) )),
  inference(superposition,[],[f537,f108])).
fof(f2318,plain,(
  ( ! [X59,X60,X58] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),X60) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X58),X59))) | s('type.semanticPrimitives.error_result'(X58),sK5(X60,'type.semanticPrimitives.error_result'(X58))) = s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rraise_1'(s(X58,sK5(X59,X58)))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),X60) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X58)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X60,'type.semanticPrimitives.error_result'(X58))))) | s('type.semanticPrimitives.error_result'(X58),X59) = s('type.semanticPrimitives.error_result'(X58),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X59,X58))))) )),
  inference(superposition,[],[f537,f95])).
fof(f2315,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),X54) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rraise_1'(s(X50,happ(s(fun(X51,X50),X52),s(X51,X53))))))) | s('type.semanticPrimitives.error_result'(X50),sK5(X54,'type.semanticPrimitives.error_result'(X50))) = s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X51,X50),X52),s('type.semanticPrimitives.error_result'(X51),'const.semanticPrimitives.Rraise_1'(s(X51,X53))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),X54) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X50)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X54,'type.semanticPrimitives.error_result'(X50)))))) )),
  inference(superposition,[],[f537,f108])).
fof(f2314,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X45)),X49) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X45)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X45),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X48))))) | s('type.semanticPrimitives.error_result'(X45),sK5(X49,'type.semanticPrimitives.error_result'(X45))) = s('type.semanticPrimitives.error_result'(X45),'const.evalProps.map_error_result_2'(s(fun(X46,X45),X47),s('type.semanticPrimitives.error_result'(X46),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X48))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X45)),X49) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X45)),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X49,'type.semanticPrimitives.error_result'(X45)))))) )),
  inference(superposition,[],[f537,f107])).
fof(f2296,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X2) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X1))) | s(X0,sK5(X2,X0)) = s(X0,sK5('const.semanticPrimitives.Rraise_1'(s(X0,X1)),X0)) | s('type.semanticPrimitives.error_result'(X0),X2) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X2,X0))))) )),
  inference(superposition,[],[f537,f596])).
fof(f2339,plain,(
  ( ! [X30,X31,X32] : (s('type.semanticPrimitives.error_result'(X30),X31) = s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),X31)),'type.semanticPrimitives.error_result'(X30)),X30)))) | s('type.semanticPrimitives.error_result'(X30),X31) != s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rraise_1'(s(X30,X32))) | s(X30,X32) = s(X30,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),X31)),'type.semanticPrimitives.error_result'(X30)),X30))) )),
  inference(forward_demodulation,[],[f2295,f596])).
fof(f2295,plain,(
  ( ! [X30,X31,X32] : (s('type.semanticPrimitives.error_result'(X30),X31) != s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rraise_1'(s(X30,X32))) | s(X30,X32) = s(X30,sK5(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),X31)),'type.semanticPrimitives.error_result'(X30)),X30)) | s('type.semanticPrimitives.error_result'(X30),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),X31)),'type.semanticPrimitives.error_result'(X30))) = s('type.semanticPrimitives.error_result'(X30),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X30),X31)),'type.semanticPrimitives.error_result'(X30)),X30))))) )),
  inference(superposition,[],[f537,f596])).
fof(f2337,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,X20))) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))) | s(X16,X20) = s(X16,sK5('const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19)))),X16))) )),
  inference(subsumption_resolution,[],[f2289,f102])).
fof(f2289,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,X20))) != s('type.semanticPrimitives.error_result'(X16),'const.evalProps.map_error_result_2'(s(fun(X17,X16),X18),s('type.semanticPrimitives.error_result'(X17),'const.semanticPrimitives.Rraise_1'(s(X17,X19))))) | s(X16,X20) = s(X16,sK5('const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19)))),X16)) | s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19))))) = s('type.semanticPrimitives.error_result'(X16),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X16,happ(s(fun(X17,X16),X18),s(X17,X19)))),X16))))) )),
  inference(superposition,[],[f537,f108])).
fof(f2336,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'(X13),X14) != s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X15))) | s(X13,X15) = s(X13,sK5('const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13))),X13)) | s('type.semanticPrimitives.error_result'(X13),X14) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X14,X13))))) )),
  inference(subsumption_resolution,[],[f2288,f102])).
fof(f2288,plain,(
  ( ! [X14,X15,X13] : (s('type.semanticPrimitives.error_result'(X13),X14) != s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,X15))) | s(X13,X15) = s(X13,sK5('const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13))),X13)) | s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13)))) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(X13,sK5(X14,X13))),X13)))) | s('type.semanticPrimitives.error_result'(X13),X14) = s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X14,X13))))) )),
  inference(superposition,[],[f537,f95])).
fof(f2335,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X9))) != s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))))) | s(X5,X9) = s(X5,sK5('const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8)))),X5))) )),
  inference(subsumption_resolution,[],[f2285,f576])).
fof(f2285,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X9))) != s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,happ(s(fun(X6,X5),X7),s(X6,X8))))) | s(X5,X9) = s(X5,sK5('const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8)))),X5)) | s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))))) = s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8)))),X5))))) )),
  inference(superposition,[],[f537,f108])).
fof(f2283,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,sK5('const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))),X0))) )),
  inference(unit_resulting_resolution,[],[f108,f576,f537])).
fof(f2204,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) )),
  inference(subsumption_resolution,[],[f2203,f115])).
fof(f2203,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) )),
  inference(forward_demodulation,[],[f2032,f596])).
fof(f2032,plain,(
  ( ! [X1] : (p(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) | s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool))) )),
  inference(superposition,[],[f85,f1429])).
fof(f1987,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool))) )),
  inference(superposition,[],[f1429,f1429])).
fof(f1286,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X1)),bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f687])).
fof(f1277,plain,(
  ( ! [X64,X65,X63] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X64))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X65))) | s(bool,X65) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X63)),bool)) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f100,f687])).
fof(f1373,plain,(
  ( ! [X61,X62] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1276,f102])).
fof(f1276,plain,(
  ( ! [X61,X62] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X62))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rraise_1'(s(bool,X61)),bool)))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f95,f687])).
fof(f1242,plain,(
  ( ! [X80,X81,X79] : (s(bool,X81) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X80)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X79))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f101,f687])).
fof(f1196,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X29)),bool)) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X28)),bool))) )),
  inference(superposition,[],[f687,f596])).
fof(f1046,plain,(
  ( ! [X26,X24,X25] : (s('type.semanticPrimitives.error_result'(X24),X25) = s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X24),X25)),'type.semanticPrimitives.error_result'(X24)),X24)))) | s('type.semanticPrimitives.error_result'(X24),X25) != s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X26)))) )),
  inference(forward_demodulation,[],[f1035,f596])).
fof(f1035,plain,(
  ( ! [X26,X24,X25] : (s('type.semanticPrimitives.error_result'(X24),X25) != s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X26))) | s('type.semanticPrimitives.error_result'(X24),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X24),X25)),'type.semanticPrimitives.error_result'(X24))) = s('type.semanticPrimitives.error_result'(X24),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X24),X25)),'type.semanticPrimitives.error_result'(X24)),X24))))) )),
  inference(superposition,[],[f539,f596])).
fof(f1042,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))),X0))))) )),
  inference(forward_demodulation,[],[f1021,f107])).
fof(f1021,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3)))),X0))))) )),
  inference(unit_resulting_resolution,[],[f107,f539])).
fof(f1039,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(forward_demodulation,[],[f1024,f596])).
fof(f1024,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)))),'type.semanticPrimitives.error_result'(X0)),X0))))) )),
  inference(unit_resulting_resolution,[],[f596,f539])).
fof(f1012,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f457,f83])).
fof(f969,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),X1))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f433,f95])).
fof(f947,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2)))))))) )),
  inference(unit_resulting_resolution,[],[f433,f100])).
fof(f886,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.error_result'(X9),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))) != s('type.semanticPrimitives.error_result'(X9),'const.evalProps.map_error_result_2'(s(fun(X5,X9),X11),s('type.semanticPrimitives.error_result'(X5),'const.evalProps.map_error_result_2'(s(fun(X6,X5),X7),s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8)))))))) )),
  inference(superposition,[],[f576,f108])).
fof(f863,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3))))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5('const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))),X0))))) )),
  inference(unit_resulting_resolution,[],[f576,f95])).
fof(f864,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X2,X0),X3),s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,X4)))))))) )),
  inference(unit_resulting_resolution,[],[f576,f100])).
fof(f855,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f844,f76])).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f789,f125])).
fof(f830,plain,(
  ( ! [X26,X27,X25] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X25))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.abort'),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X27))) | s('type.semanticPrimitives.abort',X27) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X25)),X26))) )),
  inference(superposition,[],[f100,f796])).
fof(f824,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0))) != s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X3))) | s('type.semanticPrimitives.abort',X3) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)),X1))) )),
  inference(superposition,[],[f98,f796])).
fof(f800,plain,(
  ( ! [X6,X4,X5] : (s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X5))) != s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X6))) | s('type.semanticPrimitives.abort',X6) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X5)),X4))) )),
  inference(superposition,[],[f98,f531])).
fof(f796,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.abort',X0) = s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X0)),X1))) )),
  inference(unit_resulting_resolution,[],[f531,f98])).
fof(f715,plain,(
  ( ! [X54,X55,X53] : (s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X53))) != s('type.semanticPrimitives.error_result'(X54),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X55))) | s('type.semanticPrimitives.abort',X55) = s('type.semanticPrimitives.abort',sK5('const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.abort',X53)),'type.semanticPrimitives.abort'))) )),
  inference(superposition,[],[f98,f596])).
fof(f689,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f596,f125])).
fof(f687,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f596,f125])).
fof(f656,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK7(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f109,f125])).
fof(f654,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK7(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f109,f125])).
fof(f619,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X7))) != s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X8))) | s(X6,X8) = s(X6,sK5('const.semanticPrimitives.Rraise_1'(s(X6,X7)),X6))) )),
  inference(superposition,[],[f100,f530])).
fof(f615,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X4)),bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f530,f125])).
fof(f586,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),X17))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X13)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X13),'const.semanticPrimitives.Rraise_1'(s(X13,happ(s(fun(X14,X13),X15),s(X14,X16))))))) | s('type.semanticPrimitives.error_result'(X13),X17) = s('type.semanticPrimitives.error_result'(X13),'const.evalProps.map_error_result_2'(s(fun(X14,X13),X15),s('type.semanticPrimitives.error_result'(X14),'const.semanticPrimitives.Rraise_1'(s(X14,X16)))))) )),
  inference(superposition,[],[f100,f108])).
fof(f584,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.semanticPrimitives.error_result'(X3),X7) != s('type.semanticPrimitives.error_result'(X3),'const.semanticPrimitives.Rraise_1'(s(X3,happ(s(fun(X4,X3),X5),s(X4,X6))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X3)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X3),X7))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X3)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X4,X3),X5),s('type.semanticPrimitives.error_result'(X4),'const.semanticPrimitives.Rraise_1'(s(X4,X6)))))))) )),
  inference(superposition,[],[f101,f108])).
fof(f580,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),X35))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X31)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X31),'const.evalProps.map_error_result_2'(s(fun(X32,X31),X33),s('type.semanticPrimitives.error_result'(X32),'const.semanticPrimitives.Rraise_1'(s(X32,X34))))))) | s('type.semanticPrimitives.error_result'(X31),X35) = s('type.semanticPrimitives.error_result'(X31),'const.semanticPrimitives.Rraise_1'(s(X31,happ(s(fun(X32,X31),X33),s(X32,X34)))))) )),
  inference(superposition,[],[f100,f108])).
fof(f578,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.semanticPrimitives.error_result'(X21),X25) != s('type.semanticPrimitives.error_result'(X21),'const.evalProps.map_error_result_2'(s(fun(X22,X21),X23),s('type.semanticPrimitives.error_result'(X22),'const.semanticPrimitives.Rraise_1'(s(X22,X24))))) | s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),X25))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X21)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X21),'const.semanticPrimitives.Rraise_1'(s(X21,happ(s(fun(X22,X21),X23),s(X22,X24)))))))) )),
  inference(superposition,[],[f101,f108])).
fof(f574,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X4))) != s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3))))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X3)))) )),
  inference(superposition,[],[f100,f108])).
fof(f573,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.semanticPrimitives.error_result'(X8),'const.semanticPrimitives.Rraise_1'(s(X8,happ(s(fun(X4,X8),X9),s(X4,happ(s(fun(X5,X4),X6),s(X5,X7))))))) = s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X4,X8),X9),s('type.semanticPrimitives.error_result'(X4),'const.evalProps.map_error_result_2'(s(fun(X5,X4),X6),s('type.semanticPrimitives.error_result'(X5),'const.semanticPrimitives.Rraise_1'(s(X5,X7)))))))) )),
  inference(superposition,[],[f108,f108])).
fof(f572,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X0,X2),X3),s('type.semanticPrimitives.error_result'(X0),X1))) = s('type.semanticPrimitives.error_result'(X2),'const.semanticPrimitives.Rraise_1'(s(X2,happ(s(fun(X0,X2),X3),s(X0,sK5(X1,X0)))))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f108,f95])).
fof(f558,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))))))) = s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rraise_1'(s(X1,X3)))))))) )),
  inference(unit_resulting_resolution,[],[f108,f101])).
fof(f542,plain,(
  ( ! [X17,X15,X16] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X15)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X15),X16))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X15)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X15),X17))) | s('type.semanticPrimitives.error_result'(X15),X17) = s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rraise_1'(s(X15,sK5(X16,X15)))) | s('type.semanticPrimitives.error_result'(X15),X16) = s('type.semanticPrimitives.error_result'(X15),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X16,X15))))) )),
  inference(superposition,[],[f100,f95])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) != s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK5(X1,X0)) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0))))) )),
  inference(superposition,[],[f100,f95])).
fof(f531,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6('const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1)),X0))))) )),
  inference(unit_resulting_resolution,[],[f102,f95])).
fof(f504,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),X14))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(X10)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(X10),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X13))))) | s('type.semanticPrimitives.error_result'(X10),X14) = s('type.semanticPrimitives.error_result'(X10),'const.evalProps.map_error_result_2'(s(fun(X11,X10),X12),s('type.semanticPrimitives.error_result'(X11),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X13)))))) )),
  inference(superposition,[],[f100,f107])).
fof(f481,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X18))) | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f101,f125])).
fof(f478,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,X11) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X9))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f101,f125])).
fof(f476,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X3))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f101,f125])).
fof(f439,plain,(
  ( ! [X10,X11,X9] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X10))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X11))) | s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f100,f125])).
fof(f437,plain,(
  ( ! [X4,X5,X3] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X4))) | s(bool,X3) = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f100,f125])).
fof(f422,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f273,f76])).
fof(f273,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f106,f125])).
fof(f157,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f154,f115])).
fof(f154,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f103,f83])).
fof(f116,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f115,f85])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.map_error_result_Rtype_error',aHL_EXT)).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.abort',X1) != s('type.semanticPrimitives.abort',X2) | s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X1))) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f86,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f64])).
fof(f15528,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f115,f15062])).
fof(f15527,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f434,f15466])).
fof(f15526,plain,(
  $false | $spl80),
  inference(trivial_inequality_removal,[],[f15525])).
fof(f15525,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | $spl80),
  inference(forward_demodulation,[],[f457,f15466])).
fof(f15522,plain,(
  $false | $spl80),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964,f14965,f14881,f14974,f14878,f14986,f15009,f14729,f15011,f15012,f15013,f15014,f15015,f15016,f15017,f15018,f15019,f14749,f15020,f15022,f15024,f15025,f15026,f15027,f15028,f15039,f15052,f15079,f14826,f15080,f14833,f15082,f14837,f15083,f15084,f15085,f15086,f14844,f15062,f2576,f15456,f15465,f15466,f15496,f15499,f15509,f15521])).
fof(f15510,plain,(
  $false | $spl80),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701,f12125,f12127,f12131,f12132,f12149,f12151,f12152,f12154,f12155,f12156,f12157,f12167,f12169,f12170,f12174,f12175,f12176,f12180,f12189,f12190,f12191,f12192,f12198,f12200,f12203,f12204,f12205,f12206,f12208,f12212,f12213,f12420,f12421,f12423,f12424,f12969,f13014,f13008,f13004,f13226,f13238,f13241,f13250,f13588,f13589,f13590,f13591,f13592,f13593,f13594,f13595,f13596,f13599,f13600,f13602,f13620,f13624,f13631,f13648,f13649,f13650,f13651,f13652,f13653,f13697,f13784,f13785,f13786,f13788,f13703,f13705,f13707,f13708,f13720,f13936,f14001,f14002,f14263,f14264,f14265,f14266,f14248,f14908,f14907,f14905,f14601,f14934,f14964,f14965,f14881,f14974,f14878,f14986,f15009,f14729,f15011,f15012,f15013,f15014,f15015,f15016,f15017,f15018,f15019,f14749,f15020,f15022,f15024,f15025,f15026,f15027,f15028,f15039,f15052,f15079,f14826,f15080,f14833,f15082,f14837,f15083,f15084,f15085,f15086,f14844,f15062,f2576,f15456,f15465,f15466,f15496,f15499,f15509])).
fof(f15505,plain,(
  $false | $spl80),
  inference(trivial_inequality_removal,[],[f15504])).
fof(f15504,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | $spl80),
  inference(forward_demodulation,[],[f1011,f15466])).
fof(f15500,plain,(
  $false | ($spl11 | $spl80)),
  inference(subsumption_resolution,[],[f1595,f15456])).
fof(f15489,plain,(
  $false | ($spl13 | $spl80)),
  inference(subsumption_resolution,[],[f3163,f15456])).
fof(f15452,plain,(
  $false | ($spl4 | $spl13)),
  inference(subsumption_resolution,[],[f15177,f1567])).
fof(f15177,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl13)),
  inference(backward_demodulation,[],[f1567,f3163])).
fof(f15447,plain,(
  $false | ($spl4 | $spl11)),
  inference(subsumption_resolution,[],[f15135,f1567])).
fof(f15135,plain,(
  s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | ($spl4 | $spl11)),
  inference(backward_demodulation,[],[f1567,f1595])).
fof(f15436,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15090,f76])).
fof(f15090,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f1567,f115])).
fof(f15435,plain,(
  $false | ($spl4 | $spl80)),
  inference(subsumption_resolution,[],[f15434,f15062])).
fof(f15434,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f15088,f1567])).
fof(f15088,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1567,f89])).
fof(f15433,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f15091])).
fof(f15091,plain,(
  s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | $spl4),
  inference(backward_demodulation,[],[f1567,f434])).
fof(f15432,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f15094])).
fof(f15094,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | $spl4),
  inference(backward_demodulation,[],[f1567,f457])).
fof(f15431,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f15119])).
fof(f15119,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | $spl4),
  inference(backward_demodulation,[],[f1567,f1011])).
fof(f14993,plain,(
  ( ! [X118] : (s(bool,X118) = s(bool,g(s(bool,X118))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X118)))) ) | $spl74),
  inference(cnf_transformation,[],[f14993_D])).
fof(f14993_D,plain,(
  ( ! [X118] : (s(bool,X118) = s(bool,g(s(bool,X118))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X118)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f15078,plain,(
  ( ! [X156,X157] : (s(bool,X156) = s(bool,g(s(bool,X156))) | p(s(bool,X157)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X156)))) ) | $spl13),
  inference(forward_demodulation,[],[f15077,f596])).
fof(f15077,plain,(
  ( ! [X156,X157] : (p(s(bool,X157)) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X156)),bool)) = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X156)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X156)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14822,f85])).
fof(f14822,plain,(
  ( ! [X156,X157] : (p(s(bool,X157)) | s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X156)),bool)) = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X156)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X156))) | s(bool,'T_0') = s(bool,X157)) ) | $spl13),
  inference(superposition,[],[f14021,f613])).
fof(f15076,plain,(
  ( ! [X154] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X154))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X154)))) ) | $spl84),
  inference(cnf_transformation,[],[f15076_D])).
fof(f15076_D,plain,(
  ( ! [X154] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X154))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X154)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f15075,plain,(
  ( ! [X154,X155] : (s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,X154))))) | p(s(bool,X155)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X154)))) ) | $spl13),
  inference(forward_demodulation,[],[f15074,f596])).
fof(f15074,plain,(
  ( ! [X154,X155] : (p(s(bool,X155)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X154)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X154)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14821,f85])).
fof(f14821,plain,(
  ( ! [X154,X155] : (p(s(bool,X155)) | s(bool,'F_0') = s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X154)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X154))) | s(bool,'T_0') = s(bool,X155)) ) | $spl13),
  inference(superposition,[],[f13357,f613])).
fof(f15073,plain,(
  ( ! [X152] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X152))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152)))) ) | $spl82),
  inference(cnf_transformation,[],[f15073_D])).
fof(f15073_D,plain,(
  ( ! [X152] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X152))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f15072,plain,(
  ( ! [X152,X153] : (s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,X152))))) | p(s(bool,X153)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152)))) ) | $spl13),
  inference(forward_demodulation,[],[f15071,f596])).
fof(f15071,plain,(
  ( ! [X152,X153] : (p(s(bool,X153)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X152)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14820,f85])).
fof(f14820,plain,(
  ( ! [X152,X153] : (p(s(bool,X153)) | s(bool,'T_0') != s(bool,g(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X152)),bool)))))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X152))) | s(bool,'T_0') = s(bool,X153)) ) | $spl13),
  inference(superposition,[],[f13122,f613])).
fof(f14995,plain,(
  ( ! [X120] : (s(bool,'F_0') = s(bool,g(s(bool,X120))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X120)))) ) | $spl76),
  inference(cnf_transformation,[],[f14995_D])).
fof(f14995_D,plain,(
  ( ! [X120] : (s(bool,'F_0') = s(bool,g(s(bool,X120))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X120)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f15064,plain,(
  ( ! [X144,X145] : (s(bool,'F_0') = s(bool,g(s(bool,X144))) | p(s(bool,X145)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X144)))) ) | $spl13),
  inference(forward_demodulation,[],[f15063,f596])).
fof(f15063,plain,(
  ( ! [X144,X145] : (p(s(bool,X145)) | s(bool,'F_0') = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X144)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X144)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14817,f85])).
fof(f14817,plain,(
  ( ! [X144,X145] : (p(s(bool,X145)) | s(bool,'F_0') = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X144)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X144))) | s(bool,'T_0') = s(bool,X145)) ) | $spl13),
  inference(superposition,[],[f12893,f613])).
fof(f14990,plain,(
  ( ! [X114] : (s(bool,'T_0') != s(bool,g(s(bool,X114))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114)))) ) | $spl72),
  inference(cnf_transformation,[],[f14990_D])).
fof(f14990_D,plain,(
  ( ! [X114] : (s(bool,'T_0') != s(bool,g(s(bool,X114))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f15061,plain,(
  ( ! [X140,X141] : (s(bool,'T_0') != s(bool,g(s(bool,X140))) | p(s(bool,X141)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X140)))) ) | $spl13),
  inference(forward_demodulation,[],[f15060,f596])).
fof(f15060,plain,(
  ( ! [X140,X141] : (p(s(bool,X141)) | s(bool,'T_0') != s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X140)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X140)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14815,f85])).
fof(f14815,plain,(
  ( ! [X140,X141] : (p(s(bool,X141)) | s(bool,'T_0') != s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X140)),bool)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X140))) | s(bool,'T_0') = s(bool,X141)) ) | $spl13),
  inference(superposition,[],[f12891,f613])).
fof(f14984,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X95)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95)))) ) | $spl70),
  inference(cnf_transformation,[],[f14984_D])).
fof(f14984_D,plain,(
  ( ! [X95] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X95)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f15005,plain,(
  ( ! [X202,X203] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X203)),bool)) | s(bool,'T_0') = s(bool,X202) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X203)))) ) | $spl11),
  inference(subsumption_resolution,[],[f14722,f1646])).
fof(f1646,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f1595,f83])).
fof(f14722,plain,(
  ( ! [X202,X203] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X203)),bool)) | s(bool,'T_0') = s(bool,X202) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X203))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X202)),bool))) ) | $spl11),
  inference(superposition,[],[f1646,f613])).
fof(f15001,plain,(
  ( ! [X191,X190] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X191)),bool)) | s(bool,'T_0') = s(bool,X190) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X191)))) ) | $spl13),
  inference(subsumption_resolution,[],[f14717,f11835])).
fof(f11835,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl13),
  inference(superposition,[],[f3163,f83])).
fof(f14717,plain,(
  ( ! [X191,X190] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X191)),bool)) | s(bool,'T_0') = s(bool,X190) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X191))) | s(bool,'T_0') = s(bool,g(s(bool,X190)))) ) | $spl13),
  inference(superposition,[],[f11835,f613])).
fof(f14997,plain,(
  ( ! [X124] : (~p(s(bool,g(s(bool,X124)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124)))) ) | $spl78),
  inference(cnf_transformation,[],[f14997_D])).
fof(f14997_D,plain,(
  ( ! [X124] : (~p(s(bool,g(s(bool,X124)))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f14996,plain,(
  ( ! [X123,X124] : (~p(s(bool,g(s(bool,X124)))) | s(bool,'T_0') = s(bool,X123) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124)))) ) | $spl13),
  inference(forward_demodulation,[],[f14860,f596])).
fof(f14860,plain,(
  ( ! [X123,X124] : (~p(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X124)),bool))))) | s(bool,'T_0') = s(bool,X123) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124)))) ) | $spl13),
  inference(duplicate_literal_removal,[],[f14706])).
fof(f14706,plain,(
  ( ! [X123,X124] : (~p(s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X124)),bool))))) | s(bool,'T_0') = s(bool,X123) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X124))) | s(bool,'T_0') = s(bool,X123)) ) | $spl13),
  inference(superposition,[],[f12892,f613])).
fof(f12892,plain,(
  ( ! [X7] : (~p(s(bool,g(s(bool,X7)))) | s(bool,'T_0') = s(bool,X7)) ) | $spl13),
  inference(subsumption_resolution,[],[f12807,f83])).
fof(f12807,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,X7) | ~p(s(bool,g(s(bool,X7)))) | s(bool,'T_0') = s(bool,X7)) ) | $spl13),
  inference(superposition,[],[f89,f11907])).
fof(f14994,plain,(
  ( ! [X120,X119] : (s(bool,'F_0') = s(bool,g(s(bool,X120))) | s(bool,'T_0') = s(bool,X119) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X120)))) ) | ($spl13 | $spl37)),
  inference(forward_demodulation,[],[f14861,f596])).
fof(f14861,plain,(
  ( ! [X120,X119] : (s(bool,'F_0') = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X120)),bool)))) | s(bool,'T_0') = s(bool,X119) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X120)))) ) | ($spl13 | $spl37)),
  inference(duplicate_literal_removal,[],[f14704])).
fof(f14704,plain,(
  ( ! [X120,X119] : (s(bool,'F_0') = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X120)),bool)))) | s(bool,'T_0') = s(bool,X119) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X120))) | s(bool,'T_0') = s(bool,X119)) ) | ($spl13 | $spl37)),
  inference(superposition,[],[f12416,f613])).
fof(f12416,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,g(s(bool,X19))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl13 | $spl37)),
  inference(subsumption_resolution,[],[f12395,f4193])).
fof(f4193,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl37),
  inference(unit_resulting_resolution,[],[f76,f116,f3496,f422])).
fof(f12395,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,g(s(bool,X19)))) ) | ($spl13 | $spl37)),
  inference(superposition,[],[f11835,f6707])).
fof(f6707,plain,(
  ( ! [X44] : (s(bool,X44) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X44)) ) | $spl37),
  inference(subsumption_resolution,[],[f6682,f76])).
fof(f6682,plain,(
  ( ! [X44] : (~p(s(bool,'T_0')) | s(bool,X44) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X44)) ) | $spl37),
  inference(superposition,[],[f789,f4193])).
fof(f14992,plain,(
  ( ! [X118,X117] : (s(bool,X118) = s(bool,g(s(bool,X118))) | s(bool,'T_0') = s(bool,X117) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X118)))) ) | $spl13),
  inference(forward_demodulation,[],[f14862,f596])).
fof(f14862,plain,(
  ( ! [X118,X117] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X118)),bool)) = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X118)),bool)))) | s(bool,'T_0') = s(bool,X117) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X118)))) ) | $spl13),
  inference(duplicate_literal_removal,[],[f14703])).
fof(f14703,plain,(
  ( ! [X118,X117] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X118)),bool)) = s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X118)),bool)))) | s(bool,'T_0') = s(bool,X117) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X118))) | s(bool,'T_0') = s(bool,X117)) ) | $spl13),
  inference(superposition,[],[f11907,f613])).
fof(f14989,plain,(
  ( ! [X114,X113] : (s(bool,'T_0') != s(bool,g(s(bool,X114))) | s(bool,'T_0') = s(bool,X113) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114)))) ) | $spl13),
  inference(forward_demodulation,[],[f14864,f596])).
fof(f14864,plain,(
  ( ! [X114,X113] : (s(bool,'T_0') != s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool)))) | s(bool,'T_0') = s(bool,X113) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114)))) ) | $spl13),
  inference(duplicate_literal_removal,[],[f14701])).
fof(f14701,plain,(
  ( ! [X114,X113] : (s(bool,'T_0') != s(bool,g(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X114)),bool)))) | s(bool,'T_0') = s(bool,X113) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X114))) | s(bool,'T_0') = s(bool,X113)) ) | $spl13),
  inference(superposition,[],[f11835,f613])).
fof(f14988,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X106)),bool)) | s(bool,'T_0') = s(bool,X105) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X106)))) ) | $spl37),
  inference(forward_demodulation,[],[f14865,f4193])).
fof(f14865,plain,(
  ( ! [X105,X106] : (s(bool,f(s(bool,'T_0'))) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X106)),bool)) | s(bool,'T_0') = s(bool,X105) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X106)))) ) | $spl37),
  inference(duplicate_literal_removal,[],[f14697])).
fof(f14697,plain,(
  ( ! [X105,X106] : (s(bool,f(s(bool,'T_0'))) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X106)),bool)) | s(bool,'T_0') = s(bool,X105) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X106))) | s(bool,'T_0') = s(bool,X105)) ) | $spl37),
  inference(superposition,[],[f4535,f613])).
fof(f4535,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl37),
  inference(subsumption_resolution,[],[f4487,f76])).
fof(f4487,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl37),
  inference(superposition,[],[f127,f4193])).
fof(f14987,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X104)),bool)) | s(bool,'T_0') = s(bool,X103) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X104)))) ) | ($spl32 | $spl37)),
  inference(forward_demodulation,[],[f14866,f4221])).
fof(f4221,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl32 | $spl37)),
  inference(backward_demodulation,[],[f4193,f3362])).
fof(f3362,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl32),
  inference(cnf_transformation,[],[f3362_D])).
fof(f3362_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f14866,plain,(
  ( ! [X103,X104] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X104)),bool)) | s(bool,'T_0') = s(bool,X103) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X104)))) ) | ($spl32 | $spl37)),
  inference(duplicate_literal_removal,[],[f14696])).
fof(f14696,plain,(
  ( ! [X103,X104] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X104)),bool)) | s(bool,'T_0') = s(bool,X103) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X104))) | s(bool,'T_0') = s(bool,X103)) ) | ($spl32 | $spl37)),
  inference(superposition,[],[f4314,f613])).
fof(f4314,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl32 | $spl37)),
  inference(subsumption_resolution,[],[f4268,f76])).
fof(f4268,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl32 | $spl37)),
  inference(superposition,[],[f127,f4221])).
fof(f14986,plain,(
  ( ! [X101,X102] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X102)),bool)) | s(bool,'T_0') = s(bool,X101) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X102)))) )),
  inference(forward_demodulation,[],[f14867,f596])).
fof(f14867,plain,(
  ( ! [X101,X102] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X102)),bool)) | s(bool,'T_0') = s(bool,X101) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X102)))) )),
  inference(duplicate_literal_removal,[],[f14695])).
fof(f14695,plain,(
  ( ! [X101,X102] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X102)),bool)) | s(bool,'T_0') = s(bool,X101) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X102))) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f2823,f613])).
fof(f14985,plain,(
  ( ! [X99,X100] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X100)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))) | s(bool,'T_0') = s(bool,X99)) ) | $spl11),
  inference(subsumption_resolution,[],[f14694,f1646])).
fof(f14694,plain,(
  ( ! [X99,X100] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X100)),bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X99)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X100))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f2204,f613])).
fof(f14983,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X95)),bool)) | s(bool,'T_0') = s(bool,X94) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95)))) ) | $spl11),
  inference(forward_demodulation,[],[f14869,f596])).
fof(f14869,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X95)),bool))),bool)) | s(bool,'T_0') = s(bool,X94) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95)))) ) | $spl11),
  inference(duplicate_literal_removal,[],[f14692])).
fof(f14692,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X95)),bool))),bool)) | s(bool,'T_0') = s(bool,X94) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X95))) | s(bool,'T_0') = s(bool,X94)) ) | $spl11),
  inference(superposition,[],[f1646,f613])).
fof(f14976,plain,(
  ( ! [X59,X60] : (s(bool,X60) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X59)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X59))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60)))) ) | $spl68),
  inference(cnf_transformation,[],[f14976_D])).
fof(f14976_D,plain,(
  ( ! [X59,X60] : (s(bool,X60) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X59)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X59))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f14878,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X59)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) | s(bool,'T_0') = s(bool,X58) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X59)))) )),
  inference(duplicate_literal_removal,[],[f14679])).
fof(f14679,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X59)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) | s(bool,'T_0') = s(bool,X58) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X59))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f475,f613])).
fof(f14975,plain,(
  ( ! [X50,X51] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) | s(bool,'F_0') = s(bool,X51)) ) | $spl66),
  inference(cnf_transformation,[],[f14975_D])).
fof(f14975_D,plain,(
  ( ! [X50,X51] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) | s(bool,'F_0') = s(bool,X51)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f14974,plain,(
  ( ! [X50,X51,X49] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50)))) )),
  inference(forward_demodulation,[],[f14879,f596])).
fof(f14879,plain,(
  ( ! [X50,X51,X49] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X50)),bool)))) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50)))) )),
  inference(duplicate_literal_removal,[],[f14676])).
fof(f14676,plain,(
  ( ! [X50,X51,X49] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X51))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X50)),bool)))) | s(bool,'F_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X50))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f436,f613])).
fof(f14915,plain,(
  ( ! [X15,X16] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X15)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16)))) ) | $spl64),
  inference(cnf_transformation,[],[f14915_D])).
fof(f14915_D,plain,(
  ( ! [X15,X16] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X15)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14908,plain,(
  ( ! [X14,X15,X16] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X15)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15)))) )),
  inference(duplicate_literal_removal,[],[f14582])).
fof(f14582,plain,(
  ( ! [X14,X15,X16] : (s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X15)),bool)) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X16)),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X16))) | s(bool,'T_0') = s(bool,X14) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f613,f613])).
fof(f12111,plain,(
  ( ! [X177,X176] : (s(fun(fun(bool,bool),bool),X176) = s(fun(fun(bool,bool),bool),X177) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK7(X176,X177,bool,fun(bool,bool)))))) ) | $spl60),
  inference(cnf_transformation,[],[f12111_D])).
fof(f12111_D,plain,(
  ( ! [X177,X176] : (s(fun(fun(bool,bool),bool),X176) = s(fun(fun(bool,bool),bool),X177) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK7(X176,X177,bool,fun(bool,bool)))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f12112,plain,(
  ( ! [X178] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X178)))) ) | $spl62),
  inference(cnf_transformation,[],[f12112_D])).
fof(f12112_D,plain,(
  ( ! [X178] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X178)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f12108,plain,(
  ( ! [X177,X178,X176] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X178))) | s(fun(fun(bool,bool),bool),X176) = s(fun(fun(bool,bool),bool),X177) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK7(X176,X177,bool,fun(bool,bool)))))) ) | ($spl14 | $spl32)),
  inference(superposition,[],[f663,f11702])).
fof(f11702,plain,(
  ( ! [X163,X162] : (s(fun(bool,bool),X162) = s(fun(bool,bool),X163)) ) | ($spl14 | $spl32)),
  inference(global_subsumption,[],[f75,f74,f76,f83,f86,f85,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f433,f437,f439,f476,f478,f481,f504,f531,f530,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f619,f654,f656,f687,f689,f691,f715,f796,f800,f824,f830,f855,f864,f863,f885,f886,f947,f969,f1012,f1039,f1042,f1046,f1196,f1242,f1373,f1277,f1286,f1535,f1525,f1987,f2204,f2283,f2335,f2336,f2337,f2339,f2296,f2314,f2315,f2318,f2319,f2325,f2327,f2328,f2350,f2481,f2482,f2563,f2565,f3092,f3095,f2648,f3108,f3109,f3117,f2687,f2701,f3131,f2726,f2787,f2789,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3165,f3235,f3202,f3339,f3348,f3349,f3350,f3360,f3330,f3320,f3322,f3327,f3365,f3382,f3383,f3384,f3387,f3388,f3392,f3395,f3401,f3402,f3406,f3411,f3412,f3422,f3424,f3427,f3428,f3436,f3437,f3607,f3610,f3612,f2576,f1341,f2798,f2794,f2731,f2694,f2646,f2642,f2574,f2570,f2310,f2305,f2299,f2298,f534,f532,f3380,f3377,f3374,f3373,f3372,f3368,f3367,f3362,f3638,f3628,f2797,f2645,f2573,f1249,f661,f659,f3459,f3446,f3445,f3438,f3432,f3431,f3430,f3429,f3426,f3414,f3400,f3397,f3390,f3389,f3386,f3379,f3366,f3137,f3113,f2700,f2344,f2343,f2342,f2338,f2331,f2323,f2322,f2321,f2320,f2317,f2309,f2280,f2279,f2278,f2274,f2273,f2245,f2242,f2241,f2239,f2159,f2116,f2068,f1946,f1944,f1938,f1922,f1920,f1915,f1874,f1872,f1869,f1868,f1867,f1866,f1865,f1863,f1773,f1516,f1476,f1429,f1413,f1410,f1390,f1335,f1327,f1168,f1166,f1131,f1128,f1118,f1115,f1011,f1006,f930,f928,f924,f923,f915,f913,f909,f907,f906,f900,f894,f860,f792,f791,f790,f789,f773,f663,f658,f617,f536,f480,f475,f457,f441,f436,f434,f115,f90,f89,f6066,f6065,f6062,f6105,f6106,f6107,f6108,f6113,f6165,f6121,f6123,f6128,f6129,f6130,f6137,f6166,f6145,f6147,f6152,f6153,f6154,f6305,f6308,f6314,f6322,f6594,f6595,f6533,f6534,f6535,f6539,f6545,f6546,f6547,f6551,f6552,f6553,f6556,f6557,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6606,f6579,f6581,f6586,f6588,f6589,f6769,f6735,f6738,f6739,f6740,f6741,f6742,f6745,f7427,f7391,f7392,f7393,f7395,f7432,f7398,f7399,f7400,f7401,f7402,f7404,f7409,f7746,f7619,f7668,f7669,f7670,f7671,f7672,f7674,f7680,f7681,f7682,f7683,f7684,f7685,f7686,f7687,f7688,f7759,f7761,f7693,f7763,f7695,f7703,f7705,f7706,f7764,f7765,f7714,f7766,f7716,f7717,f7719,f7724,f7726,f7727,f8318,f8319,f8322,f8323,f8324,f8325,f8326,f8330,f8331,f8332,f8336,f8344,f8345,f8347,f8364,f8365,f8367,f8368,f8369,f8370,f8371,f8372,f8373,f8379,f8380,f8716,f8717,f8720,f8721,f8722,f8723,f8724,f8725,f8730,f8731,f8734,f8735,f8736,f8740,f8748,f8749,f8751,f8769,f8770,f8772,f8773,f8774,f8775,f8776,f8777,f8778,f8784,f8785,f9193,f9194,f9198,f9199,f9200,f9201,f9205,f9206,f9207,f9211,f9219,f9220,f9222,f9241,f9242,f9244,f9245,f9246,f9247,f9248,f9249,f9250,f9303,f9304,f9309,f9312,f9269,f9270,f9463,f9464,f9468,f9469,f9470,f9471,f9472,f9477,f9478,f9481,f9482,f9483,f9487,f9495,f9496,f9498,f9517,f9518,f9520,f9521,f9522,f9523,f9524,f9525,f9526,f9530,f9531,f9532,f9536,f9545,f9546,f10068,f10069,f10070,f10081,f10083,f10084,f9996,f10085,f10003,f10087,f10007,f10088,f10089,f10090,f10091,f11701])).
fof(f3422,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s('type.semanticPrimitives.error_result'(X108),'const.semanticPrimitives.Rraise_1'(s(X108,X109))) != s('type.semanticPrimitives.error_result'(X108),'const.evalProps.map_error_result_2'(s(fun(bool,X108),X110),s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | s(X108,X109) = s(X108,happ(s(fun(bool,X108),X110),s(bool,sK7(X106,X107,bool,bool)))) | s(fun(bool,bool),X106) = s(fun(bool,bool),X107)) ) | $spl14),
  inference(superposition,[],[f574,f3165])).
fof(f3382,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.evalProps.map_error_result_2'(s(fun(X105,bool),X106),s('type.semanticPrimitives.error_result'(X105),'const.semanticPrimitives.Rraise_1'(s(X105,X107))))) | s(bool,sK7(X103,X104,bool,bool)) = s(bool,happ(s(fun(X105,bool),X106),s(X105,X107))) | s(fun(bool,bool),X103) = s(fun(bool,bool),X104)) ) | $spl14),
  inference(superposition,[],[f574,f3165])).
fof(f3327,plain,(
  ( ! [X123,X121,X122] : (s('type.semanticPrimitives.error_result'(bool),X123) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s(bool,sK5(X123,bool)) = s(bool,sK7(X121,X122,bool,bool)) | s('type.semanticPrimitives.error_result'(bool),X123) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X123,bool)))) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122)) ) | $spl14),
  inference(superposition,[],[f537,f3165])).
fof(f3322,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,X107) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X107))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,sK7(X105,X106,bool,bool)))) | s(fun(bool,bool),X105) = s(fun(bool,bool),X106)) ) | $spl14),
  inference(superposition,[],[f101,f3165])).
fof(f3320,plain,(
  ( ! [X101,X99,X100] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X101))) | s(bool,X101) = s(bool,sK7(X99,X100,bool,bool)) | s(fun(bool,bool),X99) = s(fun(bool,bool),X100)) ) | $spl14),
  inference(superposition,[],[f100,f3165])).
fof(f3360,plain,(
  ( ! [X78,X76,X77] : (s(bool,X78) = s(bool,sK7(X76,X77,bool,bool)) | s(bool,'F_0') = s(bool,X78) | s(fun(bool,bool),X76) = s(fun(bool,bool),X77)) ) | $spl14),
  inference(subsumption_resolution,[],[f3310,f76])).
fof(f3310,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'T_0')) | s(bool,X78) = s(bool,sK7(X76,X77,bool,bool)) | s(bool,'F_0') = s(bool,X78) | s(fun(bool,bool),X76) = s(fun(bool,bool),X77)) ) | $spl14),
  inference(superposition,[],[f789,f3165])).
fof(f3350,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,sK7(X35,X36,bool,bool)) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36)) ) | $spl14),
  inference(subsumption_resolution,[],[f3297,f76])).
fof(f3297,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) = s(bool,sK7(X35,X36,bool,bool)) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36)) ) | $spl14),
  inference(superposition,[],[f422,f3165])).
fof(f3349,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) != s(bool,sK7(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,X31) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30)) ) | $spl14),
  inference(subsumption_resolution,[],[f3295,f76])).
fof(f3295,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X31) != s(bool,sK7(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,X31) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30)) ) | $spl14),
  inference(superposition,[],[f127,f3165])).
fof(f3348,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X18)) | s(bool,X18) = s(bool,sK7(X16,X17,bool,bool)) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17)) ) | $spl14),
  inference(subsumption_resolution,[],[f3292,f76])).
fof(f3292,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | s(bool,X18) = s(bool,sK7(X16,X17,bool,bool)) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17)) ) | $spl14),
  inference(superposition,[],[f106,f3165])).
fof(f3339,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X0,X1,bool,bool))),bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl14),
  inference(condensation,[],[f3338])).
fof(f3338,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X21) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X19,X20,bool,bool))),bool))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3281])).
fof(f3281,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,X21) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,sK7(X19,X20,bool,bool))),bool)) | s(bool,'T_0') = s(bool,X21)) ) | $spl14),
  inference(superposition,[],[f3165,f689])).
fof(f12064,plain,(
  ( ! [X177,X178,X176] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X178))) | s(fun(fun(bool,bool),bool),X176) = s(fun(fun(bool,bool),bool),X177) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK7(X176,X177,bool,fun(bool,bool)))))) ) | ($spl14 | $spl32)),
  inference(superposition,[],[f663,f11702])).
fof(f11859,plain,(
  $false | ($spl11 | $spl13 | $spl26)),
  inference(subsumption_resolution,[],[f11858,f1595])).
fof(f11858,plain,(
  s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | ($spl13 | $spl26)),
  inference(backward_demodulation,[],[f11833,f3174])).
fof(f11549,plain,(
  ( ! [X54,X52] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X54) | s(bool,X52) = s(bool,X54)) ) | $spl58),
  inference(cnf_transformation,[],[f11549_D])).
fof(f11549_D,plain,(
  ( ! [X54,X52] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X52))) | s(bool,'F_0') = s(bool,X54) | s(bool,X52) = s(bool,X54)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f11548,plain,(
  ( ! [X54,X52,X53] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) ) | $spl12),
  inference(subsumption_resolution,[],[f10777,f83])).
fof(f10777,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') != s(bool,X53) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X52))) | s(bool,'T_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) ) | $spl12),
  inference(superposition,[],[f475,f5559])).
fof(f5559,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f3598,f3598])).
fof(f3598,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X20)) ) | $spl12),
  inference(subsumption_resolution,[],[f3549,f76])).
fof(f3549,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | s(bool,X20) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X20)) ) | $spl12),
  inference(superposition,[],[f789,f3164])).
fof(f11462,plain,(
  ( ! [X420,X421] : (s(fun(bool,bool),X420) = s(fun(bool,bool),X421) | s(bool,'F_0') = s(bool,sK7(X420,X421,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X420,X421,bool,bool))))) ) | $spl54),
  inference(cnf_transformation,[],[f11462_D])).
fof(f11462_D,plain,(
  ( ! [X420,X421] : (s(fun(bool,bool),X420) = s(fun(bool,bool),X421) | s(bool,'F_0') = s(bool,sK7(X420,X421,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X420,X421,bool,bool))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11463,plain,(
  ( ! [X422] : (s(bool,'F_0') != s(bool,f(s(bool,X422))) | s(bool,'F_0') = s(bool,X422)) ) | $spl56),
  inference(cnf_transformation,[],[f11463_D])).
fof(f11463_D,plain,(
  ( ! [X422] : (s(bool,'F_0') != s(bool,f(s(bool,X422))) | s(bool,'F_0') = s(bool,X422)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f10669,plain,(
  ( ! [X420,X422,X421] : (s(bool,'F_0') != s(bool,f(s(bool,X422))) | s(fun(bool,bool),X420) = s(fun(bool,bool),X421) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X420,X421,bool,bool)))) | s(bool,'F_0') = s(bool,sK7(X420,X421,bool,bool)) | s(bool,'F_0') = s(bool,X422)) ) | $spl12),
  inference(superposition,[],[f663,f5559])).
fof(f10513,plain,(
  ( ! [X420,X422,X421] : (s(bool,'F_0') != s(bool,f(s(bool,X422))) | s(fun(bool,bool),X420) = s(fun(bool,bool),X421) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X420,X421,bool,bool)))) | s(bool,'F_0') = s(bool,X422) | s(bool,'F_0') = s(bool,sK7(X420,X421,bool,bool))) ) | $spl12),
  inference(superposition,[],[f663,f5559])).
fof(f3935,plain,(
  $false | ($spl21 | $spl22 | $spl37)),
  inference(subsumption_resolution,[],[f3865,f3754])).
fof(f3754,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3720,f3171])).
fof(f3720,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(trivial_inequality_removal,[],[f3711])).
fof(f3711,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(superposition,[],[f3170,f125])).
fof(f3865,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl21 | $spl22 | $spl37)),
  inference(backward_demodulation,[],[f3754,f3496])).
fof(f3934,plain,(
  $false | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3846,f3754])).
fof(f3846,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f3170])).
fof(f3933,plain,(
  $false | ($spl21 | $spl22)),
  inference(global_subsumption,[],[f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914,f3922,f3821,f3930,f3932])).
fof(f3932,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61)))) ) | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3830,f3754])).
fof(f3830,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X60))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X61))) | s(bool,'T_0') != s(bool,X61)) ) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f2280])).
fof(f3930,plain,(
  ( ! [X55] : (s(bool,X55) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool))) ) | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3829,f619])).
fof(f3829,plain,(
  ( ! [X55] : (s(bool,X55) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X55)))) ) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f2279])).
fof(f3821,plain,(
  ( ! [X57,X56] : (s(bool,X56) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X57))) | s(bool,X57) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X56)),bool))) ) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f2159])).
fof(f3922,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0)))) ) | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3798,f101])).
fof(f3798,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f1413])).
fof(f3914,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3755,f3754])).
fof(f3755,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f89])).
fof(f3753,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool))) ) | $spl21),
  inference(subsumption_resolution,[],[f3709,f116])).
fof(f3709,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X8)),bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl21),
  inference(superposition,[],[f3170,f689])).
fof(f3701,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(superposition,[],[f3170,f83])).
fof(f3699,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl21),
  inference(unit_resulting_resolution,[],[f76,f3170,f789])).
fof(f3700,plain,(
  s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,f(s(bool,'F_0'))))) | $spl21),
  inference(unit_resulting_resolution,[],[f3170,f100])).
fof(f3931,plain,(
  $false | ($spl21 | $spl22)),
  inference(global_subsumption,[],[f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914,f3922,f3821,f3930])).
fof(f3927,plain,(
  $false | ($spl21 | $spl22)),
  inference(global_subsumption,[],[f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914,f3922,f3821])).
fof(f3926,plain,(
  $false | ($spl11 | $spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3803,f3754])).
fof(f3803,plain,(
  s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | ($spl11 | $spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f1595])).
fof(f3925,plain,(
  $false | ($spl6 | $spl21 | $spl22)),
  inference(global_subsumption,[],[f3924,f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914,f3922])).
fof(f3924,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'T_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33)))) ) | ($spl6 | $spl21 | $spl22)),
  inference(forward_demodulation,[],[f3802,f3754])).
fof(f3802,plain,(
  ( ! [X33] : (s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) | s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | ($spl6 | $spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f1568])).
fof(f3923,plain,(
  $false | ($spl21 | $spl22)),
  inference(global_subsumption,[],[f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914,f3922])).
fof(f3916,plain,(
  $false | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f3757,f76])).
fof(f3757,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f115])).
fof(f3915,plain,(
  $false | ($spl21 | $spl22)),
  inference(global_subsumption,[],[f75,f74,f76,f115,f83,f86,f85,f90,f89,f95,f99,f98,f101,f100,f102,f106,f103,f107,f108,f109,f116,f125,f127,f157,f199,f422,f423,f434,f433,f436,f437,f439,f441,f457,f475,f476,f478,f480,f481,f504,f531,f530,f532,f534,f536,f537,f539,f540,f542,f558,f572,f573,f574,f576,f578,f580,f584,f586,f596,f613,f615,f617,f619,f654,f656,f658,f659,f661,f663,f687,f689,f691,f715,f789,f790,f791,f773,f792,f796,f800,f824,f830,f855,f860,f864,f863,f885,f886,f894,f900,f906,f907,f909,f913,f915,f923,f924,f928,f930,f947,f969,f1006,f1011,f1012,f1039,f1042,f1046,f1115,f1118,f1128,f1131,f1166,f1168,f1341,f1335,f1196,f1327,f1242,f1249,f1373,f1277,f1286,f1410,f1390,f1413,f1429,f1535,f1476,f1516,f1525,f1773,f1863,f1865,f1866,f1867,f1868,f1869,f1872,f1874,f1915,f1920,f1922,f1938,f1944,f1946,f1987,f2204,f2239,f2068,f2241,f2242,f2245,f2116,f2273,f2274,f2278,f2279,f2159,f2280,f2283,f2335,f2336,f2337,f2338,f2339,f2296,f2298,f2299,f2342,f2305,f2343,f2309,f2310,f2344,f2314,f2315,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2325,f2327,f2328,f2331,f2350,f2481,f2482,f2563,f2565,f2570,f2573,f2574,f2576,f3092,f3095,f2642,f2645,f2646,f2648,f3108,f3109,f3113,f3117,f2687,f2694,f2700,f2701,f3131,f2726,f2731,f3137,f2787,f2789,f2794,f2797,f2798,f2800,f2823,f2811,f3155,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3235,f3202,f3365,f3366,f3367,f3368,f3372,f3373,f3374,f3445,f3377,f3446,f3379,f3380,f3383,f3384,f3386,f3387,f3388,f3389,f3390,f3392,f3395,f3397,f3400,f3401,f3402,f3406,f3411,f3412,f3414,f3459,f3424,f3426,f3427,f3428,f3429,f3430,f3431,f3432,f3436,f3437,f3438,f3171,f3170,f3607,f3610,f3612,f3628,f3638,f3700,f3699,f3701,f3753,f3754,f3914])).
fof(f3913,plain,(
  $false | ($spl21 | $spl22)),
  inference(trivial_inequality_removal,[],[f3758])).
fof(f3758,plain,(
  s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f434])).
fof(f3912,plain,(
  $false | ($spl21 | $spl22)),
  inference(trivial_inequality_removal,[],[f3761])).
fof(f3761,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f457])).
fof(f3911,plain,(
  $false | ($spl21 | $spl22)),
  inference(trivial_inequality_removal,[],[f3787])).
fof(f3787,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(bool)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))))))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f3754,f1011])).
fof(f3630,plain,(
  ( ! [X19,X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool)) | s(bool,X19) = s(bool,sK7(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X17,X18,bool,bool))))) ) | $spl52),
  inference(cnf_transformation,[],[f3630_D])).
fof(f3630_D,plain,(
  ( ! [X19,X17,X18] : (s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool)) | s(bool,X19) = s(bool,sK7(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X17,X18,bool,bool))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3496,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl37),
  inference(cnf_transformation,[],[f3496_D])).
fof(f3496_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f3612,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X17,X18,bool,bool)))) | s(bool,X19) = s(bool,sK7(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X19)),bool))) )),
  inference(superposition,[],[f663,f689])).
fof(f3629,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK7(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X2,X3,bool,bool))))) ) | $spl50),
  inference(cnf_transformation,[],[f3629_D])).
fof(f3629_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK7(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X2,X3,bool,bool))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3610,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK7(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f663,f125])).
fof(f3607,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK7(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK7(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f663,f125])).
fof(f3494,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl34),
  inference(cnf_transformation,[],[f3494_D])).
fof(f3494_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3504,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool))) ) | $spl48),
  inference(cnf_transformation,[],[f3504_D])).
fof(f3504_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3490,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X7)),bool)) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl33),
  inference(superposition,[],[f3363,f689])).
fof(f3503,plain,(
  s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'T_0')))),bool)) | $spl46),
  inference(cnf_transformation,[],[f3503_D])).
fof(f3503_D,plain,(
  s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'T_0')))),bool)) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3500,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f3500_D])).
fof(f3500_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3489,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'T_0')))),bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl33),
  inference(superposition,[],[f3363,f689])).
fof(f3501,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X5)),bool))) ) | $spl44),
  inference(cnf_transformation,[],[f3501_D])).
fof(f3501_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X5)),bool))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3499,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl41),
  inference(cnf_transformation,[],[f3499_D])).
fof(f3499_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f3488,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X5) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X5)),bool))) ) | $spl33),
  inference(superposition,[],[f3363,f689])).
fof(f3497,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f3497_D])).
fof(f3497_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3487,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl33),
  inference(superposition,[],[f3363,f125])).
fof(f3486,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl33),
  inference(superposition,[],[f3363,f125])).
fof(f3485,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl33),
  inference(superposition,[],[f3363,f125])).
fof(f3484,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f3363,f125])).
fof(f3483,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl33),
  inference(superposition,[],[f3363,f83])).
fof(f3361,plain,(
  ( ! [X97,X98] : (s(fun(bool,bool),X97) = s(fun(bool,bool),X98)) ) | $spl30),
  inference(cnf_transformation,[],[f3361_D])).
fof(f3361_D,plain,(
  ( ! [X97,X98] : (s(fun(bool,bool),X97) = s(fun(bool,bool),X98)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3363,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl33),
  inference(cnf_transformation,[],[f3363_D])).
fof(f3363_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f3330,plain,(
  ( ! [X97,X98] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X97) = s(fun(bool,bool),X98)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f3319])).
fof(f3319,plain,(
  ( ! [X97,X98] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X97) = s(fun(bool,bool),X98) | s(fun(bool,bool),X97) = s(fun(bool,bool),X98)) ) | $spl14),
  inference(superposition,[],[f109,f3165])).
fof(f3164,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl12),
  inference(cnf_transformation,[],[f3164_D])).
fof(f3164_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3165,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl14),
  inference(cnf_transformation,[],[f3165_D])).
fof(f3165_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3175,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool))) ) | $spl28),
  inference(cnf_transformation,[],[f3175_D])).
fof(f3175_D,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3162,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,'T_0') = s(bool,sK7(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X21)),bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f658,f689])).
fof(f3174,plain,(
  s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'F_0')))),bool)) | $spl26),
  inference(cnf_transformation,[],[f3174_D])).
fof(f3174_D,plain,(
  s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'F_0')))),bool)) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3171,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(cnf_transformation,[],[f3171_D])).
fof(f3171_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3161,plain,(
  ( ! [X19,X20,X18] : (s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,sK7(X19,X20,bool,bool)) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,g(s(bool,'F_0')))),bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f658,f689])).
fof(f3172,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool))) ) | $spl24),
  inference(cnf_transformation,[],[f3172_D])).
fof(f3172_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3170,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl21),
  inference(cnf_transformation,[],[f3170_D])).
fof(f3170_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f3160,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK7(X15,X16,bool,bool)) | s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,X17)),bool))) )),
  inference(superposition,[],[f658,f689])).
fof(f3168,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl18),
  inference(cnf_transformation,[],[f3168_D])).
fof(f3168_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3159,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK7(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f658,f125])).
fof(f3158,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK7(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f658,f125])).
fof(f3157,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK7(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f658,f125])).
fof(f3156,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK7(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f658,f125])).
fof(f3167,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl17),
  inference(cnf_transformation,[],[f3167_D])).
fof(f3167_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f3155,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f658,f83])).
fof(f1593,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) | $spl8),
  inference(cnf_transformation,[],[f1593_D])).
fof(f1593_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1595,plain,(
  s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | $spl11),
  inference(cnf_transformation,[],[f1595_D])).
fof(f1595_D,plain,(
  s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1525,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s(bool,X0) = s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool))) )),
  inference(equality_factoring,[],[f1327])).
fof(f1568,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33)))) ) | $spl6),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1535,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(duplicate_literal_removal,[],[f1465])).
fof(f1465,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,sK5('const.semanticPrimitives.Rraise_1'(s(bool,'F_0')),bool)) | s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f475,f1327])).
fof(f512,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f507])).
fof(f507,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f501,f498])).
fof(f498,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f467,f100])).
fof(f467,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f432,f100])).
fof(f432,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))))) | $spl1),
  inference(unit_resulting_resolution,[],[f112,f100])).
fof(f501,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(sK1,X0),X1),s('type.semanticPrimitives.error_result'(sK1),sK3)))) ) | $spl2),
  inference(superposition,[],[f107,f113])).
fof(f511,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f508])).
fof(f508,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0))),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f501,f467])).
fof(f510,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f509])).
fof(f509,plain,(
  s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))) != s('type.semanticPrimitives.error_result'('type.semanticPrimitives.error_result'(sK0)),'const.semanticPrimitives.Rraise_1'(s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f501,f432])).
fof(f111,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) | $spl0),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f113,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) | $spl2),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f74,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)))),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) = s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(sK1),sK3) = s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4)))) & (s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) | s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))) | s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))) & (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) != s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))) | s('type.semanticPrimitives.error_result'(X1),X3) != s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4)))))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))) <~> s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))) <=> s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))) = s('type.semanticPrimitives.error_result'(X0),'const.evalProps.map_error_result_2'(s(fun(X1,X0),X2),s('type.semanticPrimitives.error_result'(X1),X3))) <=> s('type.semanticPrimitives.error_result'(X1),X3) = s('type.semanticPrimitives.error_result'(X1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X4))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),X9))) <=> s('type.semanticPrimitives.error_result'(X7),X9) = s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),X9))) <=> s('type.semanticPrimitives.error_result'(X7),X9) = s('type.semanticPrimitives.error_result'(X7),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X10))))),
  file('thm.evalProps.map_error_result_Rtype_error','thm.evalProps.map_error_result_Rtype_error')).
fof(f112,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) | $spl1),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f114,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) | $spl3),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f75,plain,(
  s('type.semanticPrimitives.error_result'(sK1),sK3) != s('type.semanticPrimitives.error_result'(sK1),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) | s('type.semanticPrimitives.error_result'(sK0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK4))) != s('type.semanticPrimitives.error_result'(sK0),'const.evalProps.map_error_result_2'(s(fun(sK1,sK0),sK2),s('type.semanticPrimitives.error_result'(sK1),sK3)))),
  inference(cnf_transformation,[],[f61])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 222
% Passive clauses: 1829
% Generated clauses: 27932
% Final active clauses: 117
% Final passive clauses: 945
% Input formulas: 22
% Initial clauses: 36
% 
% Duplicate literals: 1958
% Trivial inequalities: 274
% Fw subsumption resolutions: 1862
% Fw demodulations: 3300
% Bw demodulations: 481
% Condensations: 3124
% Global subsumptions: 17
% 
% Simple tautologies: 347
% Equational tautologies: 456
% Forward subsumptions: 14457
% Fw demodulations to eq. taut.: 267
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 1860
% Factoring: 4
% Forward superposition: 6426
% Backward superposition: 7142
% Self superposition: 248
% Equality factoring: 172
% Equality resolution: 47
% 
% Split clauses: 67
% Split components: 142
% SAT solver clauses: 2018
% SAT solver unit clauses: 68
% SAT solver binary clauses: 312
% 
% Sat splits: 67
% Sat splitting refutations: 36
% 
% Memory used [KB]: 13048
% Time elapsed: 1.637 s
% ------------------------------
% ------------------------------
% Success in time 3.04 s
