Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.pegComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pegComplete
% SZS output start Proof for thm.pegComplete
fof(f3021,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3016,f26,f38,f36,f3008,f25,f37,f35])).
fof(f35,plain,(
  s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) = s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) | $spl0),
  inference(cnf_transformation,[],[f35_D])).
fof(f35_D,plain,(
  s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) = s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f37,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) | $spl2),
  inference(cnf_transformation,[],[f37_D])).
fof(f37_D,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f25,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) | s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) = s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  (s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) = s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) | s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2)))))) & (s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) != s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) | s('type.grammar.parsetree'(sK1,sK0),sK3) != s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2,X3] : ((s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) | s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2)))))) & (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) != s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) | s('type.grammar.parsetree'(X1,X0),X3) != s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2)))))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ? [X0,X1,X2,X3] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) <~> s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))))),
  inference(ennf_transformation,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1,X2,X3] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) <=> s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))))),
  inference(flattening,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1] : ! [X2,X3] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) <=> s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))) = s('type.grammar.symbol'(X7,X6),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X7,X6),X9))) <=> s('type.grammar.parsetree'(X7,X6),X9) = s('type.grammar.parsetree'(X7,X6),'const.grammar.Lf_1'(s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))) = s('type.grammar.symbol'(X7,X6),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X7,X6),X9))) <=> s('type.grammar.parsetree'(X7,X6),X9) = s('type.grammar.parsetree'(X7,X6),'const.grammar.Lf_1'(s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))))))),
  file('thm.pegComplete.ptree_head_eq_tok|split|1','thm.pegComplete.ptree_head_eq_tok|split|1')).
fof(f3008,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3007,f38])).
fof(f3007,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) | $spl0),
  inference(equality_resolution,[],[f1353])).
fof(f1353,plain,(
  ( ! [X0] : (s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) != s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,X0))) | s('type.grammar.parsetree'(sK1,sK0),sK3) = s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,X0)))))) ) | $spl0),
  inference(superposition,[],[f33,f35])).
fof(f33,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) != s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) | s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : ((s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) != s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) | s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2)))))) & (s('type.grammar.parsetree'(X1,X0),X3) != s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))) | s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3)))))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2,X3] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) <=> s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3))) <=> s('type.grammar.parsetree'(X1,X0),X3) = s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X8,X9] : (s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))) = s('type.grammar.symbol'(X7,X6),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X7,X6),X9))) <=> s('type.grammar.parsetree'(X7,X6),X9) = s('type.grammar.parsetree'(X7,X6),'const.grammar.Lf_1'(s('type.grammar.symbol'(X7,X6),'const.grammar.TOK_1'(s(X7,X8))))))),
  file('thm.pegComplete.ptree_head_eq_tok|split|1','thm.pegComplete.pegComplete_2_prove_ltb')).
fof(f36,plain,(
  s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) != s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) | $spl1),
  inference(cnf_transformation,[],[f36_D])).
fof(f36_D,plain,(
  s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) != s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f38,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) != s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) | $spl3),
  inference(cnf_transformation,[],[f38_D])).
fof(f38_D,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) != s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f26,plain,(
  s('type.grammar.parsetree'(sK1,sK0),sK3) != s('type.grammar.parsetree'(sK1,sK0),'const.grammar.Lf_1'(s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))))) | s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) != s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3)))),
  inference(cnf_transformation,[],[f21])).
fof(f3016,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3009,f36])).
fof(f3009,plain,(
  s('type.grammar.symbol'(sK1,sK0),'const.grammar.TOK_1'(s(sK1,sK2))) = s('type.grammar.symbol'(sK1,sK0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(sK1,sK0),sK3))) | $spl2),
  inference(superposition,[],[f2572,f37])).
fof(f2572,plain,(
  ( ! [X2,X0,X1] : (s('type.grammar.symbol'(X0,X1),'const.grammar.TOK_1'(s(X0,X2))) = s('type.grammar.symbol'(X0,X1),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X0,X1),'const.grammar.Lf_1'(s('type.grammar.symbol'(X0,X1),'const.grammar.TOK_1'(s(X0,X2)))))))) )),
  inference(equality_resolution,[],[f34])).
fof(f34,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.grammar.parsetree'(X1,X0),X3) != s('type.grammar.parsetree'(X1,X0),'const.grammar.Lf_1'(s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))))) | s('type.grammar.symbol'(X1,X0),'const.grammar.TOK_1'(s(X1,X2))) = s('type.grammar.symbol'(X1,X0),'const.grammar.ptree_head_1'(s('type.grammar.parsetree'(X1,X0),X3)))) )),
  inference(cnf_transformation,[],[f24])).
% SZS output end Proof for thm.pegComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 31
% Passive clauses: 572
% Generated clauses: 2878
% Final active clauses: 26
% Final passive clauses: 525
% Input formulas: 7
% Initial clauses: 10
% 
% Duplicate literals: 314
% Trivial inequalities: 36
% Fw subsumption resolutions: 184
% Fw demodulations: 8
% 
% Simple tautologies: 110
% Equational tautologies: 79
% Forward subsumptions: 1152
% 
% Binary resolution: 16
% Forward superposition: 1042
% Backward superposition: 813
% Self superposition: 223
% Equality factoring: 202
% Equality resolution: 6
% 
% Split clauses: 143
% Split components: 288
% SAT solver clauses: 440
% SAT solver unit clauses: 226
% SAT solver binary clauses: 109
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 143
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 436
% TWLsolver calls for satisfiability: 17
% 
% Memory used [KB]: 1663
% Time elapsed: 0.161 s
% ------------------------------
% ------------------------------
% Success in time 0.16 s
