Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.LeftistHeap
% (65270)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 844
% Passive clauses: 5500
% Generated clauses: 22852
% Final active clauses: 703
% Final passive clauses: 4091
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 224
% Trivial inequalities: 39
% Fw subsumption resolutions: 1463
% Fw demodulations: 1855
% Bw demodulations: 138
% 
% Simple tautologies: 196
% Equational tautologies: 26
% Forward subsumptions: 13737
% Fw demodulations to eq. taut.: 15
% 
% Binary resolution: 5717
% Factoring: 1193
% Forward superposition: 5435
% Backward superposition: 6636
% Self superposition: 74
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 191
% Split components: 383
% SAT solver clauses: 361
% SAT solver unit clauses: 9
% SAT solver binary clauses: 52
% 
% Sat splits: 191
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 361
% TWLsolver calls for satisfiability: 328
% 
% Memory used [KB]: 15607
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.LeftistHeap
% (65298)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 509
% Passive clauses: 2163
% Generated clauses: 14655
% Final active clauses: 421
% Final passive clauses: 1554
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 95
% Trivial inequalities: 23
% Fw subsumption resolutions: 675
% Fw demodulations: 1253
% Bw demodulations: 148
% Condensations: 72
% Global subsumptions: 18
% 
% Simple tautologies: 205
% Equational tautologies: 8
% Forward subsumptions: 7605
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 2154
% Factoring: 196
% Forward superposition: 4171
% Backward superposition: 3151
% Self superposition: 65
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 2843
% SAT solver unit clauses: 383
% SAT solver binary clauses: 113
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 15
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2771
% TWLsolver calls for satisfiability: 2565
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.LeftistHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.LeftistHeap
% SZS output start Proof for thm.LeftistHeap
fof(f15810,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f512,f851,f849,f511,f850,f848,f3000,f3010,f3008,f3002,f3011,f3003,f3014,f3004,f3015,f3006,f3019,f3017,f5664,f3793,f5674,f3597,f3013,f5676,f15745,f15807,f15809])).
fof(f15809,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f15808,f9738])).
fof(f9738,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X0,X1),X2),s(fun(X1,fun(X1,bool)),X3),s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0'))))) )),
  inference(forward_demodulation,[],[f9731,f706])).
fof(f706,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.LeftistHeap.heap_CASE_3'(s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0'),s(X1,X2),s(fun('type.num.num',fun(X0,fun('type.LeftistHeap.heap'(X0),fun('type.LeftistHeap.heap'(X0),X1)))),X3)))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.LeftistHeap.heap_CASE_3'(s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0'),s(X1,X2),s(fun('type.num.num',fun(X0,fun('type.LeftistHeap.heap'(X0),fun('type.LeftistHeap.heap'(X0),X1)))),X3)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.LeftistHeap.heap_CASE_3'(s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0'),s(X1,X2),s(fun('type.num.num',fun(X0,fun('type.LeftistHeap.heap'(X0),fun('type.LeftistHeap.heap'(X0),X1)))),X3)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X25,X26] : s(X6,X25) = s(X6,'const.LeftistHeap.heap_CASE_3'(s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Empty_0'),s(X6,X25),s(fun('type.num.num',fun(X7,fun('type.LeftistHeap.heap'(X7),fun('type.LeftistHeap.heap'(X7),X6)))),X26)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.LeftistHeap.heap_case_def|split|0')).
fof(f9731,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X0,X1),X2),s(fun(X1,fun(X1,bool)),X3),s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.heap_CASE_3'(s('type.LeftistHeap.heap'(X4),'const.LeftistHeap.Empty_0'),s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0'),s(fun('type.num.num',fun(X4,fun('type.LeftistHeap.heap'(X4),fun('type.LeftistHeap.heap'(X4),'type.LeftistHeap.heap'(X0))))),X5))))))) )),
  inference(unit_resulting_resolution,[],[f706,f714])).
fof(f714,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4) | ~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4))))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | (s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',sK40(X4,X3,X2,X1,X0)),s(X1,sK41(X4,X3,X2,X1,X0)),s('type.LeftistHeap.heap'(X1),sK42(X4,X3,X2,X1,X0)),s('type.LeftistHeap.heap'(X1),sK43(X4,X3,X2,X1,X0)))) & s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))) & (! [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) != s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4) | p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40,sK41,sK42,sK43])],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | (? [X9,X10,X11,X12] : s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X9),s(X1,X10),s('type.LeftistHeap.heap'(X1),X11),s('type.LeftistHeap.heap'(X1),X12))) & s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))) & (! [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) != s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4) | p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4))))))),
  inference(rectify,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | (? [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) & s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))) & (! [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) != s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4) | p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4))))))),
  inference(flattening,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | (? [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) & s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))) & ((! [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) != s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4)) | p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4))))))),
  inference(nnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) <=> (? [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) & s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4)))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) <=> (? [X5,X6,X7,X8] : s('type.LeftistHeap.heap'(X1),X4) = s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8))) & ~s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X12,X13] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X7,X6),X11),s(fun(X6,fun(X6,bool)),X12),s('type.LeftistHeap.heap'(X7),X13)))) <=> (? [X14,X15,X16,X17] : s('type.LeftistHeap.heap'(X7),X13) = s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Tree_4'(s('type.num.num',X14),s(X7,X15),s('type.LeftistHeap.heap'(X7),X16),s('type.LeftistHeap.heap'(X7),X17))) & ~s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X7),X13)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.LeftistHeap.delete_min_side_def')).
fof(f15808,plain,(
  p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0')))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f848,f851])).
fof(f15807,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f15744,f850])).
fof(f15744,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f14300,f607])).
fof(f607,plain,(
  ( ! [X0,X1] : (s('type.LeftistHeap.heap'(X0),X1) = s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Tree_4'(s('type.num.num',sK18(X1,X0)),s(X0,sK19(X1,X0)),s('type.LeftistHeap.heap'(X0),sK20(X1,X0)),s('type.LeftistHeap.heap'(X0),sK21(X1,X0)))) | s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X0),X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1] : (s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X0),X1) | s('type.LeftistHeap.heap'(X0),X1) = s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Tree_4'(s('type.num.num',sK18(X1,X0)),s(X0,sK19(X1,X0)),s('type.LeftistHeap.heap'(X0),sK20(X1,X0)),s('type.LeftistHeap.heap'(X0),sK21(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20,sK21])],[f201])).
fof(f201,plain,(
  ! [X0,X1] : (s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X0),X1) | ? [X2,X3,X4,X5] : s('type.LeftistHeap.heap'(X0),X1) = s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Tree_4'(s('type.num.num',X2),s(X0,X3),s('type.LeftistHeap.heap'(X0),X4),s('type.LeftistHeap.heap'(X0),X5))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1] : (s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X0),X1) | ? [X2,X3,X4,X5] : s('type.LeftistHeap.heap'(X0),X1) = s('type.LeftistHeap.heap'(X0),'const.LeftistHeap.Tree_4'(s('type.num.num',X2),s(X0,X3),s('type.LeftistHeap.heap'(X0),X4),s('type.LeftistHeap.heap'(X0),X5))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7] : ! [X31] : (s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X7),X31) | ? [X32,X16,X10,X33] : s('type.LeftistHeap.heap'(X7),X31) = s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Tree_4'(s('type.num.num',X32),s(X7,X16),s('type.LeftistHeap.heap'(X7),X10),s('type.LeftistHeap.heap'(X7),X33))))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.LeftistHeap.heap_nchotomy')).
fof(f14300,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.LeftistHeap.heap'(sK14),sK17) != s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Tree_4'(s('type.num.num',X0),s(sK14,X1),s('type.LeftistHeap.heap'(sK14),X2),s('type.LeftistHeap.heap'(sK14),X3)))) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f14299,f850])).
fof(f14299,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) | s('type.LeftistHeap.heap'(sK14),sK17) != s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Tree_4'(s('type.num.num',X0),s(sK14,X1),s('type.LeftistHeap.heap'(sK14),X2),s('type.LeftistHeap.heap'(sK14),X3)))) ) | $spl1),
  inference(subsumption_resolution,[],[f14291,f852])).
fof(f852,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f518,f513])).
fof(f513,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb',aHL_TRUTH)).
fof(f518,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f14291,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) | s('type.LeftistHeap.heap'(sK14),sK17) != s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Tree_4'(s('type.num.num',X0),s(sK14,X1),s('type.LeftistHeap.heap'(sK14),X2),s('type.LeftistHeap.heap'(sK14),X3)))) ) | $spl1),
  inference(superposition,[],[f715,f927])).
fof(f927,plain,(
  s(bool,'F_0') = s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17))) | $spl1),
  inference(unit_resulting_resolution,[],[f849,f565])).
fof(f565,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X35] : (s(bool,'F_0') = s(bool,X35) <=> ~p(s(bool,X35)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f715,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4) | s('type.LeftistHeap.heap'(X1),X4) != s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Tree_4'(s('type.num.num',X5),s(X1,X6),s('type.LeftistHeap.heap'(X1),X7),s('type.LeftistHeap.heap'(X1),X8)))) )),
  inference(cnf_transformation,[],[f467])).
fof(f15745,plain,(
  $false | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f850,f14300,f607])).
fof(f5676,plain,(
  ( ! [X15] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15)))))))) ) | $spl22),
  inference(cnf_transformation,[],[f5676_D])).
fof(f5676_D,plain,(
  ( ! [X15] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15)))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3013,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl11),
  inference(cnf_transformation,[],[f3013_D])).
fof(f3013_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f3597,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15)))))))) )),
  inference(superposition,[],[f568,f545])).
fof(f545,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X32] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X32)))))))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|10')).
fof(f568,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X32] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X32))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|17')).
fof(f5674,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17)))))))) ) | $spl20),
  inference(cnf_transformation,[],[f5674_D])).
fof(f5674_D,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17)))))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3793,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f568,f547])).
fof(f547,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X32] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32)))))))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|9')).
fof(f5664,plain,(
  $false | $spl11),
  inference(subsumption_resolution,[],[f5663,f513])).
fof(f5663,plain,(
  ~p(s(bool,'T_0')) | $spl11),
  inference(forward_demodulation,[],[f5662,f536])).
fof(f536,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X32] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X32)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|26')).
fof(f5662,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl11),
  inference(backward_demodulation,[],[f5600,f3121])).
fof(f3121,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | $spl11),
  inference(unit_resulting_resolution,[],[f3013,f572])).
fof(f572,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X32] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X32)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X32))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|30')).
fof(f5600,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f657,f3035])).
fof(f3035,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f3023,f855])).
fof(f855,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f548,f516])).
fof(f516,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f102])).
fof(f102,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f548,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X62] : s('type.num.num',X62) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X62)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.arithmetic.MULT_LEFT_|49|')).
fof(f3023,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(unit_resulting_resolution,[],[f855,f569])).
fof(f569,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f390])).
fof(f657,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X32,X62] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X62),s('type.num.num',X32))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X62)))))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|31')).
fof(f3017,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl16),
  inference(cnf_transformation,[],[f3017_D])).
fof(f3017_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3019,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(cnf_transformation,[],[f3019_D])).
fof(f3019_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3006,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f568,f516])).
fof(f3015,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) | $spl14),
  inference(cnf_transformation,[],[f3015_D])).
fof(f3015_D,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3004,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f568,f528])).
fof(f528,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X62] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num','const.num.|48|_0')))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.arithmetic.MULT_|48|')).
fof(f3014,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) | $spl12),
  inference(cnf_transformation,[],[f3014_D])).
fof(f3014_D,plain,(
  ( ! [X5] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3003,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) )),
  inference(superposition,[],[f568,f526])).
fof(f526,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X32] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X32)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.arithmetic.MULT|split|0')).
fof(f3011,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) ) | $spl8),
  inference(cnf_transformation,[],[f3011_D])).
fof(f3011_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3002,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f568,f531])).
fof(f531,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X32] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X32)))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|6')).
fof(f3008,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(cnf_transformation,[],[f3008_D])).
fof(f3008_D,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3010,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3000,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f568,f854])).
fof(f854,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f544,f516])).
fof(f544,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X32] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.numeral.numeral_distrib|split|11')).
fof(f848,plain,(
  p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) | $spl0),
  inference(cnf_transformation,[],[f848_D])).
fof(f848_D,plain,(
  p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f850,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(sK14),sK17) | $spl3),
  inference(cnf_transformation,[],[f850_D])).
fof(f850_D,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(sK14),sK17) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f511,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(sK14),sK17) | p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) | s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(sK14),sK17)) & (~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) | s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f382])).
fof(f382,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4)) & (~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) | s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4)))),
  inference(nnf_transformation,[],[f318])).
fof(f318,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) <~> s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) <=> s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') != s('type.LeftistHeap.heap'(X1),X4))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X1,X0),X2),s(fun(X0,fun(X0,bool)),X3),s('type.LeftistHeap.heap'(X1),X4)))) <=> ~s('type.LeftistHeap.heap'(X1),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X1),X4))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.LeftistHeap.heap'(X7),X10)))) <=> ~s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X7),X10))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(X7,X6),X8),s(fun(X6,fun(X6,bool)),X9),s('type.LeftistHeap.heap'(X7),X10)))) <=> ~s('type.LeftistHeap.heap'(X7),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(X7),X10))),
  file('thm.LeftistHeap.LeftistHeap_69_prove_ltb','thm.LeftistHeap.LeftistHeap_69_prove_ltb')).
fof(f849,plain,(
  ~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) | $spl1),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f851,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) | $spl2),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f512,plain,(
  s('type.LeftistHeap.heap'(sK14),'const.LeftistHeap.Empty_0') = s('type.LeftistHeap.heap'(sK14),sK17) | ~p(s(bool,'const.LeftistHeap.delete_min_side_3'(s(fun(sK14,sK13),sK15),s(fun(sK13,fun(sK13,bool)),sK16),s('type.LeftistHeap.heap'(sK14),sK17))))),
  inference(cnf_transformation,[],[f383])).
% SZS output end Proof for thm.LeftistHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 631
% Passive clauses: 1862
% Generated clauses: 15297
% Final active clauses: 471
% Final passive clauses: 1023
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 7
% Trivial inequalities: 29
% Fw subsumption resolutions: 132
% Fw demodulations: 2178
% Bw demodulations: 309
% Condensations: 5
% 
% Simple tautologies: 70
% Equational tautologies: 20
% Forward subsumptions: 10882
% Fw demodulations to eq. taut.: 99
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 2448
% Factoring: 1
% Forward superposition: 4671
% Backward superposition: 5069
% Self superposition: 86
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 9
% Split components: 18
% SAT solver clauses: 1886
% SAT solver unit clauses: 588
% SAT solver binary clauses: 734
% 
% Sat splits: 9
% Sat splitting refutations: 4
% 
% Memory used [KB]: 21236
% Time elapsed: 0.532 s
% ------------------------------
% ------------------------------
% Success in time 1.86 s
