Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvlSem
% (13866)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 876
% Passive clauses: 4773
% Generated clauses: 29988
% Final active clauses: 862
% Final passive clauses: 3602
% Input formulas: 516
% Initial clauses: 766
% 
% Duplicate literals: 174
% Trivial inequalities: 33
% Fw subsumption resolutions: 1411
% Fw demodulations: 2507
% Bw demodulations: 10
% 
% Simple tautologies: 250
% Equational tautologies: 15
% Forward subsumptions: 20678
% 
% Binary resolution: 4469
% Factoring: 701
% Forward superposition: 9056
% Backward superposition: 10697
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 226
% Split components: 454
% SAT solver clauses: 347
% SAT solver unit clauses: 12
% SAT solver binary clauses: 66
% 
% Sat splits: 226
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 347
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 16119
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvlSem
% (13887)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 396
% Passive clauses: 1626
% Generated clauses: 12958
% Final active clauses: 387
% Final passive clauses: 1176
% Input formulas: 516
% Initial clauses: 766
% 
% Duplicate literals: 90
% Trivial inequalities: 3
% Fw subsumption resolutions: 590
% Fw demodulations: 949
% Bw demodulations: 62
% Condensations: 83
% Global subsumptions: 24
% 
% Simple tautologies: 201
% Equational tautologies: 7
% Forward subsumptions: 7532
% 
% Binary resolution: 1593
% Factoring: 192
% Forward superposition: 3213
% Backward superposition: 3655
% Self superposition: 47
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 2646
% SAT solver unit clauses: 1145
% SAT solver binary clauses: 173
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 12
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2550
% TWLsolver calls for satisfiability: 3418
% 
% Memory used [KB]: 47589
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvlSem
% (13918)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1158
% Passive clauses: 6944
% Generated clauses: 182848
% Final active clauses: 438
% Final passive clauses: 1621
% Input formulas: 516
% Initial clauses: 766
% 
% Duplicate literals: 12
% Trivial inequalities: 52
% Fw subsumption resolutions: 1411
% Fw demodulations: 63949
% Bw demodulations: 3237
% Condensations: 27
% Global subsumptions: 1840
% 
% Simple tautologies: 124
% Equational tautologies: 37
% Forward subsumptions: 66458
% Fw demodulations to eq. taut.: 577
% Bw demodulations to eq. taut.: 21
% 
% Unit resulting resolution: 63341
% Factoring: 2
% Forward superposition: 27008
% Backward superposition: 20990
% Self superposition: 165
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 11157
% SAT solver unit clauses: 5088
% SAT solver binary clauses: 5263
% 
% Sat splits: 15
% Sat splitting refutations: 2340
% 
% Memory used [KB]: 365623
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 31 next slice time: 31
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.bvlSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvlSem
% SZS output start Proof for thm.bvlSem
fof(f11020,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f853,f1162,f1160,f1163,f1164,f1165,f1166,f1167,f1169,f1278,f1170,f1257,f1276,f1309,f1317,f1336,f1337,f1338,f1339,f1340,f1482,f1341,f1343,f1431,f1483,f1451,f1495,f1315,f1567,f1568,f1569,f1570,f1571,f1573,f1754,f1574,f1661,f1753,f1667,f1668,f1669,f1670,f1762,f1672,f1673,f1674,f1675,f1676,f1677,f1750,f1749,f1748,f1760,f1892,f1891,f1883,f1980,f1981,f1982,f1983,f1984,f1985,f1986,f1987,f1988,f1989,f1992,f1993,f2004,f2005,f2006,f2007,f2008,f2010,f2191,f2011,f2098,f2190,f2104,f2105,f2106,f2107,f2197,f2109,f2110,f2111,f2112,f2113,f2114,f2187,f2186,f2185,f2196,f2516,f2518,f2522,f2534,f2535,f2539,f2685,f2686,f3127,f3126,f3125,f3130,f3279,f3278,f3270,f3368,f3369,f3370,f3371,f3372,f3373,f3374,f3375,f3376,f3377,f3380,f3381,f3439,f3440,f3441,f3619,f3618,f2517,f3791,f3792,f3793,f3794,f3795,f3797,f3989,f3798,f3885,f3988,f3891,f3892,f3893,f3894,f3994,f3896,f3897,f3898,f3899,f3900,f3901,f3976,f3975,f3974,f3993,f4006,f4022,f4021,f4042,f4044,f4047,f4050,f4052,f4056,f4078,f4079,f4080,f4081,f4082,f4273,f4083,f4275,f4085,f4173,f4274,f4174,f4276,f4179,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4194,f4195,f4269,f4268,f4282,f4608,f4609,f4610,f4611,f4612,f4614,f4808,f4615,f4702,f4807,f4709,f4710,f1758,f4711,f4712,f1761,f4813,f4714,f4715,f4716,f4717,f4718,f4719,f4795,f4794,f4793,f4812,f4840,f4841,f4842,f4843,f4844,f5036,f4845,f5038,f4847,f4935,f5037,f4936,f5039,f4942,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4957,f4958,f4288,f5032,f5031,f5045,f5688,f5690,f5726,f5727,f5689,f5763,f5765,f5801,f5802,f5764,f6025,f6026,f6027,f6028,f6029,f6031,f6227,f6032,f6119,f6226,f6277,f6278,f6279,f6280,f6281,f6283,f6478,f6284,f6371,f6477,f6755,f7158,f6757,f7563,f7971,f7564,f6756,f8147,f8148,f8149,f8150,f8151,f8345,f8152,f8347,f8154,f8242,f8346,f8243,f8348,f8251,f8253,f8254,f8255,f8256,f8257,f8258,f8259,f8260,f8266,f8267,f8357,f8341,f8340,f8353,f8368,f8370,f8369,f8387,f8388,f8389,f8390,f8391,f8393,f8578,f8394,f8481,f8577,f8687,f8688,f8689,f8690,f8691,f8692,f8694,f8782,f8888,f8890,f8936,f8937,f8938,f8939,f8940,f8942,f8943,f9030,f9071,f9135,f9136,f9137,f9138,f9139,f9334,f9140,f9336,f9142,f9230,f9335,f9231,f9337,f9255,f9256,f9342,f9356,f9357,f9358,f9359,f9360,f8829,f9361,f9363,f9451,f8830,f9541,f9542,f8889,f9543,f9544,f9545,f9546,f9547,f9742,f9548,f9744,f9550,f9638,f9743,f9639,f9745,f1279,f9648,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f9657,f9663,f9664,f9754,f9738,f9737,f9750,f9901,f9903,f9902,f9916,f9923,f9921,f9924,f9926,f9927,f9928,f9931,f10018,f10117,f10241,f10243,f10242,f10240,f10245,f10244,f10250,f10258,f10261,f10289,f10291,f10336,f10337,f10290,f10335,f10339,f10338,f10367,f10369,f10368,f10414,f10415,f10416,f10418,f10420,f10508,f10562,f10624,f10625,f10626,f1277,f10628,f1281,f10630,f1280,f10718,f10823,f1282,f10744,f10745,f10829,f4286,f11000,f11002,f11019])).
fof(f11019,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f11018,f1144])).
fof(f1144,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb',aHL_TRUTH)).
fof(f11018,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f1162,f1254])).
fof(f1254,plain,(
  s(bool,'T_0') = s(bool,sK15)),
  inference(resolution,[],[f851,f1122])).
fof(f1122,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f847])).
fof(f847,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f851,plain,(
  p(s(bool,sK15))),
  inference(cnf_transformation,[],[f731])).
fof(f731,plain,(
  p(s(bool,sK15)) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) & (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))) | ~p(s(bool,sK15)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f666])).
fof(f666,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,X3)) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2)))))) & (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))) | ~p(s(bool,X3))))),
  inference(flattening,[],[f665])).
fof(f665,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,X3)) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2)))))) & (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))) | ~p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,X3)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2)))))) | (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))) & p(s(bool,X3)))))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,X3)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2)))))) | (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X1),s('type.bvlSem.state'(X0),X2))))) & p(s(bool,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,X9)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X8))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X8)))))) | (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X8))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X8))))) & p(s(bool,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,X9)) => (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X8))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X8)))))) | (s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X8))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X8))))) & p(s(bool,X9)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bvlSem.bvlSem_|49||49|7_prove_ltb')).
fof(f11002,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f11001,f941])).
fof(f941,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1))))))),
  inference(flattening,[],[f575])).
fof(f575,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X10,X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X10))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X10))))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bvlSem.bvlSem_|49||49|6_prove_ltb|split|0')).
fof(f11001,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | $spl1),
  inference(subsumption_resolution,[],[f10847,f1160])).
fof(f10847,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))),
  inference(resolution,[],[f1333,f1058])).
fof(f1058,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f808])).
fof(f808,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f807])).
fof(f807,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f485])).
fof(f485,axiom,(
  ! [X45,X19] : (s('type.num.num',X19) = s('type.num.num',X45) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X19)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X45))))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.arithmetic.EQ_LESS_EQ')).
fof(f1333,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))))))),
  inference(resolution,[],[f852,f897])).
fof(f897,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f434])).
fof(f434,axiom,(
  ! [X45,X19] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X45),s('type.num.num',X19)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X45)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f852,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))),
  inference(cnf_transformation,[],[f731])).
fof(f11000,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10999,f941])).
fof(f10999,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | $spl1),
  inference(subsumption_resolution,[],[f10846,f1160])).
fof(f10846,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) = s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))),
  inference(resolution,[],[f1333,f1058])).
fof(f4286,plain,(
  ( ! [X273,X274] : (p(s(bool,X273)) | ~sP11(X273,X274) | p(s(bool,X274))) ) | $spl66),
  inference(cnf_transformation,[],[f4286_D])).
fof(f4286_D,plain,(
  ( ! [X273,X274] : (p(s(bool,X273)) | ~sP11(X273,X274) | p(s(bool,X274))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f10829,plain,(
  ( ! [X524,X523] : (sP11(X523,X524) | sP10(X523,'F_0')) ) | $spl172),
  inference(cnf_transformation,[],[f10829_D])).
fof(f10829_D,plain,(
  ( ! [X524,X523] : (sP11(X523,X524) | sP10(X523,'F_0')) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f10745,plain,(
  ( ! [X528,X530,X529,X527] : (sP11(X527,X528) | sP10(X527,'F_0') | p(s(bool,X529)) | p(s(bool,X530)) | ~sP11(X529,X530)) )),
  inference(resolution,[],[f5232,f1045])).
fof(f1045,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f803])).
fof(f803,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f802])).
fof(f802,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f688,f729])).
fof(f729,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f688,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f687])).
fof(f687,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f599])).
fof(f599,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f446])).
fof(f446,axiom,(
  ! [X33,X6] : ((~(p(s(bool,X6)) | p(s(bool,X33))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X33)) => p(s(bool,'F_0')))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.sat.sat_9_prove_ltb')).
fof(f5232,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP11(X0,X2) | sP10(X0,X1)) )),
  inference(resolution,[],[f4686,f1232])).
fof(f1232,plain,(
  ( ! [X144] : (~sP10(sK15,X144) | ~p(s(bool,X144))) )),
  inference(resolution,[],[f851,f1026])).
fof(f1026,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f796])).
fof(f796,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f795])).
fof(f795,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f794])).
fof(f794,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f727])).
fof(f727,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f4686,plain,(
  ( ! [X310,X309,X311] : (sP10(sK15,X309) | sP10(X310,X309) | sP11(X310,X311)) )),
  inference(resolution,[],[f1639,f1043])).
fof(f1043,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f801])).
fof(f801,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f800])).
fof(f800,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f729])).
fof(f1639,plain,(
  ( ! [X220,X219] : (sP10(sK15,X219) | sP10(X220,X219) | p(s(bool,X220))) )),
  inference(resolution,[],[f1234,f1030])).
fof(f1030,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f796])).
fof(f1234,plain,(
  ( ! [X146] : (p(s(bool,X146)) | sP10(sK15,X146)) )),
  inference(resolution,[],[f851,f1027])).
fof(f1027,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f796])).
fof(f10744,plain,(
  ( ! [X524,X526,X525,X523] : (sP11(X523,X524) | sP10(X523,'F_0') | p(s(bool,X525)) | p(s(bool,X526)) | ~sP11(X525,X526)) )),
  inference(resolution,[],[f5232,f1041])).
fof(f1041,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f801])).
fof(f1282,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl16),
  inference(cnf_transformation,[],[f1282_D])).
fof(f1282_D,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f10823,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X6,X8)) ) | $spl170),
  inference(cnf_transformation,[],[f10823_D])).
fof(f10823_D,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X6,X8)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f10718,plain,(
  ( ! [X445,X441,X443,X444,X442] : (sP11(X441,X442) | sP10(X441,X443) | p(s(bool,X444)) | ~p(s(bool,X445)) | s(bool,X443) != s(bool,X445)) )),
  inference(resolution,[],[f5232,f1151])).
fof(f1151,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f858])).
fof(f858,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f733,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f732])).
fof(f732,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f667])).
fof(f667,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X16,X43,X23,X44] : ((s(bool,X16) = s(bool,X43) & (p(s(bool,X43)) => s(bool,X23) = s(bool,X44))) => ((p(s(bool,X16)) => p(s(bool,X23))) <=> (p(s(bool,X43)) => p(s(bool,X44)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CONG')).
fof(f1280,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | $spl12),
  inference(cnf_transformation,[],[f1280_D])).
fof(f1280_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10630,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X42,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f5232,f861])).
fof(f861,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f1281,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl14),
  inference(cnf_transformation,[],[f1281_D])).
fof(f1281_D,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f10628,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X30,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f5232,f860])).
fof(f860,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f1277,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl6),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f10626,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X18,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f5232,f859])).
fof(f859,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f10625,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X12,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f5232,f857])).
fof(f857,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f10624,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP11(X6,X7) | sP10(X6,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f5232,f855])).
fof(f855,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f733])).
fof(f10562,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl168),
  inference(cnf_transformation,[],[f10562_D])).
fof(f10562_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f10508,plain,(
  ( ! [X158,X159] : (p(s(bool,X158)) | ~p(s(bool,X159)) | s(bool,X159) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f1489,f1151])).
fof(f1489,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))),
  inference(subsumption_resolution,[],[f1438,f904])).
fof(f904,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f675])).
fof(f675,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f537])).
fof(f537,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f489])).
fof(f489,axiom,(
  ! [X45,X19] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X45),s('type.num.num',X19)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X19)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.arithmetic.LESS_IMP_LESS_OR_EQ')).
fof(f1438,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))),
  inference(superposition,[],[f852,f935])).
fof(f935,plain,(
  ( ! [X2,X0,X1] : (s('type.bvlSem.state'(X0),X2) = s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1)))))) | s('type.bvlSem.state'(X0),X2) = s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),
  inference(ennf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1)))))) => s('type.bvlSem.state'(X0),X2) = s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),
  inference(flattening,[],[f563])).
fof(f563,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X2))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X0),X1)))))) => s('type.bvlSem.state'(X0),X2) = s('type.bvlSem.state'(X0),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X0),X2),s('type.bvlSem.state'(X0),X1))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X10,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X7))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(X6),X10)))))) => s('type.bvlSem.state'(X6),X7) = s('type.bvlSem.state'(X6),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(X6),X7),s('type.bvlSem.state'(X6),X10))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bvlSem.bvlSem_|49||49|6_prove_ltb|split|1')).
fof(f10420,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f1489,f861])).
fof(f10418,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f1489,f860])).
fof(f10416,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f1489,f859])).
fof(f10415,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f1489,f857])).
fof(f10414,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f1489,f855])).
fof(f10368,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X1,X0)) ) | $spl166),
  inference(cnf_transformation,[],[f10368_D])).
fof(f10368_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X1,X0)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10369,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl144),
  inference(subsumption_resolution,[],[f10354,f1088])).
fof(f1088,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f826])).
fof(f826,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f825])).
fof(f825,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f365])).
fof(f365,axiom,(
  ! [X15] : ((p(s(bool,X15)) | p(s(bool,'F_0'))) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.OR_CLAUSES|split|3')).
fof(f10354,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl144),
  inference(resolution,[],[f10025,f1041])).
fof(f10025,plain,(
  ( ! [X269,X267,X268] : (sP10(X267,'T_0') | sP11(X268,X269) | sP10(X269,X267)) ) | $spl144),
  inference(resolution,[],[f9902,f5203])).
fof(f5203,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP11(X2,X0) | sP10(X0,X1)) )),
  inference(resolution,[],[f4685,f1232])).
fof(f4685,plain,(
  ( ! [X308,X306,X307] : (sP10(sK15,X306) | sP10(X307,X306) | sP11(X308,X307)) )),
  inference(resolution,[],[f1639,f1042])).
fof(f1042,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f801])).
fof(f10367,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl144),
  inference(subsumption_resolution,[],[f10353,f1088])).
fof(f10353,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl144),
  inference(resolution,[],[f10025,f1045])).
fof(f10338,plain,(
  ( ! [X1] : (sP11('T_0',X1)) ) | $spl162),
  inference(cnf_transformation,[],[f10338_D])).
fof(f10338_D,plain,(
  ( ! [X1] : (sP11('T_0',X1)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f10339,plain,(
  ( ! [X0] : (sP10(X0,'T_0')) ) | $spl164),
  inference(cnf_transformation,[],[f10339_D])).
fof(f10339_D,plain,(
  ( ! [X0] : (sP10(X0,'T_0')) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f10335,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | sP11('T_0',X1)) ) | $spl144),
  inference(factoring,[],[f10023])).
fof(f10023,plain,(
  ( ! [X263,X265,X264] : (sP10(X263,'T_0') | sP11(X264,X265) | sP10(X263,X264)) ) | $spl144),
  inference(resolution,[],[f9902,f4450])).
fof(f4450,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP11(X1,X2) | sP10(X0,X1)) )),
  inference(resolution,[],[f3869,f1232])).
fof(f3869,plain,(
  ( ! [X310,X309,X311] : (sP10(sK15,X309) | sP10(X309,X310) | sP11(X310,X311)) )),
  inference(resolution,[],[f1638,f1043])).
fof(f1638,plain,(
  ( ! [X218,X217] : (sP10(sK15,X217) | sP10(X217,X218) | p(s(bool,X218))) )),
  inference(resolution,[],[f1234,f1027])).
fof(f10290,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X0,X1)) ) | $spl160),
  inference(cnf_transformation,[],[f10290_D])).
fof(f10290_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | sP10(X0,X1)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f10337,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl144),
  inference(subsumption_resolution,[],[f10323,f1088])).
fof(f10323,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl144),
  inference(resolution,[],[f10023,f1041])).
fof(f10336,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl144),
  inference(subsumption_resolution,[],[f10322,f1088])).
fof(f10322,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl144),
  inference(resolution,[],[f10023,f1045])).
fof(f10291,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl144),
  inference(subsumption_resolution,[],[f10276,f1088])).
fof(f10276,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,'T_0') | sP10(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl144),
  inference(resolution,[],[f10022,f1041])).
fof(f10022,plain,(
  ( ! [X261,X260,X262] : (sP10(X260,'T_0') | sP11(X261,X262) | sP10(X260,X262)) ) | $spl144),
  inference(resolution,[],[f9902,f4421])).
fof(f4421,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP11(X2,X1) | sP10(X0,X1)) )),
  inference(resolution,[],[f3868,f1232])).
fof(f3868,plain,(
  ( ! [X308,X306,X307] : (sP10(sK15,X306) | sP10(X306,X307) | sP11(X308,X307)) )),
  inference(resolution,[],[f1638,f1042])).
fof(f10289,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl144),
  inference(subsumption_resolution,[],[f10275,f1088])).
fof(f10275,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'T_0') | sP10(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl144),
  inference(resolution,[],[f10022,f1045])).
fof(f10261,plain,(
  $false | $spl154),
  inference(subsumption_resolution,[],[f10254,f1144])).
fof(f10254,plain,(
  ~p(s(bool,'T_0')) | $spl154),
  inference(resolution,[],[f10244,f1233])).
fof(f1233,plain,(
  ( ! [X145] : (~sP10(X145,sK15) | ~p(s(bool,X145))) )),
  inference(resolution,[],[f851,f1026])).
fof(f10258,plain,(
  sP10(sK15,'T_0') | $spl158),
  inference(cnf_transformation,[],[f10258_D])).
fof(f10258_D,plain,(
  sP10(sK15,'T_0') <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f10250,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | sP10(sK15,'T_0')) ) | $spl154),
  inference(resolution,[],[f10244,f1636])).
fof(f1636,plain,(
  ( ! [X213,X214] : (~sP10(X213,X214) | ~p(s(bool,X214)) | sP10(sK15,X213)) )),
  inference(resolution,[],[f1234,f1026])).
fof(f10244,plain,(
  ( ! [X1] : (sP10('T_0',X1)) ) | $spl154),
  inference(cnf_transformation,[],[f10244_D])).
fof(f10244_D,plain,(
  ( ! [X1] : (sP10('T_0',X1)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f10245,plain,(
  ( ! [X0] : (sP11(X0,'T_0')) ) | $spl156),
  inference(cnf_transformation,[],[f10245_D])).
fof(f10245_D,plain,(
  ( ! [X0] : (sP11(X0,'T_0')) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f10240,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0') | sP10('T_0',X1)) ) | $spl112),
  inference(factoring,[],[f9647])).
fof(f9647,plain,(
  ( ! [X478,X480,X479] : (sP11(X478,X479) | sP10(X479,X480) | sP10('T_0',X480)) ) | $spl112),
  inference(resolution,[],[f5203,f8369])).
fof(f10242,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP10('T_0',X1)) ) | $spl152),
  inference(cnf_transformation,[],[f10242_D])).
fof(f10242_D,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | sP10('T_0',X1)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f10243,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,X4) | sP10('T_0',X4) | p(s(bool,X5)) | p(s(bool,X3))) ) | $spl112),
  inference(subsumption_resolution,[],[f10219,f1088])).
fof(f10219,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,X4) | sP10('T_0',X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X3))) ) | $spl112),
  inference(resolution,[],[f9647,f1041])).
fof(f10241,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1) | sP10('T_0',X1) | p(s(bool,X2)) | p(s(bool,X0))) ) | $spl112),
  inference(subsumption_resolution,[],[f10218,f1088])).
fof(f10218,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1) | sP10('T_0',X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X0))) ) | $spl112),
  inference(resolution,[],[f9647,f1045])).
fof(f10117,plain,(
  ( ! [X2,X0] : (sP10(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl150),
  inference(cnf_transformation,[],[f10117_D])).
fof(f10117_D,plain,(
  ( ! [X2,X0] : (sP10(X0,'T_0') | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f10018,plain,(
  ( ! [X251,X252,X250] : (sP10(X250,'T_0') | p(s(bool,X251)) | p(s(bool,X252)) | s(bool,X250) != s(bool,X252)) ) | $spl144),
  inference(resolution,[],[f9902,f1151])).
fof(f9931,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,'T_0') | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl144),
  inference(resolution,[],[f9902,f861])).
fof(f9928,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'T_0') | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl144),
  inference(resolution,[],[f9902,f860])).
fof(f9927,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,'T_0') | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl144),
  inference(resolution,[],[f9902,f859])).
fof(f9926,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,'T_0') | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl144),
  inference(resolution,[],[f9902,f857])).
fof(f9924,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl144),
  inference(resolution,[],[f9902,f855])).
fof(f9921,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))))) ) | $spl146),
  inference(cnf_transformation,[],[f9921_D])).
fof(f9921_D,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f9923,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))) | $spl149),
  inference(cnf_transformation,[],[f9923_D])).
fof(f9923_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f9916,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))))))) ) | $spl1),
  inference(superposition,[],[f1310,f901])).
fof(f901,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(flattening,[],[f752])).
fof(f752,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f536])).
fof(f536,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(rectify,[],[f185])).
fof(f185,axiom,(
  ! [X19,X45] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num',X45)))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X45) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19))))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.arithmetic.LE_MULT_CANCEL_LBARE|split|1')).
fof(f1310,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) | $spl1),
  inference(subsumption_resolution,[],[f1296,f1057])).
fof(f1057,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f808])).
fof(f1296,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK13))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | $spl1),
  inference(superposition,[],[f1160,f935])).
fof(f9902,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10(X0,'T_0')) ) | $spl144),
  inference(cnf_transformation,[],[f9902_D])).
fof(f9902_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10(X0,'T_0')) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f9903,plain,(
  ( ! [X2,X3] : (sP10(X2,'T_0') | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f9900,f1088])).
fof(f9900,plain,(
  ( ! [X2,X3] : (sP10(X2,'T_0') | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(resolution,[],[f9660,f1041])).
fof(f9660,plain,(
  ( ! [X514,X515] : (sP11(X514,X515) | sP10(X515,'T_0')) )),
  inference(resolution,[],[f5203,f1143])).
fof(f1143,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.TRUTH')).
fof(f9901,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f9899,f1088])).
fof(f9899,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0') | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(resolution,[],[f9660,f1045])).
fof(f9750,plain,(
  ( ! [X482,X481] : (sP11(X481,X482) | sP10(X482,'T_0')) ) | $spl140),
  inference(cnf_transformation,[],[f9750_D])).
fof(f9750_D,plain,(
  ( ! [X482,X481] : (sP11(X481,X482) | sP10(X482,'T_0')) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f9737,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP11(X58,X59) | sP10(X59,'T_0')) )),
  inference(duplicate_literal_removal,[],[f9712])).
fof(f9712,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP11(X58,X59) | sP10(X59,'T_0') | ~p(s(bool,X57))) )),
  inference(superposition,[],[f5203,f1122])).
fof(f9738,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP11(X55,X56) | sP10(X56,'T_0')) )),
  inference(duplicate_literal_removal,[],[f9711])).
fof(f9711,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP11(X55,X56) | sP10(X56,'T_0') | ~p(s(bool,X54))) )),
  inference(superposition,[],[f5203,f1124])).
fof(f1124,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f848])).
fof(f848,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f9754,plain,(
  ( ! [X520,X521] : (sP11(X520,X521) | sP10(X521,'F_0')) ) | $spl142),
  inference(cnf_transformation,[],[f9754_D])).
fof(f9754_D,plain,(
  ( ! [X520,X521] : (sP11(X520,X521) | sP10(X521,'F_0')) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f9664,plain,(
  ( ! [X524,X526,X525,X527] : (sP11(X524,X525) | sP10(X525,'F_0') | p(s(bool,X526)) | p(s(bool,X527)) | ~sP11(X526,X527)) )),
  inference(resolution,[],[f5203,f1045])).
fof(f9663,plain,(
  ( ! [X520,X522,X521,X523] : (sP11(X520,X521) | sP10(X521,'F_0') | p(s(bool,X522)) | p(s(bool,X523)) | ~sP11(X522,X523)) )),
  inference(resolution,[],[f5203,f1041])).
fof(f9657,plain,(
  ( ! [X509,X507,X508] : (sP11(X507,X508) | sP10(X508,'T_0') | ~p(s(bool,X509))) )),
  inference(resolution,[],[f5203,f1107])).
fof(f1107,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f838])).
fof(f838,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f837])).
fof(f837,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X15] : ((p(s(bool,X15)) & p(s(bool,'T_0'))) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.AND_CLAUSES|split|1')).
fof(f9656,plain,(
  ( ! [X505,X504,X506] : (sP11(X504,X505) | sP10(X505,'T_0') | ~p(s(bool,X506))) )),
  inference(resolution,[],[f5203,f1103])).
fof(f1103,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f836])).
fof(f836,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f835])).
fof(f835,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X15] : ((p(s(bool,'T_0')) & p(s(bool,X15))) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.AND_CLAUSES|split|0')).
fof(f9655,plain,(
  ( ! [X502,X501,X503] : (sP11(X501,X502) | sP10(X502,'T_0') | ~p(s(bool,X503))) )),
  inference(resolution,[],[f5203,f1100])).
fof(f1100,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f834])).
fof(f834,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f833])).
fof(f833,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f360])).
fof(f360,axiom,(
  ! [X15] : ((p(s(bool,'T_0')) | p(s(bool,X15))) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.OR_CLAUSES|split|0')).
fof(f9654,plain,(
  ( ! [X499,X500,X498] : (sP11(X498,X499) | sP10(X499,'T_0') | ~p(s(bool,X500))) )),
  inference(resolution,[],[f5203,f1096])).
fof(f1096,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f832])).
fof(f832,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f831])).
fof(f831,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f628])).
fof(f628,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f361])).
fof(f361,axiom,(
  ! [X15] : ((p(s(bool,X15)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.OR_CLAUSES|split|1')).
fof(f9653,plain,(
  ( ! [X496,X495,X497] : (sP11(X495,X496) | sP10(X496,'T_0') | ~p(s(bool,X497))) )),
  inference(resolution,[],[f5203,f1085])).
fof(f1085,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f824])).
fof(f824,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f823])).
fof(f823,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f703])).
fof(f703,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,X15))) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CLAUSES|split|3')).
fof(f9652,plain,(
  ( ! [X492,X494,X493] : (sP11(X492,X493) | sP10(X493,'T_0') | p(s(bool,X494))) )),
  inference(resolution,[],[f5203,f1084])).
fof(f1084,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f824])).
fof(f9651,plain,(
  ( ! [X490,X489,X491] : (sP11(X489,X490) | sP10(X490,'T_0') | p(s(bool,X491))) )),
  inference(resolution,[],[f5203,f1081])).
fof(f1081,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f822])).
fof(f822,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f821])).
fof(f821,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f702])).
fof(f702,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X15] : ((p(s(bool,'T_0')) => p(s(bool,X15))) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CLAUSES|split|0')).
fof(f9650,plain,(
  ( ! [X486,X488,X487] : (sP11(X486,X487) | sP10(X487,'T_0') | ~p(s(bool,X488))) )),
  inference(resolution,[],[f5203,f1079])).
fof(f1079,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f820])).
fof(f820,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f819])).
fof(f819,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f701])).
fof(f701,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X15] : ((p(s(bool,'F_0')) => p(s(bool,X15))) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f9648,plain,(
  ( ! [X482,X481,X483] : (sP11(X481,X482) | sP10(X482,'T_0') | p(s(bool,X483))) )),
  inference(resolution,[],[f5203,f1075])).
fof(f1075,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f818])).
fof(f818,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f817])).
fof(f817,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CLAUSES|split|1')).
fof(f1279,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) | $spl10),
  inference(cnf_transformation,[],[f1279_D])).
fof(f1279_D,plain,(
  ( ! [X4,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,X3)) | ~p(s(bool,X4))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f9745,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X37,X38)) ) | $spl138),
  inference(cnf_transformation,[],[f9745_D])).
fof(f9745_D,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X37,X38)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9639,plain,(
  ( ! [X447,X449,X446,X448,X450] : (sP11(X446,X447) | sP10(X447,X448) | ~p(s(bool,X449)) | p(s(bool,X450)) | s(bool,X449) != s(bool,X450)) )),
  inference(resolution,[],[f5203,f1151])).
fof(f9743,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X7,X8)) ) | $spl136),
  inference(cnf_transformation,[],[f9743_D])).
fof(f9743_D,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X7,X8)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f9638,plain,(
  ( ! [X445,X441,X443,X444,X442] : (sP11(X441,X442) | sP10(X442,X443) | p(s(bool,X444)) | ~p(s(bool,X445)) | s(bool,X443) != s(bool,X445)) )),
  inference(resolution,[],[f5203,f1151])).
fof(f9550,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X43,X44) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f5203,f861])).
fof(f9744,plain,(
  ( ! [X39,X37,X41,X38,X36] : (sP11(X36,X37) | sP10(X37,X38) | ~p(s(bool,X39)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(subsumption_resolution,[],[f9549,f1151])).
fof(f9549,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP11(X36,X37) | sP10(X37,X38) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f5203,f860])).
fof(f9548,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X31,X32) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f5203,f860])).
fof(f9742,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X1,X2)) ) | $spl134),
  inference(cnf_transformation,[],[f9742_D])).
fof(f9742_D,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X1,X2)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9547,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP11(X24,X25) | sP10(X25,X26) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X26) != s(bool,X29) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f5203,f859])).
fof(f9546,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X19,X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f5203,f859])).
fof(f9545,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X13,X14) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f5203,f857])).
fof(f9544,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP11(X6,X7) | sP10(X7,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f5203,f855])).
fof(f9543,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X1,X2) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X2) != s(bool,X5) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f5203,f855])).
fof(f8889,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | p(s(bool,X1)) | sP10(X0,X1)) ) | $spl122),
  inference(cnf_transformation,[],[f8889_D])).
fof(f8889_D,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | p(s(bool,X1)) | sP10(X0,X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f9542,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,X4) | sP10('T_0',X3) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl112),
  inference(subsumption_resolution,[],[f9519,f1088])).
fof(f9519,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,X4) | sP10('T_0',X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl112),
  inference(resolution,[],[f9239,f1041])).
fof(f9239,plain,(
  ( ! [X478,X480,X479] : (sP11(X478,X479) | sP10(X480,X478) | sP10('T_0',X480)) ) | $spl112),
  inference(resolution,[],[f4450,f8369])).
fof(f9541,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1) | sP10('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl112),
  inference(subsumption_resolution,[],[f9518,f1088])).
fof(f9518,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1) | sP10('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl112),
  inference(resolution,[],[f9239,f1045])).
fof(f8830,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl120),
  inference(cnf_transformation,[],[f8830_D])).
fof(f8830_D,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f9451,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | ~p(s(bool,X160)) | s(bool,X160) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f1151])).
fof(f9363,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X24) != s(bool,X26)) ) | $spl27),
  inference(resolution,[],[f1495,f861])).
fof(f9361,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f860])).
fof(f8829,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl118),
  inference(cnf_transformation,[],[f8829_D])).
fof(f8829_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f9360,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X15) != s(bool,X16)) ) | $spl27),
  inference(resolution,[],[f1495,f859])).
fof(f9359,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f859])).
fof(f9358,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X9) != s(bool,X11)) ) | $spl27),
  inference(resolution,[],[f1495,f857])).
fof(f9357,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f855])).
fof(f9356,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X3) != s(bool,X4)) ) | $spl27),
  inference(resolution,[],[f1495,f855])).
fof(f9342,plain,(
  ( ! [X520,X521] : (sP11(X520,X521) | sP10('F_0',X520)) ) | $spl132),
  inference(cnf_transformation,[],[f9342_D])).
fof(f9342_D,plain,(
  ( ! [X520,X521] : (sP11(X520,X521) | sP10('F_0',X520)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f9256,plain,(
  ( ! [X524,X526,X525,X527] : (sP11(X524,X525) | sP10('F_0',X524) | p(s(bool,X526)) | p(s(bool,X527)) | ~sP11(X526,X527)) )),
  inference(resolution,[],[f4450,f1045])).
fof(f9255,plain,(
  ( ! [X520,X522,X521,X523] : (sP11(X520,X521) | sP10('F_0',X520) | p(s(bool,X522)) | p(s(bool,X523)) | ~sP11(X522,X523)) )),
  inference(resolution,[],[f4450,f1041])).
fof(f9337,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X38,X36)) ) | $spl130),
  inference(cnf_transformation,[],[f9337_D])).
fof(f9337_D,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X38,X36)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f9231,plain,(
  ( ! [X447,X449,X446,X448,X450] : (sP11(X446,X447) | sP10(X448,X446) | ~p(s(bool,X449)) | p(s(bool,X450)) | s(bool,X449) != s(bool,X450)) )),
  inference(resolution,[],[f4450,f1151])).
fof(f9335,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X8,X6)) ) | $spl128),
  inference(cnf_transformation,[],[f9335_D])).
fof(f9335_D,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X8,X6)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f9230,plain,(
  ( ! [X445,X441,X443,X444,X442] : (sP11(X441,X442) | sP10(X443,X441) | p(s(bool,X444)) | ~p(s(bool,X445)) | s(bool,X443) != s(bool,X445)) )),
  inference(resolution,[],[f4450,f1151])).
fof(f9142,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X44,X42) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f4450,f861])).
fof(f9336,plain,(
  ( ! [X39,X37,X41,X38,X36] : (sP11(X36,X37) | sP10(X38,X36) | ~p(s(bool,X39)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(subsumption_resolution,[],[f9141,f1151])).
fof(f9141,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP11(X36,X37) | sP10(X38,X36) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f4450,f860])).
fof(f9140,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X32,X30) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f4450,f860])).
fof(f9334,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X2,X0)) ) | $spl126),
  inference(cnf_transformation,[],[f9334_D])).
fof(f9334_D,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X2,X0)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f9139,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP11(X24,X25) | sP10(X26,X24) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X26) != s(bool,X29) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f4450,f859])).
fof(f9138,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X20,X18) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4450,f859])).
fof(f9137,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X14,X12) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f4450,f857])).
fof(f9136,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP11(X6,X7) | sP10(X8,X6) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f4450,f855])).
fof(f9135,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X2,X0) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X2) != s(bool,X5) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f4450,f855])).
fof(f9071,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl124),
  inference(cnf_transformation,[],[f9071_D])).
fof(f9071_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f9030,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f1151])).
fof(f1494,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | $spl26),
  inference(cnf_transformation,[],[f1494_D])).
fof(f1494_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8943,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X25) != s(bool,X26)) ) | $spl26),
  inference(resolution,[],[f1494,f861])).
fof(f8942,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X22) | s(bool,X23) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f861])).
fof(f8940,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f860])).
fof(f8939,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f859])).
fof(f8938,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X10) != s(bool,X11)) ) | $spl26),
  inference(resolution,[],[f1494,f857])).
fof(f8937,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X7) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f857])).
fof(f8936,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl26),
  inference(resolution,[],[f1494,f855])).
fof(f8890,plain,(
  ( ! [X4,X5,X3] : (sP10('T_0',X3) | sP10(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl112),
  inference(subsumption_resolution,[],[f8875,f1088])).
fof(f8875,plain,(
  ( ! [X4,X5,X3] : (sP10('T_0',X3) | sP10(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) ) | $spl112),
  inference(resolution,[],[f8485,f1041])).
fof(f8485,plain,(
  ( ! [X261,X260,X262] : (sP10('T_0',X260) | sP11(X261,X262) | sP10(X260,X262)) ) | $spl112),
  inference(resolution,[],[f8369,f4421])).
fof(f8888,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl112),
  inference(subsumption_resolution,[],[f8874,f1088])).
fof(f8874,plain,(
  ( ! [X2,X0,X1] : (sP10('T_0',X0) | sP10(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl112),
  inference(resolution,[],[f8485,f1045])).
fof(f8782,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | ~p(s(bool,X160)) | s(bool,X160) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f1151])).
fof(f8694,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X24) != s(bool,X26)) ) | $spl27),
  inference(resolution,[],[f1495,f861])).
fof(f8692,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f860])).
fof(f8691,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X15) != s(bool,X16)) ) | $spl27),
  inference(resolution,[],[f1495,f859])).
fof(f8690,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f859])).
fof(f8689,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X9) != s(bool,X11)) ) | $spl27),
  inference(resolution,[],[f1495,f857])).
fof(f8688,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl27),
  inference(resolution,[],[f1495,f855])).
fof(f8687,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X3) != s(bool,X4)) ) | $spl27),
  inference(resolution,[],[f1495,f855])).
fof(f8577,plain,(
  ( ! [X2,X0] : (sP10('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl114),
  inference(cnf_transformation,[],[f8577_D])).
fof(f8577_D,plain,(
  ( ! [X2,X0] : (sP10('T_0',X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f8481,plain,(
  ( ! [X251,X252,X250] : (sP10('T_0',X250) | p(s(bool,X251)) | p(s(bool,X252)) | s(bool,X250) != s(bool,X252)) ) | $spl112),
  inference(resolution,[],[f8369,f1151])).
fof(f8394,plain,(
  ( ! [X30,X28,X31,X29] : (sP10('T_0',X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl112),
  inference(resolution,[],[f8369,f861])).
fof(f8578,plain,(
  ( ! [X4,X7] : (sP10('T_0',X4) | s(bool,X4) != s(bool,X7)) ) | $spl116),
  inference(cnf_transformation,[],[f8578_D])).
fof(f8578_D,plain,(
  ( ! [X4,X7] : (sP10('T_0',X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f8393,plain,(
  ( ! [X26,X24,X27,X25] : (sP10('T_0',X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl112),
  inference(resolution,[],[f8369,f861])).
fof(f8391,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('T_0',X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl112),
  inference(resolution,[],[f8369,f860])).
fof(f8390,plain,(
  ( ! [X14,X12,X15,X13] : (sP10('T_0',X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl112),
  inference(resolution,[],[f8369,f859])).
fof(f8389,plain,(
  ( ! [X10,X8,X11,X9] : (sP10('T_0',X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl112),
  inference(resolution,[],[f8369,f857])).
fof(f8388,plain,(
  ( ! [X6,X4,X7,X5] : (sP10('T_0',X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl112),
  inference(resolution,[],[f8369,f857])).
fof(f8387,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('T_0',X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl112),
  inference(resolution,[],[f8369,f855])).
fof(f8369,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10('T_0',X0)) ) | $spl112),
  inference(cnf_transformation,[],[f8369_D])).
fof(f8369_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP10('T_0',X0)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8370,plain,(
  ( ! [X2,X3] : (sP10('T_0',X2) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f8367,f1088])).
fof(f8367,plain,(
  ( ! [X2,X3] : (sP10('T_0',X2) | p(s(bool,'F_0')) | p(s(bool,X3)) | p(s(bool,X2))) )),
  inference(resolution,[],[f8263,f1041])).
fof(f8263,plain,(
  ( ! [X511,X512] : (sP11(X511,X512) | sP10('T_0',X512)) )),
  inference(resolution,[],[f4421,f1143])).
fof(f8368,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f8366,f1088])).
fof(f8366,plain,(
  ( ! [X0,X1] : (sP10('T_0',X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(resolution,[],[f8263,f1045])).
fof(f8353,plain,(
  ( ! [X478,X479] : (sP11(X478,X479) | sP10('T_0',X479)) ) | $spl108),
  inference(cnf_transformation,[],[f8353_D])).
fof(f8353_D,plain,(
  ( ! [X478,X479] : (sP11(X478,X479) | sP10('T_0',X479)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f8340,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP11(X58,X59) | sP10('T_0',X59)) )),
  inference(duplicate_literal_removal,[],[f8315])).
fof(f8315,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | sP11(X58,X59) | sP10('T_0',X59) | ~p(s(bool,X57))) )),
  inference(superposition,[],[f4421,f1122])).
fof(f8341,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP11(X55,X56) | sP10('T_0',X56)) )),
  inference(duplicate_literal_removal,[],[f8314])).
fof(f8314,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | sP11(X55,X56) | sP10('T_0',X56) | ~p(s(bool,X54))) )),
  inference(superposition,[],[f4421,f1124])).
fof(f8357,plain,(
  ( ! [X518,X517] : (sP11(X517,X518) | sP10('F_0',X518)) ) | $spl110),
  inference(cnf_transformation,[],[f8357_D])).
fof(f8357_D,plain,(
  ( ! [X518,X517] : (sP11(X517,X518) | sP10('F_0',X518)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f8267,plain,(
  ( ! [X524,X522,X521,X523] : (sP11(X521,X522) | sP10('F_0',X522) | p(s(bool,X523)) | p(s(bool,X524)) | ~sP11(X523,X524)) )),
  inference(resolution,[],[f4421,f1045])).
fof(f8266,plain,(
  ( ! [X520,X518,X517,X519] : (sP11(X517,X518) | sP10('F_0',X518) | p(s(bool,X519)) | p(s(bool,X520)) | ~sP11(X519,X520)) )),
  inference(resolution,[],[f4421,f1041])).
fof(f8260,plain,(
  ( ! [X505,X504,X506] : (sP11(X504,X505) | sP10('T_0',X505) | ~p(s(bool,X506))) )),
  inference(resolution,[],[f4421,f1107])).
fof(f8259,plain,(
  ( ! [X502,X501,X503] : (sP11(X501,X502) | sP10('T_0',X502) | ~p(s(bool,X503))) )),
  inference(resolution,[],[f4421,f1103])).
fof(f8258,plain,(
  ( ! [X499,X500,X498] : (sP11(X498,X499) | sP10('T_0',X499) | ~p(s(bool,X500))) )),
  inference(resolution,[],[f4421,f1100])).
fof(f8257,plain,(
  ( ! [X496,X495,X497] : (sP11(X495,X496) | sP10('T_0',X496) | ~p(s(bool,X497))) )),
  inference(resolution,[],[f4421,f1096])).
fof(f8256,plain,(
  ( ! [X492,X494,X493] : (sP11(X492,X493) | sP10('T_0',X493) | ~p(s(bool,X494))) )),
  inference(resolution,[],[f4421,f1085])).
fof(f8255,plain,(
  ( ! [X490,X489,X491] : (sP11(X489,X490) | sP10('T_0',X490) | p(s(bool,X491))) )),
  inference(resolution,[],[f4421,f1084])).
fof(f8254,plain,(
  ( ! [X486,X488,X487] : (sP11(X486,X487) | sP10('T_0',X487) | p(s(bool,X488))) )),
  inference(resolution,[],[f4421,f1081])).
fof(f8253,plain,(
  ( ! [X484,X485,X483] : (sP11(X483,X484) | sP10('T_0',X484) | ~p(s(bool,X485))) )),
  inference(resolution,[],[f4421,f1079])).
fof(f8251,plain,(
  ( ! [X478,X480,X479] : (sP11(X478,X479) | sP10('T_0',X479) | p(s(bool,X480))) )),
  inference(resolution,[],[f4421,f1075])).
fof(f8348,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X38,X37)) ) | $spl106),
  inference(cnf_transformation,[],[f8348_D])).
fof(f8348_D,plain,(
  ( ! [X37,X38,X36] : (sP11(X36,X37) | sP10(X38,X37)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8243,plain,(
  ( ! [X447,X449,X446,X448,X450] : (sP11(X446,X447) | sP10(X448,X447) | ~p(s(bool,X449)) | p(s(bool,X450)) | s(bool,X449) != s(bool,X450)) )),
  inference(resolution,[],[f4421,f1151])).
fof(f8346,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X8,X7)) ) | $spl104),
  inference(cnf_transformation,[],[f8346_D])).
fof(f8346_D,plain,(
  ( ! [X6,X10,X8,X7] : (sP11(X6,X7) | s(bool,X8) != s(bool,X10) | ~p(s(bool,X10)) | sP10(X8,X7)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8242,plain,(
  ( ! [X445,X441,X443,X444,X442] : (sP11(X441,X442) | sP10(X443,X442) | p(s(bool,X444)) | ~p(s(bool,X445)) | s(bool,X443) != s(bool,X445)) )),
  inference(resolution,[],[f4421,f1151])).
fof(f8154,plain,(
  ( ! [X47,X45,X43,X46,X44,X42] : (sP11(X42,X43) | sP10(X44,X43) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,X47)) )),
  inference(resolution,[],[f4421,f861])).
fof(f8347,plain,(
  ( ! [X39,X37,X41,X38,X36] : (sP11(X36,X37) | sP10(X38,X37) | ~p(s(bool,X39)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(subsumption_resolution,[],[f8153,f1151])).
fof(f8153,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (sP11(X36,X37) | sP10(X38,X37) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X39) != s(bool,X41)) )),
  inference(resolution,[],[f4421,f860])).
fof(f8152,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (sP11(X30,X31) | sP10(X32,X31) | p(s(bool,X33)) | ~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X32) != s(bool,X34)) )),
  inference(resolution,[],[f4421,f860])).
fof(f8345,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X2,X1)) ) | $spl102),
  inference(cnf_transformation,[],[f8345_D])).
fof(f8345_D,plain,(
  ( ! [X2,X0,X5,X1] : (sP11(X0,X1) | s(bool,X2) != s(bool,X5) | sP10(X2,X1)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f8151,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (sP11(X24,X25) | sP10(X26,X25) | ~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X26) != s(bool,X29) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f4421,f859])).
fof(f8150,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (sP11(X18,X19) | sP10(X20,X19) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4421,f859])).
fof(f8149,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (sP11(X12,X13) | sP10(X14,X13) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f4421,f857])).
fof(f8148,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sP11(X6,X7) | sP10(X8,X7) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f4421,f855])).
fof(f8147,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,X1) | sP10(X2,X1) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X2) != s(bool,X5) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f4421,f855])).
fof(f6756,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) | $spl96),
  inference(cnf_transformation,[],[f6756_D])).
fof(f6756_D,plain,(
  ( ! [X6,X7,X5] : (sP3(X6,X7,X5) | ~sP1(X7,X5,X6) | ~sP2(X6,X5,X7) | ~p(s(bool,X5))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7564,plain,(
  ( ! [X4] : (sP10(X4,sK15)) ) | $spl100),
  inference(cnf_transformation,[],[f7564_D])).
fof(f7564_D,plain,(
  ( ! [X4] : (sP10(X4,sK15)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7971,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,sK15) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f7956,f2074])).
fof(f2074,plain,(
  ( ! [X216,X215] : (~sP10(X216,X215) | ~p(s(bool,X216)) | sP10(X215,sK15)) )),
  inference(resolution,[],[f1235,f1026])).
fof(f1235,plain,(
  ( ! [X147] : (p(s(bool,X147)) | sP10(X147,sK15)) )),
  inference(resolution,[],[f851,f1030])).
fof(f7956,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,sK15) | sP10(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(resolution,[],[f6296,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f736])).
fof(f736,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f735])).
fof(f735,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f734])).
fof(f734,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f716])).
fof(f716,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f6296,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,sK15) | sP10(X85,X84) | sP0(X85,X86,X87)) )),
  inference(resolution,[],[f2076,f880])).
fof(f880,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f744])).
fof(f744,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f743])).
fof(f743,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f713])).
fof(f713,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2076,plain,(
  ( ! [X220,X219] : (sP10(X219,sK15) | sP10(X220,X219) | p(s(bool,X220))) )),
  inference(resolution,[],[f1235,f1030])).
fof(f7563,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,sK15) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f7548,f2073])).
fof(f2073,plain,(
  ( ! [X213,X214] : (~sP10(X213,X214) | ~p(s(bool,X214)) | sP10(X213,sK15)) )),
  inference(resolution,[],[f1235,f1026])).
fof(f7548,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,sK15) | sP10(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(resolution,[],[f6044,f866])).
fof(f6044,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(X84,sK15) | sP10(X84,X85) | sP0(X85,X86,X87)) )),
  inference(resolution,[],[f2075,f880])).
fof(f2075,plain,(
  ( ! [X218,X217] : (sP10(X217,sK15) | sP10(X217,X218) | p(s(bool,X218))) )),
  inference(resolution,[],[f1235,f1027])).
fof(f6757,plain,(
  ( ! [X4] : (sP10(sK15,X4)) ) | $spl98),
  inference(cnf_transformation,[],[f6757_D])).
fof(f6757_D,plain,(
  ( ! [X4] : (sP10(sK15,X4)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7158,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(sK15,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f7141,f1637])).
fof(f1637,plain,(
  ( ! [X216,X215] : (~sP10(X216,X215) | ~p(s(bool,X216)) | sP10(sK15,X215)) )),
  inference(resolution,[],[f1234,f1026])).
fof(f7141,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(sK15,X4) | sP10(X5,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(resolution,[],[f4627,f866])).
fof(f4627,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(sK15,X84) | sP10(X85,X84) | sP0(X85,X86,X87)) )),
  inference(resolution,[],[f1639,f880])).
fof(f6755,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(sK15,X4) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(subsumption_resolution,[],[f6738,f1636])).
fof(f6738,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(sK15,X4) | sP10(X4,X5) | sP3(X6,X7,X5) | ~p(s(bool,X5)) | ~sP2(X6,X5,X7) | ~sP1(X7,X5,X6)) )),
  inference(resolution,[],[f3810,f866])).
fof(f3810,plain,(
  ( ! [X87,X85,X86,X84] : (sP10(sK15,X84) | sP10(X84,X85) | sP0(X85,X86,X87)) )),
  inference(resolution,[],[f1638,f880])).
fof(f6477,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,sK15) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) | $spl92),
  inference(cnf_transformation,[],[f6477_D])).
fof(f6477_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,sK15) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6371,plain,(
  ( ! [X344,X346,X345,X347] : (sP10(X344,sK15) | sP10(X345,X344) | p(s(bool,X346)) | p(s(bool,X347)) | s(bool,X345) != s(bool,X347)) )),
  inference(resolution,[],[f2076,f1151])).
fof(f6284,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,sK15) | sP10(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f2076,f861])).
fof(f6478,plain,(
  ( ! [X6,X5,X9] : (sP10(X5,sK15) | s(bool,X6) != s(bool,X9) | sP10(X6,X5)) ) | $spl94),
  inference(cnf_transformation,[],[f6478_D])).
fof(f6478_D,plain,(
  ( ! [X6,X5,X9] : (sP10(X5,sK15) | s(bool,X6) != s(bool,X9) | sP10(X6,X5)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f6283,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP10(X30,sK15) | sP10(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f2076,f861])).
fof(f6281,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,sK15) | sP10(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f2076,f860])).
fof(f6280,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,sK15) | sP10(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f2076,f859])).
fof(f6279,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,sK15) | sP10(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f2076,f857])).
fof(f6278,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,sK15) | sP10(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) )),
  inference(resolution,[],[f2076,f857])).
fof(f6277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,sK15) | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f2076,f855])).
fof(f6226,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,sK15) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) | $spl88),
  inference(cnf_transformation,[],[f6226_D])).
fof(f6226_D,plain,(
  ( ! [X0,X3,X1] : (sP10(X0,sK15) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6119,plain,(
  ( ! [X344,X346,X345,X347] : (sP10(X344,sK15) | sP10(X344,X345) | p(s(bool,X346)) | p(s(bool,X347)) | s(bool,X345) != s(bool,X347)) )),
  inference(resolution,[],[f2075,f1151])).
fof(f6032,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(X35,sK15) | sP10(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f2075,f861])).
fof(f6227,plain,(
  ( ! [X6,X5,X9] : (sP10(X5,sK15) | s(bool,X6) != s(bool,X9) | sP10(X5,X6)) ) | $spl90),
  inference(cnf_transformation,[],[f6227_D])).
fof(f6227_D,plain,(
  ( ! [X6,X5,X9] : (sP10(X5,sK15) | s(bool,X6) != s(bool,X9) | sP10(X5,X6)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f6031,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP10(X30,sK15) | sP10(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f2075,f861])).
fof(f6029,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(X20,sK15) | sP10(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f2075,f860])).
fof(f6028,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(X15,sK15) | sP10(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f2075,f859])).
fof(f6027,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(X10,sK15) | sP10(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f2075,f857])).
fof(f6026,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(X5,sK15) | sP10(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) )),
  inference(resolution,[],[f2075,f857])).
fof(f6025,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,sK15) | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f2075,f855])).
fof(f5764,plain,(
  ( ! [X0,X1] : (sP10(X0,sK15) | p(s(bool,X1)) | sP10(X1,X0)) ) | $spl86),
  inference(cnf_transformation,[],[f5764_D])).
fof(f5764_D,plain,(
  ( ! [X0,X1] : (sP10(X0,sK15) | p(s(bool,X1)) | sP10(X1,X0)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5802,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X4,X3) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f5789,f1088])).
fof(f5789,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f5646,f1041])).
fof(f5646,plain,(
  ( ! [X70,X68,X69] : (sP10(X68,sK15) | sP10(X69,X68) | sP11(X69,X70)) )),
  inference(subsumption_resolution,[],[f5645,f1155])).
fof(f1155,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f1065])).
fof(f1065,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f811])).
fof(f811,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f614])).
fof(f614,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X16] : (s(X6,X16) = s(X6,X16) <=> p(s(bool,'T_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f5645,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'T_0')) | sP10(X68,sK15) | sP10(X69,X68) | sP11(X69,X70)) )),
  inference(forward_demodulation,[],[f5611,f1254])).
fof(f5611,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,sK15)) | sP10(X68,sK15) | sP10(X69,X68) | sP11(X69,X70)) )),
  inference(resolution,[],[f2074,f4686])).
fof(f5801,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X1,X0) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f5788,f1088])).
fof(f5788,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(resolution,[],[f5646,f1045])).
fof(f5765,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X4,X3) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f5751,f1088])).
fof(f5751,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X4,X3) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(resolution,[],[f5644,f1041])).
fof(f5644,plain,(
  ( ! [X66,X67,X65] : (sP10(X65,sK15) | sP10(X66,X65) | sP11(X67,X66)) )),
  inference(subsumption_resolution,[],[f5643,f1155])).
fof(f5643,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'T_0')) | sP10(X65,sK15) | sP10(X66,X65) | sP11(X67,X66)) )),
  inference(forward_demodulation,[],[f5610,f1254])).
fof(f5610,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,sK15)) | sP10(X65,sK15) | sP10(X66,X65) | sP11(X67,X66)) )),
  inference(resolution,[],[f2074,f4685])).
fof(f5763,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f5750,f1088])).
fof(f5750,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X1,X0) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(resolution,[],[f5644,f1045])).
fof(f5689,plain,(
  ( ! [X0,X1] : (sP10(X0,sK15) | p(s(bool,X1)) | sP10(X0,X1)) ) | $spl84),
  inference(cnf_transformation,[],[f5689_D])).
fof(f5689_D,plain,(
  ( ! [X0,X1] : (sP10(X0,sK15) | p(s(bool,X1)) | sP10(X0,X1)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5727,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X3,X4) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(subsumption_resolution,[],[f5714,f1088])).
fof(f5714,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f5640,f1041])).
fof(f5640,plain,(
  ( ! [X61,X59,X60] : (sP10(X59,sK15) | sP10(X59,X60) | sP11(X60,X61)) )),
  inference(subsumption_resolution,[],[f5639,f1155])).
fof(f5639,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'T_0')) | sP10(X59,sK15) | sP10(X59,X60) | sP11(X60,X61)) )),
  inference(forward_demodulation,[],[f5607,f1254])).
fof(f5607,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,sK15)) | sP10(X59,sK15) | sP10(X59,X60) | sP11(X60,X61)) )),
  inference(resolution,[],[f2074,f3869])).
fof(f5726,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f5713,f1088])).
fof(f5713,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(resolution,[],[f5640,f1045])).
fof(f5690,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X3,X4) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f5676,f1088])).
fof(f5676,plain,(
  ( ! [X4,X5,X3] : (sP10(X3,sK15) | sP10(X3,X4) | p(s(bool,'F_0')) | p(s(bool,X5)) | p(s(bool,X4))) )),
  inference(resolution,[],[f5638,f1041])).
fof(f5638,plain,(
  ( ! [X57,X58,X56] : (sP10(X56,sK15) | sP10(X56,X57) | sP11(X58,X57)) )),
  inference(subsumption_resolution,[],[f5637,f1155])).
fof(f5637,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'T_0')) | sP10(X56,sK15) | sP10(X56,X57) | sP11(X58,X57)) )),
  inference(forward_demodulation,[],[f5606,f1254])).
fof(f5606,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,sK15)) | sP10(X56,sK15) | sP10(X56,X57) | sP11(X58,X57)) )),
  inference(resolution,[],[f2074,f3868])).
fof(f5688,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f5675,f1088])).
fof(f5675,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,sK15) | sP10(X0,X1) | p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(resolution,[],[f5638,f1045])).
fof(f5045,plain,(
  sP10('F_0','T_0') | $spl82),
  inference(cnf_transformation,[],[f5045_D])).
fof(f5045_D,plain,(
  sP10('F_0','T_0') <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5031,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP10('F_0','T_0')) )),
  inference(duplicate_literal_removal,[],[f5006])).
fof(f5006,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP10('F_0','T_0') | ~p(s(bool,X27))) )),
  inference(superposition,[],[f4824,f1122])).
fof(f4824,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP10('F_0',X0)) )),
  inference(resolution,[],[f4720,f1232])).
fof(f4720,plain,(
  ( ! [X386] : (sP10(sK15,X386) | sP10('F_0',X386)) )),
  inference(resolution,[],[f1639,f1138])).
fof(f1138,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f656])).
fof(f656,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb',aHL_FALSITY)).
fof(f5032,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP10('F_0','T_0')) )),
  inference(duplicate_literal_removal,[],[f5005])).
fof(f5005,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP10('F_0','T_0') | ~p(s(bool,X26))) )),
  inference(superposition,[],[f4824,f1124])).
fof(f4288,plain,(
  sP10('F_0','F_0') | $spl68),
  inference(cnf_transformation,[],[f4288_D])).
fof(f4288_D,plain,(
  sP10('F_0','F_0') <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4958,plain,(
  ( ! [X277,X278] : (sP10('F_0','F_0') | p(s(bool,X277)) | p(s(bool,X278)) | ~sP11(X277,X278)) )),
  inference(resolution,[],[f4824,f1045])).
fof(f4957,plain,(
  ( ! [X275,X276] : (sP10('F_0','F_0') | p(s(bool,X275)) | p(s(bool,X276)) | ~sP11(X275,X276)) )),
  inference(resolution,[],[f4824,f1041])).
fof(f4951,plain,(
  ( ! [X274] : (sP10('F_0','T_0') | ~p(s(bool,X274))) )),
  inference(resolution,[],[f4824,f1107])).
fof(f4950,plain,(
  ( ! [X273] : (sP10('F_0','T_0') | ~p(s(bool,X273))) )),
  inference(resolution,[],[f4824,f1103])).
fof(f4949,plain,(
  ( ! [X272] : (sP10('F_0','T_0') | ~p(s(bool,X272))) )),
  inference(resolution,[],[f4824,f1100])).
fof(f4948,plain,(
  ( ! [X271] : (sP10('F_0','T_0') | ~p(s(bool,X271))) )),
  inference(resolution,[],[f4824,f1096])).
fof(f4947,plain,(
  ( ! [X270] : (sP10('F_0','T_0') | ~p(s(bool,X270))) )),
  inference(resolution,[],[f4824,f1085])).
fof(f4946,plain,(
  ( ! [X269] : (sP10('F_0','T_0') | p(s(bool,X269))) )),
  inference(resolution,[],[f4824,f1084])).
fof(f4945,plain,(
  ( ! [X268] : (sP10('F_0','T_0') | p(s(bool,X268))) )),
  inference(resolution,[],[f4824,f1081])).
fof(f4944,plain,(
  ( ! [X267] : (sP10('F_0','T_0') | ~p(s(bool,X267))) )),
  inference(resolution,[],[f4824,f1079])).
fof(f4942,plain,(
  ( ! [X266] : (sP10('F_0','T_0') | p(s(bool,X266))) )),
  inference(resolution,[],[f4824,f1075])).
fof(f5039,plain,(
  ( ! [X24] : (sP10('F_0',X24)) ) | $spl80),
  inference(cnf_transformation,[],[f5039_D])).
fof(f5039_D,plain,(
  ( ! [X24] : (sP10('F_0',X24)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4936,plain,(
  ( ! [X255,X254,X256] : (sP10('F_0',X254) | ~p(s(bool,X255)) | p(s(bool,X256)) | s(bool,X255) != s(bool,X256)) )),
  inference(resolution,[],[f4824,f1151])).
fof(f5037,plain,(
  ( ! [X6,X4] : (sP10('F_0',X4) | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6))) ) | $spl78),
  inference(cnf_transformation,[],[f5037_D])).
fof(f5037_D,plain,(
  ( ! [X6,X4] : (sP10('F_0',X4) | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4935,plain,(
  ( ! [X253,X251,X252] : (sP10('F_0',X251) | p(s(bool,X252)) | ~p(s(bool,X253)) | s(bool,X251) != s(bool,X253)) )),
  inference(resolution,[],[f4824,f1151])).
fof(f4847,plain,(
  ( ! [X30,X28,X31,X29] : (sP10('F_0',X28) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4824,f861])).
fof(f5038,plain,(
  ( ! [X24,X27,X25] : (sP10('F_0',X24) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4846,f1151])).
fof(f4846,plain,(
  ( ! [X26,X24,X27,X25] : (sP10('F_0',X24) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) )),
  inference(resolution,[],[f4824,f860])).
fof(f4845,plain,(
  ( ! [X23,X21,X22,X20] : (sP10('F_0',X20) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4824,f860])).
fof(f5036,plain,(
  ( ! [X0,X3] : (sP10('F_0',X0) | s(bool,X0) != s(bool,X3)) ) | $spl76),
  inference(cnf_transformation,[],[f5036_D])).
fof(f5036_D,plain,(
  ( ! [X0,X3] : (sP10('F_0',X0) | s(bool,X0) != s(bool,X3)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4844,plain,(
  ( ! [X19,X17,X18,X16] : (sP10('F_0',X16) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f4824,f859])).
fof(f4843,plain,(
  ( ! [X14,X12,X15,X13] : (sP10('F_0',X12) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4824,f859])).
fof(f4842,plain,(
  ( ! [X10,X8,X11,X9] : (sP10('F_0',X8) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4824,f857])).
fof(f4841,plain,(
  ( ! [X6,X4,X7,X5] : (sP10('F_0',X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X4) != s(bool,X6)) )),
  inference(resolution,[],[f4824,f855])).
fof(f4840,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('F_0',X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | s(bool,X1) != s(bool,X2)) )),
  inference(resolution,[],[f4824,f855])).
fof(f4812,plain,(
  ( ! [X362] : (sP10(sK15,X362) | sP10('F_0',X362)) ) | $spl74),
  inference(cnf_transformation,[],[f4812_D])).
fof(f4812_D,plain,(
  ( ! [X362] : (sP10(sK15,X362) | sP10('F_0',X362)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4793,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP10(sK15,X49) | sP10('F_0',X49)) )),
  inference(duplicate_literal_removal,[],[f4770])).
fof(f4770,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP10(sK15,X49) | sP10('F_0',X49) | p(s(bool,X48))) )),
  inference(superposition,[],[f1639,f1072])).
fof(f1072,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f816])).
fof(f816,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f815])).
fof(f815,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f699])).
fof(f699,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f359])).
fof(f359,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X15))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_F_EQ_F')).
fof(f4794,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP10(sK15,X47) | sP10('F_0',X47)) )),
  inference(duplicate_literal_removal,[],[f4769])).
fof(f4769,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP10(sK15,X47) | sP10('F_0',X47) | p(s(bool,X46))) )),
  inference(superposition,[],[f1639,f1118])).
fof(f1118,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f845])).
fof(f845,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f635])).
fof(f635,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f4795,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | sP10(sK15,X45) | sP10('F_0',X45)) )),
  inference(duplicate_literal_removal,[],[f4768])).
fof(f4768,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | sP10(sK15,X45) | sP10('F_0',X45) | p(s(bool,X44))) )),
  inference(superposition,[],[f1639,f1120])).
fof(f1120,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f846])).
fof(f846,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f638])).
fof(f638,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f637])).
fof(f637,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f4719,plain,(
  ( ! [X385,X384] : (sP10(sK15,X384) | sP10('F_0',X384) | p(s(bool,X385))) )),
  inference(resolution,[],[f1639,f1134])).
fof(f1134,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f712])).
fof(f712,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f651])).
fof(f651,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X15] : (p(s(bool,'F_0')) => p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.FALSITY')).
fof(f4718,plain,(
  ( ! [X383,X382] : (sP10(sK15,X382) | sP10('F_0',X382) | ~p(s(bool,X383))) )),
  inference(resolution,[],[f1639,f1132])).
fof(f1132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f709])).
fof(f709,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f648])).
fof(f648,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f647])).
fof(f647,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f490])).
fof(f490,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) => ~p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_F')).
fof(f4717,plain,(
  ( ! [X380,X381] : (sP10(sK15,X380) | sP10('F_0',X380) | p(s(bool,X381))) )),
  inference(resolution,[],[f1639,f1116])).
fof(f1116,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f844])).
fof(f844,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f843])).
fof(f843,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X15] : ((p(s(bool,'F_0')) & p(s(bool,X15))) <=> p(s(bool,'F_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f4716,plain,(
  ( ! [X379,X378] : (sP10(sK15,X378) | sP10('F_0',X378) | p(s(bool,X379))) )),
  inference(resolution,[],[f1639,f1112])).
fof(f1112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f842])).
fof(f842,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f841])).
fof(f841,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f633])).
fof(f633,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X15] : ((p(s(bool,X15)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.AND_CLAUSES|split|3')).
fof(f4715,plain,(
  ( ! [X377,X376] : (sP10(sK15,X376) | sP10('F_0',X376) | p(s(bool,X377))) )),
  inference(resolution,[],[f1639,f1093])).
fof(f1093,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f830])).
fof(f830,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f829])).
fof(f829,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f362])).
fof(f362,axiom,(
  ! [X15] : ((p(s(bool,'F_0')) | p(s(bool,X15))) <=> p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.OR_CLAUSES|split|2')).
fof(f4714,plain,(
  ( ! [X374,X375] : (sP10(sK15,X374) | sP10('F_0',X374) | p(s(bool,X375))) )),
  inference(resolution,[],[f1639,f1088])).
fof(f4813,plain,(
  ( ! [X372,X373] : (sP10(sK15,X372) | sP10('F_0',X372) | p(s(bool,X373))) )),
  inference(subsumption_resolution,[],[f4713,f1155])).
fof(f4713,plain,(
  ( ! [X372,X373] : (sP10(sK15,X372) | sP10('F_0',X372) | p(s(bool,X373)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1639,f1080])).
fof(f1080,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f820])).
fof(f1761,plain,(
  ( ! [X265] : (s(bool,'F_0') = s(bool,X265)) ) | $spl36),
  inference(cnf_transformation,[],[f1761_D])).
fof(f1761_D,plain,(
  ( ! [X265] : (s(bool,'F_0') = s(bool,X265)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4712,plain,(
  ( ! [X370,X371] : (sP10(sK15,X370) | sP10('F_0',X370) | s(bool,'F_0') = s(bool,X371)) )),
  inference(resolution,[],[f1639,f1073])).
fof(f1073,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f816])).
fof(f4711,plain,(
  ( ! [X368,X369] : (sP10(sK15,X368) | sP10('F_0',X368) | ~p(s(bool,X369))) )),
  inference(resolution,[],[f1639,f1070])).
fof(f1070,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f814])).
fof(f814,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f813])).
fof(f813,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f698])).
fof(f698,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f619])).
fof(f619,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f618])).
fof(f618,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X15] : ((p(s(bool,X15)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X15)))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.IMP_CLAUSES|split|4')).
fof(f1758,plain,(
  ( ! [X261,X260] : (sP11(X260,X261)) ) | $spl32),
  inference(cnf_transformation,[],[f1758_D])).
fof(f1758_D,plain,(
  ( ! [X261,X260] : (sP11(X260,X261)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4710,plain,(
  ( ! [X366,X365,X367] : (sP10(sK15,X365) | sP10('F_0',X365) | sP11(X366,X367)) )),
  inference(resolution,[],[f1639,f1049])).
fof(f1049,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f803])).
fof(f4709,plain,(
  ( ! [X364,X362,X363] : (sP10(sK15,X362) | sP10('F_0',X362) | sP11(X363,X364)) )),
  inference(resolution,[],[f1639,f1044])).
fof(f1044,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f801])).
fof(f4807,plain,(
  ( ! [X0,X3,X1] : (sP10(sK15,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) | $spl70),
  inference(cnf_transformation,[],[f4807_D])).
fof(f4807_D,plain,(
  ( ! [X0,X3,X1] : (sP10(sK15,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X1,X0)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4702,plain,(
  ( ! [X344,X346,X345,X347] : (sP10(sK15,X344) | sP10(X345,X344) | p(s(bool,X346)) | p(s(bool,X347)) | s(bool,X345) != s(bool,X347)) )),
  inference(resolution,[],[f1639,f1151])).
fof(f4615,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(sK15,X35) | sP10(X36,X35) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1639,f861])).
fof(f4808,plain,(
  ( ! [X6,X5,X9] : (sP10(sK15,X5) | s(bool,X6) != s(bool,X9) | sP10(X6,X5)) ) | $spl72),
  inference(cnf_transformation,[],[f4808_D])).
fof(f4808_D,plain,(
  ( ! [X6,X5,X9] : (sP10(sK15,X5) | s(bool,X6) != s(bool,X9) | sP10(X6,X5)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4614,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP10(sK15,X30) | sP10(X31,X30) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f1639,f861])).
fof(f4612,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(sK15,X20) | sP10(X21,X20) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1639,f860])).
fof(f4611,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(sK15,X15) | sP10(X16,X15) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1639,f859])).
fof(f4610,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(sK15,X10) | sP10(X11,X10) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1639,f857])).
fof(f4609,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(sK15,X5) | sP10(X6,X5) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) )),
  inference(resolution,[],[f1639,f857])).
fof(f4608,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK15,X0) | sP10(X1,X0) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1639,f855])).
fof(f4282,plain,(
  sP10('T_0','F_0') | $spl64),
  inference(cnf_transformation,[],[f4282_D])).
fof(f4282_D,plain,(
  sP10('T_0','F_0') <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4268,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP10('T_0','F_0')) )),
  inference(duplicate_literal_removal,[],[f4243])).
fof(f4243,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP10('T_0','F_0') | ~p(s(bool,X27))) )),
  inference(superposition,[],[f4004,f1122])).
fof(f4004,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP10(X0,'F_0')) )),
  inference(resolution,[],[f3902,f1232])).
fof(f3902,plain,(
  ( ! [X384] : (sP10(sK15,X384) | sP10(X384,'F_0')) )),
  inference(resolution,[],[f1638,f1138])).
fof(f4269,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP10('T_0','F_0')) )),
  inference(duplicate_literal_removal,[],[f4242])).
fof(f4242,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP10('T_0','F_0') | ~p(s(bool,X26))) )),
  inference(superposition,[],[f4004,f1124])).
fof(f4195,plain,(
  ( ! [X275,X276] : (sP10('F_0','F_0') | p(s(bool,X275)) | p(s(bool,X276)) | ~sP11(X275,X276)) )),
  inference(resolution,[],[f4004,f1045])).
fof(f4194,plain,(
  ( ! [X273,X274] : (sP10('F_0','F_0') | p(s(bool,X273)) | p(s(bool,X274)) | ~sP11(X273,X274)) )),
  inference(resolution,[],[f4004,f1041])).
fof(f4188,plain,(
  ( ! [X272] : (sP10('T_0','F_0') | ~p(s(bool,X272))) )),
  inference(resolution,[],[f4004,f1107])).
fof(f4187,plain,(
  ( ! [X271] : (sP10('T_0','F_0') | ~p(s(bool,X271))) )),
  inference(resolution,[],[f4004,f1103])).
fof(f4186,plain,(
  ( ! [X270] : (sP10('T_0','F_0') | ~p(s(bool,X270))) )),
  inference(resolution,[],[f4004,f1100])).
fof(f4185,plain,(
  ( ! [X269] : (sP10('T_0','F_0') | ~p(s(bool,X269))) )),
  inference(resolution,[],[f4004,f1096])).
fof(f4184,plain,(
  ( ! [X268] : (sP10('T_0','F_0') | ~p(s(bool,X268))) )),
  inference(resolution,[],[f4004,f1085])).
fof(f4183,plain,(
  ( ! [X267] : (sP10('T_0','F_0') | p(s(bool,X267))) )),
  inference(resolution,[],[f4004,f1084])).
fof(f4182,plain,(
  ( ! [X266] : (sP10('T_0','F_0') | p(s(bool,X266))) )),
  inference(resolution,[],[f4004,f1081])).
fof(f4181,plain,(
  ( ! [X265] : (sP10('T_0','F_0') | ~p(s(bool,X265))) )),
  inference(resolution,[],[f4004,f1079])).
fof(f4179,plain,(
  ( ! [X264] : (sP10('T_0','F_0') | p(s(bool,X264))) )),
  inference(resolution,[],[f4004,f1075])).
fof(f4276,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) | $spl62),
  inference(cnf_transformation,[],[f4276_D])).
fof(f4276_D,plain,(
  ( ! [X24] : (sP10(X24,'F_0')) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4174,plain,(
  ( ! [X255,X254,X256] : (sP10(X254,'F_0') | ~p(s(bool,X255)) | p(s(bool,X256)) | s(bool,X255) != s(bool,X256)) )),
  inference(resolution,[],[f4004,f1151])).
fof(f4274,plain,(
  ( ! [X6,X4] : (sP10(X4,'F_0') | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6))) ) | $spl60),
  inference(cnf_transformation,[],[f4274_D])).
fof(f4274_D,plain,(
  ( ! [X6,X4] : (sP10(X4,'F_0') | s(bool,X4) != s(bool,X6) | ~p(s(bool,X6))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4173,plain,(
  ( ! [X253,X251,X252] : (sP10(X251,'F_0') | p(s(bool,X252)) | ~p(s(bool,X253)) | s(bool,X251) != s(bool,X253)) )),
  inference(resolution,[],[f4004,f1151])).
fof(f4085,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,'F_0') | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4004,f861])).
fof(f4275,plain,(
  ( ! [X24,X27,X25] : (sP10(X24,'F_0') | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4084,f1151])).
fof(f4084,plain,(
  ( ! [X26,X24,X27,X25] : (sP10(X24,'F_0') | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) )),
  inference(resolution,[],[f4004,f860])).
fof(f4083,plain,(
  ( ! [X23,X21,X22,X20] : (sP10(X20,'F_0') | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4004,f860])).
fof(f4273,plain,(
  ( ! [X0,X3] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X3)) ) | $spl58),
  inference(cnf_transformation,[],[f4273_D])).
fof(f4273_D,plain,(
  ( ! [X0,X3] : (sP10(X0,'F_0') | s(bool,X0) != s(bool,X3)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4082,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,'F_0') | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) )),
  inference(resolution,[],[f4004,f859])).
fof(f4081,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,'F_0') | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4004,f859])).
fof(f4080,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,'F_0') | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4004,f857])).
fof(f4079,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,'F_0') | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X4) != s(bool,X6)) )),
  inference(resolution,[],[f4004,f855])).
fof(f4078,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,'F_0') | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | s(bool,X1) != s(bool,X2)) )),
  inference(resolution,[],[f4004,f855])).
fof(f4056,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4055,f1144])).
fof(f4055,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4037,f1254])).
fof(f4037,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(duplicate_literal_removal,[],[f4036])).
fof(f4036,plain,(
  ~p(s(bool,sK15)) | ~p(s(bool,sK15)) | $spl54),
  inference(resolution,[],[f4021,f1026])).
fof(f4052,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4051,f1144])).
fof(f4051,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4040,f1254])).
fof(f4040,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(duplicate_literal_removal,[],[f4033])).
fof(f4033,plain,(
  ~p(s(bool,sK15)) | ~p(s(bool,sK15)) | $spl54),
  inference(resolution,[],[f4021,f1033])).
fof(f1033,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f797])).
fof(f797,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f728])).
fof(f728,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f596,f727])).
fof(f596,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f595])).
fof(f595,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f453])).
fof(f453,axiom,(
  ! [X201,X200] : ((p(s(bool,X200)) <=> ~p(s(bool,X201))) <=> ((p(s(bool,X200)) | p(s(bool,X201))) & (~p(s(bool,X201)) | ~p(s(bool,X200)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.sat.sat_29_prove_ltb')).
fof(f4050,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4049,f1144])).
fof(f4049,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4048,f1254])).
fof(f4048,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(subsumption_resolution,[],[f4032,f1232])).
fof(f4032,plain,(
  ~p(s(bool,sK15)) | sP10(sK15,sK15) | $spl54),
  inference(resolution,[],[f4021,f1636])).
fof(f4047,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4046,f1144])).
fof(f4046,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4045,f1254])).
fof(f4045,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(subsumption_resolution,[],[f4031,f1232])).
fof(f4031,plain,(
  ~p(s(bool,sK15)) | sP10(sK15,sK15) | $spl54),
  inference(resolution,[],[f4021,f1637])).
fof(f4044,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4043,f1144])).
fof(f4043,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4030,f1254])).
fof(f4030,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(resolution,[],[f4021,f1233])).
fof(f4042,plain,(
  $false | $spl54),
  inference(subsumption_resolution,[],[f4041,f1144])).
fof(f4041,plain,(
  ~p(s(bool,'T_0')) | $spl54),
  inference(forward_demodulation,[],[f4029,f1254])).
fof(f4029,plain,(
  ~p(s(bool,sK15)) | $spl54),
  inference(resolution,[],[f4021,f1232])).
fof(f4021,plain,(
  sP10(sK15,sK15) | $spl54),
  inference(cnf_transformation,[],[f4021_D])).
fof(f4021_D,plain,(
  sP10(sK15,sK15) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4022,plain,(
  ( ! [X2] : (sP10(X2,'F_0') | ~p(s(bool,X2))) ) | $spl56),
  inference(cnf_transformation,[],[f4022_D])).
fof(f4022_D,plain,(
  ( ! [X2] : (sP10(X2,'F_0') | ~p(s(bool,X2))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4006,plain,(
  ( ! [X2] : (sP10(X2,'F_0') | ~p(s(bool,X2)) | sP10(sK15,sK15)) )),
  inference(resolution,[],[f3902,f1636])).
fof(f3993,plain,(
  ( ! [X360] : (sP10(sK15,X360) | sP10(X360,'F_0')) ) | $spl52),
  inference(cnf_transformation,[],[f3993_D])).
fof(f3993_D,plain,(
  ( ! [X360] : (sP10(sK15,X360) | sP10(X360,'F_0')) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3974,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP10(sK15,X49) | sP10(X49,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3952])).
fof(f3952,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | sP10(sK15,X49) | sP10(X49,'F_0') | p(s(bool,X48))) )),
  inference(superposition,[],[f1638,f1072])).
fof(f3975,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP10(sK15,X47) | sP10(X47,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3951])).
fof(f3951,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | sP10(sK15,X47) | sP10(X47,'F_0') | p(s(bool,X46))) )),
  inference(superposition,[],[f1638,f1118])).
fof(f3976,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | sP10(sK15,X45) | sP10(X45,'F_0')) )),
  inference(duplicate_literal_removal,[],[f3950])).
fof(f3950,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | sP10(sK15,X45) | sP10(X45,'F_0') | p(s(bool,X44))) )),
  inference(superposition,[],[f1638,f1120])).
fof(f3901,plain,(
  ( ! [X383,X382] : (sP10(sK15,X382) | sP10(X382,'F_0') | p(s(bool,X383))) )),
  inference(resolution,[],[f1638,f1134])).
fof(f3900,plain,(
  ( ! [X380,X381] : (sP10(sK15,X380) | sP10(X380,'F_0') | ~p(s(bool,X381))) )),
  inference(resolution,[],[f1638,f1132])).
fof(f3899,plain,(
  ( ! [X379,X378] : (sP10(sK15,X378) | sP10(X378,'F_0') | p(s(bool,X379))) )),
  inference(resolution,[],[f1638,f1116])).
fof(f3898,plain,(
  ( ! [X377,X376] : (sP10(sK15,X376) | sP10(X376,'F_0') | p(s(bool,X377))) )),
  inference(resolution,[],[f1638,f1112])).
fof(f3897,plain,(
  ( ! [X374,X375] : (sP10(sK15,X374) | sP10(X374,'F_0') | p(s(bool,X375))) )),
  inference(resolution,[],[f1638,f1093])).
fof(f3896,plain,(
  ( ! [X372,X373] : (sP10(sK15,X372) | sP10(X372,'F_0') | p(s(bool,X373))) )),
  inference(resolution,[],[f1638,f1088])).
fof(f3994,plain,(
  ( ! [X370,X371] : (sP10(sK15,X370) | sP10(X370,'F_0') | p(s(bool,X371))) )),
  inference(subsumption_resolution,[],[f3895,f1155])).
fof(f3895,plain,(
  ( ! [X370,X371] : (sP10(sK15,X370) | sP10(X370,'F_0') | p(s(bool,X371)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1638,f1080])).
fof(f3894,plain,(
  ( ! [X368,X369] : (sP10(sK15,X368) | sP10(X368,'F_0') | s(bool,'F_0') = s(bool,X369)) )),
  inference(resolution,[],[f1638,f1073])).
fof(f3893,plain,(
  ( ! [X366,X367] : (sP10(sK15,X366) | sP10(X366,'F_0') | ~p(s(bool,X367))) )),
  inference(resolution,[],[f1638,f1070])).
fof(f3892,plain,(
  ( ! [X364,X365,X363] : (sP10(sK15,X363) | sP10(X363,'F_0') | sP11(X364,X365)) )),
  inference(resolution,[],[f1638,f1049])).
fof(f3891,plain,(
  ( ! [X360,X362,X361] : (sP10(sK15,X360) | sP10(X360,'F_0') | sP11(X361,X362)) )),
  inference(resolution,[],[f1638,f1044])).
fof(f3988,plain,(
  ( ! [X0,X3,X1] : (sP10(sK15,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) | $spl48),
  inference(cnf_transformation,[],[f3988_D])).
fof(f3988_D,plain,(
  ( ! [X0,X3,X1] : (sP10(sK15,X0) | s(bool,X1) != s(bool,X3) | p(s(bool,X3)) | sP10(X0,X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3885,plain,(
  ( ! [X344,X346,X345,X347] : (sP10(sK15,X344) | sP10(X344,X345) | p(s(bool,X346)) | p(s(bool,X347)) | s(bool,X345) != s(bool,X347)) )),
  inference(resolution,[],[f1638,f1151])).
fof(f3798,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP10(sK15,X35) | sP10(X35,X36) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,X39)) )),
  inference(resolution,[],[f1638,f861])).
fof(f3989,plain,(
  ( ! [X6,X5,X9] : (sP10(sK15,X5) | s(bool,X6) != s(bool,X9) | sP10(X5,X6)) ) | $spl50),
  inference(cnf_transformation,[],[f3989_D])).
fof(f3989_D,plain,(
  ( ! [X6,X5,X9] : (sP10(sK15,X5) | s(bool,X6) != s(bool,X9) | sP10(X5,X6)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f3797,plain,(
  ( ! [X30,X33,X31,X34,X32] : (sP10(sK15,X30) | sP10(X30,X31) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X33) | s(bool,X31) != s(bool,X34)) )),
  inference(resolution,[],[f1638,f861])).
fof(f3795,plain,(
  ( ! [X24,X23,X21,X22,X20] : (sP10(sK15,X20) | sP10(X20,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X21) != s(bool,X24)) )),
  inference(resolution,[],[f1638,f860])).
fof(f3794,plain,(
  ( ! [X19,X17,X15,X18,X16] : (sP10(sK15,X15) | sP10(X15,X16) | p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X16) != s(bool,X18)) )),
  inference(resolution,[],[f1638,f859])).
fof(f3793,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP10(sK15,X10) | sP10(X10,X11) | p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f1638,f857])).
fof(f3792,plain,(
  ( ! [X6,X8,X7,X5,X9] : (sP10(sK15,X5) | sP10(X5,X6) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X9)) )),
  inference(resolution,[],[f1638,f857])).
fof(f3791,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(sK15,X0) | sP10(X0,X1) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X1) != s(bool,X3)) )),
  inference(resolution,[],[f1638,f855])).
fof(f2517,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) | $spl44),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3618,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f3549])).
fof(f3549,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X38))) ) | $spl12),
  inference(superposition,[],[f1280,f1124])).
fof(f3619,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f3548])).
fof(f3548,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f1280,f1122])).
fof(f3441,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f3419])).
fof(f3419,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl46),
  inference(superposition,[],[f3130,f919])).
fof(f919,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f377])).
fof(f377,axiom,(
  ! [X19] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.numeral.numeral_distrib|split|20')).
fof(f3440,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f3420])).
fof(f3420,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl46),
  inference(superposition,[],[f3130,f920])).
fof(f920,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f430])).
fof(f430,axiom,(
  ! [X19] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.numeral.numeral_lt|split|2')).
fof(f3439,plain,(
  $false | $spl46),
  inference(equality_resolution,[],[f3130])).
fof(f3381,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3297,f3124])).
fof(f3124,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl6),
  inference(equality_resolution,[],[f1277])).
fof(f3297,plain,(
  ( ! [X316,X317] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X316),s('type.num.num',X317))))) ) | $spl6),
  inference(resolution,[],[f3124,f1054])).
fof(f1054,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f427])).
fof(f427,axiom,(
  ! [X19,X45] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X19)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.numeral.numeral_lte|split|5')).
fof(f3380,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3296,f3124])).
fof(f3296,plain,(
  ( ! [X315,X314] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X314),s('type.num.num',X315))))) ) | $spl6),
  inference(resolution,[],[f3124,f1021])).
fof(f1021,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f791])).
fof(f791,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f592])).
fof(f592,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(flattening,[],[f591])).
fof(f591,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(rectify,[],[f429])).
fof(f429,axiom,(
  ! [X45,X19] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X19)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X45)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.arithmetic.NOT_LEQ')).
fof(f3377,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3289,f3124])).
fof(f3289,plain,(
  ( ! [X300,X299] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X299))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X300))))))) ) | $spl6),
  inference(resolution,[],[f3124,f1054])).
fof(f3376,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3288,f3124])).
fof(f3288,plain,(
  ( ! [X298,X297] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X297))),s('type.num.num',X298))))) ) | $spl6),
  inference(resolution,[],[f3124,f1021])).
fof(f3375,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3285,f3124])).
fof(f3285,plain,(
  ( ! [X290,X289] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X289),s('type.num.num',X290))))) ) | $spl6),
  inference(resolution,[],[f3124,f898])).
fof(f898,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f750])).
fof(f3374,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3281,f3124])).
fof(f3281,plain,(
  ( ! [X279,X278] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X278),s('type.num.num',X279))))) ) | $spl6),
  inference(resolution,[],[f3124,f899])).
fof(f899,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f751])).
fof(f751,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f534])).
fof(f534,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f392])).
fof(f392,axiom,(
  ! [X19,X45] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X45)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X45),s('type.num.num',X19)))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.numeral.numeral_lt|split|5')).
fof(f3373,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3275,f3124])).
fof(f3275,plain,(
  ( ! [X269,X268] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X268),s('type.num.num',X269))))) ) | $spl6),
  inference(resolution,[],[f3124,f898])).
fof(f3372,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3274,f3124])).
fof(f3274,plain,(
  ( ! [X267,X266] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X266))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X267))))))) ) | $spl6),
  inference(resolution,[],[f3124,f899])).
fof(f3371,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3272,f3124])).
fof(f3272,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(resolution,[],[f3124,f1142])).
fof(f1142,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f850])).
fof(f850,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f658])).
fof(f658,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f372])).
fof(f372,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f3370,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3271,f3124])).
fof(f3271,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(resolution,[],[f3124,f1140])).
fof(f1140,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f849])).
fof(f849,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f657])).
fof(f657,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f369])).
fof(f369,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f3369,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3268,f3124])).
fof(f3268,plain,(
  ( ! [X262] : (~p(s(bool,X262))) ) | $spl6),
  inference(resolution,[],[f3124,f1132])).
fof(f3368,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3261,f3124])).
fof(f3261,plain,(
  ( ! [X255] : (~p(s(bool,X255))) ) | $spl6),
  inference(resolution,[],[f3124,f1070])).
fof(f3270,plain,(
  $false | $spl6),
  inference(resolution,[],[f3124,f1138])).
fof(f3278,plain,(
  $false | $spl6),
  inference(resolution,[],[f3124,f910])).
fof(f910,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f543])).
fof(f543,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f494])).
fof(f494,axiom,(
  ! [X19] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.|48|_0'))))),
  file('thm.bvlSem.bvlSem_|49||49|7_prove_ltb','thm.prim_rec.NOT_LESS_|48|')).
fof(f3279,plain,(
  $false | $spl6),
  inference(resolution,[],[f3124,f852])).
fof(f3130,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl46),
  inference(cnf_transformation,[],[f3130_D])).
fof(f3130_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3125,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3055])).
fof(f3055,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl6),
  inference(superposition,[],[f1277,f1120])).
fof(f3126,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3054])).
fof(f3054,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl6),
  inference(superposition,[],[f1277,f1118])).
fof(f3127,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3052])).
fof(f3052,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl6),
  inference(superposition,[],[f1277,f1072])).
fof(f2686,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f2683])).
fof(f2683,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(superposition,[],[f1278,f1254])).
fof(f2685,plain,(
  $false | $spl8),
  inference(equality_resolution,[],[f1278])).
fof(f2539,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) )),
  inference(forward_demodulation,[],[f2538,f1254])).
fof(f2538,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,sK15) != s(bool,X98)) )),
  inference(subsumption_resolution,[],[f2288,f1151])).
fof(f2288,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,sK15) != s(bool,X98)) )),
  inference(superposition,[],[f861,f1254])).
fof(f2535,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2284,f1121])).
fof(f1121,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f847])).
fof(f2284,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(superposition,[],[f861,f1254])).
fof(f2534,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2283,f1151])).
fof(f2283,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,sK15)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(superposition,[],[f861,f1254])).
fof(f2522,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) )),
  inference(forward_demodulation,[],[f2521,f1254])).
fof(f2521,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,sK15) != s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2268,f1151])).
fof(f2268,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,sK15) != s(bool,X38)) )),
  inference(superposition,[],[f857,f1254])).
fof(f2518,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(subsumption_resolution,[],[f2264,f1151])).
fof(f2264,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,sK15)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(superposition,[],[f857,f1254])).
fof(f2516,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(subsumption_resolution,[],[f2263,f1121])).
fof(f2263,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(superposition,[],[f857,f1254])).
fof(f2196,plain,(
  sP10('F_0',sK15) | $spl42),
  inference(cnf_transformation,[],[f2196_D])).
fof(f2196_D,plain,(
  sP10('F_0',sK15) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2185,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP10('F_0',sK15)) )),
  inference(duplicate_literal_removal,[],[f2164])).
fof(f2164,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP10('F_0',sK15) | p(s(bool,X30))) )),
  inference(superposition,[],[f1235,f1072])).
fof(f2186,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP10('F_0',sK15)) )),
  inference(duplicate_literal_removal,[],[f2163])).
fof(f2163,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP10('F_0',sK15) | p(s(bool,X29))) )),
  inference(superposition,[],[f1235,f1118])).
fof(f2187,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP10('F_0',sK15)) )),
  inference(duplicate_literal_removal,[],[f2162])).
fof(f2162,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP10('F_0',sK15) | p(s(bool,X28))) )),
  inference(superposition,[],[f1235,f1120])).
fof(f2114,plain,(
  ( ! [X272] : (sP10('F_0',sK15) | p(s(bool,X272))) )),
  inference(resolution,[],[f1235,f1134])).
fof(f2113,plain,(
  ( ! [X271] : (sP10('F_0',sK15) | ~p(s(bool,X271))) )),
  inference(resolution,[],[f1235,f1132])).
fof(f2112,plain,(
  ( ! [X270] : (sP10('F_0',sK15) | p(s(bool,X270))) )),
  inference(resolution,[],[f1235,f1116])).
fof(f2111,plain,(
  ( ! [X269] : (sP10('F_0',sK15) | p(s(bool,X269))) )),
  inference(resolution,[],[f1235,f1112])).
fof(f2110,plain,(
  ( ! [X268] : (sP10('F_0',sK15) | p(s(bool,X268))) )),
  inference(resolution,[],[f1235,f1093])).
fof(f2109,plain,(
  ( ! [X267] : (sP10('F_0',sK15) | p(s(bool,X267))) )),
  inference(resolution,[],[f1235,f1088])).
fof(f2197,plain,(
  ( ! [X266] : (sP10('F_0',sK15) | p(s(bool,X266))) )),
  inference(subsumption_resolution,[],[f2108,f1155])).
fof(f2108,plain,(
  ( ! [X266] : (sP10('F_0',sK15) | p(s(bool,X266)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1235,f1080])).
fof(f2107,plain,(
  ( ! [X265] : (sP10('F_0',sK15) | s(bool,'F_0') = s(bool,X265)) )),
  inference(resolution,[],[f1235,f1073])).
fof(f2106,plain,(
  ( ! [X264] : (sP10('F_0',sK15) | ~p(s(bool,X264))) )),
  inference(resolution,[],[f1235,f1070])).
fof(f2105,plain,(
  ( ! [X263,X262] : (sP10('F_0',sK15) | sP11(X262,X263)) )),
  inference(resolution,[],[f1235,f1049])).
fof(f2104,plain,(
  ( ! [X261,X260] : (sP10('F_0',sK15) | sP11(X260,X261)) )),
  inference(resolution,[],[f1235,f1044])).
fof(f2190,plain,(
  ( ! [X2,X0] : (sP10(X0,sK15) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl38),
  inference(cnf_transformation,[],[f2190_D])).
fof(f2190_D,plain,(
  ( ! [X2,X0] : (sP10(X0,sK15) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2098,plain,(
  ( ! [X251,X252,X250] : (sP10(X250,sK15) | p(s(bool,X251)) | p(s(bool,X252)) | s(bool,X250) != s(bool,X252)) )),
  inference(resolution,[],[f1235,f1151])).
fof(f2011,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(X28,sK15) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1235,f861])).
fof(f2191,plain,(
  ( ! [X4,X7] : (sP10(X4,sK15) | s(bool,X4) != s(bool,X7)) ) | $spl40),
  inference(cnf_transformation,[],[f2191_D])).
fof(f2191_D,plain,(
  ( ! [X4,X7] : (sP10(X4,sK15) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2010,plain,(
  ( ! [X26,X24,X27,X25] : (sP10(X24,sK15) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) )),
  inference(resolution,[],[f1235,f861])).
fof(f2008,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(X16,sK15) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1235,f860])).
fof(f2007,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(X12,sK15) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1235,f859])).
fof(f2006,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(X8,sK15) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1235,f857])).
fof(f2005,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(X4,sK15) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) )),
  inference(resolution,[],[f1235,f857])).
fof(f2004,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,sK15) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1235,f855])).
fof(f1993,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1910,f1282])).
fof(f1910,plain,(
  ( ! [X316,X317] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X316),s('type.num.num',X317))))) ) | $spl16),
  inference(resolution,[],[f1282,f1054])).
fof(f1992,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1909,f1282])).
fof(f1909,plain,(
  ( ! [X315,X314] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X314),s('type.num.num',X315))))) ) | $spl16),
  inference(resolution,[],[f1282,f1021])).
fof(f1989,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1902,f1282])).
fof(f1902,plain,(
  ( ! [X300,X299] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X299))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X300))))))) ) | $spl16),
  inference(resolution,[],[f1282,f1054])).
fof(f1988,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1901,f1282])).
fof(f1901,plain,(
  ( ! [X298,X297] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X297))),s('type.num.num',X298))))) ) | $spl16),
  inference(resolution,[],[f1282,f1021])).
fof(f1987,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1898,f1282])).
fof(f1898,plain,(
  ( ! [X290,X289] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X289),s('type.num.num',X290))))) ) | $spl16),
  inference(resolution,[],[f1282,f898])).
fof(f1986,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1894,f1282])).
fof(f1894,plain,(
  ( ! [X279,X278] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X278),s('type.num.num',X279))))) ) | $spl16),
  inference(resolution,[],[f1282,f899])).
fof(f1985,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1888,f1282])).
fof(f1888,plain,(
  ( ! [X269,X268] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X268),s('type.num.num',X269))))) ) | $spl16),
  inference(resolution,[],[f1282,f898])).
fof(f1984,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1887,f1282])).
fof(f1887,plain,(
  ( ! [X267,X266] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X266))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X267))))))) ) | $spl16),
  inference(resolution,[],[f1282,f899])).
fof(f1983,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1885,f1282])).
fof(f1885,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1282,f1142])).
fof(f1982,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1884,f1282])).
fof(f1884,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1282,f1140])).
fof(f1981,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1881,f1282])).
fof(f1881,plain,(
  ( ! [X262] : (~p(s(bool,X262))) ) | $spl16),
  inference(resolution,[],[f1282,f1132])).
fof(f1980,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f1874,f1282])).
fof(f1874,plain,(
  ( ! [X255] : (~p(s(bool,X255))) ) | $spl16),
  inference(resolution,[],[f1282,f1070])).
fof(f1883,plain,(
  $false | $spl16),
  inference(resolution,[],[f1282,f1138])).
fof(f1891,plain,(
  $false | $spl16),
  inference(resolution,[],[f1282,f910])).
fof(f1892,plain,(
  $false | $spl16),
  inference(resolution,[],[f1282,f852])).
fof(f1760,plain,(
  sP10(sK15,'F_0') | $spl34),
  inference(cnf_transformation,[],[f1760_D])).
fof(f1760_D,plain,(
  sP10(sK15,'F_0') <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1748,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP10(sK15,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1727])).
fof(f1727,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP10(sK15,'F_0') | p(s(bool,X30))) )),
  inference(superposition,[],[f1234,f1072])).
fof(f1749,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP10(sK15,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1726])).
fof(f1726,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP10(sK15,'F_0') | p(s(bool,X29))) )),
  inference(superposition,[],[f1234,f1118])).
fof(f1750,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP10(sK15,'F_0')) )),
  inference(duplicate_literal_removal,[],[f1725])).
fof(f1725,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP10(sK15,'F_0') | p(s(bool,X28))) )),
  inference(superposition,[],[f1234,f1120])).
fof(f1677,plain,(
  ( ! [X272] : (sP10(sK15,'F_0') | p(s(bool,X272))) )),
  inference(resolution,[],[f1234,f1134])).
fof(f1676,plain,(
  ( ! [X271] : (sP10(sK15,'F_0') | ~p(s(bool,X271))) )),
  inference(resolution,[],[f1234,f1132])).
fof(f1675,plain,(
  ( ! [X270] : (sP10(sK15,'F_0') | p(s(bool,X270))) )),
  inference(resolution,[],[f1234,f1116])).
fof(f1674,plain,(
  ( ! [X269] : (sP10(sK15,'F_0') | p(s(bool,X269))) )),
  inference(resolution,[],[f1234,f1112])).
fof(f1673,plain,(
  ( ! [X268] : (sP10(sK15,'F_0') | p(s(bool,X268))) )),
  inference(resolution,[],[f1234,f1093])).
fof(f1672,plain,(
  ( ! [X267] : (sP10(sK15,'F_0') | p(s(bool,X267))) )),
  inference(resolution,[],[f1234,f1088])).
fof(f1762,plain,(
  ( ! [X266] : (sP10(sK15,'F_0') | p(s(bool,X266))) )),
  inference(subsumption_resolution,[],[f1671,f1155])).
fof(f1671,plain,(
  ( ! [X266] : (sP10(sK15,'F_0') | p(s(bool,X266)) | ~p(s(bool,'T_0'))) )),
  inference(resolution,[],[f1234,f1080])).
fof(f1670,plain,(
  ( ! [X265] : (sP10(sK15,'F_0') | s(bool,'F_0') = s(bool,X265)) )),
  inference(resolution,[],[f1234,f1073])).
fof(f1669,plain,(
  ( ! [X264] : (sP10(sK15,'F_0') | ~p(s(bool,X264))) )),
  inference(resolution,[],[f1234,f1070])).
fof(f1668,plain,(
  ( ! [X263,X262] : (sP10(sK15,'F_0') | sP11(X262,X263)) )),
  inference(resolution,[],[f1234,f1049])).
fof(f1667,plain,(
  ( ! [X261,X260] : (sP10(sK15,'F_0') | sP11(X260,X261)) )),
  inference(resolution,[],[f1234,f1044])).
fof(f1753,plain,(
  ( ! [X2,X0] : (sP10(sK15,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f1753_D])).
fof(f1753_D,plain,(
  ( ! [X2,X0] : (sP10(sK15,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1661,plain,(
  ( ! [X251,X252,X250] : (sP10(sK15,X250) | p(s(bool,X251)) | p(s(bool,X252)) | s(bool,X250) != s(bool,X252)) )),
  inference(resolution,[],[f1234,f1151])).
fof(f1574,plain,(
  ( ! [X30,X28,X31,X29] : (sP10(sK15,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1234,f861])).
fof(f1754,plain,(
  ( ! [X4,X7] : (sP10(sK15,X4) | s(bool,X4) != s(bool,X7)) ) | $spl30),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X4,X7] : (sP10(sK15,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1573,plain,(
  ( ! [X26,X24,X27,X25] : (sP10(sK15,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) )),
  inference(resolution,[],[f1234,f861])).
fof(f1571,plain,(
  ( ! [X19,X17,X18,X16] : (sP10(sK15,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1234,f860])).
fof(f1570,plain,(
  ( ! [X14,X12,X15,X13] : (sP10(sK15,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1234,f859])).
fof(f1569,plain,(
  ( ! [X10,X8,X11,X9] : (sP10(sK15,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1234,f857])).
fof(f1568,plain,(
  ( ! [X6,X4,X7,X5] : (sP10(sK15,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) )),
  inference(resolution,[],[f1234,f857])).
fof(f1567,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(sK15,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1234,f855])).
fof(f1315,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))))))))) ) | $spl18),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))))))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1495,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | $spl27),
  inference(cnf_transformation,[],[f1495_D])).
fof(f1495_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1451,plain,(
  ( ! [X8] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))))))))) )),
  inference(superposition,[],[f852,f901])).
fof(f1483,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl24),
  inference(cnf_transformation,[],[f1483_D])).
fof(f1483_D,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1431,plain,(
  ( ! [X158,X159] : (p(s(bool,X158)) | ~p(s(bool,X159)) | s(bool,X159) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f852,f1151])).
fof(f1343,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X23) != s(bool,X25)) )),
  inference(resolution,[],[f852,f861])).
fof(f1341,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f852,f860])).
fof(f1482,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) | $spl22),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1340,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X16) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X14) != s(bool,X15)) )),
  inference(resolution,[],[f852,f859])).
fof(f1339,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f852,f859])).
fof(f1338,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X8) != s(bool,X10)) )),
  inference(resolution,[],[f852,f857])).
fof(f1337,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14)))))) )),
  inference(resolution,[],[f852,f855])).
fof(f1336,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))))) | s(bool,X2) != s(bool,X3)) )),
  inference(resolution,[],[f852,f855])).
fof(f1317,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) | $spl21),
  inference(cnf_transformation,[],[f1317_D])).
fof(f1317_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1309,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))))))))) ) | $spl1),
  inference(superposition,[],[f1160,f901])).
fof(f1276,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK15) != s(bool,X1)) ) | $spl4),
  inference(cnf_transformation,[],[f1276_D])).
fof(f1276_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK15) != s(bool,X1)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1257,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,sK15) != s(bool,X157)) )),
  inference(resolution,[],[f851,f1151])).
fof(f1170,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK15) != s(bool,X21) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f851,f861])).
fof(f1278,plain,(
  ( ! [X5] : (s(bool,sK15) != s(bool,X5)) ) | $spl8),
  inference(cnf_transformation,[],[f1278_D])).
fof(f1278_D,plain,(
  ( ! [X5] : (s(bool,sK15) != s(bool,X5)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1169,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK15) != s(bool,X20)) )),
  inference(resolution,[],[f851,f861])).
fof(f1167,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK15) != s(bool,X14)) )),
  inference(resolution,[],[f851,f860])).
fof(f1166,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK15) != s(bool,X10)) )),
  inference(resolution,[],[f851,f859])).
fof(f1165,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK15) != s(bool,X6) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f851,f857])).
fof(f1164,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK15) != s(bool,X5)) )),
  inference(resolution,[],[f851,f857])).
fof(f1163,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK15) != s(bool,X1)) )),
  inference(resolution,[],[f851,f855])).
fof(f1160,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))) | $spl1),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1162,plain,(
  ~p(s(bool,sK15)) | $spl3),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  ~p(s(bool,sK15)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f853,plain,(
  ~p(s(bool,sK15)) | s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),sK14))) != s('type.num.num','const.bvlSem.state_clock_1'(s('type.bvlSem.state'(sK12),'const.bvlSem.check_clock_2'(s('type.bvlSem.state'(sK12),sK13),s('type.bvlSem.state'(sK12),sK14)))))),
  inference(cnf_transformation,[],[f731])).
% SZS output end Proof for thm.bvlSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 451
% Passive clauses: 3264
% Generated clauses: 9962
% Final active clauses: 437
% Final passive clauses: 2610
% Input formulas: 516
% Initial clauses: 300
% 
% Selected by SInE selection: 104
% SInE iterations: 1
% 
% Duplicate literals: 321
% Trivial inequalities: 21
% Fw subsumption resolutions: 1069
% Fw demodulations: 670
% Bw demodulations: 2
% 
% Simple tautologies: 396
% Equational tautologies: 2
% Forward subsumptions: 3730
% 
% Binary resolution: 5795
% Factoring: 74
% Forward superposition: 1577
% Backward superposition: 244
% Equality resolution: 5
% 
% Split clauses: 323
% Split components: 666
% SAT solver clauses: 369
% SAT solver unit clauses: 45
% SAT solver binary clauses: 248
% 
% Sat splits: 323
% Sat splitting refutations: 45
% 
% Memory used [KB]: 4733
% Time elapsed: 0.165 s
% ------------------------------
% ------------------------------
% Success in time 12.07 s
