Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (42961)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 182
% Passive clauses: 3795
% Generated clauses: 20984
% Final active clauses: 160
% Final passive clauses: 3418
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 568
% Trivial inequalities: 347
% Fw subsumption resolutions: 1706
% Fw demodulations: 2933
% Bw demodulations: 22
% 
% Simple tautologies: 267
% Equational tautologies: 45
% Forward subsumptions: 10774
% 
% Binary resolution: 666
% Factoring: 1
% Forward superposition: 7111
% Backward superposition: 7164
% Self superposition: 170
% Equality factoring: 70
% Equality resolution: 30
% 
% Split clauses: 261
% Split components: 524
% SAT solver clauses: 744
% SAT solver unit clauses: 9
% SAT solver binary clauses: 159
% 
% Sat splits: 261
% 
% TWLsolver clauses: 744
% TWLsolver calls for satisfiability: 197
% 
% Memory used [KB]: 20596
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (42980)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1101
% Generated clauses: 15923
% Final active clauses: 140
% Final passive clauses: 885
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 268
% Trivial inequalities: 306
% Fw subsumption resolutions: 846
% Fw demodulations: 3503
% Bw demodulations: 68
% Condensations: 799
% Global subsumptions: 8
% 
% Simple tautologies: 115
% Equational tautologies: 41
% Forward subsumptions: 6856
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 145
% Factoring: 1
% Forward superposition: 4360
% Backward superposition: 4300
% Self superposition: 151
% Equality factoring: 72
% Equality resolution: 27
% 
% Split clauses: 82
% Split components: 164
% SAT solver clauses: 2207
% SAT solver unit clauses: 152
% SAT solver binary clauses: 53
% SAT solver learnt clauses: 16
% 
% Sat splits: 82
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2175
% TWLsolver calls for satisfiability: 381
% 
% Memory used [KB]: 18549
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f28061,plain,(
  $false),
  inference(subsumption_resolution,[],[f28060,f28059])).
fof(f28059,plain,(
  ( ! [X367] : (~p(s(bool,X367))) )),
  inference(forward_demodulation,[],[f28058,f320])).
fof(f320,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X16] : ! [X22,X31,X32,X33,X34] : s(X16,X33) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),X16),X22),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|7','thm.stackSem.result_case_def|split|3')).
fof(f28058,plain,(
  ( ! [X364,X366,X368,X365,X367] : (~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X364),s(fun('type.wordSem.word_loc'(sK0),bool),X365),s(fun('type.wordSem.word_loc'(sK0),bool),X366),s(bool,X367),s(bool,X368))))) )),
  inference(subsumption_resolution,[],[f27863,f27803])).
fof(f27803,plain,(
  ( ! [X6,X10,X9] : (s(X6,X10) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),X9),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f17002,f320])).
fof(f17002,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f323,f231])).
fof(f231,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f185])).
fof(f185,plain,(
  ? [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') != s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : ~s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X0),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.stackSem.result_distinct|split|7','thm.stackSem.result_distinct|split|7')).
fof(f323,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X5),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X16,X6] : ! [X7,X22,X31,X32,X33,X34] : s(X16,happ(s(fun('type.wordSem.word_loc'(X6),X16),X32),s('type.wordSem.word_loc'(X6),X7))) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Halt_1'(s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),X16),X22),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|7','thm.stackSem.result_case_def|split|2')).
fof(f27863,plain,(
  ( ! [X364,X366,X368,X365,X367] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X366),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X364),s(fun('type.wordSem.word_loc'(sK0),bool),X365),s(fun('type.wordSem.word_loc'(sK0),bool),X366),s(bool,X367),s(bool,X368))))) )),
  inference(superposition,[],[f246,f17002])).
fof(f246,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X18] : (s(bool,'F_0') = s(bool,X18) <=> ~p(s(bool,X18)))),
  file('thm.stackSem.result_distinct|split|7','thm.bool.EQ_CLAUSES|split|3')).
fof(f28060,plain,(
  ( ! [X424,X427] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X424),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X427))))) )),
  inference(subsumption_resolution,[],[f27873,f28059])).
fof(f27873,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X424),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X427)))) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X422),s(fun('type.wordSem.word_loc'(sK0),bool),X423),s(fun('type.wordSem.word_loc'(sK0),bool),X424),s(bool,X425),s(bool,X426))))) )),
  inference(superposition,[],[f314,f17002])).
fof(f314,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) => p(s(bool,X42))))),
  file('thm.stackSem.result_distinct|split|7','==>_HL_0')).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 221
% Passive clauses: 2370
% Generated clauses: 27950
% Final active clauses: 156
% Final passive clauses: 1678
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 773
% Trivial inequalities: 109
% Fw subsumption resolutions: 1522
% Fw demodulations: 4590
% Bw demodulations: 488
% Condensations: 1332
% Global subsumptions: 13
% 
% Simple tautologies: 195
% Equational tautologies: 180
% Forward subsumptions: 16491
% Fw demodulations to eq. taut.: 78
% Bw demodulations to eq. taut.: 24
% 
% Unit resulting resolution: 1452
% Factoring: 3
% Forward superposition: 10169
% Backward superposition: 6742
% Self superposition: 193
% Equality factoring: 76
% Equality resolution: 21
% 
% Split clauses: 62
% Split components: 135
% SAT solver clauses: 2586
% SAT solver unit clauses: 581
% SAT solver binary clauses: 803
% 
% Sat splits: 62
% Sat splitting refutations: 77
% 
% Memory used [KB]: 29807
% Time elapsed: 1.260 s
% ------------------------------
% ------------------------------
% Success in time 2.6 s
