Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.interp
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.interp
% SZS output start Proof for thm.interp
fof(f1525,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f1524])).
fof(f1524,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3)))),
  inference(forward_demodulation,[],[f1523,f48])).
fof(f48,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK5(X2,X1,X0)),s(X1,sK6(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK5(X2,X1,X0)),s(X1,sK6(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : ! [X11] : ? [X12,X13] : s('type.pair.prod'(X6,X7),X11) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13)))),
  file('thm.interp.interp_|49||49|_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f1523,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),'const.pair.,_2'(s(sK1,sK5(sK3,sK2,sK1)),s(sK2,sK6(sK3,sK2,sK1))))))),
  inference(superposition,[],[f1196,f53])).
fof(f53,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X6,X7] : ! [X10,X11,X14] : s(X8,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,X8)),X10),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X11),s(X7,X14))))) = s(X8,happ(s(fun(X7,X8),happ(s(fun(X6,fun(X7,X8)),X10),s(X6,X11))),s(X7,X14)))),
  file('thm.interp.interp_|49||49|_prove_ltb','thm.pair.UNCURRY_DEF')).
fof(f1196,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,happ(s(fun(sK2,sK0),happ(s(fun(sK1,fun(sK2,sK0)),sK4),s(sK1,sK5(sK3,sK2,sK1)))),s(sK2,sK6(sK3,sK2,sK1))))),
  inference(forward_demodulation,[],[f1195,f298])).
fof(f298,plain,(
  ( ! [X4,X5,X3] : (s(X3,sK5(X5,X4,X3)) = s(X3,'const.pair.FST_1'(s('type.pair.prod'(X3,X4),X5)))) )),
  inference(superposition,[],[f49,f48])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X6] : ! [X11,X14] : s(X6,X11) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X11),s(X7,X14)))))),
  file('thm.interp.interp_|49||49|_prove_ltb','thm.pair.FST')).
fof(f1195,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,happ(s(fun(sK2,sK0),happ(s(fun(sK1,fun(sK2,sK0)),sK4),s(sK1,'const.pair.FST_1'(s('type.pair.prod'(sK1,sK2),sK3))))),s(sK2,sK6(sK3,sK2,sK1))))),
  inference(forward_demodulation,[],[f40,f297])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (s(X1,sK6(X2,X1,X0)) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))) )),
  inference(superposition,[],[f50,f48])).
fof(f50,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : ! [X11,X14] : s(X7,X14) = s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X11),s(X7,X14)))))),
  file('thm.interp.interp_|49||49|_prove_ltb','thm.pair.SND')).
fof(f40,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,happ(s(fun(sK2,sK0),happ(s(fun(sK1,fun(sK2,sK0)),sK4),s(sK1,'const.pair.FST_1'(s('type.pair.prod'(sK1,sK2),sK3))))),s(sK2,'const.pair.SND_1'(s('type.pair.prod'(sK1,sK2),sK3)))))),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  s(sK0,'const.pair.UNCURRY_2'(s(fun(sK1,fun(sK2,sK0)),sK4),s('type.pair.prod'(sK1,sK2),sK3))) != s(sK0,happ(s(fun(sK2,sK0),happ(s(fun(sK1,fun(sK2,sK0)),sK4),s(sK1,'const.pair.FST_1'(s('type.pair.prod'(sK1,sK2),sK3))))),s(sK2,'const.pair.SND_1'(s('type.pair.prod'(sK1,sK2),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f34])).
fof(f34,plain,(
  ? [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X4),s('type.pair.prod'(X1,X2),X3))) != s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X3))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X3)))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X4),s('type.pair.prod'(X1,X2),X3))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X3))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X3)))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X4),s('type.pair.prod'(X1,X2),X3))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X2),X3))))),s(X2,'const.pair.SND_1'(s('type.pair.prod'(X1,X2),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10] : s(X6,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,X6)),X10),s('type.pair.prod'(X7,X8),X9))) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X10),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X8),X9))))),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X7,X8),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10] : s(X6,'const.pair.UNCURRY_2'(s(fun(X7,fun(X8,X6)),X10),s('type.pair.prod'(X7,X8),X9))) = s(X6,happ(s(fun(X8,X6),happ(s(fun(X7,fun(X8,X6)),X10),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X8),X9))))),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X7,X8),X9)))))),
  file('thm.interp.interp_|49||49|_prove_ltb','thm.interp.interp_|49||49|_prove_ltb')).
% SZS output end Proof for thm.interp
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 32
% Passive clauses: 389
% Generated clauses: 1465
% Final active clauses: 30
% Final passive clauses: 336
% Input formulas: 13
% Initial clauses: 14
% 
% Duplicate literals: 67
% Trivial inequalities: 4
% Fw subsumption resolutions: 10
% Fw demodulations: 180
% 
% Simple tautologies: 16
% Equational tautologies: 15
% Forward subsumptions: 798
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 4
% Forward superposition: 592
% Backward superposition: 508
% Self superposition: 58
% Equality factoring: 26
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 2174
% Time elapsed: 0.104 s
% ------------------------------
% ------------------------------
% Success in time 0.12 s
