Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.patSem
% (13759)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 145
% Passive clauses: 2608
% Generated clauses: 14249
% Final active clauses: 140
% Final passive clauses: 2324
% Input formulas: 69
% Initial clauses: 146
% 
% Duplicate literals: 561
% Trivial inequalities: 371
% Fw subsumption resolutions: 1420
% Fw demodulations: 1906
% Bw demodulations: 4
% 
% Simple tautologies: 388
% Equational tautologies: 55
% Forward subsumptions: 6285
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 237
% Factoring: 1
% Forward superposition: 4639
% Backward superposition: 4654
% Self superposition: 124
% Equality factoring: 86
% Equality resolution: 39
% 
% Split clauses: 161
% Split components: 334
% SAT solver clauses: 346
% SAT solver unit clauses: 9
% SAT solver binary clauses: 77
% 
% Sat splits: 161
% 
% TWLsolver clauses: 346
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.patSem
% (13780)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 153
% Passive clauses: 1350
% Generated clauses: 14975
% Final active clauses: 145
% Final passive clauses: 1158
% Input formulas: 69
% Initial clauses: 146
% 
% Duplicate literals: 464
% Trivial inequalities: 460
% Fw subsumption resolutions: 1022
% Fw demodulations: 1797
% Bw demodulations: 39
% Condensations: 746
% Global subsumptions: 8
% 
% Simple tautologies: 306
% Equational tautologies: 47
% Forward subsumptions: 6628
% 
% Binary resolution: 173
% Factoring: 1
% Forward superposition: 4906
% Backward superposition: 3637
% Self superposition: 111
% Equality factoring: 60
% Equality resolution: 39
% 
% Split clauses: 60
% Split components: 120
% SAT solver clauses: 2214
% SAT solver unit clauses: 155
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 60
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2180
% TWLsolver calls for satisfiability: 653
% 
% Memory used [KB]: 17014
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.patSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.patSem
% SZS output start Proof for thm.patSem
fof(f101695,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f255,f422,f420,f254,f421,f419,f424,f1377,f1388,f2179,f2185,f2182,f4913,f4920,f4954,f4956,f4958,f4962,f4964,f4966,f4968,f4972,f4974,f4976,f4978,f4981,f4983,f4985,f4987,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5024,f5026,f5056,f5060,f5062,f5154,f5156,f5158,f5160,f5255,f5256,f5257,f5258,f5259,f5712,f5713,f5716,f5717,f9246,f9361,f9363,f9365,f9366,f9362,f9372,f9373,f9364,f5632,f5652,f4606,f4626,f9409,f1387,f9410,f9632,f9634,f9641,f9642,f9676,f9678,f9680,f9682,f13372,f2180,f13379,f13373,f2186,f13419,f13436,f13434,f16769,f16844,f17185,f9311,f17812,f17889,f9310,f101574,f101589,f101592,f101643,f101670])).
fof(f101670,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f101669,f6126])).
fof(f6126,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',sK1))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X3),s('type.num.num',sK0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f422,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.num.num',X2) = s('type.num.num',X5)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5))) & (s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3) | s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5))) & ((s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3) | s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5)) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5)))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) <=> (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3) & s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4) & s('type.num.num',X2) = s('type.num.num',X5)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17,X18,X19,X20,X21] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X16),s('type.list.list'('type.patLang.exp'),X17),s('type.num.num',X18))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patLang.exp'),X20),s('type.num.num',X21))) <=> (s('type.list.list'('type.patSem.v'),X16) = s('type.list.list'('type.patSem.v'),X19) & s('type.list.list'('type.patLang.exp'),X17) = s('type.list.list'('type.patLang.exp'),X20) & s('type.num.num',X18) = s('type.num.num',X21)))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_|49||49||split|3')).
fof(f101669,plain,(
  ( ! [X430,X429] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X429),s('type.list.list'('type.patLang.exp'),X430),s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X429),s('type.list.list'('type.patLang.exp'),X430),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f90467,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.list.list'('type.patLang.exp'),X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X22,X23] : s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X22),s('type.num.num',X23))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X8))),s('type.list.list'('type.patLang.exp'),X22))),s('type.num.num',X23)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_3_const.patSem.Recclosure')).
fof(f90467,plain,(
  ( ! [X430,X429] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X429),s('type.list.list'('type.patLang.exp'),X430),s('type.num.num',sK1))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X429))),s('type.list.list'('type.patLang.exp'),X430))),s('type.num.num',sK0)))) ) | $spl0),
  inference(superposition,[],[f355,f84869])).
fof(f84869,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f74954,f322])).
fof(f322,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.patSem.v'),X0),X7)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X9)),X11),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X9)),X12),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.patSem.v'),X9),X15)))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_case_def|split|4')).
fof(f74954,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK1))) = s(X0,'const.patSem.v_CASE_7'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),X0)),X3),s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp',X0)),X4),s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num',X0))),X5),s(fun('type.num.num',X0),X1),s(fun('type.list.list'('type.patSem.v'),X0),X6)))) ) | $spl0),
  inference(superposition,[],[f322,f419])).
fof(f101643,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f101642,f1842])).
fof(f1842,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK0),s('type.list.list'('type.patSem.v'),X0))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.patSem.v'),X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f422,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.num.num',X0) = s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3))) & (s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3))) & ((s('type.num.num',X0) != s('type.num.num',X2) | s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3)) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3)))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) <=> (s('type.num.num',X0) = s('type.num.num',X2) & s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16,X17,X19,X20] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.patSem.v'),X17))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X19),s('type.list.list'('type.patSem.v'),X20))) <=> (s('type.num.num',X16) = s('type.num.num',X19) & s('type.list.list'('type.patSem.v'),X17) = s('type.list.list'('type.patSem.v'),X20)))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_|49||49||split|1')).
fof(f101642,plain,(
  ( ! [X191] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK0),s('type.list.list'('type.patSem.v'),X191))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.patSem.v'),X191)))) ) | $spl0),
  inference(forward_demodulation,[],[f90419,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.patSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X0))),s('type.list.list'('type.patSem.v'),X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X22] : s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.patSem.v'),X22))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',X8))),s('type.list.list'('type.patSem.v'),X22)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_2_const.patSem.Conv')).
fof(f90419,plain,(
  ( ! [X191] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK1),s('type.list.list'('type.patSem.v'),X191))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.patSem.v'),'type.patSem.v')),'const.patSem.Conv_0'),s('type.num.num',sK0))),s('type.list.list'('type.patSem.v'),X191)))) ) | $spl0),
  inference(superposition,[],[f324,f84869])).
fof(f101592,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f101591,f422])).
fof(f101591,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f90335,f304])).
fof(f304,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.patSem.v_|49||49||split|4','_i0_ax')).
fof(f90335,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f304,f84869])).
fof(f101589,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f101588,f6126])).
fof(f101588,plain,(
  ( ! [X47,X48] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X47),s('type.list.list'('type.patLang.exp'),X48),s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X47),s('type.list.list'('type.patLang.exp'),X48),s('type.num.num',sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f90333,f355])).
fof(f90333,plain,(
  ( ! [X47,X48] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X47),s('type.list.list'('type.patLang.exp'),X48),s('type.num.num',sK1))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),happ(s(fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v')),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.list.list'('type.patLang.exp'),fun('type.num.num','type.patSem.v'))),'const.patSem.Recclosure_0'),s('type.list.list'('type.patSem.v'),X47))),s('type.list.list'('type.patLang.exp'),X48))),s('type.num.num',sK0)))) ) | $spl0),
  inference(superposition,[],[f84869,f355])).
fof(f101574,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f101573,f422])).
fof(f101573,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl0),
  inference(forward_demodulation,[],[f90310,f304])).
fof(f90310,plain,(
  s('type.num.num',sK1) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f84869,f304])).
fof(f9310,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl20),
  inference(cnf_transformation,[],[f9310_D])).
fof(f9310_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f17889,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) | $spl38),
  inference(cnf_transformation,[],[f17889_D])).
fof(f17889_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f17812,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f17643])).
fof(f17643,plain,(
  ( ! [X66,X65] : (s(bool,X66) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6072,f16387])).
fof(f16387,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f16372])).
fof(f16372,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f16289,f438])).
fof(f438,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f265,f265])).
fof(f265,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.patSem.v_|49||49||split|4',aHL_BOOL_CASES)).
fof(f16289,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl10),
  inference(forward_demodulation,[],[f16207,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X37,X9] : ! [X45] : s(X9,X45) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X37),'const.sum.INL_1'(s(X9,X45)))))),
  file('thm.patSem.v_|49||49||split|4','thm.sum.OUTL')).
fof(f16207,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f9369,f8065])).
fof(f8065,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f7985,f258])).
fof(f258,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.patSem.v_|49||49||split|4',aHL_FALSITY)).
fof(f7985,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl10),
  inference(superposition,[],[f266,f2778])).
fof(f2778,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2659,f2211])).
fof(f2211,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f425,f2182,f2071])).
fof(f2071,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f2070,f304])).
fof(f2070,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f2038,f256])).
fof(f256,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.patSem.v_|49||49||split|4',aHL_TRUTH)).
fof(f2038,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f330,f465])).
fof(f465,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f265])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X34,X35] : ((p(s(bool,X34)) => p(s(bool,X35))) => ((p(s(bool,X35)) => p(s(bool,X34))) => s(bool,X34) = s(bool,X35)))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.IMP_ANTISYM_AX')).
fof(f425,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f258,f266])).
fof(f2659,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2294,f304])).
fof(f2294,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f2269,f256])).
fof(f2269,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f2071,f2211])).
fof(f266,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f9369,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2182,f2211,f2205,f2293,f2294,f2421,f2383,f2389,f2399,f2553,f2556,f2778,f2668,f2814,f2830,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f3262,f3263,f3265,f3267,f3268,f3269,f3270,f3305,f3334,f3336,f3338,f3339,f3340,f3341,f3446,f3516,f3517,f3951,f3957,f3610,f3984,f4017,f4018,f4020,f4022,f4023,f4024,f4025,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4332,f4333,f4357,f4358,f4359,f4531,f4365,f4366,f4419,f4421,f4422,f4533,f4437,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4755,f4889,f4913,f4710,f4607,f4680,f4685,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2816,f2184,f343,f342,f316,f315,f312,f4087,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4344,f4345,f4914,f5260,f4608,f4235,f4117,f4019,f3335,f3264,f2780,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f4021,f3337,f3266,f2782,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5712,f5714,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6661,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6922,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f7962,f7964,f7969,f8065,f8073,f8411,f8422,f8846,f8851,f8853,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9361,f9363,f9365,f9366,f9368])).
fof(f9368,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9367])).
fof(f9367,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f6860,f3517])).
fof(f6860,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6571,f2668])).
fof(f8952,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X66),s('type.list.list'('type.patLang.exp'),X68),s('type.num.num',X69))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X70),s('type.list.list'('type.patLang.exp'),X71),s('type.num.num',X72))) | s('type.list.list'('type.patSem.v'),X70) = s('type.list.list'('type.patSem.v'),sK4(X66,X67,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X67)))) )),
  inference(superposition,[],[f396,f829])).
fof(f8951,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X60),s('type.list.list'('type.patLang.exp'),X61),s('type.num.num',X62))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X63),s('type.list.list'('type.patLang.exp'),X64),s('type.num.num',X65))) | s('type.list.list'('type.patSem.v'),X63) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X59),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X60)))))) )),
  inference(superposition,[],[f396,f334])).
fof(f8950,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X53),s('type.list.list'('type.patLang.exp'),X54),s('type.num.num',X55))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X56),s('type.list.list'('type.patLang.exp'),X57),s('type.num.num',X58))) | s('type.list.list'('type.patSem.v'),X56) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X51,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X51,X52),s('type.list.list'('type.patSem.v'),X53)))))) )),
  inference(superposition,[],[f396,f338])).
fof(f8949,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X45),s('type.list.list'('type.patLang.exp'),X46),s('type.num.num',X47))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X48),s('type.list.list'('type.patLang.exp'),X49),s('type.num.num',X50))) | s('type.list.list'('type.patSem.v'),X48) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X44,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X45)))))) )),
  inference(superposition,[],[f396,f333])).
fof(f8948,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X37),s('type.list.list'('type.patLang.exp'),X39),s('type.num.num',X40))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X41),s('type.list.list'('type.patLang.exp'),X42),s('type.num.num',X43))) | s('type.list.list'('type.patSem.v'),X41) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X36),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X37),s(X36,X38)))))) )),
  inference(superposition,[],[f396,f337])).
fof(f8947,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X33),s('type.list.list'('type.patLang.exp'),X34),s('type.num.num',X35))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X29),s('type.num.num',X30))),s('type.list.list'('type.patLang.exp'),X31),s('type.num.num',X32))) | s('type.list.list'('type.patSem.v'),X33) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X28),s(fun('type.num.num','type.list.list'('type.patSem.v')),X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))) )),
  inference(superposition,[],[f396,f319])).
fof(f8946,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X23),s('type.num.num',X24))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X25),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X27))) | s('type.list.list'('type.patSem.v'),X25) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X21),s(fun('type.num.num','type.list.list'('type.patSem.v')),X22),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f396,f311])).
fof(f8945,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patLang.exp'),X19),s('type.num.num',X20))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X14),X15))),s('type.list.list'('type.patLang.exp'),X16),s('type.num.num',X17))) | s('type.list.list'('type.patSem.v'),X18) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X14),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X14),X15)))) )),
  inference(superposition,[],[f396,f335])).
fof(f8944,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X9),s('type.num.num',X10))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X11),s('type.list.list'('type.patLang.exp'),X12),s('type.num.num',X13))) | s('type.list.list'('type.patSem.v'),X11) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X6),'type.list.list'('type.patSem.v')),sK16(X7,X8,X6,'type.list.list'('type.patSem.v'))),s('type.list.list'(X6),'const.list.NIL_0')))) )),
  inference(superposition,[],[f396,f353])).
fof(f8943,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X3) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f396,f304])).
fof(f8853,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f8818,f256])).
fof(f8818,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X104),bool),sK16(X105,'|dollar|exists_1'(s(fun('type.list.list'(X102),bool),sK16(X103,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X102,bool))),X104,bool))))) ) | $spl10),
  inference(superposition,[],[f4680,f5714])).
fof(f8851,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(subsumption_resolution,[],[f8795,f256])).
fof(f8795,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X39,bool)))) | s(bool,'T_0') = s(bool,X41)) ) | $spl10),
  inference(superposition,[],[f440,f5714])).
fof(f8846,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(subsumption_resolution,[],[f8777,f425])).
fof(f8777,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X23),bool),sK16(X24,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X23,bool)))) | s(bool,'F_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f5714,f2668])).
fof(f8422,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X1))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2)) )),
  inference(equality_resolution,[],[f393])).
fof(f8411,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.patLang.exp',X48) != s('type.patLang.exp',X50) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X52),s('type.patLang.exp',X50))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X51),s('type.patLang.exp',sK4(X48,X49,'type.patLang.exp')))) | s('type.list.list'('type.patSem.v'),X51) != s('type.list.list'('type.patSem.v'),X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.patLang.exp',bool),X49)))) )),
  inference(superposition,[],[f393,f829])).
fof(f8073,plain,(
  ( ! [X43,X42] : (s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f8007,f258])).
fof(f8007,plain,(
  ( ! [X43,X42] : (p(s(bool,'F_0')) | s(bool,X43) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X42))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f2073,f2778])).
fof(f7969,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,sK4(X75,X76,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76)))) ) | $spl10),
  inference(superposition,[],[f2778,f829])).
fof(f7964,plain,(
  ( ! [X64,X62,X63] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X63),s('type.num.num',X64))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X62),s(fun('type.num.num',bool),X63),s('type.num.num','const.num.SUC_1'(s('type.num.num',X64)))))) ) | $spl10),
  inference(superposition,[],[f2778,f319])).
fof(f7962,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) ) | $spl10),
  inference(superposition,[],[f2778,f335])).
fof(f7346,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X1))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(equality_resolution,[],[f390])).
fof(f7335,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s('type.list.list'('type.patSem.v'),X48) != s('type.list.list'('type.patSem.v'),X50) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X52),s('type.list.list'('type.patSem.v'),X50))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X51),s('type.list.list'('type.patSem.v'),sK4(X48,X49,'type.list.list'('type.patSem.v'))))) | s('type.num.num',X51) != s('type.num.num',X52) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X49)))) )),
  inference(superposition,[],[f390,f829])).
fof(f6892,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f6571,f829])).
fof(f6891,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X116),'const.sum.INL_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f6571,f334])).
fof(f6890,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f6571,f338])).
fof(f6889,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X109),'const.sum.INL_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,X108)))))) )),
  inference(superposition,[],[f6571,f333])).
fof(f6888,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X106),'const.sum.INL_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f6571,f337])).
fof(f6887,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X102),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f6571,f319])).
fof(f6886,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6571,f311])).
fof(f6885,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f6571,f335])).
fof(f6884,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6571,f353])).
fof(f6922,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f6864,f425])).
fof(f6864,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f6571,f2668])).
fof(f6721,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f6689,f256])).
fof(f6689,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f351,f438])).
fof(f6686,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f351,f829])).
fof(f6676,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f351,f829])).
fof(f6632,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f6367,f829])).
fof(f6631,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X115))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f6367,f334])).
fof(f6630,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f6367,f338])).
fof(f6629,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,X108))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,X108)))))) )),
  inference(superposition,[],[f6367,f333])).
fof(f6628,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X106,bool),'const.sum.INR_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f6367,f337])).
fof(f6627,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f6367,f319])).
fof(f6626,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6367,f311])).
fof(f6625,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f6367,f335])).
fof(f6624,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6367,f353])).
fof(f6661,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(subsumption_resolution,[],[f6604,f425])).
fof(f6604,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,X40))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl10),
  inference(superposition,[],[f6367,f2668])).
fof(f6579,plain,(
  ( ! [X61,X62,X63] : (s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f6522,f258])).
fof(f6522,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'F_0')) | s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f2073,f2069])).
fof(f6571,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6501,f258])).
fof(f6501,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f266,f2069])).
fof(f6487,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X119),'const.sum.INL_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,sK4(X117,X118,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X118)))) )),
  inference(superposition,[],[f2069,f829])).
fof(f6485,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))))) )),
  inference(superposition,[],[f2069,f338])).
fof(f6483,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X106),'const.sum.INL_1'(s(bool,X104))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,X104),s(X103,X105)))))) )),
  inference(superposition,[],[f2069,f337])).
fof(f6482,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X102),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f2069,f319])).
fof(f6481,plain,(
  ( ! [X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X98),'const.sum.INL_1'(s(bool,X96))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2069,f311])).
fof(f6480,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),X94))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X93),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X93),X94)))) )),
  inference(superposition,[],[f2069,f335])).
fof(f6479,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X89),bool),sK16(X90,X91,X89,bool)),s('type.list.list'(X89),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2069,f353])).
fof(f6436,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f6404,f256])).
fof(f6404,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f350,f438])).
fof(f6401,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f350,f829])).
fof(f6391,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f350,f829])).
fof(f6375,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f6318,f258])).
fof(f6318,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f2073,f2067])).
fof(f6367,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6298,f258])).
fof(f6298,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f266,f2067])).
fof(f6284,plain,(
  ( ! [X116,X114,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X116,bool),'const.sum.INR_1'(s(bool,X114))))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f2067,f829])).
fof(f6282,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X110,bool),'const.sum.INR_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X108),s(bool,X109)))))) )),
  inference(superposition,[],[f2067,f338])).
fof(f6280,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,X101))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X100),'const.pair.,_2'(s(bool,X101),s(X100,X102)))))) )),
  inference(superposition,[],[f2067,f337])).
fof(f6279,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X97),s('type.num.num',X98))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X96),s(fun('type.num.num',bool),X97),s('type.num.num','const.num.SUC_1'(s('type.num.num',X98)))))) )),
  inference(superposition,[],[f2067,f319])).
fof(f6278,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X93))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2067,f311])).
fof(f6277,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X90),X91))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X90),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X90),X91)))) )),
  inference(superposition,[],[f2067,f335])).
fof(f6276,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X86),bool),sK16(X87,X88,X86,bool)),s('type.list.list'(X86),'const.list.NIL_0')))) )),
  inference(superposition,[],[f2067,f353])).
fof(f6158,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X80),s('type.list.list'('type.patLang.exp'),X81),s('type.num.num',X78))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X82),s('type.list.list'('type.patLang.exp'),X83),s('type.num.num',X84))) | s('type.num.num',X84) = s('type.num.num',sK4(X78,X79,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X79)))) )),
  inference(superposition,[],[f398,f829])).
fof(f6157,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X73),s('type.list.list'('type.patLang.exp'),X74),s('type.num.num',X72))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X75),s('type.list.list'('type.patLang.exp'),X76),s('type.num.num',X77))) | s('type.num.num',X77) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X71),'const.sum.INL_1'(s('type.num.num',X72)))))) )),
  inference(superposition,[],[f398,f334])).
fof(f6156,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65,X63] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X66),s('type.list.list'('type.patLang.exp'),X67),s('type.num.num',X65))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X68),s('type.list.list'('type.patLang.exp'),X69),s('type.num.num',X70))) | s('type.num.num',X70) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X63,'type.num.num'),'const.pair.,_2'(s(X63,X64),s('type.num.num',X65)))))) )),
  inference(superposition,[],[f398,f338])).
fof(f6155,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X58),s('type.list.list'('type.patLang.exp'),X59),s('type.num.num',X57))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X60),s('type.list.list'('type.patLang.exp'),X61),s('type.num.num',X62))) | s('type.num.num',X62) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X56,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X57)))))) )),
  inference(superposition,[],[f398,f333])).
fof(f6154,plain,(
  ( ! [X54,X52,X50,X48,X55,X53,X51,X49] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X51),s('type.list.list'('type.patLang.exp'),X52),s('type.num.num',X49))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X53),s('type.list.list'('type.patLang.exp'),X54),s('type.num.num',X55))) | s('type.num.num',X55) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X48),'const.pair.,_2'(s('type.num.num',X49),s(X48,X50)))))) )),
  inference(superposition,[],[f398,f337])).
fof(f6153,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X45),s('type.list.list'('type.patLang.exp'),X46),s('type.num.num',X47))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X43),s('type.list.list'('type.patLang.exp'),X44),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X41),s('type.num.num',X42))))) | s('type.num.num',X47) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X40),s(fun('type.num.num','type.num.num'),X41),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42)))))) )),
  inference(superposition,[],[f398,f319])).
fof(f6152,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X35),s('type.list.list'('type.patLang.exp'),X36),s('type.num.num',X33))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X37),s('type.list.list'('type.patLang.exp'),X38),s('type.num.num',X39))) | s('type.num.num',X39) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X33),s(fun('type.num.num','type.num.num'),X34),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f398,f311])).
fof(f6151,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X32] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X30),s('type.list.list'('type.patLang.exp'),X31),s('type.num.num',X32))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X28),s('type.list.list'('type.patLang.exp'),X29),s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X26),X27))))) | s('type.num.num',X32) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X26),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X26),X27)))) )),
  inference(superposition,[],[f398,f335])).
fof(f6150,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X22),s('type.num.num',X20))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X23),s('type.list.list'('type.patLang.exp'),X24),s('type.num.num',X25))) | s('type.num.num',X25) = s('type.num.num',happ(s(fun('type.list.list'(X18),'type.num.num'),sK16(X19,X20,X18,'type.num.num')),s('type.list.list'(X18),'const.list.NIL_0')))) )),
  inference(superposition,[],[f398,f353])).
fof(f5699,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,X146)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X148),bool),sK16(X149,sK4(X146,X147,bool),X148,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X147)))) )),
  inference(superposition,[],[f4680,f829])).
fof(f5698,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,X143)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X144),bool),sK16(X145,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X142),'const.sum.INL_1'(s(bool,X143)))),X144,bool))))) )),
  inference(superposition,[],[f4680,f334])).
fof(f5697,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,X139)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK16(X141,'const.pair.SND_1'(s('type.pair.prod'(X137,bool),'const.pair.,_2'(s(X137,X138),s(bool,X139)))),X140,bool))))) )),
  inference(superposition,[],[f4680,f338])).
fof(f5696,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X134)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X135),bool),sK16(X136,'const.sum.OUTR_1'(s('type.sum.sum'(X133,bool),'const.sum.INR_1'(s(bool,X134)))),X135,bool))))) )),
  inference(superposition,[],[f4680,f333])).
fof(f5695,plain,(
  ( ! [X132,X130,X128,X131,X129] : (~p(s(bool,X129)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X131),bool),sK16(X132,'const.pair.FST_1'(s('type.pair.prod'(bool,X128),'const.pair.,_2'(s(bool,X129),s(X128,X130)))),X131,bool))))) )),
  inference(superposition,[],[f4680,f337])).
fof(f5694,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun('type.num.num',bool),X124),s('type.num.num',X125)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X126),bool),sK16(X127,'const.ind_type.FCONS_3'(s(bool,X123),s(fun('type.num.num',bool),X124),s('type.num.num','const.num.SUC_1'(s('type.num.num',X125)))),X126,bool))))) )),
  inference(superposition,[],[f4680,f319])).
fof(f5693,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X121),bool),sK16(X122,'const.ind_type.FCONS_3'(s(bool,X119),s(fun('type.num.num',bool),X120),s('type.num.num','const.num.|48|_0')),X121,bool))))) )),
  inference(superposition,[],[f4680,f311])).
fof(f5692,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),X116)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X117),bool),sK16(X118,happ(s(fun('type.pair.prod'(bool,X115),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X115),X116)),X117,bool))))) )),
  inference(superposition,[],[f4680,f335])).
fof(f5691,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X113),bool),sK16(X114,happ(s(fun('type.list.list'(X110),bool),sK16(X111,X112,X110,bool)),s('type.list.list'(X110),'const.list.NIL_0')),X113,bool))))) )),
  inference(superposition,[],[f4680,f353])).
fof(f5690,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK16(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4680,f304])).
fof(f5723,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f5687,f256])).
fof(f5687,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4680,f533])).
fof(f5722,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X94),bool),sK16(X95,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X93)))),X94,bool)))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f5686,f256])).
fof(f5686,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X94),bool),sK16(X95,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X93)))),X94,bool)))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f4680,f506])).
fof(f5719,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X87),bool),sK16(X88,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'T_0',X85,bool))),X87,bool))))) )),
  inference(subsumption_resolution,[],[f5684,f256])).
fof(f5684,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X87),bool),sK16(X88,'|dollar|exists_1'(s(fun('type.list.list'(X85),bool),sK16(X86,'T_0',X85,bool))),X87,bool))))) )),
  inference(superposition,[],[f4680,f4709])).
fof(f5715,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'|dollar|exists_1'(s(fun(X75,bool),X76)),X77,bool)))) | s(X75,X79) = s(X75,sK4(X79,X76,X75))) )),
  inference(subsumption_resolution,[],[f5681,f256])).
fof(f5681,plain,(
  ( ! [X78,X76,X79,X77,X75] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'|dollar|exists_1'(s(fun(X75,bool),X76)),X77,bool)))) | s(X75,X79) = s(X75,sK4(X79,X76,X75))) )),
  inference(superposition,[],[f4680,f829])).
fof(f5714,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X73,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5680,f256])).
fof(f5680,plain,(
  ( ! [X74,X73] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X73,bool))))) ) | $spl10),
  inference(superposition,[],[f4680,f2211])).
fof(f5711,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK16(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(subsumption_resolution,[],[f5678,f256])).
fof(f5678,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X68),bool),sK16(X69,'const.ind_type.ISO_2'(s(fun(X67,X67),'_i_0_0'),s(fun(X67,X67),'_i_0_0')),X68,bool))))) )),
  inference(superposition,[],[f4680,f433])).
fof(f5710,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f5677,f256])).
fof(f5677,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4680,f465])).
fof(f5700,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5661,f256])).
fof(f5661,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4680,f438])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : ! [X31,X24] : (! [X45] : (s(X9,X24) = s(X9,X45) => p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X31)))))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.JRH_INDUCT_UTIL')).
fof(f4535,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4483,f256])).
fof(f4483,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f330,f829])).
fof(f4536,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4484,f256])).
fof(f4484,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f411,f829])).
fof(f822,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f313,f304])).
fof(f923,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK16(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f313,f353])).
fof(f1508,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f313,f335])).
fof(f823,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X6) = s(X3,sK4(X6,'const.ind_type.FCONS_3'(s(fun(X3,bool),X4),s(fun('type.num.num',fun(X3,bool)),X5),s('type.num.num','const.num.|48|_0')),X3))) )),
  inference(superposition,[],[f313,f311])).
fof(f2382,plain,(
  ( ! [X134,X132,X136,X135,X133] : (p(s(bool,'|dollar|exists_1'(s(fun(X132,bool),happ(s(fun('type.num.num',fun(X132,bool)),X134),s('type.num.num',X135)))))) | s(X132,X136) = s(X132,sK4(X136,'const.ind_type.FCONS_3'(s(fun(X132,bool),X133),s(fun('type.num.num',fun(X132,bool)),X134),s('type.num.num','const.num.SUC_1'(s('type.num.num',X135)))),X132))) )),
  inference(superposition,[],[f313,f319])).
fof(f824,plain,(
  ( ! [X10,X8,X7,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X7,bool),X9)))) | s(X7,X11) = s(X7,sK4(X11,'const.pair.FST_1'(s('type.pair.prod'(fun(X7,bool),X8),'const.pair.,_2'(s(fun(X7,bool),X9),s(X8,X10)))),X7))) )),
  inference(superposition,[],[f313,f337])).
fof(f825,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | s(X12,X15) = s(X12,sK4(X15,'const.sum.OUTR_1'(s('type.sum.sum'(X13,fun(X12,bool)),'const.sum.INR_1'(s(fun(X12,bool),X14)))),X12))) )),
  inference(superposition,[],[f313,f333])).
fof(f826,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X16,bool),X19)))) | s(X16,X20) = s(X16,sK4(X20,'const.pair.SND_1'(s('type.pair.prod'(X17,fun(X16,bool)),'const.pair.,_2'(s(X17,X18),s(fun(X16,bool),X19)))),X16))) )),
  inference(superposition,[],[f313,f338])).
fof(f827,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f313,f334])).
fof(f4442,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f313,f829])).
fof(f4406,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f266,f829])).
fof(f4407,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f267,f829])).
fof(f4408,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f270,f829])).
fof(f4409,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f271,f829])).
fof(f4410,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f327,f829])).
fof(f4412,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f330,f829])).
fof(f4413,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f411,f829])).
fof(f4418,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f440,f829])).
fof(f4420,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f487,f829])).
fof(f4425,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f2071,f829])).
fof(f4426,plain,(
  ( ! [X111,X109,X110] : (p(s(bool,X109)) | s(bool,X111) != s(bool,sK4(X109,X110,bool)) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f2073,f829])).
fof(f4548,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4532,f829])).
fof(f4440,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f312,f829])).
fof(f4443,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f413,f829])).
fof(f4783,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(fun(X57,X58),X59) != s(fun(X57,X58),X61) | s(X58,happ(s(fun(X57,X58),X61),s(X57,X62))) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f348,f829])).
fof(f4451,plain,(
  ( ! [X196,X197,X195] : (s('type.ast.lit',X195) != s('type.ast.lit',X197) | s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X197))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK4(X195,X196,'type.ast.lit')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X196)))) )),
  inference(superposition,[],[f332,f829])).
fof(f2346,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f319,f438])).
fof(f2348,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f319,f438])).
fof(f2374,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f465,f319])).
fof(f2782,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(forward_demodulation,[],[f2663,f2211])).
fof(f2663,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',X31))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))) ) | $spl10),
  inference(superposition,[],[f2294,f319])).
fof(f3266,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))))) ) | $spl10),
  inference(superposition,[],[f2668,f319])).
fof(f3337,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61))) | s(bool,X62) = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | s(bool,'F_0') = s(bool,X62)) ) | $spl10),
  inference(superposition,[],[f2668,f319])).
fof(f4021,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X75),s('type.num.num',X76))) | s(bool,X77) = s(bool,'const.ind_type.FCONS_3'(s(bool,X74),s(fun('type.num.num',bool),X75),s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl10),
  inference(superposition,[],[f3517,f319])).
fof(f4119,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f506,f319])).
fof(f4237,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X81),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X79),s('type.num.num',X80))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X78),s(fun('type.num.num',bool),X79),s('type.num.num','const.num.SUC_1'(s('type.num.num',X80)))))) )),
  inference(superposition,[],[f533,f319])).
fof(f1477,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f335,f438])).
fof(f1479,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f335,f438])).
fof(f1481,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X12),X13)))) )),
  inference(superposition,[],[f335,f265])).
fof(f2019,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f465,f335])).
fof(f2187,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0')))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29))))) )),
  inference(forward_demodulation,[],[f2165,f335])).
fof(f2165,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f312,f335])).
fof(f2780,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(forward_demodulation,[],[f2661,f2211])).
fof(f2661,plain,(
  ( ! [X26,X25] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) ) | $spl10),
  inference(superposition,[],[f2294,f335])).
fof(f3264,plain,(
  ( ! [X68,X66,X67] : (s(bool,X68) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67)))) ) | $spl10),
  inference(superposition,[],[f2668,f335])).
fof(f3335,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),X54))) | s(bool,X55) = s(bool,happ(s(fun('type.pair.prod'(bool,X53),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X53),X54))) | s(bool,'F_0') = s(bool,X55)) ) | $spl10),
  inference(superposition,[],[f2668,f335])).
fof(f4019,plain,(
  ( ! [X70,X68,X69] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),X69))) | s(bool,X70) = s(bool,happ(s(fun('type.pair.prod'(bool,X68),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X68),X69))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f3517,f335])).
fof(f4117,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f506,f335])).
fof(f4235,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f533,f335])).
fof(f4608,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4532,f335])).
fof(f5260,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) )),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f313,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f822,f823,f824,f825,f826,f827,f829,f830,f831,f832,f833,f834,f835,f902,f923,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1477,f1479,f1481,f1490,f1526,f1508,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2019,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2187,f2346,f2348,f2374,f2421,f2382,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4117,f4118,f4119,f4120,f4121,f4122,f4123,f4202,f4234,f4235,f4236,f4237,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4406,f4407,f4408,f4409,f4410,f4412,f4413,f4418,f4419,f4420,f4421,f4422,f4425,f4426,f4440,f4441,f4442,f4443,f4447,f4451,f4453,f4460,f4462,f4470,f4535,f4536,f4537,f4548,f4608,f4666,f4709,f4703,f4783,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914])).
fof(f4914,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f4902,f256])).
fof(f4902,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4608,f438])).
fof(f4345,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X36) != s(bool,f(s(X33,sK12(X34,X35,bool,X33)))) | s(fun(X33,bool),X34) = s(fun(X33,bool),X35) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(X33,sK12(X34,X35,bool,X33))))) ) | $spl10),
  inference(superposition,[],[f347,f2668])).
fof(f4344,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'F_0') != s(bool,f(s(X29,sK12(X30,X31,bool,X29)))) | s(fun(X29,bool),X30) = s(fun(X29,bool),X31) | s(bool,X32) = s(bool,g(s(X29,sK12(X30,X31,bool,X29)))) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f347,f2668])).
fof(f4338,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK12(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK12(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f347,f438])).
fof(f4337,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK12(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK12(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f347,f438])).
fof(f4336,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f347,f265])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : ! [X31] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X31),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X24] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X31),s('type.list.list'(X9),X24)))) => ! [X32] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X31),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X32),s('type.list.list'(X9),X24)))))))) => ! [X33] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X31),s('type.list.list'(X9),X33)))))),
  file('thm.patSem.v_|49||49||split|4','thm.list.list_INDUCT')).
fof(f413,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f4073,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f4052,f304])).
fof(f4052,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f304])).
fof(f4074,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f4053,f353])).
fof(f4053,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f353])).
fof(f4076,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f4055,f311])).
fof(f4055,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X10),s('type.list.list'(X9),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X9),bool),X10),s(fun('type.num.num',fun('type.list.list'(X9),bool)),X11),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f311])).
fof(f4078,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X16),bool),X17),'const.pair.,_2'(s(fun('type.list.list'(X16),bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f4057,f337])).
fof(f4057,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),X18),s('type.list.list'(X16),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X16),bool),X17),'const.pair.,_2'(s(fun('type.list.list'(X16),bool),X18),s(X17,X19)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X16),bool),X17),'const.pair.,_2'(s(fun('type.list.list'(X16),bool),X18),s(X17,X19))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f337])).
fof(f4079,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X21,fun('type.list.list'(X20),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X20),bool),X22)))),X20)))))) )),
  inference(forward_demodulation,[],[f4058,f333])).
fof(f4058,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X21,fun('type.list.list'(X20),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X20),bool),X22)))),X20))))) | ~p(s(bool,happ(s(fun('type.list.list'(X20),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X21,fun('type.list.list'(X20),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X20),bool),X22))))),s('type.list.list'(X20),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f333])).
fof(f4080,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23)))))) )),
  inference(forward_demodulation,[],[f4059,f338])).
fof(f4059,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,happ(s(fun('type.list.list'(X23),bool),X26),s('type.list.list'(X23),sK2('const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26)))),X23))))) | ~p(s(bool,happ(s(fun('type.list.list'(X23),bool),'const.pair.SND_1'(s('type.pair.prod'(X24,fun('type.list.list'(X23),bool)),'const.pair.,_2'(s(X24,X25),s(fun('type.list.list'(X23),bool),X26))))),s('type.list.list'(X23),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f338])).
fof(f4081,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27)))))) )),
  inference(forward_demodulation,[],[f4060,f334])).
fof(f4060,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29))))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f334])).
fof(f4082,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f4061,f258])).
fof(f4061,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f413,f265])).
fof(f4087,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),'const.list.NIL_0')))) | s(bool,X23) = s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),sK2(X22,X21)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(subsumption_resolution,[],[f4069,f258])).
fof(f4069,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),'const.list.NIL_0')))) | s(bool,X23) = s(bool,happ(s(fun('type.list.list'(X21),bool),X22),s('type.list.list'(X21),sK2(X22,X21)))) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f413,f2668])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ! [X31,X45] : (p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45)))) => p(s(bool,happ(s(fun(X9,bool),X31),s(X9,'const.min.@_1'(s(fun(X9,bool),X31)))))))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.SELECT_AX')).
fof(f315,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X57,X31] : (! [X45] : (p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45)))) => p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X45))))) => (? [X45] : p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45)))) => ? [X45] : p(s(bool,happ(s(fun(X9,bool),X57),s(X9,X45))))))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.MONO_EXISTS')).
fof(f316,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f342,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X37] : ! [X54,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X37),X10),s(fun(X37,X9),X54)))) => (! [X31] : (! [X45] : p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45)))) <=> ! [X45] : p(s(bool,happ(s(fun(X9,bool),X31),s(X9,happ(s(fun(X37,X9),X54),s(X37,X45))))))) & ! [X31] : (? [X45] : p(s(bool,happ(s(fun(X9,bool),X31),s(X9,X45)))) <=> ? [X45] : p(s(bool,happ(s(fun(X9,bool),X31),s(X9,happ(s(fun(X37,X9),X54),s(X37,X45))))))) & ! [X6,X56] : (s(X9,X6) = s(X9,happ(s(fun(X37,X9),X54),s(X37,X56))) <=> s(X37,X56) = s(X37,happ(s(fun(X9,X37),X10),s(X9,X6))))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.ISO_USAGE')).
fof(f343,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f2184,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2158,f256])).
fof(f2158,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f312,f438])).
fof(f2816,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(subsumption_resolution,[],[f2815,f256])).
fof(f2815,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(forward_demodulation,[],[f2700,f2211])).
fof(f2700,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl10),
  inference(superposition,[],[f312,f2294])).
fof(f4532,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4390,f266])).
fof(f4390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f314,f829])).
fof(f4628,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4595,f256])).
fof(f4595,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4532,f438])).
fof(f341,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f344,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f4077,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f4056,f319])).
fof(f4056,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X12),bool),happ(s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num',X15))),s('type.list.list'(X12),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun('type.list.list'(X12),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X12),bool),X13),s(fun('type.num.num',fun('type.list.list'(X12),bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.list.list'(X12),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f319])).
fof(f4075,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f4054,f335])).
fof(f4054,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f413,f335])).
fof(f2164,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.min.@_1'(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)))))))) )),
  inference(superposition,[],[f312,f353])).
fof(f345,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f234])).
fof(f346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f234])).
fof(f350,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0)))))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X37,X9] : ! [X10,X54] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X37),X10),s(fun(X37,X9),X54)))) <=> (! [X45] : s(X37,X45) = s(X37,happ(s(fun(X9,X37),X10),s(X9,happ(s(fun(X37,X9),X54),s(X37,X45))))) & ! [X38] : s(X9,X38) = s(X9,happ(s(fun(X37,X9),X54),s(X37,happ(s(fun(X9,X37),X10),s(X9,X38)))))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.ISO')).
fof(f351,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f352,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f242])).
fof(f395,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6))))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X37,X51,X9,X52] : ! [X53,X54,X55,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X52),X10),s(fun(X52,X9),X55)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X37,X51),X54),s(fun(X51,X37),X53))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X37),fun(X52,X51)),'_dst_x00x11_2'(s(fun(X37,X51),X54),s(fun(fun(X9,X37),fun(X52,X37)),'_dstsel_0_x01_1'(s(fun(X52,X9),X55))))),s(fun(fun(X52,X51),fun(X9,X37)),'_dst_x00x11_2'(s(fun(X51,X37),X53),s(fun(fun(X52,X51),fun(X9,X51)),'_dstsel_0_x01_1'(s(fun(X9,X52),X10)))))))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.ISO_FUN')).
fof(f4369,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f829,f304])).
fof(f4370,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f829,f353])).
fof(f4372,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f829,f311])).
fof(f4374,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23))) | s(X21,X25) = s(X21,sK4(X25,'const.pair.FST_1'(s('type.pair.prod'(fun(X21,bool),X22),'const.pair.,_2'(s(fun(X21,bool),X23),s(X22,X24)))),X21))) )),
  inference(superposition,[],[f829,f337])).
fof(f4375,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28))) | s(X26,X29) = s(X26,sK4(X29,'const.sum.OUTR_1'(s('type.sum.sum'(X27,fun(X26,bool)),'const.sum.INR_1'(s(fun(X26,bool),X28)))),X26))) )),
  inference(superposition,[],[f829,f333])).
fof(f4376,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X30,bool),X33))) | s(X30,X34) = s(X30,sK4(X34,'const.pair.SND_1'(s('type.pair.prod'(X31,fun(X30,bool)),'const.pair.,_2'(s(X31,X32),s(fun(X30,bool),X33)))),X30))) )),
  inference(superposition,[],[f829,f338])).
fof(f4377,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X37))) | s(X35,X38) = s(X35,sK4(X38,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X35,bool),X36),'const.sum.INL_1'(s(fun(X35,bool),X37)))),X35))) )),
  inference(superposition,[],[f829,f334])).
fof(f4378,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f829,f829])).
fof(f4373,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),happ(s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num',X19))))) | s(X16,X20) = s(X16,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X16,bool),X17),s(fun('type.num.num',fun(X16,bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))) )),
  inference(superposition,[],[f829,f319])).
fof(f4371,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f829,f335])).
fof(f4685,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(subsumption_resolution,[],[f4676,f258])).
fof(f4676,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'F_0')) | ~p(s(bool,X35)) | s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X33),bool),sK16(X34,X35,X33,bool)))) | s(bool,'F_0') = s(bool,X36)) ) | $spl10),
  inference(superposition,[],[f4607,f2668])).
fof(f4680,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4668,f258])).
fof(f4668,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4607,f265])).
fof(f4607,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4532,f353])).
fof(f4710,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4666,f4607])).
fof(f4889,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4842,f256])).
fof(f4842,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f440,f4709])).
fof(f4755,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f4720,f258])).
fof(f4720,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'F_0')) | s(bool,X26) = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X24),bool),sK16(X25,'T_0',X24,bool)))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f4666,f2668])).
fof(f4703,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f353,f4666,f440])).
fof(f4709,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f425,f4666,f2071])).
fof(f4666,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f256,f4607])).
fof(f4537,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4487,f256])).
fof(f4487,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f440,f829])).
fof(f4470,plain,(
  ( ! [X269,X271,X273,X270,X272] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X271),s('type.patLang.exp',X272))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X273),s('type.patLang.exp',X269))) | s('type.patLang.exp',X272) = s('type.patLang.exp',sK4(X269,X270,'type.patLang.exp')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.patLang.exp',bool),X270)))) )),
  inference(superposition,[],[f392,f829])).
fof(f4462,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X233),s('type.patLang.exp',X235))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X236),s('type.patLang.exp',X237))) | s('type.list.list'('type.patSem.v'),X236) = s('type.list.list'('type.patSem.v'),sK4(X233,X234,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X234)))) )),
  inference(superposition,[],[f391,f829])).
fof(f4460,plain,(
  ( ! [X227,X225,X223,X226,X224] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X225),s('type.list.list'('type.patSem.v'),X223))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X226),s('type.list.list'('type.patSem.v'),X227))) | s('type.list.list'('type.patSem.v'),X227) = s('type.list.list'('type.patSem.v'),sK4(X223,X224,'type.list.list'('type.patSem.v'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'('type.patSem.v'),bool),X224)))) )),
  inference(superposition,[],[f389,f829])).
fof(f4453,plain,(
  ( ! [X202,X203,X201] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X201))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X203))) | s('type.ast.lit',X203) = s('type.ast.lit',sK4(X201,X202,'type.ast.lit')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ast.lit',bool),X202)))) )),
  inference(superposition,[],[f331,f829])).
fof(f4447,plain,(
  ( ! [X177,X175,X178,X176,X179] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X175),s('type.list.list'('type.patSem.v'),X177))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X178),s('type.list.list'('type.patSem.v'),X179))) | s('type.num.num',X178) = s('type.num.num',sK4(X175,X176,'type.num.num')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),X176)))) )),
  inference(superposition,[],[f388,f829])).
fof(f4441,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f314,f829])).
fof(f4437,plain,(
  ( ! [X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,X141) = s(bool,sK4(X139,X140,bool)) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) ) | $spl10),
  inference(superposition,[],[f3517,f829])).
fof(f4533,plain,(
  ( ! [X114,X115] : (s(bool,'T_0') != s(bool,X114) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) ) | $spl10),
  inference(forward_demodulation,[],[f4428,f2211])).
fof(f4428,plain,(
  ( ! [X114,X115] : (s(bool,X114) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,sK4(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) ) | $spl10),
  inference(superposition,[],[f2293,f829])).
fof(f4422,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f533,f829])).
fof(f4421,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f506,f829])).
fof(f4419,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f465,f829])).
fof(f4366,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,sK4(X30,X31,bool))) ) | $spl10),
  inference(superposition,[],[f829,f2668])).
fof(f4365,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,X29) = s(bool,sK4(X27,X28,bool)) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f829,f2668])).
fof(f4531,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X26))) | s(bool,'F_0') = s(bool,sK4(X25,X26,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f4364,f2211])).
fof(f4364,plain,(
  ( ! [X26,X25] : (s(bool,X25) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X26))) | s(bool,'F_0') = s(bool,sK4(X25,X26,bool))) ) | $spl10),
  inference(superposition,[],[f829,f2294])).
fof(f4359,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f829,f438])).
fof(f4358,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f829,f438])).
fof(f4357,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f829,f265])).
fof(f4333,plain,(
  ( ! [X35,X33,X36,X34] : (s(X35,f(s(bool,X36))) != s(X35,g(s(bool,X36))) | s(fun(bool,X35),X33) = s(fun(bool,X35),X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK12(X33,X34,X35,bool))) ) | $spl10),
  inference(superposition,[],[f347,f2668])).
fof(f4332,plain,(
  ( ! [X30,X31,X29,X32] : (s(X31,f(s(bool,'F_0'))) != s(X31,g(s(bool,'F_0'))) | s(fun(bool,X31),X29) = s(fun(bool,X31),X30) | s(bool,X32) = s(bool,sK12(X29,X30,X31,bool)) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f347,f2668])).
fof(f4326,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK12(X7,X8,X9,bool))) )),
  inference(superposition,[],[f347,f438])).
fof(f4325,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK12(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f347,f438])).
fof(f4324,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,X2,bool))) )),
  inference(superposition,[],[f347,f265])).
fof(f4321,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4264,f256])).
fof(f4264,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f440,f533])).
fof(f4241,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X95),'const.sum.INL_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,X94)))))) )),
  inference(superposition,[],[f533,f334])).
fof(f4240,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X89,bool),'const.pair.,_2'(s(X89,X90),s(bool,X91)))))) )),
  inference(superposition,[],[f533,f338])).
fof(f4239,plain,(
  ( ! [X88,X87,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X88),'const.sum.INL_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X86,bool),'const.sum.INR_1'(s(bool,X87)))))) )),
  inference(superposition,[],[f533,f333])).
fof(f4238,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X85),'const.sum.INL_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X83),s(X82,X84)))))) )),
  inference(superposition,[],[f533,f337])).
fof(f4236,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X77),'const.sum.INL_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X75),s(fun('type.num.num',bool),X76),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f533,f311])).
fof(f4234,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK16(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f533,f353])).
fof(f4202,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4146,f256])).
fof(f4146,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f440,f506])).
fof(f4123,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f506,f334])).
fof(f4122,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X86,bool),'const.pair.,_2'(s(X86,X87),s(bool,X88)))))) )),
  inference(superposition,[],[f506,f338])).
fof(f4121,plain,(
  ( ! [X85,X83,X84] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X85,bool),'const.sum.INR_1'(s(bool,X84))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X83,bool),'const.sum.INR_1'(s(bool,X84)))))) )),
  inference(superposition,[],[f506,f333])).
fof(f4120,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X79),'const.pair.,_2'(s(bool,X80),s(X79,X81)))))) )),
  inference(superposition,[],[f506,f337])).
fof(f4118,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f506,f311])).
fof(f4116,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X68,bool),'const.sum.INR_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK16(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f506,f353])).
fof(f4025,plain,(
  ( ! [X90,X91,X89] : (s(bool,'T_0') != s(bool,X90) | s(bool,X91) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X90))))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f3517,f334])).
fof(f4024,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') != s(bool,X87) | s(bool,X88) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87))))) | s(bool,'F_0') = s(bool,X88)) ) | $spl10),
  inference(superposition,[],[f3517,f338])).
fof(f4023,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X82,bool),'const.sum.INR_1'(s(bool,X83))))) | s(bool,'F_0') = s(bool,X84)) ) | $spl10),
  inference(superposition,[],[f3517,f333])).
fof(f4022,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X81) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X79),s(X78,X80))))) | s(bool,'F_0') = s(bool,X81)) ) | $spl10),
  inference(superposition,[],[f3517,f337])).
fof(f4020,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X73) = s(bool,'const.ind_type.FCONS_3'(s(bool,X71),s(fun('type.num.num',bool),X72),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f3517,f311])).
fof(f4018,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) = s(bool,happ(s(fun('type.list.list'(X64),bool),sK16(X65,X66,X64,bool)),s('type.list.list'(X64),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl10),
  inference(superposition,[],[f3517,f353])).
fof(f4017,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,X63) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X62))) | s(bool,'F_0') = s(bool,X63)) ) | $spl10),
  inference(superposition,[],[f3517,f304])).
fof(f3984,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f3517,f438])).
fof(f3610,plain,(
  ( ! [X173,X174,X172] : (~p(s(bool,happ(s(fun(bool,bool),X173),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173)))) | s(bool,X174) = s(bool,sK4(X172,X173,bool)) | s(bool,'F_0') = s(bool,X174)) ) | $spl10),
  inference(superposition,[],[f314,f2668])).
fof(f3957,plain,(
  ( ! [X92,X90,X93,X91] : (s(X90,X92) = s(X90,sK4(X92,X91,X90)) | s(bool,X93) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))) | s(bool,'F_0') = s(bool,X93)) ) | $spl10),
  inference(subsumption_resolution,[],[f3584,f258])).
fof(f3584,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,'F_0')) | s(X90,X92) = s(X90,sK4(X92,X91,X90)) | s(bool,X93) = s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))) | s(bool,'F_0') = s(bool,X93)) ) | $spl10),
  inference(superposition,[],[f313,f2668])).
fof(f3951,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3552,f258])).
fof(f3552,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl10),
  inference(superposition,[],[f327,f2668])).
fof(f3517,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f425,f2668])).
fof(f3516,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl10),
  inference(superposition,[],[f270,f2668])).
fof(f3446,plain,(
  ( ! [X185,X186,X184] : (~p(s(bool,happ(s(fun(bool,bool),X185),s(bool,X186)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X185)))) | s(bool,'F_0') = s(bool,X186) | s(bool,'F_0') = s(bool,sK4(X184,X185,bool))) ) | $spl10),
  inference(superposition,[],[f314,f2668])).
fof(f3341,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X74),'const.sum.INL_1'(s(bool,X75))))) | s(bool,'F_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f2668,f334])).
fof(f3340,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl10),
  inference(superposition,[],[f2668,f338])).
fof(f3339,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') = s(bool,X68) | s(bool,X69) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl10),
  inference(superposition,[],[f2668,f333])).
fof(f3338,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X66) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X63),'const.pair.,_2'(s(bool,X64),s(X63,X65))))) | s(bool,'F_0') = s(bool,X66)) ) | $spl10),
  inference(superposition,[],[f2668,f337])).
fof(f3336,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.ind_type.FCONS_3'(s(bool,X56),s(fun('type.num.num',bool),X57),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl10),
  inference(superposition,[],[f2668,f311])).
fof(f3334,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') = s(bool,X51) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(X49),bool),sK16(X50,X51,X49,bool)),s('type.list.list'(X49),'const.list.NIL_0'))) | s(bool,'F_0') = s(bool,X52)) ) | $spl10),
  inference(superposition,[],[f2668,f353])).
fof(f3305,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f2668,f2668])).
fof(f3270,plain,(
  ( ! [X88,X87,X89] : (s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X87),'const.sum.INL_1'(s(bool,X88)))))) ) | $spl10),
  inference(superposition,[],[f2668,f334])).
fof(f3269,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,X85) = s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X85)))))) ) | $spl10),
  inference(superposition,[],[f2668,f338])).
fof(f3268,plain,(
  ( ! [X80,X81,X82] : (s(bool,X81) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X80,bool),'const.sum.INR_1'(s(bool,X81)))))) ) | $spl10),
  inference(superposition,[],[f2668,f333])).
fof(f3267,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X77) = s(bool,X79) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X77),s(X76,X78)))))) ) | $spl10),
  inference(superposition,[],[f2668,f337])).
fof(f3265,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f2668,f311])).
fof(f3263,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X62),bool),sK16(X63,X64,X62,bool)),s('type.list.list'(X62),'const.list.NIL_0')))) ) | $spl10),
  inference(superposition,[],[f2668,f353])).
fof(f3262,plain,(
  ( ! [X61,X60] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60)))) ) | $spl10),
  inference(superposition,[],[f2668,f304])).
fof(f3203,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X44),s('type.patLang.exp',X45))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X46),s('type.patLang.exp',X47))) | s('type.list.list'('type.patSem.v'),X46) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X44)))))) )),
  inference(superposition,[],[f391,f334])).
fof(f3202,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X39),s('type.patLang.exp',X40))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X41),s('type.patLang.exp',X42))) | s('type.list.list'('type.patSem.v'),X41) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.patSem.v'),X39)))))) )),
  inference(superposition,[],[f391,f338])).
fof(f3201,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X33),s('type.patLang.exp',X34))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X35),s('type.patLang.exp',X36))) | s('type.list.list'('type.patSem.v'),X35) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X32,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X33)))))) )),
  inference(superposition,[],[f391,f333])).
fof(f3200,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X27),s('type.patLang.exp',X29))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X30),s('type.patLang.exp',X31))) | s('type.list.list'('type.patSem.v'),X30) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X26),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X27),s(X26,X28)))))) )),
  inference(superposition,[],[f391,f337])).
fof(f3199,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X24),s('type.patLang.exp',X25))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num',X22))),s('type.patLang.exp',X23))) | s('type.list.list'('type.patSem.v'),X24) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X20),s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f391,f319])).
fof(f3198,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X15),s('type.patLang.exp',X17))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X18),s('type.patLang.exp',X19))) | s('type.list.list'('type.patSem.v'),X18) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X15),s(fun('type.num.num','type.list.list'('type.patSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f391,f311])).
fof(f3197,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X14))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11))),s('type.patLang.exp',X12))) | s('type.list.list'('type.patSem.v'),X13) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X10),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11)))) )),
  inference(superposition,[],[f391,f335])).
fof(f3196,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X6),s('type.patLang.exp',X7))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X9))) | s('type.list.list'('type.patSem.v'),X8) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.patSem.v')),sK16(X5,X6,X4,'type.list.list'('type.patSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f391,f353])).
fof(f3195,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X2) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f391,f304])).
fof(f2862,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X45),s('type.list.list'('type.patSem.v'),X44))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X46),s('type.list.list'('type.patSem.v'),X47))) | s('type.list.list'('type.patSem.v'),X47) = s('type.list.list'('type.patSem.v'),'const.sum.OUTL_1'(s('type.sum.sum'('type.list.list'('type.patSem.v'),X43),'const.sum.INL_1'(s('type.list.list'('type.patSem.v'),X44)))))) )),
  inference(superposition,[],[f389,f334])).
fof(f2861,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X40),s('type.list.list'('type.patSem.v'),X39))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X41),s('type.list.list'('type.patSem.v'),X42))) | s('type.list.list'('type.patSem.v'),X42) = s('type.list.list'('type.patSem.v'),'const.pair.SND_1'(s('type.pair.prod'(X37,'type.list.list'('type.patSem.v')),'const.pair.,_2'(s(X37,X38),s('type.list.list'('type.patSem.v'),X39)))))) )),
  inference(superposition,[],[f389,f338])).
fof(f2860,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X34),s('type.list.list'('type.patSem.v'),X33))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X35),s('type.list.list'('type.patSem.v'),X36))) | s('type.list.list'('type.patSem.v'),X36) = s('type.list.list'('type.patSem.v'),'const.sum.OUTR_1'(s('type.sum.sum'(X32,'type.list.list'('type.patSem.v')),'const.sum.INR_1'(s('type.list.list'('type.patSem.v'),X33)))))) )),
  inference(superposition,[],[f389,f333])).
fof(f2859,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X29),s('type.list.list'('type.patSem.v'),X27))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X30),s('type.list.list'('type.patSem.v'),X31))) | s('type.list.list'('type.patSem.v'),X31) = s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X26),'const.pair.,_2'(s('type.list.list'('type.patSem.v'),X27),s(X26,X28)))))) )),
  inference(superposition,[],[f389,f337])).
fof(f2858,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X24),s('type.list.list'('type.patSem.v'),X25))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X23),s('type.list.list'('type.patSem.v'),happ(s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num',X22))))) | s('type.list.list'('type.patSem.v'),X25) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X20),s(fun('type.num.num','type.list.list'('type.patSem.v')),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f389,f319])).
fof(f2857,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X15))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) | s('type.list.list'('type.patSem.v'),X19) = s('type.list.list'('type.patSem.v'),'const.ind_type.FCONS_3'(s('type.list.list'('type.patSem.v'),X15),s(fun('type.num.num','type.list.list'('type.patSem.v')),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f389,f311])).
fof(f2856,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X13),s('type.list.list'('type.patSem.v'),X14))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.patSem.v'),'const.pair.FST_1'(s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11))))) | s('type.list.list'('type.patSem.v'),X14) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.pair.prod'('type.list.list'('type.patSem.v'),X10),'type.list.list'('type.patSem.v')),'const.pair.FST_0'),s('type.pair.prod'('type.list.list'('type.patSem.v'),X10),X11)))) )),
  inference(superposition,[],[f389,f335])).
fof(f2855,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X7),s('type.list.list'('type.patSem.v'),X6))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X8),s('type.list.list'('type.patSem.v'),X9))) | s('type.list.list'('type.patSem.v'),X9) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'(X4),'type.list.list'('type.patSem.v')),sK16(X5,X6,X4,'type.list.list'('type.patSem.v'))),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f389,f353])).
fof(f2854,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.list.list'('type.patSem.v'),X3) = s('type.list.list'('type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),'type.list.list'('type.patSem.v')),'_i_0_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(superposition,[],[f389,f304])).
fof(f2830,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(forward_demodulation,[],[f2731,f2211])).
fof(f2731,plain,(
  ( ! [X101,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | s(bool,'F_0') = s(bool,sK4(X100,X101,bool))) ) | $spl10),
  inference(superposition,[],[f314,f2294])).
fof(f2814,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2813,f256])).
fof(f2813,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(forward_demodulation,[],[f2699,f2211])).
fof(f2699,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37)))) | s(bool,'F_0') = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK4(X38,X37,X36))))) ) | $spl10),
  inference(superposition,[],[f314,f2294])).
fof(f2668,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2294,f2294])).
fof(f2556,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X6),s('type.list.list'('type.patSem.v'),X7))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.list.list'('type.patSem.v'),X8))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X5)))))) )),
  inference(superposition,[],[f388,f264])).
fof(f2553,plain,(
  ( ! [X1] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X1))))) = s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X1)))))))) )),
  inference(superposition,[],[f264,f263])).
fof(f2399,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X212),s('type.patLang.exp',X213))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X214),s('type.patLang.exp',happ(s(fun('type.num.num','type.patLang.exp'),X210),s('type.num.num',X211))))) | s('type.patLang.exp',X213) = s('type.patLang.exp','const.ind_type.FCONS_3'(s('type.patLang.exp',X209),s(fun('type.num.num','type.patLang.exp'),X210),s('type.num.num','const.num.SUC_1'(s('type.num.num',X211)))))) )),
  inference(superposition,[],[f392,f319])).
fof(f2389,plain,(
  ( ! [X165,X163,X166,X164] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X166))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',happ(s(fun('type.num.num','type.ast.lit'),X164),s('type.num.num',X165))))) | s('type.ast.lit',X166) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X163),s(fun('type.num.num','type.ast.lit'),X164),s('type.num.num','const.num.SUC_1'(s('type.num.num',X165)))))) )),
  inference(superposition,[],[f331,f319])).
fof(f2383,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X140),s('type.list.list'('type.patSem.v'),X141))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X138),s('type.num.num',X139))),s('type.list.list'('type.patSem.v'),X142))) | s('type.num.num',X140) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X137),s(fun('type.num.num','type.num.num'),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) )),
  inference(superposition,[],[f388,f319])).
fof(f2421,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130)))))) | ~p(s(bool,happ(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130))),s(X127,sK4(X131,'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))),X127)))))) )),
  inference(forward_demodulation,[],[f2381,f319])).
fof(f2381,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,happ(s(fun(X127,bool),happ(s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num',X130))),s(X127,sK4(X131,'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))),X127))))) | p(s(bool,'|dollar|exists_1'(s(fun(X127,bool),'const.ind_type.FCONS_3'(s(fun(X127,bool),X128),s(fun('type.num.num',fun(X127,bool)),X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))))))))) )),
  inference(superposition,[],[f314,f319])).
fof(f2293,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f2265,f256])).
fof(f2265,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f440,f2211])).
fof(f2205,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f353,f2182,f440])).
fof(f2111,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2090,f256])).
fof(f2090,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'const.ind_type.ISO_2'(s(fun(X27,X27),'_i_0_0'),s(fun(X27,X27),'_i_0_0'))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f2071,f433])).
fof(f2110,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2089,f256])).
fof(f2089,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f2071,f465])).
fof(f2105,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2081,f256])).
fof(f2081,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2071,f438])).
fof(f2074,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f2042,f256])).
fof(f2042,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f440,f465])).
fof(f2073,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2072,f304])).
fof(f2072,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2039,f256])).
fof(f2039,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f411,f465])).
fof(f2069,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f2024,f304])).
fof(f2024,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38)))))) )),
  inference(superposition,[],[f465,f334])).
fof(f2068,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(forward_demodulation,[],[f2023,f304])).
fof(f2023,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X34,bool),'const.pair.,_2'(s(X34,X35),s(bool,X36)))))) )),
  inference(superposition,[],[f465,f338])).
fof(f2067,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(forward_demodulation,[],[f2022,f304])).
fof(f2022,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33)))))) )),
  inference(superposition,[],[f465,f333])).
fof(f2066,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X29),'const.pair.,_2'(s(bool,X30),s(X29,X31)))))) )),
  inference(forward_demodulation,[],[f2021,f304])).
fof(f2021,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X29),'const.pair.,_2'(s(bool,X30),s(X29,X31)))))) )),
  inference(superposition,[],[f465,f337])).
fof(f2065,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2020,f304])).
fof(f2020,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X27),s(fun('type.num.num',bool),X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f465,f311])).
fof(f2064,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f2018,f304])).
fof(f2018,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f465,f353])).
fof(f1984,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X39),s('type.patLang.exp',X38))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X40),s('type.patLang.exp',X41))) | s('type.patLang.exp',X41) = s('type.patLang.exp','const.sum.OUTL_1'(s('type.sum.sum'('type.patLang.exp',X37),'const.sum.INL_1'(s('type.patLang.exp',X38)))))) )),
  inference(superposition,[],[f392,f334])).
fof(f1983,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X34),s('type.patLang.exp',X33))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X35),s('type.patLang.exp',X36))) | s('type.patLang.exp',X36) = s('type.patLang.exp','const.pair.SND_1'(s('type.pair.prod'(X31,'type.patLang.exp'),'const.pair.,_2'(s(X31,X32),s('type.patLang.exp',X33)))))) )),
  inference(superposition,[],[f392,f338])).
fof(f1982,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X28),s('type.patLang.exp',X27))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X29),s('type.patLang.exp',X30))) | s('type.patLang.exp',X30) = s('type.patLang.exp','const.sum.OUTR_1'(s('type.sum.sum'(X26,'type.patLang.exp'),'const.sum.INR_1'(s('type.patLang.exp',X27)))))) )),
  inference(superposition,[],[f392,f333])).
fof(f1981,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X23),s('type.patLang.exp',X21))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X24),s('type.patLang.exp',X25))) | s('type.patLang.exp',X25) = s('type.patLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.patLang.exp',X20),'const.pair.,_2'(s('type.patLang.exp',X21),s(X20,X22)))))) )),
  inference(superposition,[],[f392,f337])).
fof(f1980,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X17),s('type.patLang.exp',X15))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X18),s('type.patLang.exp',X19))) | s('type.patLang.exp',X19) = s('type.patLang.exp','const.ind_type.FCONS_3'(s('type.patLang.exp',X15),s(fun('type.num.num','type.patLang.exp'),X16),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f392,f311])).
fof(f1979,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X14))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X12),s('type.patLang.exp','const.pair.FST_1'(s('type.pair.prod'('type.patLang.exp',X10),X11))))) | s('type.patLang.exp',X14) = s('type.patLang.exp',happ(s(fun('type.pair.prod'('type.patLang.exp',X10),'type.patLang.exp'),'const.pair.FST_0'),s('type.pair.prod'('type.patLang.exp',X10),X11)))) )),
  inference(superposition,[],[f392,f335])).
fof(f1978,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X7),s('type.patLang.exp',X6))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X9))) | s('type.patLang.exp',X9) = s('type.patLang.exp',happ(s(fun('type.list.list'(X4),'type.patLang.exp'),sK16(X5,X6,X4,'type.patLang.exp')),s('type.list.list'(X4),'const.list.NIL_0')))) )),
  inference(superposition,[],[f392,f353])).
fof(f1977,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.patLang.exp',X3) = s('type.patLang.exp',happ(s(fun('type.patLang.exp','type.patLang.exp'),'_i_0_0'),s('type.patLang.exp',X0)))) )),
  inference(superposition,[],[f392,f304])).
fof(f1851,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X42),s('type.list.list'('type.patSem.v'),X43))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X44),s('type.list.list'('type.patSem.v'),X45))) | s('type.num.num',X44) = s('type.num.num','const.sum.OUTL_1'(s('type.sum.sum'('type.num.num',X41),'const.sum.INL_1'(s('type.num.num',X42)))))) )),
  inference(superposition,[],[f388,f334])).
fof(f1850,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X37),s('type.list.list'('type.patSem.v'),X38))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X39),s('type.list.list'('type.patSem.v'),X40))) | s('type.num.num',X39) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X35,'type.num.num'),'const.pair.,_2'(s(X35,X36),s('type.num.num',X37)))))) )),
  inference(superposition,[],[f388,f338])).
fof(f1849,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X31),s('type.list.list'('type.patSem.v'),X32))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X33),s('type.list.list'('type.patSem.v'),X34))) | s('type.num.num',X33) = s('type.num.num','const.sum.OUTR_1'(s('type.sum.sum'(X30,'type.num.num'),'const.sum.INR_1'(s('type.num.num',X31)))))) )),
  inference(superposition,[],[f388,f333])).
fof(f1848,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X25),s('type.list.list'('type.patSem.v'),X27))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X28),s('type.list.list'('type.patSem.v'),X29))) | s('type.num.num',X28) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X24),'const.pair.,_2'(s('type.num.num',X25),s(X24,X26)))))) )),
  inference(superposition,[],[f388,f337])).
fof(f1847,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X19),s('type.list.list'('type.patSem.v'),X21))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X22),s('type.list.list'('type.patSem.v'),X23))) | s('type.num.num',X22) = s('type.num.num','const.ind_type.FCONS_3'(s('type.num.num',X19),s(fun('type.num.num','type.num.num'),X20),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f388,f311])).
fof(f1846,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X14),X15))),s('type.list.list'('type.patSem.v'),X16))) | s('type.num.num',X17) = s('type.num.num',happ(s(fun('type.pair.prod'('type.num.num',X14),'type.num.num'),'const.pair.FST_0'),s('type.pair.prod'('type.num.num',X14),X15)))) )),
  inference(superposition,[],[f388,f335])).
fof(f1845,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X10),s('type.list.list'('type.patSem.v'),X11))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X12),s('type.list.list'('type.patSem.v'),X13))) | s('type.num.num',X12) = s('type.num.num',happ(s(fun('type.list.list'(X8),'type.num.num'),sK16(X9,X10,X8,'type.num.num')),s('type.list.list'(X8),'const.list.NIL_0')))) )),
  inference(superposition,[],[f388,f353])).
fof(f1844,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X4),s('type.list.list'('type.patSem.v'),X5))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X6),s('type.list.list'('type.patSem.v'),X7))) | s('type.num.num',X6) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'_i_0_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f388,f304])).
fof(f1512,plain,(
  ( ! [X99,X97,X98] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X99))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X97),X98))))) | s('type.ast.lit',X99) = s('type.ast.lit',happ(s(fun('type.pair.prod'('type.ast.lit',X97),'type.ast.lit'),'const.pair.FST_0'),s('type.pair.prod'('type.ast.lit',X97),X98)))) )),
  inference(superposition,[],[f331,f335])).
fof(f1526,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1507,f335])).
fof(f1507,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f314,f335])).
fof(f1490,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f314,f335])).
fof(f1385,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1372,f256])).
fof(f1372,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f314,f438])).
fof(f1371,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f314,f265])).
fof(f1369,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f314,f438])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f314,f438])).
fof(f1384,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26)))))) )),
  inference(forward_demodulation,[],[f1366,f334])).
fof(f1366,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28))))))))) )),
  inference(superposition,[],[f314,f334])).
fof(f1383,plain,(
  ( ! [X24,X23,X21,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24)))) | ~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21)))))) )),
  inference(forward_demodulation,[],[f1365,f338])).
fof(f1365,plain,(
  ( ! [X24,X23,X21,X25,X22] : (~p(s(bool,happ(s(fun(X21,bool),X24),s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))))) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24))))))))) )),
  inference(superposition,[],[f314,f338])).
fof(f1382,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | ~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK4(X20,'const.sum.OUTR_1'(s('type.sum.sum'(X18,fun(X17,bool)),'const.sum.INR_1'(s(fun(X17,bool),X19)))),X17)))))) )),
  inference(forward_demodulation,[],[f1364,f333])).
fof(f1364,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK4(X20,'const.sum.OUTR_1'(s('type.sum.sum'(X18,fun(X17,bool)),'const.sum.INR_1'(s(fun(X17,bool),X19)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X18,fun(X17,bool)),'const.sum.INR_1'(s(fun(X17,bool),X19))))))))) )),
  inference(superposition,[],[f314,f333])).
fof(f1381,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f1363,f337])).
fof(f1363,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f314,f337])).
fof(f1380,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1362,f311])).
fof(f1362,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f314,f311])).
fof(f1379,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1361,f353])).
fof(f1361,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f314,f353])).
fof(f1378,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1360,f304])).
fof(f1360,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f314,f304])).
fof(f1353,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1110,f256])).
fof(f1110,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f440,f438])).
fof(f1352,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1109,f256])).
fof(f1109,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f411,f438])).
fof(f1351,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1108,f256])).
fof(f1108,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f330,f438])).
fof(f999,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f438,f334])).
fof(f998,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f438,f338])).
fof(f997,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X28,bool),'const.sum.INR_1'(s(bool,X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f438,f333])).
fof(f996,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,X25),s(X24,X26))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f438,f337])).
fof(f995,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f438,f311])).
fof(f994,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK16(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f438,f353])).
fof(f993,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f438,f304])).
fof(f983,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f438,f438])).
fof(f966,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43)))))) )),
  inference(superposition,[],[f438,f334])).
fof(f965,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X39),s(bool,X40)))))) )),
  inference(superposition,[],[f438,f338])).
fof(f964,plain,(
  ( ! [X37,X35,X36] : (s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,X36)))))) )),
  inference(superposition,[],[f438,f333])).
fof(f963,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))))) )),
  inference(superposition,[],[f438,f337])).
fof(f962,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f438,f311])).
fof(f961,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK16(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f438,f353])).
fof(f960,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f438,f304])).
fof(f927,plain,(
  ( ! [X103,X105,X106,X104] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X105))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X106))) | s('type.ast.lit',X106) = s('type.ast.lit',happ(s(fun('type.list.list'(X103),'type.ast.lit'),sK16(X104,X105,X103,'type.ast.lit')),s('type.list.list'(X103),'const.list.NIL_0')))) )),
  inference(superposition,[],[f331,f353])).
fof(f902,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,X2,X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f353,f265])).
fof(f835,plain,(
  ( ! [X17,X18,X16] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X18))) | s('type.ast.lit',X18) = s('type.ast.lit','const.sum.OUTL_1'(s('type.sum.sum'('type.ast.lit',X16),'const.sum.INL_1'(s('type.ast.lit',X17)))))) )),
  inference(superposition,[],[f331,f334])).
fof(f834,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X14))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X15))) | s('type.ast.lit',X15) = s('type.ast.lit','const.pair.SND_1'(s('type.pair.prod'(X12,'type.ast.lit'),'const.pair.,_2'(s(X12,X13),s('type.ast.lit',X14)))))) )),
  inference(superposition,[],[f331,f338])).
fof(f833,plain,(
  ( ! [X10,X11,X9] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X10))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X11))) | s('type.ast.lit',X11) = s('type.ast.lit','const.sum.OUTR_1'(s('type.sum.sum'(X9,'type.ast.lit'),'const.sum.INR_1'(s('type.ast.lit',X10)))))) )),
  inference(superposition,[],[f331,f333])).
fof(f832,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X6))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X8))) | s('type.ast.lit',X8) = s('type.ast.lit','const.pair.FST_1'(s('type.pair.prod'('type.ast.lit',X5),'const.pair.,_2'(s('type.ast.lit',X6),s(X5,X7)))))) )),
  inference(superposition,[],[f331,f337])).
fof(f831,plain,(
  ( ! [X4,X2,X3] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X4))) | s('type.ast.lit',X4) = s('type.ast.lit','const.ind_type.FCONS_3'(s('type.ast.lit',X2),s(fun('type.num.num','type.ast.lit'),X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f331,f311])).
fof(f830,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X1) = s('type.ast.lit',happ(s(fun('type.ast.lit','type.ast.lit'),'_i_0_0'),s('type.ast.lit',X0)))) )),
  inference(superposition,[],[f331,f304])).
fof(f829,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f828,f258])).
fof(f828,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f313,f265])).
fof(f796,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f783,f256])).
fof(f783,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f440,f433])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) )),
  inference(superposition,[],[f338,f265])).
fof(f679,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) )),
  inference(superposition,[],[f337,f265])).
fof(f626,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f311,f265])).
fof(f533,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f334,f265])).
fof(f506,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f333,f265])).
fof(f487,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f482,f258])).
fof(f482,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f327,f265])).
fof(f440,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f265])).
fof(f433,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f259,f267])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f253])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patLang.exp'),X1) = s('type.list.list'('type.patLang.exp'),X4)) )),
  inference(cnf_transformation,[],[f253])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'('type.patLang.exp'),X1) != s('type.list.list'('type.patLang.exp'),X4) | s('type.num.num',X2) != s('type.num.num',X5) | s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X0),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X5))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f253])).
fof(f394,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X25,X26,X46,X27] : ! [X28,X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X46,X25),X28),s(fun(X26,X46),happ(s(fun(X27,fun(X26,X46)),X29),s(X27,X30))))) = s(fun(X26,X25),happ(s(fun(X27,fun(X26,X25)),'_dst_x00x11_2'(s(fun(X46,X25),X28),s(fun(X27,fun(X26,X46)),X29))),s(X27,X30)))),
  file('thm.patSem.v_|49||49||split|4','_dst_x00x11_ax')).
fof(f391,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3))) & (s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2) | s('type.patLang.exp',X1) != s('type.patLang.exp',X3) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : ((s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3))) & ((s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2) | s('type.patLang.exp',X1) != s('type.patLang.exp',X3)) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3)))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) <=> (s('type.list.list'('type.patSem.v'),X0) = s('type.list.list'('type.patSem.v'),X2) & s('type.patLang.exp',X1) = s('type.patLang.exp',X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X16,X17,X19,X20] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X16),s('type.patLang.exp',X17))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X19),s('type.patLang.exp',X20))) <=> (s('type.list.list'('type.patSem.v'),X16) = s('type.list.list'('type.patSem.v'),X19) & s('type.patLang.exp',X17) = s('type.patLang.exp',X20)))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_|49||49||split|2')).
fof(f392,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.patLang.exp',X1) = s('type.patLang.exp',X3)) )),
  inference(cnf_transformation,[],[f251])).
fof(f393,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patLang.exp',X1) != s('type.patLang.exp',X3) | s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X3))) | s('type.list.list'('type.patSem.v'),X0) != s('type.list.list'('type.patSem.v'),X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.list.list'('type.patSem.v'),X1) = s('type.list.list'('type.patSem.v'),X3)) )),
  inference(cnf_transformation,[],[f249])).
fof(f390,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'('type.patSem.v'),X1) != s('type.list.list'('type.patSem.v'),X3) | s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.num.num',X0) != s('type.num.num',X2)) )),
  inference(cnf_transformation,[],[f249])).
fof(f411,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f371])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X45,X49,X38,X50] : ((s(bool,X45) = s(bool,X49) & (p(s(bool,X49)) => s(bool,X38) = s(bool,X50))) => ((p(s(bool,X45)) => p(s(bool,X38))) <=> (p(s(bool,X49)) => p(s(bool,X50)))))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.IMP_CONG')).
fof(f363,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X25,X26,X27] : ! [X28,X29,X30] : s(X25,happ(s(fun(X26,X25),X28),s(X26,happ(s(fun(X27,X26),X29),s(X27,X30))))) = s(X25,happ(s(fun(X27,X25),'_dst_x0x1_2'(s(fun(X26,X25),X28),s(fun(X27,X26),X29))),s(X27,X30)))),
  file('thm.patSem.v_|49||49||split|4','_dst_x0x1_ax')).
fof(f362,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X25,X26,X27] : ! [X29,X30] : s(fun(X26,X25),'_dst_x0x1_2'(s(fun(X27,X25),X30),s(fun(X26,X27),X29))) = s(fun(X26,X25),happ(s(fun(fun(X27,X25),fun(X26,X25)),'_dstsel_0_x01_1'(s(fun(X26,X27),X29))),s(fun(X27,X25),X30)))),
  file('thm.patSem.v_|49||49||split|4','_dstsel_0_x01_ax')).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X37,X9] : ! [X39,X11] : ? [X40] : (s(X37,X39) = s(X37,happ(s(fun('type.list.list'(X9),X37),X40),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X16,X17] : s(X37,happ(s(fun('type.list.list'(X9),X37),X40),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X16),s('type.list.list'(X9),X17))))) = s(X37,happ(s(fun(X37,X37),happ(s(fun('type.list.list'(X9),fun(X37,X37)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X37,X37))),X11),s(X9,X16))),s('type.list.list'(X9),X17))),s(X37,happ(s(fun('type.list.list'(X9),X37),X40),s('type.list.list'(X9),X17))))))),
  file('thm.patSem.v_|49||49||split|4','thm.list.list_Axiom')).
fof(f354,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f348,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X37,X9] : ! [X10,X54] : (s(fun(X9,X37),X10) = s(fun(X9,X37),X54) <=> ! [X45] : s(X37,happ(s(fun(X9,X37),X10),s(X9,X45))) = s(X37,happ(s(fun(X9,X37),X54),s(X9,X45))))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.FUN_EQ_THM')).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f238])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.patSem.v_|49||49||split|4',aHL_EXT)).
fof(f338,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X37] : ! [X45,X38] : s(X37,X38) = s(X37,'const.pair.SND_1'(s('type.pair.prod'(X9,X37),'const.pair.,_2'(s(X9,X45),s(X37,X38)))))),
  file('thm.patSem.v_|49||49||split|4','thm.pair.SND')).
fof(f337,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X37,X9] : ! [X45,X38] : s(X9,X45) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X37),'const.pair.,_2'(s(X9,X45),s(X37,X38)))))),
  file('thm.patSem.v_|49||49||split|4','thm.pair.FST')).
fof(f336,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X37,X9] : ! [X41] : ? [X10] : ! [X42,X43,X44] : s(X37,happ(s(fun('type.ind_type.recspace'(X9),X37),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X42),s(X9,X43),s(fun('type.num.num','type.ind_type.recspace'(X9)),X44))))) = s(X37,happ(s(fun(fun('type.num.num',X37),X37),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X37),X37)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X37),X37))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X37),X37)))),X41),s('type.num.num',X42))),s(X9,X43))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X44))),s(fun('type.num.num',X37),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X37),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X44)))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.CONSTR_REC')).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9,X37] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X37),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X37),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X37),X8)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_1_const.pair.FST')).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X9,X37] : ! [X45] : s(X37,X45) = s(X37,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X37),'const.sum.INR_1'(s(X37,X45)))))),
  file('thm.patSem.v_|49||49||split|4','thm.sum.OUTR')).
fof(f331,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | s('type.ast.lit',X0) = s('type.ast.lit',X1)) & (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) <=> s('type.ast.lit',X0) = s('type.ast.lit',X1))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X6))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) <=> s('type.ast.lit',X6) = s('type.ast.lit',X7))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_|49||49||split|0')).
fof(f332,plain,(
  ( ! [X0,X1] : (s('type.ast.lit',X0) != s('type.ast.lit',X1) | s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f230])).
fof(f327,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f326,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0))),s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X1)))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0))),s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X1)))))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X16,X17] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X16),s('type.patLang.exp',X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X16))),s('type.num.num','const.patLang.exp_size_1'(s('type.patLang.exp',X17)))))))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_size_def|split|2')).
fof(f325,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X1)))))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X0),s('type.list.list'('type.patSem.v'),X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X1)))))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X16,X17] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X16),s('type.list.list'('type.patSem.v'),X17))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X17)))))))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_size_def|split|1')).
fof(f323,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.patLang.exp',X1)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X0),s('type.patLang.exp',X1))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X0))),s('type.patLang.exp',X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8,X22] : s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X8),s('type.patLang.exp',X22))) = s('type.patSem.v',happ(s(fun('type.patLang.exp','type.patSem.v'),happ(s(fun('type.list.list'('type.patSem.v'),fun('type.patLang.exp','type.patSem.v')),'const.patSem.Closure_0'),s('type.list.list'('type.patSem.v'),X8))),s('type.patLang.exp',X22)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_2_const.patSem.Closure')).
fof(f321,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9] : ! [X8,X22,X23] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X22),s('type.num.num',X23))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X22))),s('type.num.num',X23)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_3_const.ind_type.FCONS')).
fof(f320,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X9] : ! [X8,X22,X23] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X22),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X22))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X23)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_1_3_const.ind_type.CONSTR')).
fof(f319,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X6,X10,X36] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X36))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.FCONS|split|1')).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.FCONS|split|0')).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.patSem.v_|49||49||split|4','thm.bool.EQ_CLAUSES|split|2')).
fof(f271,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f267,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f197])).
fof(f264,plain,(
  ( ! [X0] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X0)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.patSem.v|49|_size_1'(s('type.list.list'('type.patSem.v'),X6)))))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_size_def|split|5')).
fof(f263,plain,(
  ( ! [X0] : (s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.patSem.v_size_1'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_size_def|split|4')).
fof(f262,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X8))) = s('type.patSem.v',happ(s(fun('type.list.list'('type.patSem.v'),'type.patSem.v'),'const.patSem.Vectorv_0'),s('type.list.list'('type.patSem.v'),X8)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_1_const.patSem.Vectorv')).
fof(f261,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),'const.patSem.Loc_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),'const.patSem.Loc_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X8))) = s('type.patSem.v',happ(s(fun('type.num.num','type.patSem.v'),'const.patSem.Loc_0'),s('type.num.num',X8)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_1_const.patSem.Loc')).
fof(f260,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X8))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X8)))),
  file('thm.patSem.v_|49||49||split|4','ARITY_0_1_const.patSem.Litv')).
fof(f259,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.patSem.v_|49||49||split|4','thm.ind_type.ISO_REFL')).
fof(f6072,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f5439,f5332])).
fof(f5332,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f5294,f258])).
fof(f5294,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'))) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f4920,f2668])).
fof(f5439,plain,(
  ( ! [X23,X22] : (s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(subsumption_resolution,[],[f5394,f256])).
fof(f5394,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X22),bool),'const.pair.FST_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl14),
  inference(superposition,[],[f440,f5284])).
fof(f5284,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f425,f4920,f2071])).
fof(f9311,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl22),
  inference(cnf_transformation,[],[f9311_D])).
fof(f9311_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f17185,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f17025])).
fof(f17025,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X103,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6072,f10201])).
fof(f10201,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10191])).
fof(f10191,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9376,f438])).
fof(f9376,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2182,f2211,f2205,f2293,f2294,f2421,f2383,f2389,f2399,f2553,f2556,f2778,f2668,f2814,f2830,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f3262,f3263,f3265,f3267,f3268,f3269,f3270,f3305,f3334,f3336,f3338,f3339,f3340,f3341,f3446,f3516,f3517,f3951,f3957,f3610,f3984,f4017,f4018,f4020,f4022,f4023,f4024,f4025,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4332,f4333,f4357,f4358,f4359,f4531,f4365,f4366,f4419,f4421,f4422,f4533,f4437,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4755,f4889,f4913,f4710,f4607,f4680,f4685,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2816,f2184,f343,f342,f316,f315,f312,f4087,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4344,f4345,f4914,f5260,f4608,f4235,f4117,f4019,f3335,f3264,f2780,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f4021,f3337,f3266,f2782,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5712,f5714,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6661,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6922,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f7962,f7964,f7969,f8065,f8073,f8411,f8422,f8846,f8851,f8853,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9361,f9363,f9365,f9366,f9368,f9369,f9372,f9373,f9375])).
fof(f9375,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl10),
  inference(condensation,[],[f9374])).
fof(f9374,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(subsumption_resolution,[],[f6600,f3517])).
fof(f6600,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6367,f2668])).
fof(f16844,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) | $spl36),
  inference(cnf_transformation,[],[f16844_D])).
fof(f16844_D,plain,(
  ( ! [X103,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f16769,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104)) ) | ($spl10 | $spl14)),
  inference(duplicate_literal_removal,[],[f16613])).
fof(f16613,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f6072,f10034])).
fof(f10034,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f10024])).
fof(f10024,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f9369,f438])).
fof(f13434,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl32),
  inference(cnf_transformation,[],[f13434_D])).
fof(f13434_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13436,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl34),
  inference(cnf_transformation,[],[f13436_D])).
fof(f13436_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13419,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl30),
  inference(superposition,[],[f13373,f304])).
fof(f2186,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl12),
  inference(cnf_transformation,[],[f2186_D])).
fof(f2186_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f13373,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl30),
  inference(cnf_transformation,[],[f13373_D])).
fof(f13373_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f13379,plain,(
  ( ! [X59,X60] : (p(s(bool,sK5('_i_0_0',X60,bool))) | ~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool)))))) )),
  inference(forward_demodulation,[],[f13348,f304])).
fof(f13348,plain,(
  ( ! [X59,X60] : (~p(s(bool,X59)) | p(s(bool,happ(s(fun(bool,bool),X60),s(bool,sK6('_i_0_0',X60,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X60,bool)))))) )),
  inference(superposition,[],[f315,f304])).
fof(f2180,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2180_D])).
fof(f2180_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13372,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f13371,f304])).
fof(f13371,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f13335,f256])).
fof(f13335,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f315,f465])).
fof(f9682,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2421,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4419,f4421,f4422,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914,f5260,f4608,f4235,f4117,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f8411,f8422,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9365,f9366,f9372,f9373,f5632,f5652,f4606,f4626,f2360,f9409,f9310,f9633,f9675,f9677,f9679,f9681])).
fof(f9681,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9589,f256])).
fof(f9589,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9310,f351])).
fof(f9679,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f9588,f256])).
fof(f9588,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl20),
  inference(backward_demodulation,[],[f9310,f350])).
fof(f9677,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(subsumption_resolution,[],[f9587,f256])).
fof(f9587,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl20),
  inference(backward_demodulation,[],[f9310,f346])).
fof(f9675,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(subsumption_resolution,[],[f9586,f256])).
fof(f9586,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl20),
  inference(backward_demodulation,[],[f9310,f345])).
fof(f9633,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f9412,f9310])).
fof(f9412,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f9310,f270])).
fof(f2360,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f265,f319])).
fof(f9680,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2421,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4419,f4421,f4422,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914,f5260,f4608,f4235,f4117,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f8411,f8422,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9365,f9366,f9372,f9373,f5632,f5652,f4606,f4626,f2360,f9409,f9310,f9633,f9675,f9677,f9679])).
fof(f9678,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2421,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4419,f4421,f4422,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914,f5260,f4608,f4235,f4117,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f8411,f8422,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9365,f9366,f9372,f9373,f5632,f5652,f4606,f4626,f2360,f9409,f9310,f9633,f9675,f9677])).
fof(f9676,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2421,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4419,f4421,f4422,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914,f5260,f4608,f4235,f4117,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f8411,f8422,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9365,f9366,f9372,f9373,f5632,f5652,f4606,f4626,f2360,f9409,f9310,f9633,f9675])).
fof(f9642,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9548,f9310])).
fof(f9548,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9310,f9376])).
fof(f9641,plain,(
  $false | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9546,f9310])).
fof(f9546,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'T_0')))))) ) | ($spl10 | $spl20)),
  inference(backward_demodulation,[],[f9310,f9369])).
fof(f9634,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f267,f266,f271,f270,f304,f311,f314,f319,f320,f321,f322,f323,f324,f325,f326,f330,f327,f332,f331,f333,f334,f335,f336,f337,f338,f347,f349,f348,f354,f353,f355,f362,f363,f411,f390,f389,f388,f393,f392,f391,f394,f399,f398,f397,f396,f425,f433,f438,f440,f465,f487,f506,f533,f626,f679,f736,f796,f829,f830,f831,f832,f833,f834,f835,f902,f927,f960,f961,f962,f963,f964,f965,f966,f983,f993,f994,f995,f996,f997,f998,f999,f1351,f1352,f1353,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1367,f1369,f1371,f1385,f1377,f1490,f1526,f1512,f1844,f1845,f1846,f1847,f1848,f1849,f1850,f1851,f1977,f1978,f1979,f1980,f1981,f1982,f1983,f1984,f2064,f2065,f2066,f2067,f2068,f2069,f2071,f2073,f2074,f2105,f2110,f2111,f2179,f2185,f2421,f2383,f2389,f2399,f2553,f2556,f2854,f2855,f2856,f2857,f2858,f2859,f2860,f2861,f2862,f3195,f3196,f3197,f3198,f3199,f3200,f3201,f3202,f3203,f4116,f4118,f4120,f4121,f4122,f4123,f4202,f4234,f4236,f4238,f4239,f4240,f4241,f4321,f4324,f4325,f4326,f4357,f4358,f4359,f4419,f4421,f4422,f4441,f4447,f4453,f4460,f4462,f4470,f4537,f4666,f4709,f4703,f4889,f4913,f4710,f4607,f4680,f4371,f4373,f4378,f4377,f4376,f4375,f4374,f4372,f4370,f4369,f395,f352,f351,f350,f346,f345,f2164,f4075,f4077,f344,f341,f4628,f4532,f2184,f343,f342,f316,f315,f312,f4082,f4081,f4080,f4079,f4078,f4076,f4074,f4073,f413,f306,f4336,f4337,f4338,f4914,f5260,f4608,f4235,f4117,f2187,f2019,f1481,f1479,f1477,f4237,f4119,f2374,f2348,f2346,f4451,f4783,f4443,f4440,f4548,f4426,f4425,f4420,f4418,f4413,f4412,f4410,f4409,f4408,f4407,f4406,f4442,f827,f826,f825,f824,f2382,f823,f1508,f923,f822,f4536,f4535,f313,f5700,f5710,f5711,f5715,f5719,f5722,f5723,f5690,f5691,f5692,f5693,f5694,f5695,f5696,f5697,f5698,f5699,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6276,f6277,f6278,f6279,f6280,f6282,f6284,f6367,f6375,f6391,f6401,f6436,f6479,f6480,f6481,f6482,f6483,f6485,f6487,f6571,f6579,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6676,f6686,f6721,f6884,f6885,f6886,f6887,f6888,f6889,f6890,f6891,f6892,f7335,f7346,f8411,f8422,f8943,f8944,f8945,f8946,f8947,f8948,f8949,f8950,f8951,f8952,f9365,f9366,f9372,f9373,f5632,f5652,f4606,f4626,f2360,f9409,f9310,f9633])).
fof(f9632,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f9411,f256])).
fof(f9411,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f9310,f258])).
fof(f9410,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl28),
  inference(cnf_transformation,[],[f9410_D])).
fof(f9410_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1387,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1387_D])).
fof(f1387_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f9409,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f2030,f256])).
fof(f2030,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f314,f465])).
fof(f4626,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4593,f256])).
fof(f4593,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4532,f465])).
fof(f4606,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4532,f304])).
fof(f5652,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f5619,f256])).
fof(f5619,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4532,f465])).
fof(f5632,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4532,f304])).
fof(f9364,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f9364_D])).
fof(f9364_D,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9373,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f6617,f6367])).
fof(f6617,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f6367,f2067])).
fof(f9372,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f6619,f6571])).
fof(f6619,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6367,f2069])).
fof(f9362,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl24),
  inference(cnf_transformation,[],[f9362_D])).
fof(f9362_D,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9366,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f6877,f6367])).
fof(f6877,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X70,bool),'const.sum.INR_1'(s(bool,X71))))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f6571,f2067])).
fof(f9365,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f6879,f6571])).
fof(f6879,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6571,f2069])).
fof(f9363,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(subsumption_resolution,[],[f8020,f8065])).
fof(f8020,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f6367,f2778])).
fof(f9361,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(subsumption_resolution,[],[f8021,f8065])).
fof(f8021,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f6571,f2778])).
fof(f9246,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9113])).
fof(f9113,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2478,f6804])).
fof(f6804,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f6790])).
fof(f6790,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f6659,f438])).
fof(f6659,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f6587,f1402])).
fof(f1402,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1388,f271])).
fof(f6587,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1403,f6367])).
fof(f1403,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1388,f266])).
fof(f2478,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2477,f258])).
fof(f2477,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2443,f1402])).
fof(f2443,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2073,f1426])).
fof(f1426,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1404,f256])).
fof(f1404,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1388,f438])).
fof(f5717,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X81,bool))))) ) | $spl18),
  inference(cnf_transformation,[],[f5717_D])).
fof(f5717_D,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X81,bool))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5716,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,X80)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X81,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f5682,f270])).
fof(f5682,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,X80)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X81),bool),sK16(X82,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X81,bool)))) | s(bool,'F_0') = s(bool,X80)) ) | $spl6),
  inference(superposition,[],[f4680,f2112])).
fof(f2112,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl6),
  inference(subsumption_resolution,[],[f2091,f256])).
fof(f2091,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl6),
  inference(superposition,[],[f2071,f1603])).
fof(f1603,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1550,f267])).
fof(f1550,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1403,f313])).
fof(f5713,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X71),bool),sK16(X72,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X71,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f5713_D])).
fof(f5713_D,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X71),bool),sK16(X72,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X71,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f5712,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X71),bool),sK16(X72,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X71,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f5679,f270])).
fof(f5679,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X71),bool),sK16(X72,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X71,bool)))) | s(bool,'F_0') = s(bool,X70)) ) | $spl10),
  inference(superposition,[],[f4680,f2294])).
fof(f5259,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4965,f258])).
fof(f4965,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f425,f2186,f327])).
fof(f5258,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4967,f2186])).
fof(f4967,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1403,f2186,f327])).
fof(f5257,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4975,f258])).
fof(f4975,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f425,f2186,f327])).
fof(f5256,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4977,f2186])).
fof(f4977,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1403,f2186,f327])).
fof(f5255,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5247,f258])).
fof(f5247,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(backward_demodulation,[],[f5038,f4710])).
fof(f5038,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(unit_resulting_resolution,[],[f334,f2186,f2073])).
fof(f5160,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5159,f425])).
fof(f5159,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5041,f433])).
fof(f5041,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f433,f2186,f2073])).
fof(f5158,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5157,f425])).
fof(f5157,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f5042,f1603])).
fof(f5042,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1603,f2186,f2073])).
fof(f5156,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5155,f425])).
fof(f5155,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl12),
  inference(forward_demodulation,[],[f5043,f4709])).
fof(f5043,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f4709,f2186,f2073])).
fof(f5154,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f5153,f425])).
fof(f5153,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f5044,f2211])).
fof(f5044,plain,(
  s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f2211,f2186,f2073])).
fof(f5062,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f2186,f4607])).
fof(f5060,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1403,f2186,f313])).
fof(f5056,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f433,f425,f2186,f2073])).
fof(f5026,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1403,f1403,f2186,f487])).
fof(f5024,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f425,f425,f2186,f487])).
fof(f5002,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f334,f2186,f411])).
fof(f5000,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f338,f2186,f411])).
fof(f4998,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f333,f2186,f411])).
fof(f4996,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f337,f2186,f411])).
fof(f4994,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f311,f2186,f411])).
fof(f4992,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f353,f2186,f411])).
fof(f4990,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f304,f2186,f411])).
fof(f4987,plain,(
  $false | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f256,f2211,f2186,f411])).
fof(f4985,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f4709,f2186,f411])).
fof(f4983,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f256,f1603,f2186,f411])).
fof(f4981,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f433,f2186,f411])).
fof(f4978,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2186,f1403,f2186,f327])).
fof(f4976,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f425,f2186,f327])).
fof(f4974,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2186,f1403,f2186,f327])).
fof(f4972,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2186,f425,f2186,f327])).
fof(f4968,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2186,f1403,f2186,f327])).
fof(f4966,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f425,f2186,f327])).
fof(f4964,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2186,f1403,f2186,f327])).
fof(f4962,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2186,f425,f2186,f327])).
fof(f4958,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f425,f2186,f271])).
fof(f4956,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f433,f2186,f266])).
fof(f4954,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f2186])).
fof(f4920,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl14),
  inference(cnf_transformation,[],[f4920_D])).
fof(f4920_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4913,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4608,f337])).
fof(f2182,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2182_D])).
fof(f2182_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2185,plain,(
  ( ! [X24] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f2163,f304])).
fof(f2163,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f312,f304])).
fof(f2179,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2178,f304])).
fof(f2178,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2156,f256])).
fof(f2156,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f312,f465])).
fof(f1388,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1388_D])).
fof(f1388_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1377,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f314,f304])).
fof(f424,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f423])).
fof(f423,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f420,f421])).
fof(f419,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) | $spl0),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f421,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f254,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) = s('type.num.num',sK1)) & (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) | s('type.num.num',sK0) != s('type.num.num',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f195])).
fof(f195,plain,(
  ? [X0,X1] : ((s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ? [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X6))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X6))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.patSem.v_|49||49||split|4','thm.patSem.v_|49||49||split|4')).
fof(f420,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) | $spl1),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f422,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f255,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK0))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK1)))),
  inference(cnf_transformation,[],[f196])).
% SZS output end Proof for thm.patSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 183
% Passive clauses: 1844
% Generated clauses: 101763
% Final active clauses: 133
% Final passive clauses: 1201
% Input formulas: 69
% Initial clauses: 146
% 
% Duplicate literals: 766
% Trivial inequalities: 177
% Fw subsumption resolutions: 946
% Fw demodulations: 58337
% Bw demodulations: 342
% Condensations: 971
% Global subsumptions: 8
% 
% Simple tautologies: 288
% Equational tautologies: 149
% Forward subsumptions: 16604
% Fw demodulations to eq. taut.: 21597
% Bw demodulations to eq. taut.: 58
% 
% Unit resulting resolution: 27899
% Factoring: 1
% Forward superposition: 6153
% Backward superposition: 5259
% Self superposition: 147
% Equality factoring: 80
% Equality resolution: 36
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 1956
% SAT solver unit clauses: 194
% SAT solver binary clauses: 730
% 
% Sat splits: 26
% Sat splitting refutations: 50
% 
% Memory used [KB]: 45031
% Time elapsed: 3.461 s
% ------------------------------
% ------------------------------
% Success in time 4.83 s
