Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.ml_translator
% (23414)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 131
% Passive clauses: 1624
% Generated clauses: 16503
% Final active clauses: 120
% Final passive clauses: 1368
% Input formulas: 16
% Initial clauses: 37
% 
% Duplicate literals: 958
% Trivial inequalities: 271
% Fw subsumption resolutions: 2828
% Fw demodulations: 1732
% Bw demodulations: 7
% 
% Simple tautologies: 321
% Equational tautologies: 56
% Forward subsumptions: 6896
% 
% Binary resolution: 441
% Forward superposition: 5118
% Backward superposition: 4535
% Self superposition: 185
% Equality factoring: 159
% Equality resolution: 28
% 
% Split clauses: 1258
% Split components: 2643
% SAT solver clauses: 1910
% SAT solver unit clauses: 62
% SAT solver binary clauses: 329
% 
% Sat splits: 1257
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1899
% TWLsolver calls for satisfiability: 168
% 
% Memory used [KB]: 8059
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ml_translator
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_translator
% SZS output start Proof for thm.ml_translator
fof(f2644,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2639,f61,f113,f130,f132,f59,f111,f66,f117,f2568,f2585,f79,f127,f2528,f2529,f63,f115,f139,f144,f60,f109,f70,f119,f1902,f1918,f76,f121,f78,f124])).
fof(f124,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) | $spl24),
  inference(cnf_transformation,[],[f124_D])).
fof(f124_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))) | ! [X3] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))) => ! [X3] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0)))))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10,X7,X6] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X7),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X10)))))))) => ! [X19] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'('type.one.one'),X19),s('type.ast.exp',X7),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X19),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X10)))))))))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.ml_translator.evaluate_empty_state_IMP')).
fof(f121,plain,(
  ( ! [X4,X0,X3,X1] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),sK5(X1,X0)),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))) ) | $spl18),
  inference(cnf_transformation,[],[f121_D])).
fof(f121_D,plain,(
  ( ! [X4,X0,X3,X1] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),sK5(X1,X0)),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f76,plain,(
  ( ! [X4,X0,X3,X1] : (p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),sK5(X1,X0)),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4))))))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ! [X5] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X5),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),sK3(X5,X1,X0)),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4(X5,X1,X0)))))))))) & (! [X3,X4] : ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),sK5(X1,X0)),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5])],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ! [X5] : ? [X6,X7] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X5),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X6),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X7))))))))) & (? [X2] : ! [X3,X4] : ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X2),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))))),
  inference(rectify,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ! [X2] : ? [X3,X4] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X2),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4))))))))) & (? [X2] : ! [X3,X4] : ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X2),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : (p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) <=> ! [X2] : ? [X3,X4] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),X2),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X3),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : (p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7)))) <=> ! [X20] : ? [X21,X18] : p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'('type.one.one'),X20),s('type.ast.exp',X7),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),X21),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X18)))))))))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.ml_translator.always_evaluates_def')).
fof(f1918,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X2))))))))) ) | $spl246),
  inference(cnf_transformation,[],[f1918_D])).
fof(f1918_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X2))))))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f1902,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X2)))))))) | p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1))))) ) | ($spl18 | $spl24)),
  inference(resolution,[],[f124,f121])).
fof(f119,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f70,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.bool.EQ_CLAUSES|split|3')).
fof(f109,plain,(
  ~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1)))) | $spl1),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  ~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f60,plain,(
  ~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1))))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2)))) & ~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f35])).
fof(f35,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) & ~p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) => p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7),s(fun('type.semanticPrimitives.v',bool),X8)))) => p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7),s(fun('type.semanticPrimitives.v',bool),X8)))) => p(s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7)))))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.ml_translator.Eval_IMP_always_evaluates')).
fof(f144,plain,(
  s(bool,'F_0') = s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1))) | $spl36),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  s(bool,'F_0') = s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f139,plain,(
  s(bool,'F_0') = s(bool,'const.ml_translator.always_evaluates_2'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1))) | ($spl1 | $spl14)),
  inference(resolution,[],[f119,f109])).
fof(f115,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f63,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ml_translator.Eval_IMP_always_evaluates',aHL_FALSITY)).
fof(f2529,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) | $spl384),
  inference(cnf_transformation,[],[f2529_D])).
fof(f2529_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f2528,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) | ($spl7 | $spl36 | $spl246)),
  inference(subsumption_resolution,[],[f2515,f115])).
fof(f2515,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',sK1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X0))))))))) ) | ($spl36 | $spl246)),
  inference(superposition,[],[f1918,f144])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK7(X2,X1,X0))))))))) | ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))) ) | $spl30),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK7(X2,X1,X0))))))))) | ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f79,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK7(X2,X1,X0))))))))) | ~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK7(X2,X1,X0))))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',sK7(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | ? [X4] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X4)))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(rectify,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) | ? [X3] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3)))))) & (! [X3] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))) | p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2))))))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.ast.exp',X1),s(fun('type.semanticPrimitives.v',bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X0),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X1),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X2),s('type.semanticPrimitives.v',X3))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.ast.exp',X7),s(fun('type.semanticPrimitives.v',bool),X8)))) <=> ? [X18] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,'F_0'),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.ast.exp',X7),s('type.pair.prod'('type.semanticPrimitives.state'('type.one.one'),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'('type.one.one'),'const.ml_translator.empty_state_0'),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X18)))))))) & p(s(bool,happ(s(fun('type.semanticPrimitives.v',bool),X8),s('type.semanticPrimitives.v',X18))))))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.ml_translator.Eval_def')).
fof(f2585,plain,(
  ( ! [X0] : (~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),X0))))) ) | $spl386),
  inference(cnf_transformation,[],[f2585_D])).
fof(f2585_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),X0))))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f2568,plain,(
  ( ! [X0] : (~p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),X0))))) ) | ($spl30 | $spl384)),
  inference(resolution,[],[f2529,f127])).
fof(f117,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f66,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.bool.EQ_CLAUSES|split|1')).
fof(f111,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f59,plain,(
  p(s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2))))),
  inference(cnf_transformation,[],[f42])).
fof(f132,plain,(
  s(bool,'T_0') = s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2))) | $spl34),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  s(bool,'T_0') = s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f130,plain,(
  s(bool,'T_0') = s(bool,'const.ml_translator.Eval_3'(s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK0),s('type.ast.exp',sK1),s(fun('type.semanticPrimitives.v',bool),sK2))) | ($spl2 | $spl10)),
  inference(resolution,[],[f117,f111])).
fof(f113,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f61,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ml_translator.Eval_IMP_always_evaluates','thm.bool.TRUTH')).
fof(f2639,plain,(
  $false | ($spl4 | $spl34 | $spl386)),
  inference(subsumption_resolution,[],[f2628,f113])).
fof(f2628,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl386)),
  inference(superposition,[],[f2585,f132])).
% SZS output end Proof for thm.ml_translator
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 81
% Passive clauses: 217
% Generated clauses: 2611
% Final active clauses: 40
% Final passive clauses: 85
% Input formulas: 16
% Initial clauses: 37
% 
% Duplicate literals: 154
% Trivial inequalities: 47
% Fw subsumption resolutions: 308
% Fw demodulations: 164
% Bw demodulations: 64
% Condensations: 142
% 
% Simple tautologies: 87
% Equational tautologies: 20
% Forward subsumptions: 1040
% 
% Binary resolution: 140
% Forward superposition: 776
% Backward superposition: 422
% Self superposition: 36
% Equality factoring: 34
% Equality resolution: 9
% 
% Split clauses: 99
% Split components: 213
% SAT solver clauses: 450
% SAT solver unit clauses: 32
% SAT solver binary clauses: 27
% SAT solver learnt clauses: 1
% 
% Sat splits: 99
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 449
% TWLsolver calls for satisfiability: 237
% 
% Memory used [KB]: 2814
% Time elapsed: 0.087 s
% ------------------------------
% ------------------------------
% Success in time 0.71 s
