Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% (104142)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 3945
% Generated clauses: 13385
% Final active clauses: 119
% Final passive clauses: 3582
% Input formulas: 31
% Initial clauses: 80
% 
% Duplicate literals: 634
% Trivial inequalities: 174
% Fw subsumption resolutions: 2036
% Fw demodulations: 425
% Bw demodulations: 5
% 
% Simple tautologies: 287
% Equational tautologies: 78
% Forward subsumptions: 4686
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 492
% Factoring: 2
% Forward superposition: 4987
% Backward superposition: 3959
% Self superposition: 224
% Equality factoring: 172
% Equality resolution: 32
% 
% Split clauses: 513
% Split components: 1067
% SAT solver clauses: 1260
% SAT solver unit clauses: 10
% SAT solver binary clauses: 180
% 
% Sat splits: 513
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1253
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (104176)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 1539
% Generated clauses: 12488
% Final active clauses: 146
% Final passive clauses: 1304
% Input formulas: 31
% Initial clauses: 80
% 
% Duplicate literals: 493
% Trivial inequalities: 152
% Fw subsumption resolutions: 1121
% Fw demodulations: 1008
% Bw demodulations: 31
% Condensations: 789
% Global subsumptions: 9
% 
% Simple tautologies: 184
% Equational tautologies: 44
% Forward subsumptions: 4808
% Fw demodulations to eq. taut.: 47
% Bw demodulations to eq. taut.: 39
% 
% Binary resolution: 480
% Factoring: 2
% Forward superposition: 3095
% Backward superposition: 3377
% Self superposition: 121
% Equality factoring: 80
% Equality resolution: 35
% 
% Split clauses: 142
% Split components: 304
% SAT solver clauses: 2317
% SAT solver unit clauses: 80
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 20
% SAT solver learnt literals: 4
% 
% Sat splits: 142
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 2279
% TWLsolver calls for satisfiability: 455
% 
% Memory used [KB]: 13816
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f25776,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f130,f228,f227,f129,f230,f128,f232,f226,f238,f259,f263,f261,f268,f624,f636,f625,f647,f650,f648,f765,f766,f760,f754,f778,f767,f780,f781,f782,f787,f789,f793,f795,f797,f832,f833,f834,f835,f836,f837,f838,f839,f1229,f1224,f1223,f1232,f1233,f1215,f1209,f1242,f1244,f1234,f1252,f1230,f1254,f1243,f1255,f1231,f1468,f1470,f1494,f1498,f1499,f1502,f1503,f1545,f1546,f2085,f2087,f2062,f2061,f2088,f2089,f1996,f1997,f1999,f2002,f2086,f1930,f1929,f2112,f2253,f2255,f2231,f2230,f2256,f2257,f2229,f2170,f2090,f2171,f2091,f2173,f2092,f2174,f2254,f2095,f2271,f2258,f2464,f2465,f2527,f2661,f2529,f2530,f2664,f2660,f3171,f3202,f3167,f3209,f3255,f3257,f3258,f3260,f3263,f3261,f3264,f3265,f2876,f2877,f3268,f2882,f2636,f2585,f2576,f3271,f3275,f3276,f2516,f2519,f241,f3278,f3280,f2490,f2493,f3272,f3281,f3393,f3398,f3552,f3553,f3555,f7098,f7101,f7106,f7109,f7114,f7073,f7115,f7099,f3270,f7161,f7162,f7185,f7187,f7163,f7699,f7700,f9143,f9307,f9144,f9645,f9646,f9647,f9648,f9649,f9664,f9650,f9667,f9668,f9651,f9669,f9660,f9652,f9666,f9670,f9653,f9654,f9655,f9674,f9673,f9656,f9663,f9661,f9671,f9763,f9797,f9795,f9812,f9813,f9814,f9815,f9816,f9817,f9834,f9818,f9835,f9819,f9836,f9820,f9821,f9822,f9823,f9842,f9979,f9980,f9981,f9983,f9984,f9985,f9987,f9990,f9991,f9992,f9996,f9999,f10003,f10005,f10008,f10011,f10354,f10355,f10356,f10363,f10380,f10381,f11027,f11028,f11029,f11036,f11053,f11069,f11071,f11072,f11076,f11078,f11080,f11082,f11084,f11086,f11088,f11089,f11091,f11097,f11099,f11101,f11110,f11112,f11114,f11116,f11117,f11118,f11119,f11121,f11123,f11125,f11127,f11129,f11132,f11134,f11135,f11193,f11199,f11201,f11202,f11207,f11523,f11524,f11525,f11526,f11527,f11531,f11532,f11533,f11534,f11535,f11540,f11541,f11542,f11580,f11543,f11555,f9840,f9827,f11556,f9833,f11557,f9829,f9830,f11562,f11563,f11564,f11584,f9839,f11565,f9828,f11566,f9832,f9837,f11613,f11617,f11619,f12280,f12282,f12281,f12279,f12283,f12409,f12410,f12461,f12460,f12833,f12835,f12838,f13190,f13191,f13196,f13203,f12777,f13267,f13293,f13249,f13307,f13308,f13310,f9826,f13311,f13312,f3266,f12794,f13313,f13318,f13319,f13675,f13676,f13679,f13682,f13684,f13729,f13797,f13798,f13804,f13806,f13807,f14021,f14022,f14023,f14025,f14026,f14028,f14030,f14031,f14032,f14033,f14034,f14037,f14040,f14374,f14375,f14376,f14383,f14400,f14043,f9831,f3267,f16933,f16999,f17036,f17037,f18721,f18722,f18748,f18747,f19344,f19345,f19286,f19346,f21369,f21372,f21373,f21375,f21370,f649,f21689,f21692,f21755,f21751,f21699,f21700,f21762,f21765,f21766,f21767,f21759,f21768,f21708,f21709,f21770,f21769,f21771,f21773,f21772,f21774,f21776,f21775,f21717,f21777,f21778,f21780,f21779,f21783,f21784,f21786,f21787,f21752,f21788,f21789,f21760,f18716,f18715,f21813,f21812,f22706,f22769,f260,f23412,f23413,f23415,f23416,f23419,f23420,f23425,f23426,f239,f24044,f24046,f24045,f25774,f25766,f25775])).
fof(f25775,plain,(
  $false | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f25767])).
fof(f25767,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f25683,f1476])).
fof(f1476,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),X4))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))) | s('type.semanticPrimitives.result'(sK0,sK3),X4) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) ) | $spl0),
  inference(superposition,[],[f187,f226])).
fof(f187,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9,X6] : ! [X14,X25] : (s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,X14))) = s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,X25))) <=> s(X6,X14) = s(X6,X25))),
  file('thm.evalProps.map_result_Rerr','thm.semanticPrimitives.result_|49||49||split|0')).
fof(f25683,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(X0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,X0),X1),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f25682,f24405])).
fof(f24405,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rval_1'(s(sK1,X2)))) ) | $spl0),
  inference(superposition,[],[f191,f24379])).
fof(f24379,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | $spl0),
  inference(subsumption_resolution,[],[f24192,f191])).
fof(f24192,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rval_1'(s(sK0,happ(s(fun(sK1,sK0),sK5),s(sK1,sK10(sK7,sK2,sK1)))))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | $spl0),
  inference(superposition,[],[f1395,f226])).
fof(f1395,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'(X3,X4),'const.semanticPrimitives.Rval_1'(s(X3,happ(s(fun(X0,X3),X5),s(X0,sK10(X2,X1,X0)))))) = s('type.semanticPrimitives.result'(X3,X4),'const.evalProps.map_result_3'(s(fun(X0,X3),X5),s(fun(X1,X4),X6),s('type.semanticPrimitives.result'(X0,X1),X2))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0))))) )),
  inference(superposition,[],[f206,f185])).
fof(f185,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X9] : ! [X26] : (? [X14] : s('type.semanticPrimitives.result'(X6,X9),X26) = s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,X14))) | ? [X13] : s('type.semanticPrimitives.result'(X6,X9),X26) = s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X13))))),
  file('thm.evalProps.map_result_Rerr','thm.semanticPrimitives.result_nchotomy')).
fof(f206,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'(X2,X1),'const.semanticPrimitives.Rval_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.semanticPrimitives.result'(X2,X1),'const.evalProps.map_result_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.semanticPrimitives.result'(X3,X0),'const.semanticPrimitives.Rval_1'(s(X3,X6)))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.semanticPrimitives.result'(X2,X1),'const.semanticPrimitives.Rval_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.semanticPrimitives.result'(X2,X1),'const.evalProps.map_result_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.semanticPrimitives.result'(X3,X0),'const.semanticPrimitives.Rval_1'(s(X3,X6)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s('type.semanticPrimitives.result'(X2,X1),'const.semanticPrimitives.Rval_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.semanticPrimitives.result'(X2,X1),'const.evalProps.map_result_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.semanticPrimitives.result'(X3,X0),'const.semanticPrimitives.Rval_1'(s(X3,X6)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8,X9,X6,X7] : ! [X11,X10,X24] : s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,happ(s(fun(X7,X6),X11),s(X7,X24))))) = s('type.semanticPrimitives.result'(X6,X9),'const.evalProps.map_result_3'(s(fun(X7,X6),X11),s(fun(X8,X9),X10),s('type.semanticPrimitives.result'(X7,X8),'const.semanticPrimitives.Rval_1'(s(X7,X24)))))),
  file('thm.evalProps.map_result_Rerr','thm.evalProps.map_result_def|split|0')).
fof(f191,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ! [X2,X3] : ~s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X9] : ! [X25,X14] : ~s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,X14))) = s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X25)))),
  file('thm.evalProps.map_result_Rerr','thm.semanticPrimitives.result_distinct')).
fof(f25682,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(X0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,X0),X1),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rval_1'(s(sK1,sK10(sK7,sK2,sK1))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f24737,f1443])).
fof(f1443,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s('type.semanticPrimitives.result'(X7,X8),'const.evalProps.map_result_3'(s(fun(X4,X7),X10),s(fun(X5,X8),X9),s('type.semanticPrimitives.result'(X4,X5),X6))) = s('type.semanticPrimitives.result'(X7,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),'const.evalProps.map_error_result_2'(s(fun(X5,X8),X9),s('type.semanticPrimitives.error_result'(X5),sK11(X6,X5,X4)))))) | s('type.semanticPrimitives.result'(X4,X5),X6) = s('type.semanticPrimitives.result'(X4,X5),'const.semanticPrimitives.Rval_1'(s(X4,sK10(X6,X5,X4))))) )),
  inference(superposition,[],[f207,f185])).
fof(f207,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X3,X2),X5),s('type.semanticPrimitives.error_result'(X3),X6))))) = s('type.semanticPrimitives.result'(X1,X2),'const.evalProps.map_result_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X3,X2),X5),s('type.semanticPrimitives.error_result'(X3),X6))))) = s('type.semanticPrimitives.result'(X1,X2),'const.evalProps.map_result_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),'const.evalProps.map_error_result_2'(s(fun(X3,X2),X5),s('type.semanticPrimitives.error_result'(X3),X6))))) = s('type.semanticPrimitives.result'(X1,X2),'const.evalProps.map_result_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7,X6,X9,X8] : ! [X11,X10,X13] : s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),'const.evalProps.map_error_result_2'(s(fun(X8,X9),X10),s('type.semanticPrimitives.error_result'(X8),X13))))) = s('type.semanticPrimitives.result'(X6,X9),'const.evalProps.map_result_3'(s(fun(X7,X6),X11),s(fun(X8,X9),X10),s('type.semanticPrimitives.result'(X7,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),X13)))))),
  file('thm.evalProps.map_result_Rerr','thm.evalProps.map_result_def|split|1')).
fof(f24737,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1))))))) ) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f24399,f189])).
fof(f189,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))) | s('type.semanticPrimitives.error_result'(X1),X2) = s('type.semanticPrimitives.error_result'(X1),X3)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3] : ((s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))) | s('type.semanticPrimitives.error_result'(X1),X2) = s('type.semanticPrimitives.error_result'(X1),X3)) & (s('type.semanticPrimitives.error_result'(X1),X2) != s('type.semanticPrimitives.error_result'(X1),X3) | s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3)))))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))) <=> s('type.semanticPrimitives.error_result'(X1),X2) = s('type.semanticPrimitives.error_result'(X1),X3))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))) <=> s('type.semanticPrimitives.error_result'(X1),X2) = s('type.semanticPrimitives.error_result'(X1),X3))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X9] : ! [X14,X25] : (s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X14))) = s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X25))) <=> s('type.semanticPrimitives.error_result'(X9),X14) = s('type.semanticPrimitives.error_result'(X9),X25))),
  file('thm.evalProps.map_result_Rerr','thm.semanticPrimitives.result_|49||49||split|1')).
fof(f24399,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f24379,f228])).
fof(f25766,plain,(
  $false | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f226,f25683])).
fof(f25774,plain,(
  $false | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f25773])).
fof(f25773,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) | ($spl0 | $spl2)),
  inference(superposition,[],[f25683,f226])).
fof(f24045,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | $spl216),
  inference(cnf_transformation,[],[f24045_D])).
fof(f24045_D,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f24046,plain,(
  ( ! [X45,X44] : (s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45)))) ) | $spl218),
  inference(cnf_transformation,[],[f24046_D])).
fof(f24046_D,plain,(
  ( ! [X45,X44] : (s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f24044,plain,(
  ( ! [X45,X43,X44] : (s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45))) | s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f24043,f188])).
fof(f188,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f24043,plain,(
  ( ! [X45,X43,X44] : (s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45))) | s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | ($spl16 | $spl106)),
  inference(forward_demodulation,[],[f23715,f12264])).
fof(f12264,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl106),
  inference(condensation,[],[f12263])).
fof(f12263,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl106),
  inference(subsumption_resolution,[],[f12236,f134])).
fof(f134,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.map_result_Rerr',aHL_BOOL_CASES)).
fof(f12236,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl106),
  inference(superposition,[],[f9667,f134])).
fof(f23715,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,X43) | s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45))) = s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,f(s(bool,X43)))) ) | ($spl16 | $spl106)),
  inference(superposition,[],[f346,f23414])).
fof(f23414,plain,(
  ( ! [X77] : (s(bool,X77) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X77)))) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f23387,f131])).
fof(f131,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.map_result_Rerr',aHL_TRUTH)).
fof(f23387,plain,(
  ( ! [X77] : (~p(s(bool,'T_0')) | s(bool,X77) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X77)))) ) | ($spl16 | $spl106)),
  inference(superposition,[],[f23229,f12264])).
fof(f23229,plain,(
  ( ! [X39,X41] : (~p(s(bool,X41)) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,f(s(bool,X39)))) ) | ($spl16 | $spl106)),
  inference(forward_demodulation,[],[f23228,f1561])).
fof(f1561,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK10('const.semanticPrimitives.Rval_1'(s(X0,X1)),X2,X0))) )),
  inference(unit_resulting_resolution,[],[f1277,f187])).
fof(f1277,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10('const.semanticPrimitives.Rval_1'(s(X0,X2)),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f191,f185])).
fof(f23228,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X41)) | s(bool,X41) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X39)),X40,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X39)))) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f23094,f131])).
fof(f23094,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | s(bool,X41) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X39)),X40,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X39)))) ) | ($spl16 | $spl106)),
  inference(superposition,[],[f178,f12383])).
fof(f12383,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X13)),X14,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X13)))) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f12303,f840])).
fof(f840,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl16),
  inference(unit_resulting_resolution,[],[f133,f625])).
fof(f133,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.evalProps.map_result_Rerr',aHL_FALSITY)).
fof(f12303,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X13)),X14,bool))) ) | $spl106),
  inference(superposition,[],[f12264,f1636])).
fof(f1636,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) )),
  inference(superposition,[],[f1561,f233])).
fof(f233,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f134,f134])).
fof(f178,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X15,X16] : ((p(s(bool,X15)) => p(s(bool,X16))) => ((p(s(bool,X16)) => p(s(bool,X15))) => s(bool,X15) = s(bool,X16)))),
  file('thm.evalProps.map_result_Rerr','thm.bool.IMP_ANTISYM_AX')).
fof(f346,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,X8) | s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rval_1'(s(bool,X7))) = s('type.semanticPrimitives.result'(bool,X9),'const.semanticPrimitives.Rval_1'(s(bool,X8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f188,f134])).
fof(f239,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f23426,plain,(
  ( ! [X110,X108] : (s(bool,X110) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X108,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X110)))) ) | $spl214),
  inference(cnf_transformation,[],[f23426_D])).
fof(f23426_D,plain,(
  ( ! [X110,X108] : (s(bool,X110) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X108,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X110)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f23425,plain,(
  ( ! [X109,X110,X108] : (s(bool,X110) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X108,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X110))) | s(bool,'F_0') = s(bool,X109)) ) | ($spl14 | $spl16 | $spl86 | $spl106)),
  inference(subsumption_resolution,[],[f23400,f261])).
fof(f23400,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X109)) | s(bool,X110) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X108,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X110))) | s(bool,'F_0') = s(bool,X109)) ) | ($spl16 | $spl86 | $spl106)),
  inference(superposition,[],[f23229,f7099])).
fof(f23420,plain,(
  ( ! [X85,X83] : (s(bool,X85) = s(bool,f(s(fun(bool,bool),X83))) | s(bool,'T_0') = s(bool,f(s(bool,X85)))) ) | $spl212),
  inference(cnf_transformation,[],[f23420_D])).
fof(f23420_D,plain,(
  ( ! [X85,X83] : (s(bool,X85) = s(bool,f(s(fun(bool,bool),X83))) | s(bool,'T_0') = s(bool,f(s(bool,X85)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f23419,plain,(
  ( ! [X85,X83,X84] : (s(bool,X85) = s(bool,f(s(fun(bool,bool),X83))) | s(bool,'T_0') = s(bool,f(s(bool,X85))) | s(bool,'F_0') = s(bool,X84)) ) | ($spl14 | $spl16 | $spl106 | $spl168)),
  inference(subsumption_resolution,[],[f23391,f261])).
fof(f23391,plain,(
  ( ! [X85,X83,X84] : (~p(s(bool,X84)) | s(bool,X85) = s(bool,f(s(fun(bool,bool),X83))) | s(bool,'T_0') = s(bool,f(s(bool,X85))) | s(bool,'F_0') = s(bool,X84)) ) | ($spl16 | $spl106 | $spl168)),
  inference(superposition,[],[f23229,f18831])).
fof(f18831,plain,(
  ( ! [X21,X22] : (s(bool,X22) = s(bool,f(s(fun(bool,bool),X21))) | s(bool,'F_0') = s(bool,X22)) ) | $spl168),
  inference(trivial_inequality_removal,[],[f18823])).
fof(f18823,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X22) = s(bool,f(s(fun(bool,bool),X21))) | s(bool,'F_0') = s(bool,X22)) ) | $spl168),
  inference(superposition,[],[f18748,f7068])).
fof(f7068,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f7067])).
fof(f7067,plain,(
  ( ! [X21,X19,X22] : (s(bool,X21) = s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X21)) )),
  inference(duplicate_literal_removal,[],[f6834])).
fof(f6834,plain,(
  ( ! [X21,X19,X22] : (s(bool,X21) = s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X19)) )),
  inference(superposition,[],[f1850,f1850])).
fof(f1850,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X24)),X25,bool)) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1846,f131])).
fof(f1846,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X24)),X25,bool)) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f1813,f1634])).
fof(f1634,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X3)),X4,bool)) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f1561,f134])).
fof(f1813,plain,(
  ( ! [X10,X8] : (~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s(bool,'F_0') = s(bool,X8)) )),
  inference(forward_demodulation,[],[f1812,f1561])).
fof(f1812,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1785,f131])).
fof(f1785,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f178,f1634])).
fof(f23416,plain,(
  ( ! [X79] : (s(bool,X79) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X79)))) ) | $spl210),
  inference(cnf_transformation,[],[f23416_D])).
fof(f23416_D,plain,(
  ( ! [X79] : (s(bool,X79) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X79)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f23415,plain,(
  ( ! [X78,X79] : (s(bool,X79) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X79))) | s(bool,'F_0') = s(bool,X78)) ) | ($spl14 | $spl16 | $spl106 | $spl123 | $spl142)),
  inference(subsumption_resolution,[],[f23388,f261])).
fof(f23388,plain,(
  ( ! [X78,X79] : (~p(s(bool,X78)) | s(bool,X79) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X79))) | s(bool,'F_0') = s(bool,X78)) ) | ($spl16 | $spl106 | $spl123 | $spl142)),
  inference(superposition,[],[f23229,f14015])).
fof(f14015,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X15)) ) | ($spl123 | $spl142)),
  inference(trivial_inequality_removal,[],[f14014])).
fof(f14014,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,f(s(bool,'T_0')))) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f7068])).
fof(f13812,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(backward_demodulation,[],[f9839,f9797])).
fof(f23413,plain,(
  ( ! [X76] : (s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X76)))) ) | $spl208),
  inference(cnf_transformation,[],[f23413_D])).
fof(f23413_D,plain,(
  ( ! [X76] : (s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X76)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f23412,plain,(
  ( ! [X76,X75] : (s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X76))) | s(bool,'F_0') = s(bool,X75)) ) | ($spl14 | $spl16 | $spl106 | $spl150)),
  inference(subsumption_resolution,[],[f23386,f261])).
fof(f23386,plain,(
  ( ! [X76,X75] : (~p(s(bool,X75)) | s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X76))) | s(bool,'F_0') = s(bool,X75)) ) | ($spl16 | $spl106 | $spl150)),
  inference(superposition,[],[f23229,f12283])).
fof(f260,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f22769,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) != s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,X47)))) ) | $spl206),
  inference(cnf_transformation,[],[f22769_D])).
fof(f22769_D,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) != s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,X47)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f22706,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X47) != s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl74 | $spl154)),
  inference(duplicate_literal_removal,[],[f22473])).
fof(f22473,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X47) != s(bool,happ(s(fun(bool,bool),X46),s(bool,'F_0'))) | s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X48),'const.semanticPrimitives.Rval_1'(s(bool,X47))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | ($spl74 | $spl154)),
  inference(superposition,[],[f345,f15857])).
fof(f15857,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl74 | $spl154)),
  inference(subsumption_resolution,[],[f15790,f131])).
fof(f15790,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl74 | $spl154)),
  inference(superposition,[],[f15738,f233])).
fof(f15738,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | ($spl74 | $spl154)),
  inference(forward_demodulation,[],[f15683,f1561])).
fof(f15683,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool)))))) ) | ($spl74 | $spl154)),
  inference(unit_resulting_resolution,[],[f3266,f13813])).
fof(f13813,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl154),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2527,f2529,f2530,f2705,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2576,f2490,f2493,f4947,f5015,f5270,f5271,f5419,f5421,f5527,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f2543,f6685,f304,f320,f321,f345,f346,f641,f961,f964,f1104,f1110,f1278,f1634,f1933,f9527,f9545,f9547,f9549,f651,f417,f639,f959,f2311,f2316,f5212,f6183,f12461,f6128,f628,f5430,f5659,f5525,f5428,f1280,f1279,f6682,f5102,f2322,f1408,f2323,f1458,f2319,f1449,f2321,f1401,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f2326,f1763,f6230,f2333,f2332,f2308,f1698,f1686,f1664,f974,f174,f173,f172,f6184,f6161,f2310,f474,f2314,f962,f957,f637,f453,f5657,f5523,f5426,f4952,f1108,f1106,f9881,f9844,f9630,f9594,f9590,f9583,f9577,f9559,f9550,f9526,f9177,f9136,f8466,f8343,f8341,f8340,f8339,f8335,f8330,f8328,f8327,f8326,f8226,f8223,f8220,f8219,f8162,f8158,f8157,f8156,f8079,f7301,f7300,f7244,f7243,f7242,f7241,f7239,f7238,f7237,f7160,f7078,f7027,f7020,f6941,f6933,f6801,f6198,f6196,f6152,f6131,f6130,f6116,f5531,f5060,f4170,f4168,f4167,f4166,f4161,f4150,f4148,f4147,f4146,f4143,f2840,f2814,f2753,f2731,f2716,f2714,f2713,f2709,f2706,f2695,f2688,f2687,f2628,f2590,f2484,f2482,f2446,f2413,f2410,f2400,f2397,f2381,f2377,f2375,f2374,f2373,f2369,f2368,f2365,f2361,f2359,f2358,f2357,f2354,f2353,f2352,f2346,f2340,f2317,f2309,f2298,f2286,f1956,f1903,f1892,f1850,f1819,f1818,f1817,f1816,f1815,f1813,f1797,f1777,f133,f303])).
fof(f303,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK9('F_0',X3,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f173,f134])).
fof(f1777,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X27)),X29,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X27)),X28,bool))) )),
  inference(superposition,[],[f1634,f1561])).
fof(f1797,plain,(
  ( ! [X47,X45,X46,X44] : (s('type.semanticPrimitives.result'(bool,X46),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X46),'const.semanticPrimitives.Rval_1'(s(bool,X47))) | s(bool,X47) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X44)),X45,bool)) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f187,f1634])).
fof(f1815,plain,(
  ( ! [X13,X11] : (p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s(bool,'F_0') = s(bool,X11)) )),
  inference(forward_demodulation,[],[f1814,f1561])).
fof(f1814,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,X13) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X11)),X12,bool)) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1786,f131])).
fof(f1786,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | p(s(bool,X13)) | s(bool,X13) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X11)),X12,bool)) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f180,f1634])).
fof(f1816,plain,(
  ( ! [X35,X36] : (s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1794,f191])).
fof(f1794,plain,(
  ( ! [X35,X36] : (s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X35))) = s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X36),sK11('const.semanticPrimitives.Rval_1'(s(bool,X35)),X36,bool)))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f185,f1634])).
fof(f1817,plain,(
  ( ! [X39,X37,X40] : (s(bool,X37) != s(bool,X40) | ~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s(bool,'F_0') = s(bool,X37)) )),
  inference(forward_demodulation,[],[f1795,f1561])).
fof(f1795,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X40)))) | s(bool,X40) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool)) | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f172,f1634])).
fof(f1818,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X41)),X42,bool),X43,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X41)) )),
  inference(forward_demodulation,[],[f1796,f1561])).
fof(f1796,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X41)),X42,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X41)),X42,bool),X43,bool)) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f173,f1634])).
fof(f1819,plain,(
  ( ! [X54,X52,X55] : (s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,X52))) = s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,X54))) | s(bool,'T_0') != s(bool,X54) | s(bool,'F_0') = s(bool,X52)) )),
  inference(forward_demodulation,[],[f1799,f1561])).
fof(f1799,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,X54))) = s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X52)),X53,bool)))) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f188,f1634])).
fof(f1892,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X25)),X26,bool))) )),
  inference(superposition,[],[f1867,f1561])).
fof(f1903,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X29) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X27)),X28,bool)) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1867,f1561])).
fof(f1956,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') = s(bool,sK9('F_0',X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1946,f131])).
fof(f1946,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK9('F_0',X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f303,f233])).
fof(f2286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,X2) = s(bool,sK9(X2,sK10('const.semanticPrimitives.Rval_1'(s(fun(bool,bool),X0)),X1,fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f304,f1561])).
fof(f2298,plain,(
  ( ! [X4,X2,X3] : (s(bool,X3) = s(bool,sK9(X3,X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2294,f131])).
fof(f2294,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,sK9(X3,X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f304,f233])).
fof(f2309,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,sK9(X4,X5,bool)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f651,f134])).
fof(f2317,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X28),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool),X28,bool)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f651,f1634])).
fof(f2340,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1816,f233])).
fof(f2346,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X1))) = s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f1816,f1816])).
fof(f2352,plain,(
  ( ! [X23,X21,X22] : (s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X23),'const.semanticPrimitives.Rval_1'(s(bool,X21))) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X21)),X22,bool))) )),
  inference(superposition,[],[f1816,f1561])).
fof(f2353,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.result'(bool,X4),X5) = s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK10(X5,X4,bool)) | s('type.semanticPrimitives.result'(bool,X4),X5) = s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X4),sK11(X5,X4,bool))))) )),
  inference(superposition,[],[f1816,f185])).
fof(f2354,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s('type.semanticPrimitives.result'(X4,X5),'const.semanticPrimitives.Rval_1'(s(X4,happ(s(fun(bool,X4),X6),s(bool,'T_0'))))) = s('type.semanticPrimitives.result'(X4,X5),'const.evalProps.map_result_3'(s(fun(bool,X4),X6),s(fun(X2,X5),X7),s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,X3))))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f206,f1816])).
fof(f2357,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X17),bool),X19),s('type.semanticPrimitives.result'(bool,X17),'const.semanticPrimitives.Rval_1'(s(bool,X18)))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X17),bool),X19),s('type.semanticPrimitives.result'(bool,X17),X20)))) | s('type.semanticPrimitives.result'(bool,X17),X20) != s('type.semanticPrimitives.result'(bool,X17),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f172,f1816])).
fof(f2358,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X21),bool),X23),s('type.semanticPrimitives.result'(bool,X21),'const.semanticPrimitives.Rval_1'(s(bool,X22)))))) | s('type.semanticPrimitives.result'(bool,X21),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X21),sK9('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X23,'type.semanticPrimitives.result'(bool,X21))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f173,f1816])).
fof(f2359,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X24),X26),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X24),X27))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X24),X26),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X24),'const.semanticPrimitives.Rval_1'(s(bool,X25))))) | s('type.semanticPrimitives.result'(bool,X24),X27) = s('type.semanticPrimitives.result'(bool,X24),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f187,f1816])).
fof(f2361,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.semanticPrimitives.result'(bool,X32),X34) != s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,X33))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X32),X35),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X32),X34))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X32),X35),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f188,f1816])).
fof(f2365,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X51),bool),X53),s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,X52)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X51),bool),X53),s('type.semanticPrimitives.result'(bool,X51),sK9('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X53,'type.semanticPrimitives.result'(bool,X51))))) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f651,f1816])).
fof(f2368,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s('type.semanticPrimitives.result'(X4,X5),'const.semanticPrimitives.Rval_1'(s(X4,happ(s(fun(bool,X4),X6),s(bool,X3))))) = s('type.semanticPrimitives.result'(X4,X5),'const.evalProps.map_result_3'(s(fun(bool,X4),X6),s(fun(X2,X5),X7),s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f206,f1816])).
fof(f2369,plain,(
  ( ! [X10,X8,X9] : (s('type.semanticPrimitives.result'(bool,X8),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X8),'const.semanticPrimitives.Rval_1'(s(bool,X10))) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f187,f1816])).
fof(f2373,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X25),bool),X27),s('type.semanticPrimitives.result'(bool,X25),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X25),bool),X27),s('type.semanticPrimitives.result'(bool,X25),X28)))) | s('type.semanticPrimitives.result'(bool,X25),X28) != s('type.semanticPrimitives.result'(bool,X25),'const.semanticPrimitives.Rval_1'(s(bool,X26))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f172,f1816])).
fof(f2374,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X29),bool),X31),s('type.semanticPrimitives.result'(bool,X29),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) | s('type.semanticPrimitives.result'(bool,X29),'const.semanticPrimitives.Rval_1'(s(bool,X30))) = s('type.semanticPrimitives.result'(bool,X29),sK9('const.semanticPrimitives.Rval_1'(s(bool,X30)),X31,'type.semanticPrimitives.result'(bool,X29))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f173,f1816])).
fof(f2375,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X32),X34),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X32),X35))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X32),X34),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) | s('type.semanticPrimitives.result'(bool,X32),X35) = s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,X33))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f187,f1816])).
fof(f2377,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.semanticPrimitives.result'(bool,X40),X42) != s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X40),X43),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X40),X42))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X40),X43),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X41))))) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f188,f1816])).
fof(f2381,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X59),bool),X61),s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X59),bool),X61),s('type.semanticPrimitives.result'(bool,X59),sK9('const.semanticPrimitives.Rval_1'(s(bool,X60)),X61,'type.semanticPrimitives.result'(bool,X59))))) | s(bool,'F_0') = s(bool,X60)) )),
  inference(superposition,[],[f651,f1816])).
fof(f2397,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.semanticPrimitives.result'(bool,X34),'const.semanticPrimitives.Rval_1'(s(bool,X32))) != s('type.semanticPrimitives.result'(bool,X34),'const.semanticPrimitives.Rval_1'(s(bool,X35))) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X32)),X33,bool))) )),
  inference(superposition,[],[f320,f1561])).
fof(f2400,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'(bool,X6),X7) != s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,sK10(X7,X6,bool)) | s('type.semanticPrimitives.result'(bool,X6),X7) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),sK11(X7,X6,bool))))) )),
  inference(superposition,[],[f320,f185])).
fof(f2410,plain,(
  ( ! [X37,X35,X36,X34] : (s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X34))) != s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X37))) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X34)),X35,bool)) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f320,f1561])).
fof(f2413,plain,(
  ( ! [X6,X8,X7] : (s('type.semanticPrimitives.result'(bool,X6),X7) != s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,sK10(X7,X6,bool)) | s(bool,'T_0') = s(bool,X8) | s('type.semanticPrimitives.result'(bool,X6),X7) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),sK11(X7,X6,bool))))) )),
  inference(superposition,[],[f320,f185])).
fof(f2446,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X32) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X30)),X31,bool)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2433,f131])).
fof(f2433,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X32) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X30)),X31,bool)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f626,f1634])).
fof(f2482,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,X30))) | s(bool,X33) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X30)),X31,bool)) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f321,f1561])).
fof(f2484,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.result'(bool,X5),X6) != s('type.semanticPrimitives.result'(bool,X5),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X7) = s(bool,sK10(X6,X5,bool)) | s(bool,'T_0') = s(bool,X7) | s('type.semanticPrimitives.result'(bool,X5),X6) = s('type.semanticPrimitives.result'(bool,X5),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X5),sK11(X6,X5,bool))))) )),
  inference(superposition,[],[f321,f185])).
fof(f2590,plain,(
  ( ! [X121,X118,X120,X119,X117] : (s('type.semanticPrimitives.result'(bool,X119),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X119),'const.semanticPrimitives.Rval_1'(s(bool,X120))) | s(bool,X120) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X117)),X118,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X117)),X121,bool))) )),
  inference(superposition,[],[f187,f1777])).
fof(f2628,plain,(
  ( ! [X125,X123,X126,X124,X122] : (s('type.semanticPrimitives.result'(bool,X124),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X124),'const.semanticPrimitives.Rval_1'(s(bool,X125))) | s(bool,X125) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X122)),X123,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X122)),X126,bool))) )),
  inference(superposition,[],[f187,f1777])).
fof(f2687,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(bool,X67) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X66,bool)) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X69,bool))) )),
  inference(subsumption_resolution,[],[f2578,f131])).
fof(f2578,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | s(bool,X67) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X66,bool)) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X69,bool))) )),
  inference(superposition,[],[f626,f1777])).
fof(f2688,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X81) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X79)),X80,bool)) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X79)),X82,bool))) )),
  inference(subsumption_resolution,[],[f2581,f131])).
fof(f2581,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | s(bool,X81) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X79)),X80,bool)) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X79)),X82,bool))) )),
  inference(superposition,[],[f1813,f1777])).
fof(f2695,plain,(
  ( ! [X127,X130,X131,X129] : (s('type.semanticPrimitives.result'(bool,X130),'const.semanticPrimitives.Rval_1'(s(bool,X127))) = s('type.semanticPrimitives.result'(bool,X130),'const.semanticPrimitives.Rval_1'(s(bool,X129))) | s(bool,'T_0') != s(bool,X129) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X127)),X131,bool))) )),
  inference(forward_demodulation,[],[f2592,f1561])).
fof(f2592,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'T_0') != s(bool,X129) | s('type.semanticPrimitives.result'(bool,X130),'const.semanticPrimitives.Rval_1'(s(bool,X129))) = s('type.semanticPrimitives.result'(bool,X130),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X127)),X128,bool)))) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X127)),X131,bool))) )),
  inference(superposition,[],[f188,f1777])).
fof(f2706,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,X89) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X88,bool)) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X90,bool))) )),
  inference(subsumption_resolution,[],[f2619,f133])).
fof(f2619,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,'F_0')) | s(bool,X89) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X88,bool)) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X90,bool))) )),
  inference(superposition,[],[f1815,f1777])).
fof(f2709,plain,(
  ( ! [X111,X112,X110] : (s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rval_1'(s(bool,X110))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X110)),X112,bool))) )),
  inference(subsumption_resolution,[],[f2625,f191])).
fof(f2625,plain,(
  ( ! [X111,X112,X110] : (s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rval_1'(s(bool,X110))) | s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rval_1'(s(bool,X110))) = s('type.semanticPrimitives.result'(bool,X111),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X111),sK11('const.semanticPrimitives.Rval_1'(s(bool,X110)),X111,bool)))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X110)),X112,bool))) )),
  inference(superposition,[],[f185,f1777])).
fof(f2713,plain,(
  ( ! [X118,X120,X119] : (s(bool,'T_0') = s(bool,X118) | s(bool,X118) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool),X120,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f2712,f1561])).
fof(f2712,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X118) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool),X120,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X121,bool))) )),
  inference(forward_demodulation,[],[f2627,f1561])).
fof(f2627,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,'F_0')))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool),X120,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X121,bool))) )),
  inference(superposition,[],[f173,f1777])).
fof(f2714,plain,(
  ( ! [X134,X132,X136,X135] : (s('type.semanticPrimitives.result'(bool,X135),'const.semanticPrimitives.Rval_1'(s(bool,X132))) = s('type.semanticPrimitives.result'(bool,X135),'const.semanticPrimitives.Rval_1'(s(bool,X134))) | s(bool,'F_0') != s(bool,X134) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X132)),X136,bool))) )),
  inference(forward_demodulation,[],[f2630,f1561])).
fof(f2630,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,'F_0') != s(bool,X134) | s('type.semanticPrimitives.result'(bool,X135),'const.semanticPrimitives.Rval_1'(s(bool,X134))) = s('type.semanticPrimitives.result'(bool,X135),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X132)),X133,bool)))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X132)),X136,bool))) )),
  inference(superposition,[],[f188,f1777])).
fof(f2716,plain,(
  ( ! [X156,X157,X155] : (s(bool,'T_0') = s(bool,X155) | ~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X157),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X155)),X156,bool),X157,bool))))) )),
  inference(forward_demodulation,[],[f2634,f1561])).
fof(f2634,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X157),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X155)),X156,bool),X157,bool)))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X155)),X158,bool))) )),
  inference(superposition,[],[f651,f1777])).
fof(f2731,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) != s(bool,X44) | s('type.semanticPrimitives.result'(bool,X45),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X45),'const.semanticPrimitives.Rval_1'(s(bool,X44))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X42)),X43,bool))) )),
  inference(superposition,[],[f345,f1561])).
fof(f2753,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f2752])).
fof(f2752,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f2751])).
fof(f2751,plain,(
  ( ! [X10,X8,X7,X9] : (s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,X7))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2734,f233])).
fof(f2734,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X8) != s(bool,X9) | s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,X7))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f345,f233])).
fof(f2814,plain,(
  ( ! [X116,X114,X115,X113] : (s('type.semanticPrimitives.result'(bool,X115),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X115),'const.semanticPrimitives.Rval_1'(s(bool,X116))) | s(bool,X116) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X113)),X114,bool)) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f187,f2543])).
fof(f2840,plain,(
  ( ! [X88,X87,X89] : (s(bool,X89) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X88,bool)) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2805,f133])).
fof(f2805,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,'F_0')) | s(bool,X89) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X87)),X88,bool)) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f1815,f2543])).
fof(f4143,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.semanticPrimitives.result'(X19,X20),'const.semanticPrimitives.Rval_1'(s(X19,happ(s(fun(bool,X19),X21),s(bool,'F_0'))))) = s('type.semanticPrimitives.result'(X19,X20),'const.evalProps.map_result_3'(s(fun(bool,X19),X21),s(fun(X17,X20),X22),s('type.semanticPrimitives.result'(bool,X17),'const.semanticPrimitives.Rval_1'(s(bool,X18))))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f206,f2753])).
fof(f4146,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X32),bool),X34),s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,X33)))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X32),bool),X34),s('type.semanticPrimitives.result'(bool,X32),X35)))) | s('type.semanticPrimitives.result'(bool,X32),X35) != s('type.semanticPrimitives.result'(bool,X32),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f172,f2753])).
fof(f4147,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X36),bool),X38),s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,X37)))))) | s('type.semanticPrimitives.result'(bool,X36),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X36),sK9('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X38,'type.semanticPrimitives.result'(bool,X36))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f173,f2753])).
fof(f4148,plain,(
  ( ! [X39,X41,X42,X40] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X39),X41),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X39),X42))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X39),X41),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X39),'const.semanticPrimitives.Rval_1'(s(bool,X40))))) | s('type.semanticPrimitives.result'(bool,X39),X42) = s('type.semanticPrimitives.result'(bool,X39),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f187,f2753])).
fof(f4150,plain,(
  ( ! [X47,X50,X48,X49] : (s('type.semanticPrimitives.result'(bool,X47),X49) != s('type.semanticPrimitives.result'(bool,X47),'const.semanticPrimitives.Rval_1'(s(bool,X48))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X47),X50),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X47),X49))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X47),X50),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X47),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f188,f2753])).
fof(f4161,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.semanticPrimitives.result'(X19,X20),'const.semanticPrimitives.Rval_1'(s(X19,happ(s(fun(bool,X19),X21),s(bool,X18))))) = s('type.semanticPrimitives.result'(X19,X20),'const.evalProps.map_result_3'(s(fun(bool,X19),X21),s(fun(X17,X20),X22),s('type.semanticPrimitives.result'(bool,X17),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f206,f2753])).
fof(f4166,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X40),bool),X42),s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X40),bool),X42),s('type.semanticPrimitives.result'(bool,X40),X43)))) | s('type.semanticPrimitives.result'(bool,X40),X43) != s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X41))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f172,f2753])).
fof(f4167,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X44),bool),X46),s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))))) | s('type.semanticPrimitives.result'(bool,X44),'const.semanticPrimitives.Rval_1'(s(bool,X45))) = s('type.semanticPrimitives.result'(bool,X44),sK9('const.semanticPrimitives.Rval_1'(s(bool,X45)),X46,'type.semanticPrimitives.result'(bool,X44))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f173,f2753])).
fof(f4168,plain,(
  ( ! [X47,X50,X48,X49] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X47),X49),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X47),X50))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X47),X49),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X47),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))))) | s('type.semanticPrimitives.result'(bool,X47),X50) = s('type.semanticPrimitives.result'(bool,X47),'const.semanticPrimitives.Rval_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f187,f2753])).
fof(f4170,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.semanticPrimitives.result'(bool,X55),X57) != s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X55),X58),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X55),X57))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X55),X58),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X55),'const.semanticPrimitives.Rval_1'(s(bool,X56))))) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f188,f2753])).
fof(f5060,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rval_1'(s(bool,X1)))) )),
  inference(condensation,[],[f5059])).
fof(f5059,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,X67) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X66,bool)) | s('type.semanticPrimitives.result'(bool,X68),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X68),'const.semanticPrimitives.Rval_1'(s(bool,X67))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X65)) )),
  inference(duplicate_literal_removal,[],[f4899])).
fof(f4899,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,X67) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X65)),X66,bool)) | s('type.semanticPrimitives.result'(bool,X68),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X68),'const.semanticPrimitives.Rval_1'(s(bool,X67))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f345,f1635])).
fof(f5531,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s('type.semanticPrimitives.result'(bool,X15),'const.semanticPrimitives.Rval_1'(s(bool,X14))) = s('type.semanticPrimitives.result'(bool,X15),'const.semanticPrimitives.Rval_1'(s(bool,X16))) | s(bool,'F_0') = s(bool,X16) | s(bool,X14) = s(bool,X17) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X17)),X18,bool))) )),
  inference(superposition,[],[f1816,f1636])).
fof(f6116,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,X110) = s(bool,X111) | s(bool,'F_0') = s(bool,X111) | s(bool,X110) = s(bool,X112) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X112)),X113,bool))) )),
  inference(subsumption_resolution,[],[f5596,f131])).
fof(f5596,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | s(bool,X110) = s(bool,X111) | s(bool,'F_0') = s(bool,X111) | s(bool,X110) = s(bool,X112) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X112)),X113,bool))) )),
  inference(superposition,[],[f1813,f1636])).
fof(f6130,plain,(
  ( ! [X280,X282,X281,X283] : (s(bool,'F_0') = s(bool,sK9('F_0',X280,bool)) | s(bool,'T_0') = s(bool,X281) | s(bool,X282) = s(bool,happ(s(fun(bool,bool),X280),s(bool,X281))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X282)),X283,bool))) )),
  inference(subsumption_resolution,[],[f5636,f131])).
fof(f5636,plain,(
  ( ! [X280,X282,X281,X283] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK9('F_0',X280,bool)) | s(bool,'T_0') = s(bool,X281) | s(bool,X282) = s(bool,happ(s(fun(bool,bool),X280),s(bool,X281))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X282)),X283,bool))) )),
  inference(superposition,[],[f303,f1636])).
fof(f6131,plain,(
  ( ! [X285,X287,X284,X286] : (s(bool,X285) = s(bool,sK9(X285,X284,bool)) | s(bool,'T_0') = s(bool,X285) | s(bool,X286) = s(bool,happ(s(fun(bool,bool),X284),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X286)),X287,bool))) )),
  inference(subsumption_resolution,[],[f5637,f131])).
fof(f5637,plain,(
  ( ! [X285,X287,X284,X286] : (~p(s(bool,'T_0')) | s(bool,X285) = s(bool,sK9(X285,X284,bool)) | s(bool,'T_0') = s(bool,X285) | s(bool,X286) = s(bool,happ(s(fun(bool,bool),X284),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X286)),X287,bool))) )),
  inference(superposition,[],[f304,f1636])).
fof(f6152,plain,(
  ( ! [X154,X152,X155,X153] : (s(bool,X154) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X152)),X153,bool)) | s(bool,'F_0') = s(bool,X154) | s(bool,X152) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) )),
  inference(subsumption_resolution,[],[f5697,f131])).
fof(f5697,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'T_0')) | s(bool,X154) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X152)),X153,bool)) | s(bool,'F_0') = s(bool,X154) | s(bool,X152) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f1813,f1636])).
fof(f6196,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X113),bool),X115),s('type.semanticPrimitives.result'(bool,X113),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X113),bool),X115),s('type.semanticPrimitives.result'(bool,X113),sK9('const.semanticPrimitives.Rval_1'(s(bool,X114)),X115,'type.semanticPrimitives.result'(bool,X113))))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f651,f2753])).
fof(f6198,plain,(
  ( ! [X121,X120,X119] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X119),bool),X121),s('type.semanticPrimitives.result'(bool,X119),'const.semanticPrimitives.Rval_1'(s(bool,X120)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(bool,X119),bool),X121),s('type.semanticPrimitives.result'(bool,X119),sK9('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X121,'type.semanticPrimitives.result'(bool,X119))))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f651,f2753])).
fof(f6801,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,X112) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X114)),X115,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X112)),X113,bool)) | s(bool,'F_0') = s(bool,X114)) )),
  inference(superposition,[],[f1850,f1561])).
fof(f6933,plain,(
  ( ! [X420,X422,X418,X421,X419] : (s(bool,X421) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X419)),X420,bool)) | s('type.semanticPrimitives.result'(bool,X422),'const.semanticPrimitives.Rval_1'(s(bool,X418))) = s('type.semanticPrimitives.result'(bool,X422),'const.semanticPrimitives.Rval_1'(s(bool,X421))) | s(bool,'F_0') = s(bool,X418) | s(bool,'F_0') = s(bool,X419)) )),
  inference(superposition,[],[f188,f1850])).
fof(f6941,plain,(
  ( ! [X466,X469,X465,X467,X468] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X468)),X469,bool)) != s(bool,f(s(X465,sK12(X466,X467,bool,X465)))) | s(fun(X465,bool),X466) = s(fun(X465,bool),X467) | s(bool,'F_0') = s(bool,g(s(X465,sK12(X466,X467,bool,X465)))) | s(bool,'F_0') = s(bool,X468)) )),
  inference(superposition,[],[f186,f1850])).
fof(f7020,plain,(
  ( ! [X308,X310,X309,X311,X307] : (s('type.semanticPrimitives.result'(bool,X310),'const.semanticPrimitives.Rval_1'(s(bool,X309))) != s('type.semanticPrimitives.result'(bool,X310),'const.semanticPrimitives.Rval_1'(s(bool,X311))) | s(bool,X311) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X307)),X308,bool)) | s(bool,'F_0') = s(bool,X309) | s(bool,'F_0') = s(bool,X307)) )),
  inference(superposition,[],[f187,f1850])).
fof(f7027,plain,(
  ( ! [X346,X349,X347,X348] : (~p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X348)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X349),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X346)),X347,bool),X349,bool)))) | s(bool,'F_0') = s(bool,X348) | s(bool,'F_0') = s(bool,X346)) )),
  inference(superposition,[],[f651,f1850])).
fof(f7078,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X3,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f7077])).
fof(f7077,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X40)),X41,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X42)),X43,bool)) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f6788])).
fof(f6788,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X40)),X41,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X42)),X43,bool)) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f1850,f1850])).
fof(f7160,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X70)),X71,bool)) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X72)) )),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f346,f345,f321,f320,f133,f628,f1100,f1102,f1104,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1634,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1777,f1813,f1815,f1816,f1797,f1819,f1849,f1850,f1867,f1892,f1903,f2298,f2340,f2346,f2352,f2354,f2359,f2361,f2368,f2369,f2375,f2377,f2397,f2410,f2446,f2482,f2527,f2529,f2530,f2543,f2687,f2688,f2590,f2695,f2705,f2706,f2709,f2628,f2714,f2731,f2753,f2838,f2840,f2814,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2576,f2490,f2493,f2484,f2413,f2400,f2353,f1280,f1279,f1278,f1956,f1763,f2326,f2373,f2374,f2381,f2357,f2358,f2365,f1408,f2322,f1458,f2323,f1449,f2319,f1401,f2321,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f2716,f2713,f2309,f2286,f304,f2317,f2310,f1818,f474,f2316,f2314,f2311,f1933,f1817,f964,f962,f961,f959,f957,f641,f639,f637,f453,f417,f303,f2333,f2332,f2308,f1698,f1686,f1664,f974,f651,f174,f173,f172,f1106,f1108,f1110,f4143,f4146,f4147,f4148,f4150,f4161,f4166,f4167,f4168,f4170,f5102,f5060,f4947,f4952,f5212,f5015,f5270,f5271,f5419,f5421,f5426,f5428,f5430,f5523,f5525,f5527,f5531,f5567,f6116,f6128,f6130,f6131,f5650,f5652,f5657,f5659,f5661,f6144,f6152,f6159,f6161,f5717,f5744,f5732,f6183,f6184,f6196,f6198,f6230,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6685,f6682,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7078,f6801,f7106,f7114,f7073,f7068,f7159])).
fof(f7237,plain,(
  ( ! [X399,X403,X400,X402] : (~p(s(bool,happ(s(fun(bool,bool),X402),s(bool,X400)))) | p(s(bool,happ(s(fun(bool,bool),X402),s(bool,X403)))) | s(bool,X399) != s(bool,X403) | s(bool,'F_0') = s(bool,X399) | s(bool,'F_0') = s(bool,X400)) )),
  inference(forward_demodulation,[],[f6929,f1561])).
fof(f6929,plain,(
  ( ! [X399,X401,X403,X400,X402] : (~p(s(bool,happ(s(fun(bool,bool),X402),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X400)),X401,bool))))) | p(s(bool,happ(s(fun(bool,bool),X402),s(bool,X403)))) | s(bool,X399) != s(bool,X403) | s(bool,'F_0') = s(bool,X399) | s(bool,'F_0') = s(bool,X400)) )),
  inference(superposition,[],[f172,f1850])).
fof(f7238,plain,(
  ( ! [X405,X407,X404] : (~p(s(bool,happ(s(fun(bool,bool),X407),s(bool,X405)))) | s(bool,X404) = s(bool,sK9(X404,X407,bool)) | s(bool,'F_0') = s(bool,X404) | s(bool,'F_0') = s(bool,X405)) )),
  inference(forward_demodulation,[],[f6930,f1561])).
fof(f6930,plain,(
  ( ! [X405,X407,X404,X406] : (~p(s(bool,happ(s(fun(bool,bool),X407),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X405)),X406,bool))))) | s(bool,X404) = s(bool,sK9(X404,X407,bool)) | s(bool,'F_0') = s(bool,X404) | s(bool,'F_0') = s(bool,X405)) )),
  inference(superposition,[],[f173,f1850])).
fof(f7239,plain,(
  ( ! [X412,X408,X409,X411] : (s('type.semanticPrimitives.result'(bool,X411),'const.semanticPrimitives.Rval_1'(s(bool,X409))) != s('type.semanticPrimitives.result'(bool,X411),'const.semanticPrimitives.Rval_1'(s(bool,X412))) | s(bool,X408) = s(bool,X412) | s(bool,'F_0') = s(bool,X408) | s(bool,'F_0') = s(bool,X409)) )),
  inference(forward_demodulation,[],[f6931,f1561])).
fof(f6931,plain,(
  ( ! [X412,X408,X410,X409,X411] : (s('type.semanticPrimitives.result'(bool,X411),'const.semanticPrimitives.Rval_1'(s(bool,X412))) != s('type.semanticPrimitives.result'(bool,X411),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X409)),X410,bool)))) | s(bool,X408) = s(bool,X412) | s(bool,'F_0') = s(bool,X408) | s(bool,'F_0') = s(bool,X409)) )),
  inference(superposition,[],[f187,f1850])).
fof(f7241,plain,(
  ( ! [X447,X448,X450] : (~p(s(bool,happ(s(fun(bool,bool),X450),s(bool,X448)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,sK9(X447,X450,bool)))) | s(bool,'F_0') = s(bool,X447) | s(bool,'F_0') = s(bool,X448)) )),
  inference(forward_demodulation,[],[f6938,f1561])).
fof(f6938,plain,(
  ( ! [X447,X449,X448,X450] : (~p(s(bool,happ(s(fun(bool,bool),X450),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X448)),X449,bool))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X450),s(bool,sK9(X447,X450,bool)))) | s(bool,'F_0') = s(bool,X447) | s(bool,'F_0') = s(bool,X448)) )),
  inference(superposition,[],[f651,f1850])).
fof(f7242,plain,(
  ( ! [X478,X480,X479] : (p(s(bool,happ(s(fun(bool,bool),X479),s(bool,X480)))) | ~p(s(bool,happ(s(fun(bool,bool),X479),s(bool,X478)))) | s(bool,'F_0') = s(bool,sK9(X478,X479,bool)) | s(bool,'F_0') = s(bool,X480)) )),
  inference(forward_demodulation,[],[f6942,f1561])).
fof(f6942,plain,(
  ( ! [X478,X480,X479,X481] : (p(s(bool,happ(s(fun(bool,bool),X479),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X480)),X481,bool))))) | ~p(s(bool,happ(s(fun(bool,bool),X479),s(bool,X478)))) | s(bool,'F_0') = s(bool,sK9(X478,X479,bool)) | s(bool,'F_0') = s(bool,X480)) )),
  inference(superposition,[],[f174,f1850])).
fof(f7243,plain,(
  ( ! [X484,X482,X483] : (s('type.semanticPrimitives.result'(bool,X483),X482) = s('type.semanticPrimitives.result'(bool,X483),'const.semanticPrimitives.Rval_1'(s(bool,X484))) | s('type.semanticPrimitives.result'(bool,X483),X482) = s('type.semanticPrimitives.result'(bool,X483),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X483),sK11(X482,X483,bool)))) | s(bool,'F_0') = s(bool,sK10(X482,X483,bool)) | s(bool,'F_0') = s(bool,X484)) )),
  inference(forward_demodulation,[],[f6943,f1561])).
fof(f6943,plain,(
  ( ! [X484,X482,X485,X483] : (s('type.semanticPrimitives.result'(bool,X483),X482) = s('type.semanticPrimitives.result'(bool,X483),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X484)),X485,bool)))) | s('type.semanticPrimitives.result'(bool,X483),X482) = s('type.semanticPrimitives.result'(bool,X483),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X483),sK11(X482,X483,bool)))) | s(bool,'F_0') = s(bool,sK10(X482,X483,bool)) | s(bool,'F_0') = s(bool,X484)) )),
  inference(superposition,[],[f185,f1850])).
fof(f7244,plain,(
  ( ! [X492,X490,X493,X491] : (s(X492,f(s(bool,X493))) != s(X492,g(s(bool,X493))) | s(fun(bool,X492),X490) = s(fun(bool,X492),X491) | s(bool,'F_0') = s(bool,sK12(X490,X491,X492,bool)) | s(bool,'F_0') = s(bool,X493)) )),
  inference(forward_demodulation,[],[f6945,f1561])).
fof(f6945,plain,(
  ( ! [X492,X494,X490,X493,X491] : (s(X492,f(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X493)),X494,bool)))) != s(X492,g(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X493)),X494,bool)))) | s(fun(bool,X492),X490) = s(fun(bool,X492),X491) | s(bool,'F_0') = s(bool,sK12(X490,X491,X492,bool)) | s(bool,'F_0') = s(bool,X493)) )),
  inference(superposition,[],[f186,f1850])).
fof(f7300,plain,(
  ( ! [X304,X306,X303,X305] : (s(bool,X303) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X303)),X304,bool),X306,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X306),s(bool,X305)))) | s(bool,'F_0') = s(bool,X305) | s(bool,'F_0') = s(bool,X303)) )),
  inference(forward_demodulation,[],[f7019,f1561])).
fof(f7019,plain,(
  ( ! [X304,X306,X303,X305] : (~p(s(bool,happ(s(fun(bool,bool),X306),s(bool,X305)))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X303)),X304,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X303)),X304,bool),X306,bool)) | s(bool,'F_0') = s(bool,X305) | s(bool,'F_0') = s(bool,X303)) )),
  inference(superposition,[],[f173,f1850])).
fof(f7301,plain,(
  ( ! [X321,X320,X317,X319] : (s('type.semanticPrimitives.result'(bool,X321),'const.semanticPrimitives.Rval_1'(s(bool,X317))) = s('type.semanticPrimitives.result'(bool,X321),'const.semanticPrimitives.Rval_1'(s(bool,X320))) | s(bool,X319) != s(bool,X320) | s(bool,'F_0') = s(bool,X319) | s(bool,'F_0') = s(bool,X317)) )),
  inference(forward_demodulation,[],[f7022,f1561])).
fof(f7022,plain,(
  ( ! [X321,X318,X320,X317,X319] : (s(bool,X319) != s(bool,X320) | s('type.semanticPrimitives.result'(bool,X321),'const.semanticPrimitives.Rval_1'(s(bool,X320))) = s('type.semanticPrimitives.result'(bool,X321),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X317)),X318,bool)))) | s(bool,'F_0') = s(bool,X319) | s(bool,'F_0') = s(bool,X317)) )),
  inference(superposition,[],[f188,f1850])).
fof(f8079,plain,(
  ( ! [X114,X112,X113] : (s(bool,X112) = s(bool,X114) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X112)),X113,bool)) | s(bool,'F_0') = s(bool,X114)) )),
  inference(superposition,[],[f7068,f1561])).
fof(f8156,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X5)))) | s(bool,X7) = s(bool,sK9(X7,X6,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,X5) = s(bool,X8) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f304,f7068])).
fof(f8157,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,X9))) != s('type.semanticPrimitives.result'(bool,X10),'const.semanticPrimitives.Rval_1'(s(bool,X11))) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,X9) = s(bool,X13) | s(bool,'F_0') = s(bool,X13)) )),
  inference(superposition,[],[f321,f7068])).
fof(f8158,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X14) != s(bool,X15) | s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X15))) = s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,X14) = s(bool,X18) | s(bool,'F_0') = s(bool,X18)) )),
  inference(superposition,[],[f346,f7068])).
fof(f8162,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.semanticPrimitives.result'(bool,X31),'const.semanticPrimitives.Rval_1'(s(bool,X30))) = s('type.semanticPrimitives.result'(bool,X31),'const.semanticPrimitives.Rval_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f2753,f7068])).
fof(f8219,plain,(
  ( ! [X243,X241,X242,X240] : (~p(s(bool,happ(s(fun(bool,bool),X241),s(bool,X240)))) | p(s(bool,happ(s(fun(bool,bool),X241),s(bool,X242)))) | s(bool,'F_0') != s(bool,X242) | s(bool,X240) = s(bool,X243) | s(bool,'F_0') = s(bool,X243)) )),
  inference(superposition,[],[f172,f7068])).
fof(f8220,plain,(
  ( ! [X246,X244,X245] : (~p(s(bool,happ(s(fun(bool,bool),X245),s(bool,X244)))) | s(bool,'F_0') = s(bool,sK9('F_0',X245,bool)) | s(bool,X244) = s(bool,X246) | s(bool,'F_0') = s(bool,X246)) )),
  inference(superposition,[],[f173,f7068])).
fof(f8223,plain,(
  ( ! [X255,X257,X256,X258] : (s(bool,X255) != s(bool,X256) | s('type.semanticPrimitives.result'(bool,X257),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X257),'const.semanticPrimitives.Rval_1'(s(bool,X256))) | s(bool,X255) = s(bool,X258) | s(bool,'F_0') = s(bool,X258)) )),
  inference(superposition,[],[f188,f7068])).
fof(f8226,plain,(
  ( ! [X277,X279,X278] : (~p(s(bool,happ(s(fun(bool,bool),X278),s(bool,X277)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X278),s(bool,sK9('F_0',X278,bool)))) | s(bool,X277) = s(bool,X279) | s(bool,'F_0') = s(bool,X279)) )),
  inference(superposition,[],[f651,f7068])).
fof(f8326,plain,(
  ( ! [X333,X331,X332,X334] : (~p(s(bool,happ(s(fun(bool,bool),X332),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X332),s(bool,X333)))) | s(bool,X331) != s(bool,X333) | s(bool,X331) = s(bool,X334) | s(bool,'F_0') = s(bool,X334)) )),
  inference(superposition,[],[f172,f7068])).
fof(f8327,plain,(
  ( ! [X335,X337,X336] : (~p(s(bool,happ(s(fun(bool,bool),X336),s(bool,'F_0')))) | s(bool,X335) = s(bool,sK9(X335,X336,bool)) | s(bool,X335) = s(bool,X337) | s(bool,'F_0') = s(bool,X337)) )),
  inference(superposition,[],[f173,f7068])).
fof(f8328,plain,(
  ( ! [X341,X339,X340,X338] : (s('type.semanticPrimitives.result'(bool,X339),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X339),'const.semanticPrimitives.Rval_1'(s(bool,X340))) | s(bool,X338) = s(bool,X340) | s(bool,X338) = s(bool,X341) | s(bool,'F_0') = s(bool,X341)) )),
  inference(superposition,[],[f187,f7068])).
fof(f8330,plain,(
  ( ! [X346,X349,X347,X348] : (s(bool,'F_0') != s(bool,X347) | s('type.semanticPrimitives.result'(bool,X348),'const.semanticPrimitives.Rval_1'(s(bool,X346))) = s('type.semanticPrimitives.result'(bool,X348),'const.semanticPrimitives.Rval_1'(s(bool,X347))) | s(bool,X346) = s(bool,X349) | s(bool,'F_0') = s(bool,X349)) )),
  inference(superposition,[],[f188,f7068])).
fof(f8335,plain,(
  ( ! [X372,X370,X371] : (~p(s(bool,happ(s(fun(bool,bool),X371),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X371),s(bool,sK9(X370,X371,bool)))) | s(bool,X370) = s(bool,X372) | s(bool,'F_0') = s(bool,X372)) )),
  inference(superposition,[],[f651,f7068])).
fof(f8339,plain,(
  ( ! [X391,X393,X392,X394] : (s(bool,'F_0') != s(bool,f(s(X391,sK12(X392,X393,bool,X391)))) | s(fun(X391,bool),X392) = s(fun(X391,bool),X393) | s(bool,X394) = s(bool,g(s(X391,sK12(X392,X393,bool,X391)))) | s(bool,'F_0') = s(bool,X394)) )),
  inference(superposition,[],[f186,f7068])).
fof(f8340,plain,(
  ( ! [X401,X403,X402] : (p(s(bool,happ(s(fun(bool,bool),X402),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X402),s(bool,X401)))) | s(bool,X403) = s(bool,sK9(X401,X402,bool)) | s(bool,'F_0') = s(bool,X403)) )),
  inference(superposition,[],[f174,f7068])).
fof(f8341,plain,(
  ( ! [X405,X404,X406] : (s('type.semanticPrimitives.result'(bool,X405),X404) = s('type.semanticPrimitives.result'(bool,X405),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.result'(bool,X405),X404) = s('type.semanticPrimitives.result'(bool,X405),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X405),sK11(X404,X405,bool)))) | s(bool,X406) = s(bool,sK10(X404,X405,bool)) | s(bool,'F_0') = s(bool,X406)) )),
  inference(superposition,[],[f185,f7068])).
fof(f8343,plain,(
  ( ! [X412,X410,X413,X411] : (s(X412,f(s(bool,'F_0'))) != s(X412,g(s(bool,'F_0'))) | s(fun(bool,X412),X410) = s(fun(bool,X412),X411) | s(bool,X413) = s(bool,sK12(X410,X411,X412,bool)) | s(bool,'F_0') = s(bool,X413)) )),
  inference(superposition,[],[f186,f7068])).
fof(f8466,plain,(
  ( ! [X436,X438,X437,X439] : (s(bool,X439) != s(bool,f(s(X436,sK12(X437,X438,bool,X436)))) | s(fun(X436,bool),X437) = s(fun(X436,bool),X438) | s(bool,'F_0') = s(bool,g(s(X436,sK12(X437,X438,bool,X436)))) | s(bool,'F_0') = s(bool,X439)) )),
  inference(superposition,[],[f186,f7068])).
fof(f9136,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f8230,f133])).
fof(f8230,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f175,f7068])).
fof(f9177,plain,(
  ( ! [X255,X253,X254,X252] : (~p(s(bool,happ(s(fun(X252,bool),X253),s(X252,X254)))) | s(bool,X255) = s(bool,happ(s(fun(X252,bool),X253),s(X252,sK9(X254,X253,X252)))) | s(bool,'F_0') = s(bool,X255)) )),
  inference(subsumption_resolution,[],[f8304,f133])).
fof(f8304,plain,(
  ( ! [X255,X253,X254,X252] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X252,bool),X253),s(X252,X254)))) | s(bool,X255) = s(bool,happ(s(fun(X252,bool),X253),s(X252,sK9(X254,X253,X252)))) | s(bool,'F_0') = s(bool,X255)) )),
  inference(superposition,[],[f174,f7068])).
fof(f9526,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X2,X3,bool)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f641,f233])).
fof(f9550,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X81)),X82,bool),X83,bool)) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f641,f1636])).
fof(f9559,plain,(
  ( ! [X116,X114,X115] : (~p(s(bool,happ(s(fun(bool,bool),X116),s(bool,X114)))) | p(s(bool,happ(s(fun(bool,bool),X116),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X114)),X115,bool),X116,bool))) )),
  inference(superposition,[],[f641,f1561])).
fof(f9577,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK9(X2,sK10('const.semanticPrimitives.Rval_1'(s(fun(bool,bool),X0)),X1,fun(bool,bool)),bool))) )),
  inference(forward_demodulation,[],[f9524,f1561])).
fof(f9524,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),sK10('const.semanticPrimitives.Rval_1'(s(fun(bool,bool),X0)),X1,fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X2,sK10('const.semanticPrimitives.Rval_1'(s(fun(bool,bool),X0)),X1,fun(bool,bool)),bool))) )),
  inference(superposition,[],[f641,f1561])).
fof(f9583,plain,(
  ( ! [X30,X33,X31] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X30,X33,bool)) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f9582,f172])).
fof(f9582,plain,(
  ( ! [X30,X33,X31] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X30,X33,bool)) | s(bool,'F_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X31)) )),
  inference(forward_demodulation,[],[f9534,f1561])).
fof(f9534,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X31)),X32,bool))))) | p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X30,X33,bool)) | s(bool,'F_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f641,f1850])).
fof(f9590,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X95)))) | p(s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X93)),X94,bool),X96,bool)) | s(bool,'F_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f9553,f172])).
fof(f9553,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),X96),s(bool,X95)))) | p(s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X93)),X94,bool),X96,bool)) | s(bool,'F_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f641,f1850])).
fof(f9594,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X3,X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f9561,f131])).
fof(f9561,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X3,X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f641,f233])).
fof(f9630,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK12(X2,X3,X1,bool)) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f1104,f7068])).
fof(f9844,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X4) | s(bool,X0) = s(bool,X4)) )),
  inference(superposition,[],[f961,f7068])).
fof(f9881,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f9860,f131])).
fof(f9860,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f961,f233])).
fof(f1106,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s(X4,sK12(X5,X6,bool,X4)))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,g(s(X4,sK12(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f186,f233])).
fof(f1108,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) != s(bool,f(s(X12,sK12(X13,X14,bool,X12)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,g(s(X12,sK12(X13,X14,bool,X12)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f186,f233])).
fof(f4952,plain,(
  ( ! [X360,X362,X359,X361,X363] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X362)),X363,bool)) != s(bool,f(s(X359,sK12(X360,X361,bool,X359)))) | s(fun(X359,bool),X360) = s(fun(X359,bool),X361) | s(bool,'T_0') = s(bool,X362) | s(bool,'T_0') = s(bool,g(s(X359,sK12(X360,X361,bool,X359))))) )),
  inference(superposition,[],[f186,f1635])).
fof(f5426,plain,(
  ( ! [X412,X410,X413,X409,X411] : (s(bool,X412) != s(bool,f(s(X409,sK12(X410,X411,bool,X409)))) | s(fun(X409,bool),X410) = s(fun(X409,bool),X411) | s(bool,'T_0') = s(bool,X412) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(X409,sK12(X410,X411,bool,X409))))),X413,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f5523,plain,(
  ( ! [X408,X405,X407,X404,X406] : (s(bool,X407) != s(bool,f(s(X404,sK12(X405,X406,bool,X404)))) | s(fun(X404,bool),X405) = s(fun(X404,bool),X406) | s(bool,'T_0') = s(bool,g(s(X404,sK12(X405,X406,bool,X404)))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X407)),X408,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f5657,plain,(
  ( ! [X389,X391,X387,X388,X390] : (s(bool,'T_0') != s(bool,f(s(X387,sK12(X388,X389,bool,X387)))) | s(fun(X387,bool),X388) = s(fun(X387,bool),X389) | s(bool,X390) = s(bool,g(s(X387,sK12(X388,X389,bool,X387)))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X390)),X391,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f453,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,X40)))) | s(bool,'T_0') = s(bool,sK9('T_0',X41,bool)) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f173,f233])).
fof(f637,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK9(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f174,f233])).
fof(f957,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) | s(bool,X4) != s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f172,f233])).
fof(f962,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X23)))) | p(s(bool,happ(s(fun(bool,bool),X24),s(bool,X25)))) | s(bool,'T_0') != s(bool,X25) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f172,f233])).
fof(f2314,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X19),s(bool,sK9('T_0',X19,bool)))) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f651,f233])).
fof(f474,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0')))) | s(bool,X53) = s(bool,sK9(X53,X54,bool)) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f173,f233])).
fof(f2310,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,sK9(X6,X7,bool)))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f651,f233])).
fof(f6161,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,X234) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X234)),X235,bool),X236,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,'T_0')))) | s(bool,X234) = s(bool,X237) | s(bool,'T_0') = s(bool,X237)) )),
  inference(forward_demodulation,[],[f5716,f1561])).
fof(f5716,plain,(
  ( ! [X237,X235,X236,X234] : (~p(s(bool,happ(s(fun(bool,bool),X236),s(bool,'T_0')))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X234)),X235,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X234)),X235,bool),X236,bool)) | s(bool,X234) = s(bool,X237) | s(bool,'T_0') = s(bool,X237)) )),
  inference(superposition,[],[f173,f1636])).
fof(f6184,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X62),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X60)),X61,bool),X62,bool)))) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f651,f1636])).
fof(f172,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK9(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X23,X14] : (? [X18] : (s(X6,X14) = s(X6,X18) & p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X18))))) <=> p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X14)))))),
  file('thm.evalProps.map_result_Rerr','thm.bool.UNWIND_THM2')).
fof(f173,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK9(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f117])).
fof(f174,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f974,plain,(
  ( ! [X6,X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X5)))) | s(X2,X4) != s(X2,X5) | s(bool,X6) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f967,f131])).
fof(f967,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X5)))) | s(X2,X4) != s(X2,X5) | s(bool,X6) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f172,f233])).
fof(f1664,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(X93,bool),X94),s(X93,X96)))) | s(X93,X96) = s(X93,sK9(X96,sK10('const.semanticPrimitives.Rval_1'(s(fun(X93,bool),X94)),X95,fun(X93,bool)),X93))) )),
  inference(superposition,[],[f173,f1561])).
fof(f1686,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X9) = s(X8,sK9(sK10('const.semanticPrimitives.Rval_1'(s(X8,X9)),X10,X8),X11,X8)) | ~p(s(bool,happ(s(fun(X8,bool),X11),s(X8,X9))))) )),
  inference(forward_demodulation,[],[f1640,f1561])).
fof(f1640,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X11),s(X8,X9)))) | s(X8,sK10('const.semanticPrimitives.Rval_1'(s(X8,X9)),X10,X8)) = s(X8,sK9(sK10('const.semanticPrimitives.Rval_1'(s(X8,X9)),X10,X8),X11,X8))) )),
  inference(superposition,[],[f173,f1561])).
fof(f1698,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,happ(s(fun(X89,bool),X90),s(X89,X92)))) | p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK9(X92,sK10('const.semanticPrimitives.Rval_1'(s(fun(X89,bool),X90)),X91,fun(X89,bool)),X89)))))) )),
  inference(forward_demodulation,[],[f1663,f1561])).
fof(f1663,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(X89,bool),X90),s(X89,sK9(X92,sK10('const.semanticPrimitives.Rval_1'(s(fun(X89,bool),X90)),X91,fun(X89,bool)),X89))))) | ~p(s(bool,happ(s(fun(X89,bool),sK10('const.semanticPrimitives.Rval_1'(s(fun(X89,bool),X90)),X91,fun(X89,bool))),s(X89,X92))))) )),
  inference(superposition,[],[f174,f1561])).
fof(f2308,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X3),s(X0,sK9(sK10('const.semanticPrimitives.Rval_1'(s(X0,X1)),X2,X0),X3,X0))))) )),
  inference(superposition,[],[f651,f1561])).
fof(f2332,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X3,sK10('const.semanticPrimitives.Rval_1'(s(fun(X0,bool),X1)),X2,fun(X0,bool)),X0)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f2307,f1561])).
fof(f2307,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),sK10('const.semanticPrimitives.Rval_1'(s(fun(X0,bool),X1)),X2,fun(X0,bool))),s(X0,sK9(X3,sK10('const.semanticPrimitives.Rval_1'(s(fun(X0,bool),X1)),X2,fun(X0,bool)),X0))))) )),
  inference(superposition,[],[f651,f1561])).
fof(f2333,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK9(X5,X4,X3)))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2328,f131])).
fof(f2328,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,sK9(X5,X4,X3)))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f651,f233])).
fof(f6230,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),X35),s(X34,sK9(X36,X35,X34)))) | s(bool,X37) = s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool))) )),
  inference(subsumption_resolution,[],[f6212,f131])).
fof(f6212,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),X35),s(X34,sK9(X36,X35,X34)))) | s(bool,X37) = s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool))) )),
  inference(superposition,[],[f651,f1636])).
fof(f1763,plain,(
  ( ! [X45,X43,X46,X44] : (s('type.semanticPrimitives.error_result'(X43),X44) = s('type.semanticPrimitives.error_result'(X43),sK9(sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X43),X44)),X43,X45),X46,'type.semanticPrimitives.error_result'(X43))) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X43),bool),X46),s('type.semanticPrimitives.error_result'(X43),X44))))) )),
  inference(forward_demodulation,[],[f1751,f1712])).
fof(f1751,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X43),bool),X46),s('type.semanticPrimitives.error_result'(X43),X44)))) | s('type.semanticPrimitives.error_result'(X43),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X43),X44)),X43,X45)) = s('type.semanticPrimitives.error_result'(X43),sK9(sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X43),X44)),X43,X45),X46,'type.semanticPrimitives.error_result'(X43)))) )),
  inference(superposition,[],[f173,f1712])).
fof(f2326,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X72),bool),X75),s('type.semanticPrimitives.error_result'(X72),X73)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(X72),bool),X75),s('type.semanticPrimitives.error_result'(X72),sK9(sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X72),X73)),X72,X74),X75,'type.semanticPrimitives.error_result'(X72)))))) )),
  inference(superposition,[],[f651,f1712])).
fof(f1290,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X15,X16),bool),X18),s('type.semanticPrimitives.result'(X15,X16),X17)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X15,X16),bool),X18),s('type.semanticPrimitives.result'(X15,X16),X19)))) | s('type.semanticPrimitives.result'(X15,X16),X19) != s('type.semanticPrimitives.result'(X15,X16),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X16),sK11(X17,X16,X15)))) | s('type.semanticPrimitives.result'(X15,X16),X17) = s('type.semanticPrimitives.result'(X15,X16),'const.semanticPrimitives.Rval_1'(s(X15,sK10(X17,X16,X15))))) )),
  inference(superposition,[],[f172,f185])).
fof(f1291,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X20,X21),bool),X23),s('type.semanticPrimitives.result'(X20,X21),X22)))) | s('type.semanticPrimitives.result'(X20,X21),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X21),sK11(X22,X21,X20)))) = s('type.semanticPrimitives.result'(X20,X21),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X21),sK11(X22,X21,X20))),X23,'type.semanticPrimitives.result'(X20,X21))) | s('type.semanticPrimitives.result'(X20,X21),X22) = s('type.semanticPrimitives.result'(X20,X21),'const.semanticPrimitives.Rval_1'(s(X20,sK10(X22,X21,X20))))) )),
  inference(superposition,[],[f173,f185])).
fof(f1303,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X17,X18),bool),X20),s('type.semanticPrimitives.result'(X17,X18),X19)))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X17,X18),bool),X20),s('type.semanticPrimitives.result'(X17,X18),X21)))) | s('type.semanticPrimitives.result'(X17,X18),X21) != s('type.semanticPrimitives.result'(X17,X18),'const.semanticPrimitives.Rval_1'(s(X17,sK10(X19,X18,X17)))) | s('type.semanticPrimitives.result'(X17,X18),X19) = s('type.semanticPrimitives.result'(X17,X18),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X18),sK11(X19,X18,X17))))) )),
  inference(superposition,[],[f172,f185])).
fof(f1304,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X22,X23),bool),X25),s('type.semanticPrimitives.result'(X22,X23),X24)))) | s('type.semanticPrimitives.result'(X22,X23),'const.semanticPrimitives.Rval_1'(s(X22,sK10(X24,X23,X22)))) = s('type.semanticPrimitives.result'(X22,X23),sK9('const.semanticPrimitives.Rval_1'(s(X22,sK10(X24,X23,X22))),X25,'type.semanticPrimitives.result'(X22,X23))) | s('type.semanticPrimitives.result'(X22,X23),X24) = s('type.semanticPrimitives.result'(X22,X23),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X23),sK11(X24,X23,X22))))) )),
  inference(superposition,[],[f173,f185])).
fof(f1400,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X40,X41),bool),X47),s('type.semanticPrimitives.result'(X40,X41),'const.evalProps.map_result_3'(s(fun(X42,X40),X43),s(fun(X45,X41),X46),s('type.semanticPrimitives.result'(X42,X45),'const.semanticPrimitives.Rval_1'(s(X42,X44)))))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X40,X41),bool),X47),s('type.semanticPrimitives.result'(X40,X41),X48)))) | s('type.semanticPrimitives.result'(X40,X41),X48) != s('type.semanticPrimitives.result'(X40,X41),'const.semanticPrimitives.Rval_1'(s(X40,happ(s(fun(X42,X40),X43),s(X42,X44)))))) )),
  inference(superposition,[],[f172,f206])).
fof(f1407,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X0,X1),bool),X7),s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,happ(s(fun(X2,X0),X3),s(X2,X6)))))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X0,X1),bool),X7),s('type.semanticPrimitives.result'(X0,X1),X8)))) | s('type.semanticPrimitives.result'(X0,X1),X8) != s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X4,X1),X5),s('type.semanticPrimitives.result'(X2,X4),'const.semanticPrimitives.Rval_1'(s(X2,X6)))))) )),
  inference(superposition,[],[f172,f206])).
fof(f1448,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X40,X41),bool),X47),s('type.semanticPrimitives.result'(X40,X41),'const.evalProps.map_result_3'(s(fun(X45,X40),X46),s(fun(X42,X41),X43),s('type.semanticPrimitives.result'(X45,X42),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X42),X44)))))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X40,X41),bool),X47),s('type.semanticPrimitives.result'(X40,X41),X48)))) | s('type.semanticPrimitives.result'(X40,X41),X48) != s('type.semanticPrimitives.result'(X40,X41),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X41),'const.evalProps.map_error_result_2'(s(fun(X42,X41),X43),s('type.semanticPrimitives.error_result'(X42),X44)))))) )),
  inference(superposition,[],[f172,f207])).
fof(f1457,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X13,X11,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X5,X6),bool),X12),s('type.semanticPrimitives.result'(X5,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X9,X6),X10),s('type.semanticPrimitives.error_result'(X9),X11)))))))) | p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X5,X6),bool),X12),s('type.semanticPrimitives.result'(X5,X6),X13)))) | s('type.semanticPrimitives.result'(X5,X6),X13) != s('type.semanticPrimitives.result'(X5,X6),'const.evalProps.map_result_3'(s(fun(X7,X5),X8),s(fun(X9,X6),X10),s('type.semanticPrimitives.result'(X7,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11)))))) )),
  inference(superposition,[],[f172,f207])).
fof(f2318,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X29,X30),bool),X32),s('type.semanticPrimitives.result'(X29,X30),X31)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X29,X30),bool),X32),s('type.semanticPrimitives.result'(X29,X30),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X30),sK11(X31,X30,X29))),X32,'type.semanticPrimitives.result'(X29,X30))))) | s('type.semanticPrimitives.result'(X29,X30),X31) = s('type.semanticPrimitives.result'(X29,X30),'const.semanticPrimitives.Rval_1'(s(X29,sK10(X31,X30,X29))))) )),
  inference(superposition,[],[f651,f185])).
fof(f2320,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X41,X42),bool),X44),s('type.semanticPrimitives.result'(X41,X42),X43)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X41,X42),bool),X44),s('type.semanticPrimitives.result'(X41,X42),sK9('const.semanticPrimitives.Rval_1'(s(X41,sK10(X43,X42,X41))),X44,'type.semanticPrimitives.result'(X41,X42))))) | s('type.semanticPrimitives.result'(X41,X42),X43) = s('type.semanticPrimitives.result'(X41,X42),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X42),sK11(X43,X42,X41))))) )),
  inference(superposition,[],[f651,f185])).
fof(f1401,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X49,X50),bool),X56),s('type.semanticPrimitives.result'(X49,X50),'const.evalProps.map_result_3'(s(fun(X51,X49),X52),s(fun(X54,X50),X55),s('type.semanticPrimitives.result'(X51,X54),'const.semanticPrimitives.Rval_1'(s(X51,X53)))))))) | s('type.semanticPrimitives.result'(X49,X50),'const.semanticPrimitives.Rval_1'(s(X49,happ(s(fun(X51,X49),X52),s(X51,X53))))) = s('type.semanticPrimitives.result'(X49,X50),sK9('const.semanticPrimitives.Rval_1'(s(X49,happ(s(fun(X51,X49),X52),s(X51,X53)))),X56,'type.semanticPrimitives.result'(X49,X50)))) )),
  inference(superposition,[],[f173,f206])).
fof(f2321,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X45,X46),bool),X52),s('type.semanticPrimitives.result'(X45,X46),'const.evalProps.map_result_3'(s(fun(X47,X45),X48),s(fun(X50,X46),X51),s('type.semanticPrimitives.result'(X47,X50),'const.semanticPrimitives.Rval_1'(s(X47,X49)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X45,X46),bool),X52),s('type.semanticPrimitives.result'(X45,X46),sK9('const.semanticPrimitives.Rval_1'(s(X45,happ(s(fun(X47,X45),X48),s(X47,X49)))),X52,'type.semanticPrimitives.result'(X45,X46)))))) )),
  inference(superposition,[],[f651,f206])).
fof(f1449,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X49,X50),bool),X56),s('type.semanticPrimitives.result'(X49,X50),'const.evalProps.map_result_3'(s(fun(X54,X49),X55),s(fun(X51,X50),X52),s('type.semanticPrimitives.result'(X54,X51),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X51),X53)))))))) | s('type.semanticPrimitives.result'(X49,X50),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X51,X50),X52),s('type.semanticPrimitives.error_result'(X51),X53))))) = s('type.semanticPrimitives.result'(X49,X50),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X50),'const.evalProps.map_error_result_2'(s(fun(X51,X50),X52),s('type.semanticPrimitives.error_result'(X51),X53)))),X56,'type.semanticPrimitives.result'(X49,X50)))) )),
  inference(superposition,[],[f173,f207])).
fof(f2319,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X33,X34),bool),X40),s('type.semanticPrimitives.result'(X33,X34),'const.evalProps.map_result_3'(s(fun(X38,X33),X39),s(fun(X35,X34),X36),s('type.semanticPrimitives.result'(X38,X35),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X35),X37)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X33,X34),bool),X40),s('type.semanticPrimitives.result'(X33,X34),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X34),'const.evalProps.map_error_result_2'(s(fun(X35,X34),X36),s('type.semanticPrimitives.error_result'(X35),X37)))),X40,'type.semanticPrimitives.result'(X33,X34)))))) )),
  inference(superposition,[],[f651,f207])).
fof(f1458,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X14,X15),bool),X21),s('type.semanticPrimitives.result'(X14,X15),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X15),'const.evalProps.map_error_result_2'(s(fun(X18,X15),X19),s('type.semanticPrimitives.error_result'(X18),X20)))))))) | s('type.semanticPrimitives.result'(X14,X15),'const.evalProps.map_result_3'(s(fun(X16,X14),X17),s(fun(X18,X15),X19),s('type.semanticPrimitives.result'(X16,X18),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X18),X20))))) = s('type.semanticPrimitives.result'(X14,X15),sK9('const.evalProps.map_result_3'(s(fun(X16,X14),X17),s(fun(X18,X15),X19),s('type.semanticPrimitives.result'(X16,X18),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X18),X20)))),X21,'type.semanticPrimitives.result'(X14,X15)))) )),
  inference(superposition,[],[f173,f207])).
fof(f2323,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X61,X62),bool),X68),s('type.semanticPrimitives.result'(X61,X62),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X62),'const.evalProps.map_error_result_2'(s(fun(X65,X62),X66),s('type.semanticPrimitives.error_result'(X65),X67)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X61,X62),bool),X68),s('type.semanticPrimitives.result'(X61,X62),sK9('const.evalProps.map_result_3'(s(fun(X63,X61),X64),s(fun(X65,X62),X66),s('type.semanticPrimitives.result'(X63,X65),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X65),X67)))),X68,'type.semanticPrimitives.result'(X61,X62)))))) )),
  inference(superposition,[],[f651,f207])).
fof(f1408,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9,X16] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X9,X10),bool),X16),s('type.semanticPrimitives.result'(X9,X10),'const.semanticPrimitives.Rval_1'(s(X9,happ(s(fun(X11,X9),X12),s(X11,X15)))))))) | s('type.semanticPrimitives.result'(X9,X10),'const.evalProps.map_result_3'(s(fun(X11,X9),X12),s(fun(X13,X10),X14),s('type.semanticPrimitives.result'(X11,X13),'const.semanticPrimitives.Rval_1'(s(X11,X15))))) = s('type.semanticPrimitives.result'(X9,X10),sK9('const.evalProps.map_result_3'(s(fun(X11,X9),X12),s(fun(X13,X10),X14),s('type.semanticPrimitives.result'(X11,X13),'const.semanticPrimitives.Rval_1'(s(X11,X15)))),X16,'type.semanticPrimitives.result'(X9,X10)))) )),
  inference(superposition,[],[f173,f206])).
fof(f2322,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(X53,X54),bool),X60),s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rval_1'(s(X53,happ(s(fun(X55,X53),X56),s(X55,X59)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(X53,X54),bool),X60),s('type.semanticPrimitives.result'(X53,X54),sK9('const.evalProps.map_result_3'(s(fun(X55,X53),X56),s(fun(X57,X54),X58),s('type.semanticPrimitives.result'(X55,X57),'const.semanticPrimitives.Rval_1'(s(X55,X59)))),X60,'type.semanticPrimitives.result'(X53,X54)))))) )),
  inference(superposition,[],[f651,f206])).
fof(f5102,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X3,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f5101])).
fof(f5101,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X21,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X23,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(duplicate_literal_removal,[],[f4834])).
fof(f4834,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X21,bool)) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X23,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f1635,f1635])).
fof(f6682,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X21,bool)) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X23,bool)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(duplicate_literal_removal,[],[f6607])).
fof(f6607,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X21,bool)) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X23,bool)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f5732,f1635])).
fof(f1279,plain,(
  ( ! [X4,X2,X3] : (s('type.semanticPrimitives.result'(bool,X3),X2) = s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.result'(bool,X3),X2) = s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),sK11(X2,X3,bool)))) | s(bool,X4) = s(bool,sK10(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f185,f233])).
fof(f1280,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.result'(bool,X6),X5) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X7))) | s('type.semanticPrimitives.result'(bool,X6),X5) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),sK11(X5,X6,bool)))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK10(X5,X6,bool))) )),
  inference(superposition,[],[f185,f233])).
fof(f5428,plain,(
  ( ! [X428,X426,X425,X427] : (s('type.semanticPrimitives.result'(bool,X426),X425) = s('type.semanticPrimitives.result'(bool,X426),'const.semanticPrimitives.Rval_1'(s(bool,X427))) | s('type.semanticPrimitives.result'(bool,X426),X425) = s('type.semanticPrimitives.result'(bool,X426),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X426),sK11(X425,X426,bool)))) | s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,sK10(X425,X426,bool))),X428,bool))) )),
  inference(superposition,[],[f185,f1636])).
fof(f5525,plain,(
  ( ! [X420,X422,X421,X423] : (s('type.semanticPrimitives.result'(bool,X421),X420) = s('type.semanticPrimitives.result'(bool,X421),'const.semanticPrimitives.Rval_1'(s(bool,X422))) | s('type.semanticPrimitives.result'(bool,X421),X420) = s('type.semanticPrimitives.result'(bool,X421),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X421),sK11(X420,X421,bool)))) | s(bool,'T_0') = s(bool,sK10(X420,X421,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X422)),X423,bool))) )),
  inference(superposition,[],[f185,f1636])).
fof(f5659,plain,(
  ( ! [X405,X403,X404,X402] : (s('type.semanticPrimitives.result'(bool,X403),X402) = s('type.semanticPrimitives.result'(bool,X403),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.result'(bool,X403),X402) = s('type.semanticPrimitives.result'(bool,X403),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X403),sK11(X402,X403,bool)))) | s(bool,X404) = s(bool,sK10(X402,X403,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X404)),X405,bool))) )),
  inference(superposition,[],[f185,f1636])).
fof(f5430,plain,(
  ( ! [X436,X434,X437,X433,X435] : (s(X435,f(s(bool,X436))) != s(X435,g(s(bool,X436))) | s(fun(bool,X435),X433) = s(fun(bool,X435),X434) | s(bool,'T_0') = s(bool,X436) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,sK12(X433,X434,X435,bool))),X437,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f628,plain,(
  ( ! [X30,X31,X29,X32] : (s(X29,X31) = s(X29,sK9(X31,X30,X29)) | s(bool,X32) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f467,f131])).
fof(f467,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | s(X29,X31) = s(X29,sK9(X31,X30,X29)) | s(bool,X32) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f173,f233])).
fof(f6128,plain,(
  ( ! [X269,X271,X267,X268,X270] : (s(X267,X269) = s(X267,sK9(X269,X268,X267)) | s(bool,X270) = s(bool,happ(s(fun(X267,bool),X268),s(X267,X269))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X270)),X271,bool))) )),
  inference(subsumption_resolution,[],[f5633,f131])).
fof(f5633,plain,(
  ( ! [X269,X271,X267,X268,X270] : (~p(s(bool,'T_0')) | s(X267,X269) = s(X267,sK9(X269,X268,X267)) | s(bool,X270) = s(bool,happ(s(fun(X267,bool),X268),s(X267,X269))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X270)),X271,bool))) )),
  inference(superposition,[],[f173,f1636])).
fof(f6183,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,X58)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X56)),X57,bool),X59,bool)))) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f651,f1635])).
fof(f5212,plain,(
  ( ! [X231,X233,X230,X232] : (s(bool,X230) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X230)),X231,bool),X233,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X233),s(bool,X232)))) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X232)) )),
  inference(forward_demodulation,[],[f5014,f1561])).
fof(f5014,plain,(
  ( ! [X231,X233,X230,X232] : (~p(s(bool,happ(s(fun(bool,bool),X233),s(bool,X232)))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X230)),X231,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X230)),X231,bool),X233,bool)) | s(bool,'T_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f173,f1635])).
fof(f2316,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(bool,bool),X25),s(bool,X24)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,sK9('F_0',X25,bool)))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f651,f134])).
fof(f2311,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,sK9(X9,X11,bool)))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f651,f233])).
fof(f959,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X13)))) | p(s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) | s(bool,X12) != s(bool,X15) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f172,f233])).
fof(f639,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'T_0') = s(bool,sK9(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f174,f233])).
fof(f417,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X57),s(bool,X56)))) | s(bool,X55) = s(bool,sK9(X55,X57,bool)) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f173,f233])).
fof(f651,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK9(X18,X17,X16))))) )),
  inference(subsumption_resolution,[],[f646,f133])).
fof(f646,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,sK9(X18,X17,X16))))) )),
  inference(superposition,[],[f174,f134])).
fof(f9549,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X80),s(bool,X79)))) | p(s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X77)),X78,bool),X80,bool)) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f641,f1635])).
fof(f9547,plain,(
  ( ! [X72,X73] : (~p(s(bool,happ(s(fun(bool,bool),X73),s(bool,X72)))) | p(s(bool,happ(s(fun(bool,bool),X73),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9('F_0',X73,bool)) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f641,f134])).
fof(f9545,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X67),s(bool,X66)))) | p(s(bool,happ(s(fun(bool,bool),X67),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9('F_0',X67,bool)) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) = s(bool,X68)) )),
  inference(superposition,[],[f641,f7068])).
fof(f9527,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X5,X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f641,f233])).
fof(f1933,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK9('F_0',sK10('const.semanticPrimitives.Rval_1'(s(fun(bool,bool),X0)),X1,fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f303,f1561])).
fof(f1278,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),X0) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X0,X1,bool)))) | s('type.semanticPrimitives.result'(bool,X1),X0) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X0,X1,bool))) )),
  inference(superposition,[],[f185,f134])).
fof(f1110,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,f(s(X20,sK12(X21,X22,bool,X20)))) | s(fun(X20,bool),X21) = s(fun(X20,bool),X22) | s(bool,'T_0') = s(bool,g(s(X20,sK12(X21,X22,bool,X20))))) )),
  inference(superposition,[],[f186,f134])).
fof(f1104,plain,(
  ( ! [X21,X22,X20] : (s(X22,f(s(bool,'F_0'))) != s(X22,g(s(bool,'F_0'))) | s(fun(bool,X22),X20) = s(fun(bool,X22),X21) | s(bool,'T_0') = s(bool,sK12(X20,X21,X22,bool))) )),
  inference(superposition,[],[f186,f134])).
fof(f964,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X31)))) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X33)))) | s(bool,'F_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f172,f134])).
fof(f961,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X21),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X22)))) | s(bool,X20) != s(bool,X22) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f172,f134])).
fof(f641,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(X12,X13,bool))) )),
  inference(superposition,[],[f174,f134])).
fof(f321,plain,(
  ( ! [X8,X7,X9] : (s('type.semanticPrimitives.result'(bool,X8),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X8),'const.semanticPrimitives.Rval_1'(s(bool,X9))) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f187,f134])).
fof(f320,plain,(
  ( ! [X6,X4,X5] : (s('type.semanticPrimitives.result'(bool,X5),'const.semanticPrimitives.Rval_1'(s(bool,X4))) != s('type.semanticPrimitives.result'(bool,X5),'const.semanticPrimitives.Rval_1'(s(bool,X6))) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f187,f134])).
fof(f304,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | s(bool,X4) = s(bool,sK9(X4,X5,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f173,f134])).
fof(f6685,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f6602])).
fof(f6602,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5732,f134])).
fof(f2543,plain,(
  ( ! [X23,X25] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X23)),X25,bool)) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f1777,f1561])).
fof(f9548,plain,(
  ( ! [X76,X74,X75] : (~p(s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X76),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X74)),X75,bool),X76,bool)) | s(bool,'F_0') = s(bool,X74)) )),
  inference(superposition,[],[f641,f1634])).
fof(f9396,plain,(
  ( ! [X14,X12,X10,X8,X19,X17,X15,X13,X11,X9,X18,X16] : (s('type.semanticPrimitives.result'(X15,X16),'const.semanticPrimitives.Rval_1'(s(X15,X17))) != s('type.semanticPrimitives.result'(X15,X16),'const.evalProps.map_result_3'(s(fun(X8,X15),X18),s(fun(X9,X16),X19),s('type.semanticPrimitives.result'(X8,X9),'const.evalProps.map_result_3'(s(fun(X13,X8),X14),s(fun(X10,X9),X11),s('type.semanticPrimitives.result'(X13,X10),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X10),X12)))))))) )),
  inference(superposition,[],[f1447,f207])).
fof(f9395,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'(X3,X4),'const.semanticPrimitives.Rval_1'(s(X3,X5))) != s('type.semanticPrimitives.result'(X3,X4),'const.evalProps.map_result_3'(s(fun(X0,X3),X6),s(fun(X1,X4),X7),s('type.semanticPrimitives.result'(X0,X1),X2))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f1447,f185])).
fof(f9159,plain,(
  ( ! [X109,X107,X108] : (s(bool,X107) != s(bool,X108) | s(bool,'F_0') = s(bool,X108) | s(bool,X107) = s(bool,X109) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f8261,f133])).
fof(f8261,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,'F_0')) | s(bool,X107) != s(bool,X108) | s(bool,'F_0') = s(bool,X108) | s(bool,X107) = s(bool,X109) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f1815,f7068])).
fof(f8006,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X3) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f7068,f7068])).
fof(f7885,plain,(
  ( ! [X14,X12,X10,X8,X19,X17,X15,X13,X11,X9,X18,X16] : (s('type.semanticPrimitives.result'(X15,X16),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X16),X17))) != s('type.semanticPrimitives.result'(X15,X16),'const.evalProps.map_result_3'(s(fun(X8,X15),X18),s(fun(X9,X16),X19),s('type.semanticPrimitives.result'(X8,X9),'const.evalProps.map_result_3'(s(fun(X10,X8),X11),s(fun(X13,X9),X14),s('type.semanticPrimitives.result'(X10,X13),'const.semanticPrimitives.Rval_1'(s(X10,X12)))))))) )),
  inference(superposition,[],[f1399,f206])).
fof(f7884,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'(X3,X4),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X4),X5))) != s('type.semanticPrimitives.result'(X3,X4),'const.evalProps.map_result_3'(s(fun(X0,X3),X6),s(fun(X1,X4),X7),s('type.semanticPrimitives.result'(X0,X1),X2))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0))))) )),
  inference(superposition,[],[f1399,f185])).
fof(f7838,plain,(
  ( ! [X14,X12,X10,X8,X19,X17,X15,X13,X11,X9,X18,X16] : (s('type.semanticPrimitives.result'(X8,X9),'const.evalProps.map_result_3'(s(fun(X15,X8),X16),s(fun(X17,X9),X18),s('type.semanticPrimitives.result'(X15,X17),'const.semanticPrimitives.Rval_1'(s(X15,X19))))) != s('type.semanticPrimitives.result'(X8,X9),'const.evalProps.map_result_3'(s(fun(X13,X8),X14),s(fun(X10,X9),X11),s('type.semanticPrimitives.result'(X13,X10),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X10),X12)))))) )),
  inference(superposition,[],[f1399,f207])).
fof(f7752,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.result'(X4,X5),X11) != s('type.semanticPrimitives.result'(X4,X5),'const.evalProps.map_result_3'(s(fun(X9,X4),X10),s(fun(X6,X5),X7),s('type.semanticPrimitives.result'(X9,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),X8))))) | s('type.semanticPrimitives.result'(X4,X5),X11) = s('type.semanticPrimitives.result'(X4,X5),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X5),sK11(X11,X5,X4))))) )),
  inference(superposition,[],[f1299,f207])).
fof(f7757,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.semanticPrimitives.result'(X53,X54),X55) = s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X54),sK11(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54)),X54,X53)))) | s('type.semanticPrimitives.result'(X53,X54),X55) != s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X54),X57)))) )),
  inference(forward_demodulation,[],[f7747,f1561])).
fof(f7747,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.semanticPrimitives.result'(X53,X54),X55) != s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X54),X57))) | s('type.semanticPrimitives.result'(X53,X54),sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54))) = s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X54),sK11(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54)),X54,X53))))) )),
  inference(superposition,[],[f1299,f1561])).
fof(f7734,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X4,X1),X5),s('type.semanticPrimitives.result'(X2,X4),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X4),X6))))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11('const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X4,X1),X5),s('type.semanticPrimitives.result'(X2,X4),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X4),X6)))),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f207,f1299])).
fof(f7754,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),X3,'type.semanticPrimitives.result'(X0,X1)),X1,X0))))) )),
  inference(forward_demodulation,[],[f7736,f1561])).
fof(f7736,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),X3,'type.semanticPrimitives.result'(X0,X1))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),X3,'type.semanticPrimitives.result'(X0,X1)),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f1561,f1299])).
fof(f7159,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X72) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X70)),X71,bool)) | s('type.semanticPrimitives.result'(bool,X73),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X73),'const.semanticPrimitives.Rval_1'(s(bool,X69))) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f6869,f188])).
fof(f6869,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X72) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X70)),X71,bool)) | s('type.semanticPrimitives.result'(bool,X73),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X73),'const.semanticPrimitives.Rval_1'(s(bool,X69))) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X70)) )),
  inference(superposition,[],[f345,f1850])).
fof(f6768,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X5,X6),X12),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X5,X6),'const.semanticPrimitives.Rval_1'(s(X5,X13))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X5,X6),X12),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X5,X6),'const.evalProps.map_result_3'(s(fun(X10,X5),X11),s(fun(X7,X6),X8),s('type.semanticPrimitives.result'(X10,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X9)))))))) )),
  inference(superposition,[],[f316,f207])).
fof(f6767,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),X2))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f316,f185])).
fof(f6761,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X5,X6),X12),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X5,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),X13))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X5,X6),X12),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X5,X6),'const.evalProps.map_result_3'(s(fun(X7,X5),X8),s(fun(X10,X6),X11),s('type.semanticPrimitives.result'(X7,X10),'const.semanticPrimitives.Rval_1'(s(X7,X9)))))))) )),
  inference(superposition,[],[f316,f206])).
fof(f6760,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),X2))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X4))))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0))))) )),
  inference(superposition,[],[f316,f185])).
fof(f6725,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),X3),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f316,f187])).
fof(f6647,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X108) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X106)),X109,bool)) | s(bool,X108) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X106)),X107,bool)) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f5732,f1561])).
fof(f6622,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,X102) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X104)),X105,bool)) | s(bool,X104) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X102)),X103,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X102)),X103,bool))) )),
  inference(superposition,[],[f5732,f1561])).
fof(f6523,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.result'(X4,X5),X11) != s('type.semanticPrimitives.result'(X4,X5),'const.evalProps.map_result_3'(s(fun(X6,X4),X7),s(fun(X9,X5),X10),s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rval_1'(s(X6,X8))))) | s('type.semanticPrimitives.result'(X4,X5),X11) = s('type.semanticPrimitives.result'(X4,X5),'const.semanticPrimitives.Rval_1'(s(X4,sK10(X11,X5,X4))))) )),
  inference(superposition,[],[f1288,f206])).
fof(f6522,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) != s('type.semanticPrimitives.result'(X0,X1),X3) | s('type.semanticPrimitives.result'(X0,X1),X3) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(X3,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0))))) )),
  inference(superposition,[],[f1288,f185])).
fof(f6538,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.semanticPrimitives.result'(X53,X54),X55) = s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rval_1'(s(X53,sK10(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54)),X54,X53)))) | s('type.semanticPrimitives.result'(X53,X54),X55) != s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rval_1'(s(X53,X57)))) )),
  inference(forward_demodulation,[],[f6483,f1561])).
fof(f6483,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.semanticPrimitives.result'(X53,X54),X55) != s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rval_1'(s(X53,X57))) | s('type.semanticPrimitives.result'(X53,X54),sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54))) = s('type.semanticPrimitives.result'(X53,X54),'const.semanticPrimitives.Rval_1'(s(X53,sK10(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X53,X54),X55)),X56,'type.semanticPrimitives.result'(X53,X54)),X54,X53))))) )),
  inference(superposition,[],[f1288,f1561])).
fof(f6469,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X4,X1),X5),s('type.semanticPrimitives.result'(X2,X4),'const.semanticPrimitives.Rval_1'(s(X2,X6))))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10('const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X4,X1),X5),s('type.semanticPrimitives.result'(X2,X4),'const.semanticPrimitives.Rval_1'(s(X2,X6)))),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f206,f1288])).
fof(f6537,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X2)))),X3,'type.semanticPrimitives.result'(X0,X1)),X1,X0))))) )),
  inference(forward_demodulation,[],[f6470,f1561])).
fof(f6470,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X2)))),X3,'type.semanticPrimitives.result'(X0,X1))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK10(sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X2)))),X3,'type.semanticPrimitives.result'(X0,X1)),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f1561,f1288])).
fof(f6377,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X6)),X8,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X7)),X9,bool)) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f6299,f233])).
fof(f6376,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X3,bool)) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X4)),X5,bool))) )),
  inference(superposition,[],[f6299,f1636])).
fof(f6262,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,X6) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X5)),X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,X5) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f5744,f233])).
fof(f6261,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X3)),X4,bool))) )),
  inference(superposition,[],[f5744,f1636])).
fof(f6299,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) )),
  inference(condensation,[],[f6298])).
fof(f6298,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X24,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X23)),X25,bool))) )),
  inference(duplicate_literal_removal,[],[f6244])).
fof(f6244,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X24,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X23)),X25,bool))) )),
  inference(superposition,[],[f5744,f1636])).
fof(f5732,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f1636])).
fof(f5744,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X1,bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f5726])).
fof(f5726,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X1,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f1636])).
fof(f5717,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s('type.semanticPrimitives.result'(bool,X240),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X240),'const.semanticPrimitives.Rval_1'(s(bool,X241))) | s(bool,X241) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X238)),X239,bool)) | s(bool,X238) = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) )),
  inference(superposition,[],[f187,f1636])).
fof(f6159,plain,(
  ( ! [X227,X228,X226] : (s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rval_1'(s(bool,X226))) | s(bool,X226) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f5714,f191])).
fof(f5714,plain,(
  ( ! [X227,X228,X226] : (s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rval_1'(s(bool,X226))) | s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rval_1'(s(bool,X226))) = s('type.semanticPrimitives.result'(bool,X227),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X227),sK11('const.semanticPrimitives.Rval_1'(s(bool,X226)),X227,bool)))) | s(bool,X226) = s(bool,X228) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f185,f1636])).
fof(f6144,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool)) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,X121) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f5690,f131])).
fof(f5690,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | s(bool,X120) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X118)),X119,bool)) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,X121) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f626,f1636])).
fof(f5661,plain,(
  ( ! [X412,X414,X410,X413,X411] : (s(X412,f(s(bool,'T_0'))) != s(X412,g(s(bool,'T_0'))) | s(fun(bool,X412),X410) = s(fun(bool,X412),X411) | s(bool,X413) = s(bool,sK12(X410,X411,X412,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X413)),X414,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f5652,plain,(
  ( ! [X358,X360,X357,X359,X361] : (s(bool,'T_0') != s(bool,X358) | s('type.semanticPrimitives.result'(bool,X359),'const.semanticPrimitives.Rval_1'(s(bool,X357))) = s('type.semanticPrimitives.result'(bool,X359),'const.semanticPrimitives.Rval_1'(s(bool,X358))) | s(bool,X357) = s(bool,X360) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X360)),X361,bool))) )),
  inference(superposition,[],[f188,f1636])).
fof(f5650,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s('type.semanticPrimitives.result'(bool,X348),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X348),'const.semanticPrimitives.Rval_1'(s(bool,X349))) | s(bool,X347) = s(bool,X349) | s(bool,X347) = s(bool,X350) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X350)),X351,bool))) )),
  inference(superposition,[],[f187,f1636])).
fof(f5567,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s(bool,X184) != s(bool,X185) | s('type.semanticPrimitives.result'(bool,X186),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X186),'const.semanticPrimitives.Rval_1'(s(bool,X185))) | s(bool,X184) = s(bool,X187) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X187)),X188,bool))) )),
  inference(superposition,[],[f188,f1636])).
fof(f5527,plain,(
  ( ! [X428,X430,X432,X429,X431] : (s(X430,f(s(bool,X431))) != s(X430,g(s(bool,X431))) | s(fun(bool,X430),X428) = s(fun(bool,X430),X429) | s(bool,'T_0') = s(bool,sK12(X428,X429,X430,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X431)),X432,bool))) )),
  inference(superposition,[],[f186,f1636])).
fof(f5421,plain,(
  ( ! [X379,X380,X382,X378,X381] : (s(bool,X379) != s(bool,X380) | s('type.semanticPrimitives.result'(bool,X381),'const.semanticPrimitives.Rval_1'(s(bool,X378))) = s('type.semanticPrimitives.result'(bool,X381),'const.semanticPrimitives.Rval_1'(s(bool,X380))) | s(bool,'T_0') = s(bool,X379) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X378)),X382,bool))) )),
  inference(superposition,[],[f188,f1636])).
fof(f5419,plain,(
  ( ! [X372,X368,X370,X369,X371] : (s('type.semanticPrimitives.result'(bool,X370),'const.semanticPrimitives.Rval_1'(s(bool,X369))) != s('type.semanticPrimitives.result'(bool,X370),'const.semanticPrimitives.Rval_1'(s(bool,X371))) | s(bool,X368) = s(bool,X371) | s(bool,'T_0') = s(bool,X369) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X368)),X372,bool))) )),
  inference(superposition,[],[f187,f1636])).
fof(f5271,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X6) = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f1636,f233])).
fof(f5270,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X2)),X3,bool)) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X4)),X5,bool))) )),
  inference(superposition,[],[f1636,f1636])).
fof(f5015,plain,(
  ( ! [X237,X235,X238,X236,X234] : (s('type.semanticPrimitives.result'(bool,X237),'const.semanticPrimitives.Rval_1'(s(bool,X236))) != s('type.semanticPrimitives.result'(bool,X237),'const.semanticPrimitives.Rval_1'(s(bool,X238))) | s(bool,X238) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X234)),X235,bool)) | s(bool,'T_0') = s(bool,X234) | s(bool,'T_0') = s(bool,X236)) )),
  inference(superposition,[],[f187,f1635])).
fof(f4947,plain,(
  ( ! [X327,X329,X331,X328,X330] : (s(bool,X330) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X328)),X329,bool)) | s('type.semanticPrimitives.result'(bool,X331),'const.semanticPrimitives.Rval_1'(s(bool,X327))) = s('type.semanticPrimitives.result'(bool,X331),'const.semanticPrimitives.Rval_1'(s(bool,X330))) | s(bool,'T_0') = s(bool,X328) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f188,f1635])).
fof(f3269,plain,(
  ( ! [X26,X24,X23] : (s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26) | s(bool,X23) != s(bool,X24) | s(bool,X23) = s(bool,X26)) )),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f172,f346,f345,f321,f320,f133,f303,f453,f641,f639,f637,f474,f417,f304,f173,f174,f628,f651,f957,f959,f961,f962,f964,f974,f1100,f1102,f1104,f1106,f1108,f1110,f1277,f1276,f1278,f1279,f1280,f1286,f1288,f1290,f1291,f1292,f1294,f1297,f1299,f1303,f1304,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1400,f1401,f1402,f1404,f1407,f1408,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1448,f1449,f1450,f1452,f1457,f1458,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1634,f1635,f1636,f1686,f1641,f1698,f1664,f1673,f1712,f1717,f1744,f1763,f1752,f1777,f1813,f1815,f1816,f1817,f1818,f1797,f1819,f1849,f1850,f1867,f1892,f1903,f1933,f1956,f2286,f2298,f2332,f2308,f2309,f2310,f2311,f2314,f2316,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2326,f2333,f2340,f2346,f2352,f2353,f2354,f2357,f2358,f2359,f2361,f2365,f2368,f2369,f2373,f2374,f2375,f2377,f2381,f2397,f2400,f2410,f2413,f2446,f2482,f2484,f2527,f2529,f2530,f2543,f2687,f2688,f2590,f2695,f2705,f2706,f2709,f2713,f2628,f2714,f2716,f2731,f2753,f2838,f2840,f2814,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738])).
fof(f2738,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X23) != s(bool,X24) | s('type.semanticPrimitives.result'(bool,X25),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X25),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f345,f233])).
fof(f3254,plain,(
  ( ! [X41,X38,X40] : (s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X38))) = s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X41))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f3253,f1819])).
fof(f3253,plain,(
  ( ! [X41,X38,X40] : (s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X38))) = s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X41))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f2900,f1561])).
fof(f2900,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,X41))) = s('type.semanticPrimitives.result'(bool,X40),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X38)),X39,bool)))) | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f346,f1634])).
fof(f2909,plain,(
  ( ! [X30,X28,X31,X32] : (s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,X28))) = s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,X31))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X28)),X32,bool))) )),
  inference(forward_demodulation,[],[f2903,f1561])).
fof(f2903,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,X31))) = s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X28)),X29,bool)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X28)),X32,bool))) )),
  inference(trivial_inequality_removal,[],[f2898])).
fof(f2898,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'F_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,X31))) = s('type.semanticPrimitives.result'(bool,X30),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X28)),X29,bool)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X28)),X32,bool))) )),
  inference(superposition,[],[f346,f1777])).
fof(f2907,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f2891])).
fof(f2891,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f346,f134])).
fof(f2838,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X71) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X69)),X70,bool)) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f2800,f133])).
fof(f2800,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,'F_0')) | s(bool,X71) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X69)),X70,bool)) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f627,f2543])).
fof(f2705,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,X76) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X74)),X75,bool)) | s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X74)),X78,bool))) )),
  inference(subsumption_resolution,[],[f2616,f133])).
fof(f2616,plain,(
  ( ! [X78,X76,X74,X77,X75] : (p(s(bool,'F_0')) | s(bool,X76) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X74)),X75,bool)) | s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X74)),X78,bool))) )),
  inference(superposition,[],[f627,f1777])).
fof(f1867,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1855,f133])).
fof(f1855,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1815,f134])).
fof(f1849,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1838,f131])).
fof(f1838,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1813,f233])).
fof(f1752,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.error_result'(X47),X50),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.error_result'(X47),X48))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.error_result'(X47),X50),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.error_result'(X47),X51))) | s('type.semanticPrimitives.error_result'(X47),X51) = s('type.semanticPrimitives.error_result'(X47),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X47),X48)),X47,X49))) )),
  inference(superposition,[],[f187,f1712])).
fof(f1744,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.semanticPrimitives.result'(X8,X5),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X5),X6))) != s('type.semanticPrimitives.result'(X8,X5),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X5),X9))) | s('type.semanticPrimitives.error_result'(X5),X9) = s('type.semanticPrimitives.error_result'(X5),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X5),X6)),X5,X7))) )),
  inference(superposition,[],[f189,f1712])).
fof(f1717,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.semanticPrimitives.result'(X10,X11),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X11),X12))) != s('type.semanticPrimitives.result'(X10,X11),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X11),X13))) | s('type.semanticPrimitives.error_result'(X11),X13) = s('type.semanticPrimitives.error_result'(X11),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X11),X12)),X11,X10))) )),
  inference(superposition,[],[f189,f1276])).
fof(f1712,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X0),X1)),X0,X2))) )),
  inference(unit_resulting_resolution,[],[f1276,f189])).
fof(f1673,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s('type.semanticPrimitives.result'(X152,X149),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X149),X150))) != s('type.semanticPrimitives.result'(X152,X149),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X149),X153))) | s('type.semanticPrimitives.error_result'(X149),X153) = s('type.semanticPrimitives.error_result'(X149),sK10('const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.error_result'(X149),X150)),X151,'type.semanticPrimitives.error_result'(X149)))) )),
  inference(superposition,[],[f189,f1561])).
fof(f1641,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.semanticPrimitives.result'(X12,X15),'const.semanticPrimitives.Rval_1'(s(X12,X13))) != s('type.semanticPrimitives.result'(X12,X15),'const.semanticPrimitives.Rval_1'(s(X12,X16))) | s(X12,X16) = s(X12,sK10('const.semanticPrimitives.Rval_1'(s(X12,X13)),X14,X12))) )),
  inference(superposition,[],[f187,f1561])).
fof(f1635,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X5)),X6,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1561,f233])).
fof(f1587,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.semanticPrimitives.result'(X10,X11),'const.semanticPrimitives.Rval_1'(s(X10,X12))) != s('type.semanticPrimitives.result'(X10,X11),'const.semanticPrimitives.Rval_1'(s(X10,X13))) | s(X10,X13) = s(X10,sK10('const.semanticPrimitives.Rval_1'(s(X10,X12)),X11,X10))) )),
  inference(superposition,[],[f187,f1277])).
fof(f1583,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X5))) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X7))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X5)),X6,bool))) )),
  inference(superposition,[],[f1277,f233])).
fof(f1582,plain,(
  ( ! [X4,X2,X3] : (s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X3),'const.semanticPrimitives.Rval_1'(s(bool,X2))) | s(bool,X4) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X2)),X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1277,f233])).
fof(f1581,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) )),
  inference(superposition,[],[f1277,f134])).
fof(f1461,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40] : (s('type.semanticPrimitives.result'(X40,X41),X47) != s('type.semanticPrimitives.result'(X40,X41),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X41),'const.evalProps.map_error_result_2'(s(fun(X44,X41),X45),s('type.semanticPrimitives.error_result'(X44),X46))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X40,X41),X48),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X40,X41),X47))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X40,X41),X48),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X40,X41),'const.evalProps.map_result_3'(s(fun(X42,X40),X43),s(fun(X44,X41),X45),s('type.semanticPrimitives.result'(X42,X44),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X44),X46)))))))) )),
  inference(superposition,[],[f188,f207])).
fof(f1459,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25,X22] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X22,X23),X29),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X22,X23),X30))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X22,X23),X29),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X22,X23),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X23),'const.evalProps.map_error_result_2'(s(fun(X26,X23),X27),s('type.semanticPrimitives.error_result'(X26),X28))))))) | s('type.semanticPrimitives.result'(X22,X23),X30) = s('type.semanticPrimitives.result'(X22,X23),'const.evalProps.map_result_3'(s(fun(X24,X22),X25),s(fun(X26,X23),X27),s('type.semanticPrimitives.result'(X24,X26),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X26),X28)))))) )),
  inference(superposition,[],[f187,f207])).
fof(f1452,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X75,X82] : (s('type.semanticPrimitives.result'(X75,X76),X82) != s('type.semanticPrimitives.result'(X75,X76),'const.evalProps.map_result_3'(s(fun(X80,X75),X81),s(fun(X77,X76),X78),s('type.semanticPrimitives.result'(X80,X77),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X77),X79))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X75,X76),X83),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X75,X76),X82))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X75,X76),X83),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X75,X76),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X76),'const.evalProps.map_error_result_2'(s(fun(X77,X76),X78),s('type.semanticPrimitives.error_result'(X77),X79)))))))) )),
  inference(superposition,[],[f188,f207])).
fof(f1450,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X57,X58),X64),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X57,X58),X65))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X57,X58),X64),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X57,X58),'const.evalProps.map_result_3'(s(fun(X62,X57),X63),s(fun(X59,X58),X60),s('type.semanticPrimitives.result'(X62,X59),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X59),X61))))))) | s('type.semanticPrimitives.result'(X57,X58),X65) = s('type.semanticPrimitives.result'(X57,X58),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X58),'const.evalProps.map_error_result_2'(s(fun(X59,X58),X60),s('type.semanticPrimitives.error_result'(X59),X61)))))) )),
  inference(superposition,[],[f187,f207])).
fof(f1447,plain,(
  ( ! [X28,X26,X24,X23,X21,X27,X25,X22] : (s('type.semanticPrimitives.result'(X21,X22),'const.semanticPrimitives.Rval_1'(s(X21,X28))) != s('type.semanticPrimitives.result'(X21,X22),'const.evalProps.map_result_3'(s(fun(X26,X21),X27),s(fun(X23,X22),X24),s('type.semanticPrimitives.result'(X26,X23),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X23),X25)))))) )),
  inference(superposition,[],[f191,f207])).
fof(f1445,plain,(
  ( ! [X6,X12,X10,X8,X7,X5,X11,X9] : (s('type.semanticPrimitives.result'(X5,X6),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X6),X12))) != s('type.semanticPrimitives.result'(X5,X6),'const.evalProps.map_result_3'(s(fun(X10,X5),X11),s(fun(X7,X6),X8),s('type.semanticPrimitives.result'(X10,X7),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X7),X9))))) | s('type.semanticPrimitives.error_result'(X6),X12) = s('type.semanticPrimitives.error_result'(X6),'const.evalProps.map_error_result_2'(s(fun(X7,X6),X8),s('type.semanticPrimitives.error_result'(X7),X9)))) )),
  inference(superposition,[],[f189,f207])).
fof(f1444,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (s('type.semanticPrimitives.result'(X18,X19),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X19),'const.evalProps.map_error_result_2'(s(fun(X12,X19),X20),s('type.semanticPrimitives.error_result'(X12),'const.evalProps.map_error_result_2'(s(fun(X13,X12),X14),s('type.semanticPrimitives.error_result'(X13),X15))))))) = s('type.semanticPrimitives.result'(X18,X19),'const.evalProps.map_result_3'(s(fun(X11,X18),X21),s(fun(X12,X19),X20),s('type.semanticPrimitives.result'(X11,X12),'const.evalProps.map_result_3'(s(fun(X16,X11),X17),s(fun(X13,X12),X14),s('type.semanticPrimitives.result'(X16,X13),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X13),X15)))))))) )),
  inference(superposition,[],[f207,f207])).
fof(f1437,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X5,X0),X6),s(fun(X2,X1),X3),s('type.semanticPrimitives.result'(X5,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X4))))) = s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X7,X0),X8),s(fun(X2,X1),X3),s('type.semanticPrimitives.result'(X7,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X4)))))) )),
  inference(superposition,[],[f207,f207])).
fof(f1436,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(X3,X1),X4),s('type.semanticPrimitives.error_result'(X3),X5))))))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X6,X0),X7),s(fun(X3,X1),X4),s('type.semanticPrimitives.result'(X6,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X5)))))))) )),
  inference(unit_resulting_resolution,[],[f207,f188])).
fof(f1411,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X42,X40] : (s('type.semanticPrimitives.result'(X35,X36),X42) != s('type.semanticPrimitives.result'(X35,X36),'const.semanticPrimitives.Rval_1'(s(X35,happ(s(fun(X37,X35),X38),s(X37,X41))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X35,X36),X43),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X35,X36),X42))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X35,X36),X43),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X35,X36),'const.evalProps.map_result_3'(s(fun(X37,X35),X38),s(fun(X39,X36),X40),s('type.semanticPrimitives.result'(X37,X39),'const.semanticPrimitives.Rval_1'(s(X37,X41)))))))) )),
  inference(superposition,[],[f188,f206])).
fof(f1409,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X17,X18),X24),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X17,X18),X25))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X17,X18),X24),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X17,X18),'const.semanticPrimitives.Rval_1'(s(X17,happ(s(fun(X19,X17),X20),s(X19,X23))))))) | s('type.semanticPrimitives.result'(X17,X18),X25) = s('type.semanticPrimitives.result'(X17,X18),'const.evalProps.map_result_3'(s(fun(X19,X17),X20),s(fun(X21,X18),X22),s('type.semanticPrimitives.result'(X19,X21),'const.semanticPrimitives.Rval_1'(s(X19,X23)))))) )),
  inference(superposition,[],[f187,f206])).
fof(f1404,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X75,X82] : (s('type.semanticPrimitives.result'(X75,X76),X82) != s('type.semanticPrimitives.result'(X75,X76),'const.evalProps.map_result_3'(s(fun(X77,X75),X78),s(fun(X80,X76),X81),s('type.semanticPrimitives.result'(X77,X80),'const.semanticPrimitives.Rval_1'(s(X77,X79))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X75,X76),X83),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X75,X76),X82))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X75,X76),X83),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X75,X76),'const.semanticPrimitives.Rval_1'(s(X75,happ(s(fun(X77,X75),X78),s(X77,X79)))))))) )),
  inference(superposition,[],[f188,f206])).
fof(f1402,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X57,X58),X64),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X57,X58),X65))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X57,X58),X64),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X57,X58),'const.evalProps.map_result_3'(s(fun(X59,X57),X60),s(fun(X62,X58),X63),s('type.semanticPrimitives.result'(X59,X62),'const.semanticPrimitives.Rval_1'(s(X59,X61))))))) | s('type.semanticPrimitives.result'(X57,X58),X65) = s('type.semanticPrimitives.result'(X57,X58),'const.semanticPrimitives.Rval_1'(s(X57,happ(s(fun(X59,X57),X60),s(X59,X61)))))) )),
  inference(superposition,[],[f187,f206])).
fof(f1399,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s('type.semanticPrimitives.result'(X16,X17),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X17),X23))) != s('type.semanticPrimitives.result'(X16,X17),'const.evalProps.map_result_3'(s(fun(X18,X16),X19),s(fun(X21,X17),X22),s('type.semanticPrimitives.result'(X18,X21),'const.semanticPrimitives.Rval_1'(s(X18,X20)))))) )),
  inference(superposition,[],[f191,f206])).
fof(f1397,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X7))) != s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X5,X1),X6),s('type.semanticPrimitives.result'(X2,X5),'const.semanticPrimitives.Rval_1'(s(X2,X4))))) | s(X0,X7) = s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))) )),
  inference(superposition,[],[f187,f206])).
fof(f1396,plain,(
  ( ! [X14,X12,X10,X8,X17,X7,X15,X13,X11,X9,X16] : (s('type.semanticPrimitives.result'(X14,X15),'const.semanticPrimitives.Rval_1'(s(X14,happ(s(fun(X7,X14),X16),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))))) = s('type.semanticPrimitives.result'(X14,X15),'const.evalProps.map_result_3'(s(fun(X7,X14),X16),s(fun(X8,X15),X17),s('type.semanticPrimitives.result'(X7,X8),'const.evalProps.map_result_3'(s(fun(X9,X7),X10),s(fun(X12,X8),X13),s('type.semanticPrimitives.result'(X9,X12),'const.semanticPrimitives.Rval_1'(s(X9,X11)))))))) )),
  inference(superposition,[],[f206,f206])).
fof(f1372,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X5,X1),X6),s('type.semanticPrimitives.result'(X2,X5),'const.semanticPrimitives.Rval_1'(s(X2,X4))))) = s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X2,X0),X3),s(fun(X7,X1),X8),s('type.semanticPrimitives.result'(X2,X7),'const.semanticPrimitives.Rval_1'(s(X2,X4)))))) )),
  inference(superposition,[],[f206,f206])).
fof(f1371,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,happ(s(fun(X3,X0),X4),s(X3,X5))))))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(X3,X0),X4),s(fun(X6,X1),X7),s('type.semanticPrimitives.result'(X3,X6),'const.semanticPrimitives.Rval_1'(s(X3,X5)))))))) )),
  inference(unit_resulting_resolution,[],[f206,f188])).
fof(f1307,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.semanticPrimitives.result'(X36,X37),X38) != s('type.semanticPrimitives.result'(X36,X37),X39) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X36,X37),X40),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X36,X37),X39))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X36,X37),X40),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X36,X37),'const.semanticPrimitives.Rval_1'(s(X36,sK10(X38,X37,X36)))))) | s('type.semanticPrimitives.result'(X36,X37),X38) = s('type.semanticPrimitives.result'(X36,X37),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X37),sK11(X38,X37,X36))))) )),
  inference(superposition,[],[f188,f185])).
fof(f1305,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X26,X27),X29),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X26,X27),X28))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X26,X27),X29),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X26,X27),X30))) | s('type.semanticPrimitives.result'(X26,X27),X30) = s('type.semanticPrimitives.result'(X26,X27),'const.semanticPrimitives.Rval_1'(s(X26,sK10(X28,X27,X26)))) | s('type.semanticPrimitives.result'(X26,X27),X28) = s('type.semanticPrimitives.result'(X26,X27),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X27),sK11(X28,X27,X26))))) )),
  inference(superposition,[],[f187,f185])).
fof(f1299,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.semanticPrimitives.result'(X8,X9),X10) != s('type.semanticPrimitives.result'(X8,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X11))) | s('type.semanticPrimitives.result'(X8,X9),X10) = s('type.semanticPrimitives.result'(X8,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),sK11(X10,X9,X8))))) )),
  inference(superposition,[],[f191,f185])).
fof(f1297,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) | s(X0,X3) = s(X0,sK10(X2,X1,X0)) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11(X2,X1,X0))))) )),
  inference(superposition,[],[f187,f185])).
fof(f1294,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s('type.semanticPrimitives.result'(X34,X35),X36) != s('type.semanticPrimitives.result'(X34,X35),X37) | s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X34,X35),X38),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X34,X35),X37))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X34,X35),X38),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X34,X35),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X35),sK11(X36,X35,X34)))))) | s('type.semanticPrimitives.result'(X34,X35),X36) = s('type.semanticPrimitives.result'(X34,X35),'const.semanticPrimitives.Rval_1'(s(X34,sK10(X36,X35,X34))))) )),
  inference(superposition,[],[f188,f185])).
fof(f1292,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X24,X25),X27),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X24,X25),X26))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X24,X25),X27),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X24,X25),X28))) | s('type.semanticPrimitives.result'(X24,X25),X28) = s('type.semanticPrimitives.result'(X24,X25),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X25),sK11(X26,X25,X24)))) | s('type.semanticPrimitives.result'(X24,X25),X26) = s('type.semanticPrimitives.result'(X24,X25),'const.semanticPrimitives.Rval_1'(s(X24,sK10(X26,X25,X24))))) )),
  inference(superposition,[],[f187,f185])).
fof(f1288,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.semanticPrimitives.result'(X9,X10),X11) != s('type.semanticPrimitives.result'(X9,X10),'const.semanticPrimitives.Rval_1'(s(X9,X12))) | s('type.semanticPrimitives.result'(X9,X10),X11) = s('type.semanticPrimitives.result'(X9,X10),'const.semanticPrimitives.Rval_1'(s(X9,sK10(X11,X10,X9))))) )),
  inference(superposition,[],[f191,f185])).
fof(f1286,plain,(
  ( ! [X4,X2,X3,X1] : (s('type.semanticPrimitives.result'(X1,X2),X3) != s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X4))) | s('type.semanticPrimitives.error_result'(X2),X4) = s('type.semanticPrimitives.error_result'(X2),sK11(X3,X2,X1)) | s('type.semanticPrimitives.result'(X1,X2),X3) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rval_1'(s(X1,sK10(X3,X2,X1))))) )),
  inference(superposition,[],[f189,f185])).
fof(f1276,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK11('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)),X1,X0))))) )),
  inference(unit_resulting_resolution,[],[f191,f185])).
fof(f1102,plain,(
  ( ! [X14,X12,X15,X13] : (s(X14,f(s(bool,X15))) != s(X14,g(s(bool,X15))) | s(fun(bool,X14),X12) = s(fun(bool,X14),X13) | s(bool,'T_0') = s(bool,sK12(X12,X13,X14,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f186,f233])).
fof(f1100,plain,(
  ( ! [X6,X4,X7,X5] : (s(X6,f(s(bool,'T_0'))) != s(X6,g(s(bool,'T_0'))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | s(bool,X7) = s(bool,sK12(X4,X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f186,f233])).
fof(f472,plain,(
  ( ! [X47,X45,X48,X46] : (s('type.semanticPrimitives.result'(bool,X46),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X46),'const.semanticPrimitives.Rval_1'(s(bool,X47))) | s(bool,X45) = s(bool,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f187,f233])).
fof(f470,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') != s(bool,X38) | s('type.semanticPrimitives.result'(bool,X39),'const.semanticPrimitives.Rval_1'(s(bool,X37))) = s('type.semanticPrimitives.result'(bool,X39),'const.semanticPrimitives.Rval_1'(s(bool,X38))) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f188,f233])).
fof(f627,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f458,f131])).
fof(f458,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f180,f233])).
fof(f626,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f457,f131])).
fof(f457,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f178,f233])).
fof(f449,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s('type.semanticPrimitives.result'(bool,X26),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X26),'const.semanticPrimitives.Rval_1'(s(bool,X25))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f188,f233])).
fof(f415,plain,(
  ( ! [X47,X50,X48,X49] : (s('type.semanticPrimitives.result'(bool,X49),'const.semanticPrimitives.Rval_1'(s(bool,X48))) != s('type.semanticPrimitives.result'(bool,X49),'const.semanticPrimitives.Rval_1'(s(bool,X50))) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f187,f233])).
fof(f413,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s('type.semanticPrimitives.result'(bool,X42),'const.semanticPrimitives.Rval_1'(s(bool,X39))) = s('type.semanticPrimitives.result'(bool,X42),'const.semanticPrimitives.Rval_1'(s(bool,X41))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f188,f233])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f233,f233])).
fof(f316,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(X0,X1),X2),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X4)))))) )),
  inference(unit_resulting_resolution,[],[f191,f187])).
fof(f262,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f256,f133])).
fof(f256,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f175,f134])).
fof(f235,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f133,f134])).
fof(f190,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.error_result'(X1),X2) != s('type.semanticPrimitives.error_result'(X1),X3) | s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f186,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.map_result_Rerr',aHL_EXT)).
fof(f180,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X15,X16] : (s(bool,X15) = s(bool,X16) <=> ((p(s(bool,X15)) => p(s(bool,X16))) & (p(s(bool,X16)) => p(s(bool,X15)))))),
  file('thm.evalProps.map_result_Rerr','thm.bool.EQ_IMP_THM')).
fof(f175,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f81])).
fof(f345,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) != s(bool,X5) | s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X6),'const.semanticPrimitives.Rval_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f188,f134])).
fof(f21812,plain,(
  ( ! [X37,X36] : (s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'F_0') = s(bool,sK12(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X36,X37,bool,bool))))) ) | $spl202),
  inference(cnf_transformation,[],[f21812_D])).
fof(f21812_D,plain,(
  ( ! [X37,X36] : (s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'F_0') = s(bool,sK12(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X36,X37,bool,bool))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f21813,plain,(
  ( ! [X38] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(bool,'F_0') = s(bool,X38)) ) | $spl204),
  inference(cnf_transformation,[],[f21813_D])).
fof(f21813_D,plain,(
  ( ! [X38] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(bool,'F_0') = s(bool,X38)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f18715,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,f(s(bool,X35))) | s(fun(bool,bool),X33) = s(fun(bool,bool),X34) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X33,X34,bool,bool)))) | s(bool,'F_0') = s(bool,sK12(X33,X34,bool,bool)) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f1110,f7068])).
fof(f18716,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X36,X37,bool,bool)))) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,sK12(X36,X37,bool,bool))) )),
  inference(superposition,[],[f1110,f7068])).
fof(f21760,plain,(
  ( ! [X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,X48)))) | s(bool,'F_0') = s(bool,X48)) ) | $spl182),
  inference(cnf_transformation,[],[f21760_D])).
fof(f21760_D,plain,(
  ( ! [X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,X48)))) | s(bool,'F_0') = s(bool,X48)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f21789,plain,(
  ( ! [X146,X149] : (s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X146,bool),X149,bool))) ) | $spl200),
  inference(cnf_transformation,[],[f21789_D])).
fof(f21789_D,plain,(
  ( ! [X146,X149] : (s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X146,bool),X149,bool))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f21788,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'T_0') = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X146,bool),X149,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X148),s(bool,X147)))) | s(bool,'F_0') = s(bool,X147)) ) | ($spl86 | $spl120)),
  inference(forward_demodulation,[],[f21727,f1561])).
fof(f21727,plain,(
  ( ! [X146,X149,X147,X148] : (~p(s(bool,happ(s(fun(bool,bool),X148),s(bool,X147)))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X146,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X146,bool),X149,bool)) | s(bool,'F_0') = s(bool,X147)) ) | ($spl86 | $spl120)),
  inference(superposition,[],[f11459,f7099])).
fof(f11459,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,X2) = s(bool,sK9(X2,X0,bool))) ) | $spl120),
  inference(superposition,[],[f173,f9795])).
fof(f21752,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))))) ) | $spl178),
  inference(cnf_transformation,[],[f21752_D])).
fof(f21752_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f21787,plain,(
  ( ! [X142,X145,X143] : (s(bool,X142) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X142)),X143,bool),X145,bool)) | s(bool,'F_0') = s(bool,X142)) ) | $spl198),
  inference(cnf_transformation,[],[f21787_D])).
fof(f21787_D,plain,(
  ( ! [X142,X145,X143] : (s(bool,X142) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X142)),X143,bool),X145,bool)) | s(bool,'F_0') = s(bool,X142)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f21786,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,X142) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X142)),X143,bool),X145,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X144),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X142)) ) | ($spl120 | $spl138 | $spl142)),
  inference(forward_demodulation,[],[f21726,f1561])).
fof(f21726,plain,(
  ( ! [X144,X142,X145,X143] : (~p(s(bool,happ(s(fun(bool,bool),X144),s(bool,'T_0')))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X142)),X143,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X142)),X143,bool),X145,bool)) | s(bool,'F_0') = s(bool,X142)) ) | ($spl120 | $spl138 | $spl142)),
  inference(superposition,[],[f11459,f13811])).
fof(f13811,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool)) | s(bool,'F_0') = s(bool,X10)) ) | ($spl138 | $spl142)),
  inference(backward_demodulation,[],[f9839,f9836])).
fof(f21784,plain,(
  ( ! [X134,X136,X133,X137] : (s(bool,X133) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X133)),X134,bool),X136,bool)) | s(bool,'T_0') = s(bool,X137) | s(bool,X133) = s(bool,X137)) ) | $spl196),
  inference(cnf_transformation,[],[f21784_D])).
fof(f21784_D,plain,(
  ( ! [X134,X136,X133,X137] : (s(bool,X133) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X133)),X134,bool),X136,bool)) | s(bool,'T_0') = s(bool,X137) | s(bool,X133) = s(bool,X137)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f21783,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,X133) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X133)),X134,bool),X136,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X135),s(bool,'T_0')))) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | $spl120),
  inference(forward_demodulation,[],[f21724,f1561])).
fof(f21724,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(bool,bool),X135),s(bool,'T_0')))) | s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X133)),X134,bool)) = s(bool,sK9(sK10('const.semanticPrimitives.Rval_1'(s(bool,X133)),X134,bool),X136,bool)) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | $spl120),
  inference(superposition,[],[f11459,f1636])).
fof(f21779,plain,(
  ( ! [X116,X113] : (s(bool,'T_0') = s(bool,sK9(f(s(fun(bool,bool),X113)),X116,bool))) ) | $spl194),
  inference(cnf_transformation,[],[f21779_D])).
fof(f21779_D,plain,(
  ( ! [X116,X113] : (s(bool,'T_0') = s(bool,sK9(f(s(fun(bool,bool),X113)),X116,bool))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f21780,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') = s(bool,sK9(f(s(fun(bool,bool),X117)),X119,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0'))))) ) | ($spl16 | $spl120 | $spl168)),
  inference(forward_demodulation,[],[f21719,f18801])).
fof(f18801,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(fun(bool,bool),X0)))) ) | ($spl16 | $spl168)),
  inference(unit_resulting_resolution,[],[f131,f840,f18748,f626])).
fof(f21719,plain,(
  ( ! [X118,X119,X117] : (~p(s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0')))) | s(bool,f(s(fun(bool,bool),X117))) = s(bool,sK9(f(s(fun(bool,bool),X117)),X119,bool))) ) | ($spl16 | $spl120 | $spl168)),
  inference(superposition,[],[f11459,f18801])).
fof(f21778,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') = s(bool,sK9(f(s(fun(bool,bool),X113)),X116,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X115),s(bool,X114)))) | s(bool,'F_0') = s(bool,X114)) ) | ($spl16 | $spl120 | $spl168)),
  inference(forward_demodulation,[],[f21718,f18801])).
fof(f21718,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,happ(s(fun(bool,bool),X115),s(bool,X114)))) | s(bool,f(s(fun(bool,bool),X113))) = s(bool,sK9(f(s(fun(bool,bool),X113)),X116,bool)) | s(bool,'F_0') = s(bool,X114)) ) | ($spl120 | $spl168)),
  inference(superposition,[],[f11459,f18831])).
fof(f21777,plain,(
  ( ! [X112,X110] : (s(bool,f(s(bool,X110))) = s(bool,sK9(f(s(bool,X110)),X112,bool)) | s(bool,'T_0') = s(bool,X110)) ) | $spl192),
  inference(cnf_transformation,[],[f21777_D])).
fof(f21777_D,plain,(
  ( ! [X112,X110] : (s(bool,f(s(bool,X110))) = s(bool,sK9(f(s(bool,X110)),X112,bool)) | s(bool,'T_0') = s(bool,X110)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f21717,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),X111),s(bool,'T_0')))) | s(bool,f(s(bool,X110))) = s(bool,sK9(f(s(bool,X110)),X112,bool)) | s(bool,'T_0') = s(bool,X110)) ) | ($spl106 | $spl120)),
  inference(superposition,[],[f11459,f12296])).
fof(f12296,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl106),
  inference(superposition,[],[f12264,f134])).
fof(f21775,plain,(
  ( ! [X107] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'T_0')),X107,bool))) ) | $spl190),
  inference(cnf_transformation,[],[f21775_D])).
fof(f21775_D,plain,(
  ( ! [X107] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'T_0')),X107,bool))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f21776,plain,(
  ( ! [X109,X108] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'T_0')),X109,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'T_0'))))) ) | ($spl120 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f21716,f14039])).
fof(f14039,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(condensation,[],[f14038])).
fof(f14038,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl123 | $spl142)),
  inference(subsumption_resolution,[],[f14006,f134])).
fof(f14006,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f233])).
fof(f21716,plain,(
  ( ! [X109,X108] : (~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'T_0')))) | s(bool,f(s(bool,'T_0'))) = s(bool,sK9(f(s(bool,'T_0')),X109,bool))) ) | ($spl120 | $spl123 | $spl142)),
  inference(superposition,[],[f11459,f14039])).
fof(f21774,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'T_0')),X107,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X105)))) | s(bool,'F_0') = s(bool,X105)) ) | ($spl120 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f21715,f14039])).
fof(f21715,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X105)))) | s(bool,f(s(bool,'T_0'))) = s(bool,sK9(f(s(bool,'T_0')),X107,bool)) | s(bool,'F_0') = s(bool,X105)) ) | ($spl120 | $spl123 | $spl142)),
  inference(superposition,[],[f11459,f14015])).
fof(f21772,plain,(
  ( ! [X102] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'F_0')),X102,bool))) ) | $spl188),
  inference(cnf_transformation,[],[f21772_D])).
fof(f21772_D,plain,(
  ( ! [X102] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'F_0')),X102,bool))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f21773,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'F_0')),X104,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X103),s(bool,'T_0'))))) ) | ($spl106 | $spl120)),
  inference(forward_demodulation,[],[f21714,f12264])).
fof(f21714,plain,(
  ( ! [X103,X104] : (~p(s(bool,happ(s(fun(bool,bool),X103),s(bool,'T_0')))) | s(bool,f(s(bool,'F_0'))) = s(bool,sK9(f(s(bool,'F_0')),X104,bool))) ) | ($spl106 | $spl120)),
  inference(superposition,[],[f11459,f12264])).
fof(f21771,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'F_0')),X102,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,X100)))) | s(bool,'F_0') = s(bool,X100)) ) | ($spl106 | $spl120 | $spl150)),
  inference(forward_demodulation,[],[f21713,f12264])).
fof(f21713,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,happ(s(fun(bool,bool),X101),s(bool,X100)))) | s(bool,f(s(bool,'F_0'))) = s(bool,sK9(f(s(bool,'F_0')),X102,bool)) | s(bool,'F_0') = s(bool,X100)) ) | ($spl120 | $spl150)),
  inference(superposition,[],[f11459,f12283])).
fof(f21769,plain,(
  ( ! [X78] : (s(bool,'T_0') = s(bool,sK9('T_0',X78,bool))) ) | $spl184),
  inference(cnf_transformation,[],[f21769_D])).
fof(f21769_D,plain,(
  ( ! [X78] : (s(bool,'T_0') = s(bool,sK9('T_0',X78,bool))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f21770,plain,(
  ( ! [X76,X79,X77] : (s(bool,'T_0') = s(bool,X79) | ~p(s(bool,happ(s(fun(bool,bool),X77),s(bool,X76)))) | s(bool,X76) = s(bool,X79)) ) | $spl186),
  inference(cnf_transformation,[],[f21770_D])).
fof(f21770_D,plain,(
  ( ! [X76,X79,X77] : (s(bool,'T_0') = s(bool,X79) | ~p(s(bool,happ(s(fun(bool,bool),X77),s(bool,X76)))) | s(bool,X76) = s(bool,X79)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f21709,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,X85)))) | s(bool,'T_0') = s(bool,sK9('T_0',X87,bool)) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) ) | $spl120),
  inference(superposition,[],[f11459,f233])).
fof(f21708,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X82),s(bool,X81)))) | s(bool,'T_0') = s(bool,sK9('T_0',X83,bool)) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) ) | $spl120),
  inference(superposition,[],[f11459,f233])).
fof(f21768,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,X79) | ~p(s(bool,happ(s(fun(bool,bool),X77),s(bool,X76)))) | s(bool,'T_0') = s(bool,sK9('T_0',X78,bool)) | s(bool,X76) = s(bool,X79)) ) | $spl120),
  inference(forward_demodulation,[],[f21707,f1561])).
fof(f21707,plain,(
  ( ! [X80,X78,X76,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X77),s(bool,X76)))) | s(bool,'T_0') = s(bool,sK9('T_0',X78,bool)) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X79)),X80,bool))) ) | $spl120),
  inference(superposition,[],[f11459,f1636])).
fof(f21759,plain,(
  ( ! [X47,X50] : (s(bool,X47) = s(bool,sK9(X47,X50,bool)) | s(bool,'F_0') = s(bool,X47)) ) | $spl180),
  inference(cnf_transformation,[],[f21759_D])).
fof(f21759_D,plain,(
  ( ! [X47,X50] : (s(bool,X47) = s(bool,sK9(X47,X50,bool)) | s(bool,'F_0') = s(bool,X47)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f21767,plain,(
  ( ! [X74,X72,X75] : (~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,'T_0')))) | s(bool,X72) = s(bool,sK9(X72,X75,bool)) | s(bool,'F_0') = s(bool,X72)) ) | ($spl16 | $spl120 | $spl168)),
  inference(forward_demodulation,[],[f21706,f18801])).
fof(f21706,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,happ(s(fun(bool,bool),X74),s(bool,f(s(fun(bool,bool),X73)))))) | s(bool,X72) = s(bool,sK9(X72,X75,bool)) | s(bool,'F_0') = s(bool,X72)) ) | ($spl120 | $spl168)),
  inference(superposition,[],[f11459,f18831])).
fof(f21766,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,happ(s(fun(bool,bool),X70),s(bool,'T_0')))) | s(bool,X69) = s(bool,sK9(X69,X71,bool)) | s(bool,'F_0') = s(bool,X69)) ) | ($spl120 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f21705,f14039])).
fof(f21705,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,happ(s(fun(bool,bool),X70),s(bool,f(s(bool,'T_0')))))) | s(bool,X69) = s(bool,sK9(X69,X71,bool)) | s(bool,'F_0') = s(bool,X69)) ) | ($spl120 | $spl123 | $spl142)),
  inference(superposition,[],[f11459,f14015])).
fof(f21765,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X67),s(bool,'T_0')))) | s(bool,X66) = s(bool,sK9(X66,X68,bool)) | s(bool,'F_0') = s(bool,X66)) ) | ($spl106 | $spl120 | $spl150)),
  inference(forward_demodulation,[],[f21704,f12264])).
fof(f21704,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,happ(s(fun(bool,bool),X67),s(bool,f(s(bool,'F_0')))))) | s(bool,X66) = s(bool,sK9(X66,X68,bool)) | s(bool,'F_0') = s(bool,X66)) ) | ($spl120 | $spl150)),
  inference(superposition,[],[f11459,f12283])).
fof(f21762,plain,(
  ( ! [X61,X59,X62] : (~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,'T_0')))) | s(bool,X59) = s(bool,sK9(X59,X62,bool)) | s(bool,'F_0') = s(bool,X59)) ) | ($spl86 | $spl120)),
  inference(forward_demodulation,[],[f21702,f1561])).
fof(f21702,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,happ(s(fun(bool,bool),X61),s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X60,bool))))) | s(bool,X59) = s(bool,sK9(X59,X62,bool)) | s(bool,'F_0') = s(bool,X59)) ) | ($spl86 | $spl120)),
  inference(superposition,[],[f11459,f7099])).
fof(f21700,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X53),s(bool,X52)))) | s(bool,X51) = s(bool,sK9(X51,X54,bool)) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X51)) ) | $spl120),
  inference(superposition,[],[f11459,f7068])).
fof(f21699,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X49),s(bool,X48)))) | s(bool,X47) = s(bool,sK9(X47,X50,bool)) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X48)) ) | $spl120),
  inference(superposition,[],[f11459,f7068])).
fof(f21751,plain,(
  ( ! [X6,X5,X3] : (s(bool,X3) = s(bool,sK9(X3,X5,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,X3) = s(bool,X6)) ) | $spl176),
  inference(cnf_transformation,[],[f21751_D])).
fof(f21751_D,plain,(
  ( ! [X6,X5,X3] : (s(bool,X3) = s(bool,sK9(X3,X5,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,X3) = s(bool,X6)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f21755,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X37) | ~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,X34) = s(bool,sK9(X34,X36,bool)) | s(bool,X34) = s(bool,X37)) ) | $spl120),
  inference(forward_demodulation,[],[f21696,f1561])).
fof(f21696,plain,(
  ( ! [X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,X34) = s(bool,sK9(X34,X36,bool)) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool))) ) | $spl120),
  inference(superposition,[],[f11459,f1636])).
fof(f21692,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0')))) | s(bool,X15) = s(bool,sK9(X15,X17,bool)) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl120),
  inference(superposition,[],[f11459,f233])).
fof(f21689,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | s(bool,X3) = s(bool,sK9(X3,X5,bool)) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl120),
  inference(superposition,[],[f11459,f233])).
fof(f649,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl20),
  inference(cnf_transformation,[],[f649_D])).
fof(f649_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f21370,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl174),
  inference(cnf_transformation,[],[f21370_D])).
fof(f21370_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f21375,plain,(
  ( ! [X88,X89] : (s(bool,'F_0') != s(bool,f(s(bool,X89))) | p(s(bool,X88)) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl106 | $spl128 | $spl142)),
  inference(forward_demodulation,[],[f21374,f9839])).
fof(f21374,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,g(s(bool,'T_0'))) != s(bool,f(s(bool,X89))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl106 | $spl128)),
  inference(subsumption_resolution,[],[f21316,f625])).
fof(f21316,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | s(bool,g(s(bool,'T_0'))) != s(bool,f(s(bool,X89))) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X88)) ) | ($spl106 | $spl128)),
  inference(superposition,[],[f16528,f9830])).
fof(f16528,plain,(
  ( ! [X10,X9] : (p(s(bool,X10)) | s(bool,X10) != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl106),
  inference(subsumption_resolution,[],[f16440,f131])).
fof(f16440,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | s(bool,X10) != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl106),
  inference(superposition,[],[f180,f12296])).
fof(f21373,plain,(
  ( ! [X74,X73] : (p(s(bool,X73)) | s(bool,'F_0') != s(bool,f(s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f21308,f625])).
fof(f21308,plain,(
  ( ! [X74,X73] : (p(s(bool,X73)) | s(bool,'F_0') != s(bool,f(s(bool,X74))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X73)) ) | $spl106),
  inference(superposition,[],[f16528,f134])).
fof(f21372,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl106),
  inference(condensation,[],[f21371])).
fof(f21371,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,f(s(bool,X71))) | s(bool,'T_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'F_0') = s(bool,X72)) ) | $spl106),
  inference(subsumption_resolution,[],[f21307,f1813])).
fof(f21307,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,X70)) | s(bool,'F_0') != s(bool,f(s(bool,X71))) | s(bool,'T_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'F_0') = s(bool,X72)) ) | $spl106),
  inference(superposition,[],[f16528,f7068])).
fof(f21369,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl106),
  inference(condensation,[],[f21368])).
fof(f21368,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') != s(bool,f(s(bool,X68))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,X67) = s(bool,X69)) ) | $spl106),
  inference(subsumption_resolution,[],[f21306,f1813])).
fof(f21306,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,X67)) | s(bool,'F_0') != s(bool,f(s(bool,X68))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,X67) = s(bool,X69)) ) | $spl106),
  inference(superposition,[],[f16528,f7068])).
fof(f19346,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,f(s(fun(bool,bool),X50))) | s(bool,'T_0') = s(bool,X51)) ) | $spl172),
  inference(cnf_transformation,[],[f19346_D])).
fof(f19346_D,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,f(s(fun(bool,bool),X50))) | s(bool,'T_0') = s(bool,X51)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f19286,plain,(
  ( ! [X101,X102,X100] : (s(bool,X102) != s(bool,f(s(fun(bool,bool),X101))) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X102)) ) | $spl168),
  inference(duplicate_literal_removal,[],[f19136])).
fof(f19136,plain,(
  ( ! [X101,X102,X100] : (s(bool,X102) != s(bool,f(s(fun(bool,bool),X101))) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | $spl168),
  inference(superposition,[],[f1867,f18831])).
fof(f19345,plain,(
  ( ! [X52,X49] : (s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,X49)))) ) | $spl170),
  inference(cnf_transformation,[],[f19345_D])).
fof(f19345_D,plain,(
  ( ! [X52,X49] : (s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,X49)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f19344,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X51) != s(bool,f(s(fun(bool,bool),X50))) | s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,X49))) | s(bool,'T_0') = s(bool,X51)) ) | $spl168),
  inference(subsumption_resolution,[],[f19122,f188])).
fof(f19122,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,X51) != s(bool,f(s(fun(bool,bool),X50))) | s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X52),'const.semanticPrimitives.Rval_1'(s(bool,X49))) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | $spl168),
  inference(superposition,[],[f345,f18831])).
fof(f18747,plain,(
  ( ! [X48,X49] : (s(fun(fun(bool,bool),bool),X48) = s(fun(fun(bool,bool),bool),X49) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK12(X48,X49,bool,fun(bool,bool)))))) ) | $spl166),
  inference(cnf_transformation,[],[f18747_D])).
fof(f18747_D,plain,(
  ( ! [X48,X49] : (s(fun(fun(bool,bool),bool),X48) = s(fun(fun(bool,bool),bool),X49) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK12(X48,X49,bool,fun(bool,bool)))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f18748,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X50)))) ) | $spl168),
  inference(cnf_transformation,[],[f18748_D])).
fof(f18748_D,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X50)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f18722,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X53))) | s(fun(fun(bool,bool),bool),X51) = s(fun(fun(bool,bool),bool),X52) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK12(X51,X52,bool,fun(bool,bool)))))) ) | $spl120),
  inference(superposition,[],[f1110,f9795])).
fof(f18721,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X50))) | s(fun(fun(bool,bool),bool),X48) = s(fun(fun(bool,bool),bool),X49) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK12(X48,X49,bool,fun(bool,bool)))))) ) | $spl120),
  inference(superposition,[],[f1110,f9795])).
fof(f17037,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,X30)) ) | $spl164),
  inference(cnf_transformation,[],[f17037_D])).
fof(f17037_D,plain,(
  ( ! [X30,X29] : (s(bool,X29) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,X29) = s(bool,X30)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f17036,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) = s(bool,g(s(bool,'F_0'))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X28)) ) | ($spl14 | $spl114)),
  inference(subsumption_resolution,[],[f16883,f261])).
fof(f16883,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | s(bool,X29) = s(bool,g(s(bool,'F_0'))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X28)) ) | $spl114),
  inference(superposition,[],[f626,f9671])).
fof(f16999,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X55)) ) | $spl162),
  inference(cnf_transformation,[],[f16999_D])).
fof(f16999_D,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X55)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f16933,plain,(
  ( ! [X54,X55] : (s(bool,X55) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X55)) ) | $spl114),
  inference(duplicate_literal_removal,[],[f16812])).
fof(f16812,plain,(
  ( ! [X54,X55] : (s(bool,X55) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X54)) ) | $spl114),
  inference(superposition,[],[f1867,f9671])).
fof(f3267,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) | $spl76),
  inference(cnf_transformation,[],[f3267_D])).
fof(f3267_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f9831,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl130),
  inference(cnf_transformation,[],[f9831_D])).
fof(f9831_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f14043,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl123 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14009,f13811])).
fof(f14009,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f1636])).
fof(f14400,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14282,f14381])).
fof(f14381,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0)))) ) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14150,f188])).
fof(f14150,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f2753])).
fof(f14085,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f625,f647,f650,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2464,f2527,f2529,f2530,f2704,f2705,f2837,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f3264,f3265,f2876,f2877,f3268,f2882,f2738,f3269,f2636,f2576,f2490,f2493,f840,f4214,f4241,f4947,f5015,f5270,f5271,f5967,f5419,f5421,f6060,f6061,f5527,f5532,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f8159,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f9799,f9800,f9801,f9812,f9813,f9814,f9815,f9816,f9817,f9818,f9819,f9820,f9821,f9822,f9823,f964,f1634,f1933,f9527,f9547,f9549,f417,f639,f959,f2311,f2316,f5212,f6183,f9797,f9836,f9839,f13811,f13812,f133,f303,f304,f320,f321,f345,f346,f641,f961,f1104,f1110,f1278,f2543,f6685,f9545,f651,f1797,f1813,f1815,f1816,f1819,f1850,f1892,f2298,f2340,f2346,f2354,f2359,f2361,f2368,f2369,f2375,f2377,f2397,f2446,f2482,f2628,f2706,f2709,f2714,f2731,f2753,f2814,f2840,f4143,f4148,f4150,f4161,f4168,f4170,f5060,f5531,f6116,f6131,f6152,f6933,f7020,f7108,f7160,f7239,f7244,f7301,f8157,f8158,f8162,f8223,f8328,f8330,f8343,f9136,f9207,f9415,f9461,f9630,f628,f6128,f9171,f8079,f6801,f2695,f2688,f2687,f2590,f2410,f2352,f1903,f1777,f5430,f8341,f7243,f5659,f5525,f5428,f2484,f2413,f2400,f2353,f1280,f1279,f7078,f6682,f5102,f6130,f1956,f2326,f1763,f2381,f2374,f2373,f6196,f4167,f4166,f6198,f4147,f4146,f2365,f2358,f2357,f2322,f1408,f2323,f1458,f2321,f1401,f2319,f1449,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f8335,f8327,f2716,f2713,f2309,f2286,f9550,f9526,f6184,f6161,f2317,f2310,f1818,f474,f9881,f9844,f9594,f9590,f9583,f9577,f9559,f8340,f8326,f8226,f8220,f8219,f8156,f7300,f7242,f7241,f7238,f7237,f7027,f2314,f1817,f962,f957,f637,f453,f9177,f6230,f2333,f2332,f2308,f1698,f1686,f1664,f974,f174,f173,f172,f8466,f8339,f6941,f5657,f5523,f5426,f4952,f1108,f1106,f9832,f13990,f13987,f14020,f13991,f13992,f14044,f14084])).
fof(f14084,plain,(
  ( ! [X29,X27] : (s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14046,f840])).
fof(f14046,plain,(
  ( ! [X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14044,f1903])).
fof(f14044,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14043,f9832])).
fof(f13992,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s(bool,X3))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f233])).
fof(f13991,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool))) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f1636])).
fof(f14020,plain,(
  p(s(bool,f(s(bool,'T_0')))) | ($spl16 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f13986,f1561])).
fof(f13986,plain,(
  ( ! [X0] : (p(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'T_0')))),X0,bool)))) ) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f1561,f13812,f627])).
fof(f13987,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'T_0')))),X0,bool))) ) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f13812,f1636])).
fof(f13990,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'T_0')))))) ) | ($spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f13812,f187])).
fof(f9171,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X171)),X172,bool)) | s(bool,X170) = s(bool,X171) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') = s(bool,X173)) ) | $spl16),
  inference(subsumption_resolution,[],[f8279,f8159])).
fof(f8279,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X171)),X172,bool)) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X170) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') = s(bool,X173)) )),
  inference(superposition,[],[f5732,f7068])).
fof(f9461,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,X114) | s(bool,X116) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X114)),X115,bool)) | s(bool,'F_0') = s(bool,X116)) ) | $spl16),
  inference(superposition,[],[f8159,f1561])).
fof(f9415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X3)),X4,bool))) ) | $spl16),
  inference(superposition,[],[f8159,f1636])).
fof(f9207,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X178) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X177)),X179,bool)) | s(bool,X176) = s(bool,X178) | s(bool,'F_0') = s(bool,X176) | s(bool,'F_0') = s(bool,X177)) ) | $spl16),
  inference(subsumption_resolution,[],[f8396,f8159])).
fof(f8396,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,X178) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X177)),X179,bool)) | s(bool,X176) = s(bool,X178) | s(bool,'T_0') = s(bool,X178) | s(bool,'F_0') = s(bool,X176) | s(bool,'F_0') = s(bool,X177)) )),
  inference(superposition,[],[f5732,f7068])).
fof(f7108,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X54) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X53)),X55,bool)) | s(bool,'F_0') = s(bool,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X56)),X57,bool))) ) | $spl16),
  inference(subsumption_resolution,[],[f6814,f840])).
fof(f6814,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X54) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X53)),X55,bool)) | s(bool,'F_0') = s(bool,X54) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X56)),X57,bool))) )),
  inference(superposition,[],[f1850,f1636])).
fof(f9801,plain,(
  ( ! [X4,X3] : (s(bool,f(s(bool,X3))) != s(bool,g(s(bool,X3))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl123),
  inference(superposition,[],[f9797,f233])).
fof(f9800,plain,(
  ( ! [X2,X0,X1] : (s(bool,f(s(bool,X0))) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool))) ) | $spl123),
  inference(superposition,[],[f9797,f1636])).
fof(f9799,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))))) ) | $spl123),
  inference(unit_resulting_resolution,[],[f9797,f187])).
fof(f8159,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,X19) | s(bool,X19) = s(bool,X20) | s(bool,'F_0') = s(bool,X20)) ) | $spl16),
  inference(superposition,[],[f840,f7068])).
fof(f5532,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X29) != s(bool,X30) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X30)),X31,bool)) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X32)),X33,bool))) ) | $spl16),
  inference(superposition,[],[f4214,f1636])).
fof(f6061,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(bool,X226) != s(bool,X227) | s('type.semanticPrimitives.result'(bool,X228),'const.semanticPrimitives.Rval_1'(s(bool,X227))) = s('type.semanticPrimitives.result'(bool,X228),'const.semanticPrimitives.Rval_1'(s(bool,X229))) | s(bool,'T_0') = s(bool,X229) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X226)),X230,bool))) ) | $spl16),
  inference(subsumption_resolution,[],[f5484,f840])).
fof(f5484,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(bool,X226) != s(bool,X227) | s('type.semanticPrimitives.result'(bool,X228),'const.semanticPrimitives.Rval_1'(s(bool,X227))) = s('type.semanticPrimitives.result'(bool,X228),'const.semanticPrimitives.Rval_1'(s(bool,X229))) | s(bool,'T_0') = s(bool,X229) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X226)),X230,bool))) )),
  inference(superposition,[],[f346,f1636])).
fof(f6060,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s('type.semanticPrimitives.result'(bool,X222),'const.semanticPrimitives.Rval_1'(s(bool,X221))) != s('type.semanticPrimitives.result'(bool,X222),'const.semanticPrimitives.Rval_1'(s(bool,X223))) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X221)),X225,bool))) ) | $spl16),
  inference(subsumption_resolution,[],[f5483,f840])).
fof(f5483,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s('type.semanticPrimitives.result'(bool,X222),'const.semanticPrimitives.Rval_1'(s(bool,X221))) != s('type.semanticPrimitives.result'(bool,X222),'const.semanticPrimitives.Rval_1'(s(bool,X223))) | s(bool,X223) = s(bool,X224) | s(bool,'T_0') = s(bool,X224) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X221)),X225,bool))) )),
  inference(superposition,[],[f321,f1636])).
fof(f5967,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X57) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X55)),X56,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X57)),X58,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X55)),X59,bool))) ) | $spl16),
  inference(subsumption_resolution,[],[f5290,f840])).
fof(f5290,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X57) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X55)),X56,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X57)),X58,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X55)),X59,bool))) )),
  inference(superposition,[],[f1636,f1777])).
fof(f4241,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl16),
  inference(superposition,[],[f4214,f233])).
fof(f4214,plain,(
  ( ! [X42,X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X40)),X42,bool))) ) | $spl16),
  inference(superposition,[],[f2704,f1561])).
fof(f2837,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X63)),X64,bool)) | s(bool,'T_0') = s(bool,X63)) ) | $spl16),
  inference(subsumption_resolution,[],[f2798,f133])).
fof(f2798,plain,(
  ( ! [X64,X63] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X63)),X64,bool)) | s(bool,'T_0') = s(bool,X63)) ) | $spl16),
  inference(superposition,[],[f625,f2543])).
fof(f2704,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X66)),X67,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X66)),X68,bool))) ) | $spl16),
  inference(subsumption_resolution,[],[f2614,f133])).
fof(f2614,plain,(
  ( ! [X68,X66,X67] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X66)),X67,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X66)),X68,bool))) ) | $spl16),
  inference(superposition,[],[f625,f1777])).
fof(f14282,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'T_0')))))) ) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f13990])).
fof(f14383,plain,(
  $false | ($spl16 | $spl74 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14155,f14085])).
fof(f14155,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool))) ) | ($spl16 | $spl74 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f3266])).
fof(f14376,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(subsumption_resolution,[],[f14086,f131])).
fof(f14086,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f133])).
fof(f14375,plain,(
  $false | ($spl16 | $spl82 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(trivial_inequality_removal,[],[f14156])).
fof(f14156,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl16 | $spl82 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f3272])).
fof(f14374,plain,(
  $false | ($spl16 | $spl82 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(trivial_inequality_removal,[],[f14158])).
fof(f14158,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl16 | $spl82 | $spl123 | $spl131 | $spl138 | $spl142)),
  inference(backward_demodulation,[],[f14085,f3771])).
fof(f3771,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f3272,f187])).
fof(f14040,plain,(
  $false | ($spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f14039,f9832])).
fof(f14037,plain,(
  $false | ($spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f14036,f9832])).
fof(f14036,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(condensation,[],[f14035])).
fof(f14035,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'T_0')))) ) | ($spl123 | $spl142)),
  inference(subsumption_resolution,[],[f14005,f134])).
fof(f14005,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'T_0')))) ) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f233])).
fof(f14034,plain,(
  $false | ($spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f14017,f9832])).
fof(f14017,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(trivial_inequality_removal,[],[f14003])).
fof(f14003,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(superposition,[],[f13812,f134])).
fof(f14033,plain,(
  $false | ($spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13959,f9832])).
fof(f13959,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f13812,f134])).
fof(f14032,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13963,f9832])).
fof(f13963,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f13812,f7068])).
fof(f14031,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13966,f9832])).
fof(f13966,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f13812,f7068])).
fof(f14030,plain,(
  $false | ($spl16 | $spl106 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f14029,f9832])).
fof(f14029,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl106 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f13968,f12264])).
fof(f13968,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl106 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f12264,f13812,f8159])).
fof(f14028,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f14027,f9832])).
fof(f14027,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(forward_demodulation,[],[f13969,f1561])).
fof(f13969,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'T_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f1561,f13812,f8159])).
fof(f14026,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13972,f9832])).
fof(f13972,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f13812,f233])).
fof(f14025,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13974,f9832])).
fof(f13974,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f131,f840,f13812,f626])).
fof(f14023,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13980,f9832])).
fof(f13980,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f840,f13812,f233])).
fof(f14022,plain,(
  $false | ($spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13981,f9832])).
fof(f13981,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f133,f13812,f262])).
fof(f14021,plain,(
  $false | ($spl16 | $spl123 | $spl131 | $spl142)),
  inference(subsumption_resolution,[],[f13982,f9832])).
fof(f13982,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl123 | $spl142)),
  inference(unit_resulting_resolution,[],[f131,f840,f13812,f626])).
fof(f13807,plain,(
  $false | ($spl74 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f3266,f13731])).
fof(f13731,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl140 | $spl158)),
  inference(condensation,[],[f13730])).
fof(f13730,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13494,f13336])).
fof(f13336,plain,(
  ( ! [X49] : (s(bool,'F_0') = s(bool,X49)) ) | ($spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13313,f9837])).
fof(f13494,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,X0) = s(bool,X1)) ) | ($spl140 | $spl158)),
  inference(condensation,[],[f13478])).
fof(f13478,plain,(
  ( ! [X23,X19,X22,X20] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X23,bool)) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | ($spl140 | $spl158)),
  inference(backward_demodulation,[],[f13336,f6682])).
fof(f13806,plain,(
  $false | ($spl82 | $spl140 | $spl158)),
  inference(trivial_inequality_removal,[],[f13805])).
fof(f13805,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl82 | $spl140 | $spl158)),
  inference(forward_demodulation,[],[f3272,f13731])).
fof(f13804,plain,(
  $false | ($spl82 | $spl140 | $spl158)),
  inference(trivial_inequality_removal,[],[f13803])).
fof(f13803,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl82 | $spl140 | $spl158)),
  inference(forward_demodulation,[],[f3771,f13731])).
fof(f13798,plain,(
  $false | ($spl127 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f9829,f13336])).
fof(f13797,plain,(
  $false | ($spl16 | $spl128 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13796,f131])).
fof(f13796,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl128 | $spl140 | $spl158)),
  inference(forward_demodulation,[],[f13296,f13731])).
fof(f13296,plain,(
  ~p(s(bool,'F_0')) | ($spl16 | $spl128)),
  inference(global_subsumption,[],[f13295,f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2527,f2529,f2530,f2705,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2576,f2490,f2493,f4947,f5015,f5270,f5271,f5419,f5421,f5527,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f1777,f1903,f2352,f2410,f2543,f2590,f2687,f2688,f2695,f4952,f5102,f6682,f6685,f6801,f6941,f7078,f8079,f133,f303,f304,f320,f321,f345,f346,f641,f961,f964,f1104,f1110,f1278,f1634,f1797,f1813,f1815,f1816,f1817,f1818,f1819,f1850,f1892,f1933,f1956,f2286,f2298,f2309,f2317,f2340,f2346,f2353,f2354,f2357,f2358,f2359,f2361,f2365,f2368,f2369,f2373,f2374,f2375,f2377,f2381,f2397,f2400,f2413,f2446,f2482,f2484,f2628,f2706,f2709,f2713,f2714,f2716,f2731,f2753,f2814,f2840,f4143,f4146,f4147,f4148,f4150,f4161,f4166,f4167,f4168,f4170,f5060,f5531,f6116,f6130,f6131,f6152,f6196,f6198,f6933,f7020,f7027,f7160,f7237,f7238,f7239,f7241,f7242,f7243,f7244,f7300,f7301,f8156,f8157,f8158,f8162,f8219,f8220,f8223,f8226,f8326,f8327,f8328,f8330,f8335,f8339,f8340,f8341,f8343,f8466,f9136,f9177,f9526,f9527,f9545,f9547,f9549,f9550,f9559,f9577,f9583,f9590,f9594,f9630,f9844,f9881,f1106,f1108,f5426,f5523,f5657,f172,f173,f174,f651,f974,f1664,f1686,f1698,f2308,f2332,f2333,f6230,f417,f453,f637,f639,f957,f959,f962,f2311,f2314,f2316,f5212,f6183,f474,f2310,f6161,f6184,f1290,f1291,f1303,f1304,f1400,f1407,f1448,f1457,f2318,f2320,f2321,f1401,f2319,f1449,f2323,f1458,f2322,f1408,f2326,f1763,f1279,f1280,f5428,f5525,f5659,f5430,f628,f6128])).
fof(f13295,plain,(
  ( ! [X97,X96] : (~p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK9('F_0',X96,bool)) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X96),s(bool,X97)))) ) | ($spl16 | $spl128)),
  inference(forward_demodulation,[],[f12701,f13264])).
fof(f13264,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl16 | $spl128)),
  inference(condensation,[],[f13263])).
fof(f13263,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | ($spl16 | $spl128)),
  inference(condensation,[],[f13262])).
fof(f13262,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12760,f8159])).
fof(f12760,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X77)) ) | $spl128),
  inference(superposition,[],[f7068,f9830])).
fof(f12701,plain,(
  ( ! [X97,X96] : (~p(s(bool,g(s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK9('F_0',X96,bool)) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X96),s(bool,X97)))) ) | $spl128),
  inference(superposition,[],[f303,f9830])).
fof(f13729,plain,(
  $false | ($spl16 | $spl96 | $spl106 | $spl125 | $spl136 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f9826,f13559])).
fof(f13559,plain,(
  ( ! [X92] : (s(bool,'T_0') = s(bool,X92)) ) | ($spl16 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(global_subsumption,[],[f13558,f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f625,f647,f650,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2464,f2527,f2529,f2530,f2704,f2705,f2837,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f3264,f3265,f2876,f2877,f3268,f2882,f2738,f3269,f2636,f2576,f2490,f2493,f840,f4214,f4241,f4947,f5015,f5270,f5271,f5967,f5419,f5421,f6060,f6061,f5527,f5532,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f8159,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f2543,f6685,f133,f303,f304,f320,f321,f345,f346,f641,f961,f964,f1104,f1110,f1278,f1634,f1933,f9527,f9545,f9547,f9549,f9667,f651,f417,f639,f959,f2311,f2316,f5212,f6183,f12199,f12254,f12252,f12223,f12281,f12279,f12296,f12304,f13266,f6128,f628,f5430,f5659,f5525,f5428,f1280,f1279,f6682,f5102,f2322,f1408,f2323,f1458,f2319,f1449,f2321,f1401,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f2326,f1763,f6230,f2333,f2332,f2308,f1698,f1686,f1664,f974,f174,f173,f172,f6184,f6161,f2310,f474,f2314,f962,f957,f637,f453,f12383,f12262,f5657,f5523,f5426,f4952,f1108,f1106,f12404,f12391,f12294,f12264,f12261,f12221,f9881,f9844,f9837,f9630,f9594,f9590,f9583,f9577,f9559,f9550,f9526,f9461,f9415,f9207,f9177,f9171,f9136,f8466,f8343,f8341,f8340,f8339,f8335,f8330,f8328,f8327,f8326,f8226,f8223,f8220,f8219,f8162,f8158,f8157,f8156,f8079,f7301,f7300,f7244,f7243,f7242,f7241,f7239,f7238,f7237,f7160,f7108,f7078,f7027,f7020,f6941,f6933,f6801,f6198,f6196,f6152,f6131,f6130,f6116,f5531,f5060,f4170,f4168,f4167,f4166,f4161,f4150,f4148,f4147,f4146,f4143,f2840,f2814,f2753,f2731,f2716,f2714,f2713,f2709,f2706,f2695,f2688,f2687,f2628,f2590,f2484,f2482,f2446,f2413,f2410,f2400,f2397,f2381,f2377,f2375,f2374,f2373,f2369,f2368,f2365,f2361,f2359,f2358,f2357,f2354,f2353,f2352,f2346,f2340,f2317,f2309,f2298,f2286,f1956,f1903,f1892,f1850,f1819,f1818,f1817,f1816,f1815,f1813,f1797,f1777,f9835,f13336,f13520,f13553])).
fof(f13553,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl140 | $spl158)),
  inference(condensation,[],[f13552])).
fof(f13552,plain,(
  ( ! [X14,X12,X15,X13] : (s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,g(s(X12,sK12(X13,X14,bool,X12)))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13341,f134])).
fof(f13341,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X15) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,g(s(X12,sK12(X13,X14,bool,X12)))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl140 | $spl158)),
  inference(backward_demodulation,[],[f13336,f1108])).
fof(f13520,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl106 | $spl136 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13337,f13266])).
fof(f13337,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13336,f9835])).
fof(f12221,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl106),
  inference(superposition,[],[f9667,f7068])).
fof(f12261,plain,(
  ( ! [X33,X32] : (s(bool,X32) != s(bool,f(s(bool,X32))) | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X33)) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f12213,f840])).
fof(f12213,plain,(
  ( ! [X33,X32] : (s(bool,X32) != s(bool,f(s(bool,X32))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X33)) ) | $spl106),
  inference(superposition,[],[f9667,f7068])).
fof(f12294,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | $spl106),
  inference(superposition,[],[f12264,f7068])).
fof(f12391,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl106),
  inference(subsumption_resolution,[],[f12330,f131])).
fof(f12330,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl106),
  inference(superposition,[],[f626,f12264])).
fof(f12404,plain,(
  ( ! [X42] : (s(bool,'T_0') = s(bool,sK9(f(s(bool,'F_0')),X42,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0'))))) ) | $spl106),
  inference(forward_demodulation,[],[f12350,f12264])).
fof(f12350,plain,(
  ( ! [X42] : (~p(s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0')))) | s(bool,f(s(bool,'F_0'))) = s(bool,sK9(f(s(bool,'F_0')),X42,bool))) ) | $spl106),
  inference(superposition,[],[f173,f12264])).
fof(f12262,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X24,bool))) ) | ($spl16 | $spl106)),
  inference(subsumption_resolution,[],[f12230,f840])).
fof(f12230,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X22)),X24,bool))) ) | $spl106),
  inference(superposition,[],[f9667,f1636])).
fof(f13266,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X15)),X16,bool))) ) | $spl106),
  inference(global_subsumption,[],[f12304,f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2527,f2529,f2530,f2705,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2576,f2490,f2493,f4947,f5015,f5270,f5271,f5419,f5421,f5527,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f1777,f1903,f2352,f2410,f2543,f2590,f2687,f2688,f2695,f4952,f5102,f6682,f6685,f6801,f6941,f7078,f8079,f133,f303,f304,f320,f321,f345,f346,f641,f961,f964,f1104,f1110,f1278,f1634,f1797,f1813,f1815,f1816,f1817,f1818,f1819,f1850,f1892,f1933,f1956,f2286,f2298,f2309,f2317,f2340,f2346,f2353,f2354,f2357,f2358,f2359,f2361,f2365,f2368,f2369,f2373,f2374,f2375,f2377,f2381,f2397,f2400,f2413,f2446,f2482,f2484,f2628,f2706,f2709,f2713,f2714,f2716,f2731,f2753,f2814,f2840,f4143,f4146,f4147,f4148,f4150,f4161,f4166,f4167,f4168,f4170,f5060,f5531,f6116,f6130,f6131,f6152,f6196,f6198,f6933,f7020,f7027,f7160,f7237,f7238,f7239,f7241,f7242,f7243,f7244,f7300,f7301,f8156,f8157,f8158,f8162,f8219,f8220,f8223,f8226,f8326,f8327,f8328,f8330,f8335,f8339,f8340,f8341,f8343,f8466,f9136,f9177,f9526,f9527,f9545,f9547,f9549,f9550,f9559,f9577,f9583,f9590,f9594,f9630,f9844,f9881,f1106,f1108,f5426,f5523,f5657,f172,f173,f174,f651,f974,f1664,f1686,f1698,f2308,f2332,f2333,f6230,f417,f453,f637,f639,f957,f959,f962,f2311,f2314,f2316,f5212,f6183,f474,f2310,f6161,f6184,f1290,f1291,f1303,f1304,f1400,f1407,f1448,f1457,f2318,f2320,f2321,f1401,f2319,f1449,f2323,f1458,f2322,f1408,f2326,f1763,f1279,f1280,f5428,f5525,f5659,f5430,f628,f6128])).
fof(f12304,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X15)),X16,bool))) ) | $spl106),
  inference(superposition,[],[f12264,f1636])).
fof(f12223,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl106),
  inference(superposition,[],[f9667,f134])).
fof(f12252,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool)) | s(bool,X0) != s(bool,f(s(bool,'F_0')))) ) | $spl106),
  inference(condensation,[],[f12251])).
fof(f12251,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X16)),X17,bool))) ) | $spl106),
  inference(duplicate_literal_removal,[],[f12207])).
fof(f12207,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X16)),X17,bool))) ) | $spl106),
  inference(superposition,[],[f9667,f1636])).
fof(f12254,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) | $spl106),
  inference(condensation,[],[f12253])).
fof(f12253,plain,(
  ( ! [X10,X11,X9] : (s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) | $spl106),
  inference(duplicate_literal_removal,[],[f12205])).
fof(f12205,plain,(
  ( ! [X10,X11,X9] : (s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl106),
  inference(superposition,[],[f9667,f1635])).
fof(f12199,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))),X0,bool))) ) | $spl106),
  inference(unit_resulting_resolution,[],[f1561,f9667])).
fof(f13558,plain,(
  ( ! [X94,X92,X93] : (s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,X93))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl96 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13355,f2753])).
fof(f13355,plain,(
  ( ! [X94,X92,X93] : (s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,X93))) | s(bool,'T_0') != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) | ($spl96 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13336,f12556])).
fof(f12556,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') != s(bool,X93) | s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,X93))) = s('type.semanticPrimitives.result'(bool,X94),'const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,X92))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl96),
  inference(superposition,[],[f188,f10496])).
fof(f10496,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,g(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl96),
  inference(superposition,[],[f9660,f134])).
fof(f13684,plain,(
  $false | ($spl16 | $spl96 | $spl106 | $spl128 | $spl136 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13671,f131])).
fof(f13671,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl96 | $spl106 | $spl128 | $spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13559,f13296])).
fof(f13682,plain,(
  $false | ($spl16 | $spl96 | $spl106 | $spl127 | $spl136 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13650,f13559])).
fof(f13650,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl96 | $spl106 | $spl127 | $spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13559,f9829])).
fof(f13679,plain,(
  $false | ($spl16 | $spl74 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(subsumption_resolution,[],[f13591,f13559])).
fof(f13591,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool))) ) | ($spl16 | $spl74 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13559,f3266])).
fof(f13676,plain,(
  $false | ($spl16 | $spl82 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(trivial_inequality_removal,[],[f13592])).
fof(f13592,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl16 | $spl82 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13559,f3272])).
fof(f13675,plain,(
  $false | ($spl16 | $spl82 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(trivial_inequality_removal,[],[f13593])).
fof(f13593,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl16 | $spl82 | $spl96 | $spl106 | $spl136 | $spl140 | $spl158)),
  inference(backward_demodulation,[],[f13559,f3771])).
fof(f13319,plain,(
  ( ! [X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9('F_0',X22,bool))) ) | $spl160),
  inference(cnf_transformation,[],[f13319_D])).
fof(f13319_D,plain,(
  ( ! [X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9('F_0',X22,bool))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f13318,plain,(
  ( ! [X21,X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK9('F_0',X22,bool)) | s(bool,'T_0') = s(bool,X21)) ) | ($spl16 | $spl128)),
  inference(forward_demodulation,[],[f12804,f13264])).
fof(f12804,plain,(
  ( ! [X21,X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,g(s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK9('F_0',X22,bool)) | s(bool,'T_0') = s(bool,X21)) ) | $spl128),
  inference(duplicate_literal_removal,[],[f12667])).
fof(f12667,plain,(
  ( ! [X21,X22] : (~p(s(bool,happ(s(fun(bool,bool),X22),s(bool,g(s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK9('F_0',X22,bool)) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl128),
  inference(superposition,[],[f303,f9830])).
fof(f13313,plain,(
  ( ! [X49] : (s(bool,X49) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X49)) ) | $spl158),
  inference(cnf_transformation,[],[f13313_D])).
fof(f13313_D,plain,(
  ( ! [X49] : (s(bool,X49) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X49)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f12794,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48)) ) | $spl128),
  inference(duplicate_literal_removal,[],[f12677])).
fof(f12677,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | $spl128),
  inference(superposition,[],[f1867,f9830])).
fof(f3266,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X2,bool))) ) | $spl74),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X2,bool))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f13312,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X55,bool)) | s(bool,'T_0') = s(bool,X54)) ) | ($spl16 | $spl128)),
  inference(forward_demodulation,[],[f13241,f13264])).
fof(f13241,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X55,bool)) | s(bool,'T_0') = s(bool,X54)) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12680,f2837])).
fof(f12680,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X55,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X54)),X56,bool)) | s(bool,'T_0') = s(bool,X54)) ) | ($spl16 | $spl128)),
  inference(superposition,[],[f2704,f9830])).
fof(f13311,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X58,bool)) | s(bool,'T_0') = s(bool,X57)) ) | ($spl16 | $spl128)),
  inference(forward_demodulation,[],[f12793,f13264])).
fof(f12793,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X58,bool)) | s(bool,'T_0') = s(bool,X57)) ) | ($spl16 | $spl128)),
  inference(duplicate_literal_removal,[],[f12681])).
fof(f12681,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X58,bool)) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) ) | ($spl16 | $spl128)),
  inference(superposition,[],[f2837,f9830])).
fof(f9826,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl125),
  inference(cnf_transformation,[],[f9826_D])).
fof(f9826_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f13310,plain,(
  ( ! [X59] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12682,f2837])).
fof(f12682,plain,(
  ( ! [X59,X60] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X59)),X60,bool)) | s(bool,'T_0') = s(bool,X59)) ) | ($spl16 | $spl128)),
  inference(superposition,[],[f4214,f9830])).
fof(f13308,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X67)) ) | $spl128),
  inference(forward_demodulation,[],[f12791,f1561])).
fof(f12791,plain,(
  ( ! [X68,X67] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X68,bool)) | s(bool,'T_0') = s(bool,X67)) ) | $spl128),
  inference(duplicate_literal_removal,[],[f12685])).
fof(f12685,plain,(
  ( ! [X68,X67] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X68,bool)) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) ) | $spl128),
  inference(superposition,[],[f5744,f9830])).
fof(f13307,plain,(
  ( ! [X69] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl16 | $spl128)),
  inference(forward_demodulation,[],[f13306,f1561])).
fof(f13306,plain,(
  ( ! [X70,X69] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X70,bool)) | s(bool,'T_0') = s(bool,X69)) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12686,f2837])).
fof(f12686,plain,(
  ( ! [X70,X71,X69] : (s(bool,g(s(bool,'T_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X70,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X69)),X71,bool)) | s(bool,'T_0') = s(bool,X69)) ) | $spl128),
  inference(superposition,[],[f6299,f9830])).
fof(f13249,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,g(s(bool,'T_0')))) ) | ($spl16 | $spl128)),
  inference(condensation,[],[f13248])).
fof(f13248,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X80) = s(bool,X81) | s(bool,'F_0') = s(bool,X81)) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12691,f8159])).
fof(f12691,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X80) = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X80)) ) | ($spl16 | $spl128)),
  inference(superposition,[],[f8159,f9830])).
fof(f13293,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107)) ) | ($spl16 | $spl128)),
  inference(subsumption_resolution,[],[f12705,f2837])).
fof(f12705,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X107)),X108,bool))) ) | ($spl16 | $spl128)),
  inference(superposition,[],[f2837,f9830])).
fof(f13267,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | $spl156),
  inference(cnf_transformation,[],[f13267_D])).
fof(f13267_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f12777,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | $spl128),
  inference(equality_factoring,[],[f9830])).
fof(f13203,plain,(
  $false | ($spl16 | $spl74 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12923,f12840])).
fof(f12840,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12839,f2837])).
fof(f12839,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl124 | $spl128)),
  inference(forward_demodulation,[],[f12825,f9827])).
fof(f12825,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,g(s(bool,'T_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) | $spl128),
  inference(condensation,[],[f12824])).
fof(f12824,plain,(
  ( ! [X10,X11,X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) | $spl128),
  inference(duplicate_literal_removal,[],[f12611])).
fof(f12611,plain,(
  ( ! [X10,X11,X9] : (s(bool,g(s(bool,'T_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl128),
  inference(superposition,[],[f9830,f1635])).
fof(f12923,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool))) ) | ($spl16 | $spl74 | $spl124 | $spl128)),
  inference(backward_demodulation,[],[f12840,f3266])).
fof(f13196,plain,(
  $false | ($spl16 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12841,f131])).
fof(f12841,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl124 | $spl128)),
  inference(backward_demodulation,[],[f12840,f133])).
fof(f13191,plain,(
  $false | ($spl16 | $spl82 | $spl124 | $spl128)),
  inference(trivial_inequality_removal,[],[f12924])).
fof(f12924,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl16 | $spl82 | $spl124 | $spl128)),
  inference(backward_demodulation,[],[f12840,f3272])).
fof(f13190,plain,(
  $false | ($spl16 | $spl82 | $spl124 | $spl128)),
  inference(trivial_inequality_removal,[],[f12926])).
fof(f12926,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl16 | $spl82 | $spl124 | $spl128)),
  inference(backward_demodulation,[],[f12840,f3771])).
fof(f12838,plain,(
  $false | ($spl16 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12837,f840])).
fof(f12837,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl124 | $spl128)),
  inference(condensation,[],[f12836])).
fof(f12836,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl124 | $spl128)),
  inference(forward_demodulation,[],[f12830,f9827])).
fof(f12830,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl128),
  inference(duplicate_literal_removal,[],[f12606])).
fof(f12606,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl128),
  inference(superposition,[],[f9830,f134])).
fof(f12835,plain,(
  $false | ($spl16 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12834,f840])).
fof(f12834,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl16 | $spl124 | $spl128)),
  inference(forward_demodulation,[],[f12604,f9827])).
fof(f12604,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl16 | $spl128)),
  inference(unit_resulting_resolution,[],[f840,f9830])).
fof(f12833,plain,(
  $false | ($spl16 | $spl74 | $spl124 | $spl128)),
  inference(subsumption_resolution,[],[f12832,f840])).
fof(f12832,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl74 | $spl124 | $spl128)),
  inference(forward_demodulation,[],[f12831,f9827])).
fof(f12831,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl74 | $spl128)),
  inference(forward_demodulation,[],[f12605,f1561])).
fof(f12605,plain,(
  ( ! [X0] : (s(bool,g(s(bool,'T_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X0,bool))) ) | ($spl74 | $spl128)),
  inference(unit_resulting_resolution,[],[f3266,f9830])).
fof(f12460,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK9('F_0',X0,bool))) ) | $spl152),
  inference(cnf_transformation,[],[f12460_D])).
fof(f12460_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK9('F_0',X0,bool))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f12461,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl154),
  inference(cnf_transformation,[],[f12461_D])).
fof(f12461_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f12410,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,X5)))) | s(bool,'F_0') = s(bool,sK9('F_0',X3,bool)) | s(bool,'T_0') = s(bool,X5)) ) | $spl120),
  inference(superposition,[],[f303,f9795])).
fof(f12409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,'F_0') = s(bool,sK9('F_0',X0,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl120),
  inference(superposition,[],[f303,f9795])).
fof(f12283,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X27)) ) | $spl150),
  inference(cnf_transformation,[],[f12283_D])).
fof(f12283_D,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X27)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f12279,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X21)) ) | $spl106),
  inference(subsumption_resolution,[],[f12244,f134])).
fof(f12244,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X21)) ) | $spl106),
  inference(superposition,[],[f9667,f7068])).
fof(f12281,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,'F_0')))) ) | $spl106),
  inference(subsumption_resolution,[],[f12247,f134])).
fof(f12247,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X27) = s(bool,f(s(bool,'F_0')))) ) | $spl106),
  inference(superposition,[],[f9667,f7068])).
fof(f12282,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27)) ) | ($spl106 | $spl118)),
  inference(subsumption_resolution,[],[f12281,f9674])).
fof(f12280,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X21)) ) | ($spl106 | $spl118)),
  inference(subsumption_resolution,[],[f12279,f9674])).
fof(f11619,plain,(
  $false | ($spl16 | $spl122 | $spl124 | $spl126)),
  inference(subsumption_resolution,[],[f11618,f840])).
fof(f11618,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl122 | $spl124 | $spl126)),
  inference(backward_demodulation,[],[f11616,f9828])).
fof(f11616,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl122 | $spl124)),
  inference(forward_demodulation,[],[f9796,f9827])).
fof(f9796,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) | $spl122),
  inference(cnf_transformation,[],[f9796_D])).
fof(f9796_D,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(bool,'T_0'))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11617,plain,(
  $false | ($spl122 | $spl124 | $spl131)),
  inference(subsumption_resolution,[],[f11616,f9832])).
fof(f11613,plain,(
  $false | ($spl16 | $spl124 | $spl142)),
  inference(subsumption_resolution,[],[f11612,f840])).
fof(f11612,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl124 | $spl142)),
  inference(backward_demodulation,[],[f9827,f9839])).
fof(f9837,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl140),
  inference(cnf_transformation,[],[f9837_D])).
fof(f9837_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f9832,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl131),
  inference(cnf_transformation,[],[f9832_D])).
fof(f9832_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f11566,plain,(
  ( ! [X42] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X42) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X42)) ) | ($spl16 | $spl122)),
  inference(superposition,[],[f8159,f9796])).
fof(f9828,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl126),
  inference(cnf_transformation,[],[f9828_D])).
fof(f9828_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11565,plain,(
  ( ! [X41] : (s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f7068,f9796])).
fof(f9839,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl142),
  inference(cnf_transformation,[],[f9839_D])).
fof(f9839_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11584,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X13)) ) | $spl148),
  inference(cnf_transformation,[],[f11584_D])).
fof(f11584_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f11564,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X40) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f7068,f9796])).
fof(f11563,plain,(
  ( ! [X39] : (s(bool,X39) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X39)) ) | $spl122),
  inference(superposition,[],[f7068,f9796])).
fof(f11562,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X38) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X38)) ) | $spl122),
  inference(superposition,[],[f7068,f9796])).
fof(f9830,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl128),
  inference(cnf_transformation,[],[f9830_D])).
fof(f9830_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f9829,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl127),
  inference(cnf_transformation,[],[f9829_D])).
fof(f9829_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f11557,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X27) | s(bool,X27) = s(bool,g(s(bool,'T_0')))) ) | ($spl58 | $spl122)),
  inference(superposition,[],[f2465,f9796])).
fof(f9833,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl132),
  inference(cnf_transformation,[],[f9833_D])).
fof(f9833_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11556,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) ) | $spl122),
  inference(superposition,[],[f1867,f9796])).
fof(f9827,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl124),
  inference(cnf_transformation,[],[f9827_D])).
fof(f9827_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f9840,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X17)) ) | $spl144),
  inference(cnf_transformation,[],[f9840_D])).
fof(f9840_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11555,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X25) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f1867,f9796])).
fof(f11543,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl122),
  inference(superposition,[],[f233,f9796])).
fof(f11580,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl146),
  inference(cnf_transformation,[],[f11580_D])).
fof(f11580_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11542,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl122),
  inference(superposition,[],[f233,f9796])).
fof(f11541,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f233,f9796])).
fof(f11540,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl122),
  inference(superposition,[],[f233,f9796])).
fof(f11535,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl122),
  inference(superposition,[],[f134,f9796])).
fof(f11534,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f9796,f7068])).
fof(f11533,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f9796,f7068])).
fof(f11532,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X13)) ) | $spl122),
  inference(superposition,[],[f9796,f7068])).
fof(f11531,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X12) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X12)) ) | $spl122),
  inference(superposition,[],[f9796,f7068])).
fof(f11527,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl122),
  inference(superposition,[],[f9796,f233])).
fof(f11526,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl122),
  inference(superposition,[],[f9796,f233])).
fof(f11525,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl122),
  inference(superposition,[],[f9796,f233])).
fof(f11524,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl122),
  inference(superposition,[],[f9796,f233])).
fof(f11523,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl122),
  inference(superposition,[],[f9796,f134])).
fof(f11207,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11206,f131])).
fof(f11206,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f133,f11105])).
fof(f11105,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | ($spl105 | $spl106)),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2527,f2529,f2530,f2705,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2576,f2490,f2493,f4947,f5015,f5270,f5271,f5419,f5421,f5527,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f6128,f628,f5430,f5659,f5525,f5428,f1280,f1279,f1763,f2326,f1408,f2322,f1458,f2323,f1401,f2321,f1449,f2319,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f6184,f6161,f2310,f474,f6183,f5212,f2316,f2314,f2311,f962,f959,f957,f639,f637,f453,f417,f6230,f2333,f2332,f2308,f1698,f1686,f1664,f974,f651,f174,f173,f172,f5657,f5523,f5426,f1108,f1106,f9881,f9844,f9667,f9666,f9630,f9594,f9590,f9583,f9577,f9559,f9550,f9549,f9547,f9545,f9527,f9526,f9177,f9136,f8466,f8343,f8341,f8340,f8339,f8335,f8330,f8328,f8327,f8326,f8226,f8223,f8220,f8219,f8162,f8158,f8157,f8156,f7301,f7300,f7244,f7243,f7242,f7241,f7239,f7238,f7237,f7160,f7027,f7020,f6933,f6198,f6196,f6152,f6131,f6130,f6116,f5531,f5060,f4170,f4168,f4167,f4166,f4161,f4150,f4148,f4147,f4146,f4143,f2840,f2814,f2753,f2731,f2716,f2714,f2713,f2709,f2706,f2628,f2484,f2482,f2446,f2413,f2400,f2397,f2381,f2377,f2375,f2374,f2373,f2369,f2368,f2365,f2361,f2359,f2358,f2357,f2354,f2353,f2346,f2340,f2317,f2309,f2298,f2286,f1956,f1933,f1892,f1850,f1819,f1818,f1817,f1816,f1815,f1813,f1797,f1634,f1278,f1110,f1104,f964,f961,f641,f346,f345,f321,f320,f304,f303,f133,f8079,f7078,f6941,f6801,f6685,f6682,f5102,f4952,f2695,f2688,f2687,f2590,f2543,f2410,f2352,f1903,f1777,f10656,f11103,f11104,f10685])).
fof(f10685,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl105),
  inference(trivial_inequality_removal,[],[f10676])).
fof(f10676,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl105),
  inference(superposition,[],[f9666,f233])).
fof(f11104,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl105),
  inference(forward_demodulation,[],[f10684,f11103])).
fof(f10684,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl105),
  inference(trivial_inequality_removal,[],[f10679])).
fof(f10679,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X10) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl105),
  inference(superposition,[],[f9666,f1636])).
fof(f11103,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(condensation,[],[f11102])).
fof(f11102,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | $spl105),
  inference(subsumption_resolution,[],[f10682,f134])).
fof(f10682,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | $spl105),
  inference(superposition,[],[f9666,f7068])).
fof(f10656,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))))) ) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f187])).
fof(f11202,plain,(
  $false | ($spl74 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f3266,f11105])).
fof(f11201,plain,(
  $false | ($spl82 | $spl105 | $spl106)),
  inference(trivial_inequality_removal,[],[f11200])).
fof(f11200,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl82 | $spl105 | $spl106)),
  inference(forward_demodulation,[],[f3272,f11105])).
fof(f11199,plain,(
  $false | ($spl82 | $spl105 | $spl106)),
  inference(trivial_inequality_removal,[],[f11198])).
fof(f11198,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl82 | $spl105 | $spl106)),
  inference(forward_demodulation,[],[f3771,f11105])).
fof(f11193,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f9666,f11105])).
fof(f11135,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11068,f11105])).
fof(f11068,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(forward_demodulation,[],[f10653,f1561])).
fof(f10653,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f9666,f5732])).
fof(f11134,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11133,f131])).
fof(f11133,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10646,f11105])).
fof(f10646,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f131,f9666,f178])).
fof(f11132,plain,(
  $false | ($spl105 | $spl106)),
  inference(trivial_inequality_removal,[],[f11131])).
fof(f11131,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10645,f11105])).
fof(f10645,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))))) ) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f187])).
fof(f11129,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11128,f131])).
fof(f11128,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10638,f11105])).
fof(f10638,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f9666,f9666,f626])).
fof(f11127,plain,(
  $false | ($spl16 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11126,f131])).
fof(f11126,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl106)),
  inference(forward_demodulation,[],[f10637,f11105])).
fof(f10637,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f840,f840,f9666,f626])).
fof(f11125,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11124,f131])).
fof(f11124,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10636,f11105])).
fof(f10636,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f9666,f9666,f626])).
fof(f11123,plain,(
  $false | ($spl16 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11122,f131])).
fof(f11122,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl106)),
  inference(forward_demodulation,[],[f10635,f11105])).
fof(f10635,plain,(
  ~p(s(bool,'F_0')) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f840,f9666,f9666,f626])).
fof(f11121,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11120,f131])).
fof(f11120,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10632,f11105])).
fof(f10632,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f131,f9666,f178])).
fof(f11119,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11085,f11105])).
fof(f11085,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(forward_demodulation,[],[f10631,f1561])).
fof(f10631,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f5744])).
fof(f11118,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11087,f11105])).
fof(f11087,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(forward_demodulation,[],[f10629,f1561])).
fof(f10629,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f9666,f5732])).
fof(f11117,plain,(
  $false | ($spl16 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f10628,f11105])).
fof(f10628,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))),X0,bool))) ) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f9666,f2837])).
fof(f11116,plain,(
  $false | ($spl16 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f10624,f11105])).
fof(f10624,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f840,f9666,f1867])).
fof(f11114,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11113,f131])).
fof(f11113,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10614,f11105])).
fof(f10614,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f9666,f9666,f626])).
fof(f11112,plain,(
  $false | ($spl16 | $spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11111,f131])).
fof(f11111,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl106)),
  inference(forward_demodulation,[],[f10613,f11105])).
fof(f10613,plain,(
  ~p(s(bool,'F_0')) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f840,f9666,f9666,f626])).
fof(f11110,plain,(
  $false | ($spl105 | $spl106)),
  inference(subsumption_resolution,[],[f11109,f131])).
fof(f11109,plain,(
  ~p(s(bool,'T_0')) | ($spl105 | $spl106)),
  inference(forward_demodulation,[],[f10607,f11105])).
fof(f10607,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl105),
  inference(unit_resulting_resolution,[],[f9666,f235])).
fof(f11101,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11100,f131])).
fof(f11100,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10607,f10732])).
fof(f10732,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl105 | $spl110)),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f625,f647,f650,f1100,f1102,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1849,f1867,f2464,f2527,f2529,f2530,f2704,f2705,f2837,f2838,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f3264,f3265,f2876,f2877,f3268,f2882,f2738,f3269,f2636,f2576,f2490,f2493,f840,f4214,f4241,f4947,f5015,f5270,f5271,f5967,f5419,f5421,f6060,f6061,f5527,f5532,f5567,f5650,f5652,f5661,f6144,f6159,f5717,f5744,f5732,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7106,f7114,f7073,f7068,f7159,f7161,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f8159,f9143,f9159,f9307,f9395,f9396,f9548,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f6128,f628,f5430,f5659,f5525,f5428,f1280,f1279,f1763,f2326,f1408,f2322,f1458,f2323,f1401,f2321,f1449,f2319,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f6184,f6161,f2310,f474,f6183,f5212,f2316,f2314,f2311,f962,f959,f957,f639,f637,f453,f417,f6230,f2333,f2332,f2308,f1698,f1686,f1664,f974,f651,f174,f173,f172,f5657,f5523,f5426,f1108,f1106,f9881,f9844,f9669,f9666,f9630,f9594,f9590,f9583,f9577,f9559,f9550,f9549,f9547,f9545,f9527,f9526,f9461,f9415,f9207,f9177,f9136,f8466,f8343,f8341,f8340,f8339,f8335,f8330,f8328,f8327,f8326,f8226,f8223,f8220,f8219,f8162,f8158,f8157,f8156,f7301,f7300,f7244,f7243,f7242,f7241,f7239,f7238,f7237,f7160,f7108,f7027,f7020,f6933,f6198,f6196,f6152,f6131,f6130,f6116,f5531,f5060,f4170,f4168,f4167,f4166,f4161,f4150,f4148,f4147,f4146,f4143,f2840,f2814,f2753,f2731,f2716,f2714,f2713,f2709,f2706,f2628,f2484,f2482,f2446,f2413,f2400,f2397,f2381,f2377,f2375,f2374,f2373,f2369,f2368,f2365,f2361,f2359,f2358,f2357,f2354,f2353,f2346,f2340,f2317,f2309,f2298,f2286,f1956,f1933,f1892,f1850,f1819,f1818,f1817,f1816,f1815,f1813,f1797,f1634,f1278,f1110,f1104,f964,f961,f641,f346,f345,f321,f320,f304,f303,f133,f9171,f8079,f7078,f6941,f6801,f6685,f6682,f5102,f4952,f2695,f2688,f2687,f2590,f2543,f2410,f2352,f1903,f1777,f10656,f10654,f10691,f10731])).
fof(f10731,plain,(
  ( ! [X29,X27] : (s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10693,f840])).
fof(f10693,plain,(
  ( ! [X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10691,f1903])).
fof(f10691,plain,(
  ( ! [X23,X20] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X23,bool))) ) | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10690,f1634])).
fof(f10690,plain,(
  ( ! [X23,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X23,bool))) ) | ($spl16 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10654,f9669])).
fof(f10654,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl105)),
  inference(unit_resulting_resolution,[],[f840,f9666,f7068])).
fof(f11099,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11098,f131])).
fof(f11098,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10613,f10732])).
fof(f11097,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11096,f131])).
fof(f11096,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10614,f10732])).
fof(f11091,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10624,f10732])).
fof(f11089,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10628,f10732])).
fof(f11088,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11087,f10732])).
fof(f11086,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11085,f10732])).
fof(f11084,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11083,f131])).
fof(f11083,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10632,f10732])).
fof(f11082,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11081,f131])).
fof(f11081,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10635,f10732])).
fof(f11080,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11079,f131])).
fof(f11079,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10636,f10732])).
fof(f11078,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11077,f131])).
fof(f11077,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10637,f10732])).
fof(f11076,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11075,f131])).
fof(f11075,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10638,f10732])).
fof(f11072,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10645,f11034])).
fof(f11034,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0)))) ) | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10801,f188])).
fof(f10801,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f2753])).
fof(f11071,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11070,f131])).
fof(f11070,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(forward_demodulation,[],[f10646,f10732])).
fof(f11069,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f11068,f10732])).
fof(f11053,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10923,f10732])).
fof(f10923,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f9666])).
fof(f11036,plain,(
  $false | ($spl16 | $spl74 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10806,f10732])).
fof(f10806,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool))) ) | ($spl16 | $spl74 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f3266])).
fof(f11029,plain,(
  $false | ($spl16 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f10733,f131])).
fof(f10733,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f133])).
fof(f11028,plain,(
  $false | ($spl16 | $spl82 | $spl105 | $spl110)),
  inference(trivial_inequality_removal,[],[f10807])).
fof(f10807,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl16 | $spl82 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f3272])).
fof(f11027,plain,(
  $false | ($spl16 | $spl82 | $spl105 | $spl110)),
  inference(trivial_inequality_removal,[],[f10809])).
fof(f10809,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl16 | $spl82 | $spl105 | $spl110)),
  inference(backward_demodulation,[],[f10732,f3771])).
fof(f10381,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10249,f10056])).
fof(f10056,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl101 | $spl105)),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f625,f647,f650,f346,f345,f321,f320,f133,f628,f1100,f1102,f1104,f1277,f1276,f1286,f1288,f1292,f1294,f1297,f1299,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1402,f1404,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1450,f1452,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1634,f1635,f1636,f1641,f1673,f1712,f1717,f1744,f1752,f1777,f1813,f1815,f1816,f1797,f1819,f1849,f1850,f1867,f1892,f1903,f2298,f2340,f2346,f2352,f2354,f2359,f2361,f2368,f2369,f2375,f2377,f2397,f2410,f2446,f2464,f2482,f2527,f2529,f2530,f2543,f2687,f2688,f2590,f2695,f2704,f2705,f2706,f2709,f2628,f2714,f2731,f2753,f2837,f2838,f2840,f2814,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f3264,f3265,f2876,f2877,f3268,f2882,f2738,f3269,f2636,f2576,f2490,f2493,f840,f2484,f2413,f2400,f2353,f1280,f1279,f1278,f1956,f1763,f2326,f2373,f2374,f2381,f2357,f2358,f2365,f1408,f2322,f1458,f2323,f1449,f2319,f1401,f2321,f2320,f2318,f1457,f1448,f1407,f1400,f1304,f1303,f1291,f1290,f2716,f2713,f2309,f2286,f304,f2317,f2310,f1818,f474,f2316,f2314,f2311,f1933,f1817,f964,f962,f961,f959,f957,f641,f639,f637,f453,f417,f303,f2333,f2332,f2308,f1698,f1686,f1664,f974,f651,f174,f173,f172,f1106,f1108,f1110,f4143,f4146,f4147,f4148,f4150,f4161,f4166,f4167,f4168,f4170,f4214,f4241,f5102,f5060,f4947,f4952,f5212,f5015,f5270,f5271,f5967,f5419,f5421,f5426,f5428,f5430,f6060,f6061,f5523,f5525,f5527,f5531,f5532,f5567,f6116,f6128,f6130,f6131,f5650,f5652,f5657,f5659,f5661,f6144,f6152,f6159,f6161,f5717,f5744,f5732,f6183,f6184,f6196,f6198,f6230,f6299,f6261,f6262,f6376,f6377,f6537,f6469,f6538,f6522,f6523,f6685,f6682,f6622,f6647,f6725,f6760,f6761,f6767,f6768,f7078,f6801,f7106,f7108,f7114,f7073,f7068,f7159,f7160,f7161,f7237,f7238,f7239,f6933,f7241,f6941,f7242,f7243,f7244,f7300,f7020,f7301,f7027,f7754,f7734,f7757,f7752,f7838,f7884,f7885,f8006,f8079,f8156,f8157,f8158,f8159,f8162,f8219,f8220,f8223,f8226,f9136,f9143,f9159,f9171,f9177,f8326,f8327,f8328,f8330,f8335,f8339,f8340,f8341,f8343,f9207,f8466,f9307,f9395,f9396,f9415,f9461,f9577,f9526,f9527,f9583,f9545,f9547,f9548,f9549,f9550,f9590,f9559,f9594,f9630,f9645,f9646,f9647,f9648,f9649,f9650,f9651,f9652,f9653,f9654,f9655,f9656,f9666,f9663,f9844,f9881,f9945,f9941,f9978,f9904,f9946,f9948,f10004,f10015,f10055])).
fof(f10055,plain,(
  ( ! [X29,X27] : (s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10017,f840])).
fof(f10017,plain,(
  ( ! [X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl101 | $spl105)),
  inference(backward_demodulation,[],[f10015,f1903])).
fof(f10015,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10014,f9666])).
fof(f10014,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | $spl101),
  inference(subsumption_resolution,[],[f9967,f1634])).
fof(f9967,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | $spl101),
  inference(superposition,[],[f9663,f1636])).
fof(f10004,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X13)),X14,bool))) ) | ($spl16 | $spl101)),
  inference(subsumption_resolution,[],[f9955,f840])).
fof(f9955,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,f(s(bool,X13))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X13)),X14,bool))) ) | $spl101),
  inference(superposition,[],[f9663,f1636])).
fof(f9948,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl101),
  inference(superposition,[],[f9663,f134])).
fof(f9946,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | $spl101),
  inference(superposition,[],[f9663,f7068])).
fof(f9904,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))))) ) | $spl101),
  inference(unit_resulting_resolution,[],[f9663,f1816])).
fof(f9978,plain,(
  p(s(bool,f(s(bool,'F_0')))) | ($spl16 | $spl101)),
  inference(forward_demodulation,[],[f9940,f1561])).
fof(f9940,plain,(
  ( ! [X0] : (p(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))),X0,bool)))) ) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f1561,f9663,f627])).
fof(f9941,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))),X0,bool))) ) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f1636])).
fof(f9945,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,f(s(bool,'F_0')))))) ) | $spl101),
  inference(unit_resulting_resolution,[],[f9663,f187])).
fof(f10249,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f9666])).
fof(f10380,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10248,f10056])).
fof(f10248,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f9663])).
fof(f10363,plain,(
  $false | ($spl16 | $spl74 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10131,f10056])).
fof(f10131,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool))) ) | ($spl16 | $spl74 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f3266])).
fof(f10356,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10057,f131])).
fof(f10057,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f133])).
fof(f10355,plain,(
  $false | ($spl16 | $spl82 | $spl101 | $spl105)),
  inference(trivial_inequality_removal,[],[f10132])).
fof(f10132,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl16 | $spl82 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f3272])).
fof(f10354,plain,(
  $false | ($spl16 | $spl82 | $spl101 | $spl105)),
  inference(trivial_inequality_removal,[],[f10134])).
fof(f10134,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl16 | $spl82 | $spl101 | $spl105)),
  inference(backward_demodulation,[],[f10056,f3771])).
fof(f10011,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10010,f9666])).
fof(f10010,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(condensation,[],[f10009])).
fof(f10009,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl101),
  inference(subsumption_resolution,[],[f9964,f134])).
fof(f9964,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl101),
  inference(superposition,[],[f9663,f233])).
fof(f10008,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f10007,f9666])).
fof(f10007,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(condensation,[],[f10006])).
fof(f10006,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl101),
  inference(subsumption_resolution,[],[f9963,f134])).
fof(f9963,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl101),
  inference(superposition,[],[f9663,f233])).
fof(f10005,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9975,f9666])).
fof(f9975,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(trivial_inequality_removal,[],[f9961])).
fof(f9961,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(superposition,[],[f9663,f134])).
fof(f10003,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9895,f9666])).
fof(f9895,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(unit_resulting_resolution,[],[f9663,f134])).
fof(f9999,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9900,f9666])).
fof(f9900,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(unit_resulting_resolution,[],[f131,f9663,f1813])).
fof(f9996,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9995,f9666])).
fof(f9995,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(forward_demodulation,[],[f9908,f1561])).
fof(f9908,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f1850])).
fof(f9992,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9914,f9666])).
fof(f9914,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f7068])).
fof(f9991,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9917,f9666])).
fof(f9917,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f7068])).
fof(f9990,plain,(
  $false | ($spl86 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9989,f9666])).
fof(f9989,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl86 | $spl101)),
  inference(forward_demodulation,[],[f9919,f1561])).
fof(f9919,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | ($spl86 | $spl101)),
  inference(unit_resulting_resolution,[],[f9663,f7099])).
fof(f9987,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9986,f9666])).
fof(f9986,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(forward_demodulation,[],[f9921,f1561])).
fof(f9921,plain,(
  ( ! [X0] : (s(bool,f(s(bool,'F_0'))) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X0,bool))) ) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f1561,f9663,f8159])).
fof(f9985,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9924,f9666])).
fof(f9924,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f233])).
fof(f9984,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9927,f9666])).
fof(f9927,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f131,f840,f9663,f626])).
fof(f9983,plain,(
  $false | ($spl74 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9928,f9666])).
fof(f9928,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl74 | $spl101)),
  inference(unit_resulting_resolution,[],[f3266,f9663,f1636])).
fof(f9981,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9934,f9666])).
fof(f9934,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f840,f9663,f233])).
fof(f9980,plain,(
  $false | ($spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9935,f9666])).
fof(f9935,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(unit_resulting_resolution,[],[f133,f9663,f262])).
fof(f9979,plain,(
  $false | ($spl16 | $spl101 | $spl105)),
  inference(subsumption_resolution,[],[f9936,f9666])).
fof(f9936,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl101)),
  inference(unit_resulting_resolution,[],[f131,f840,f9663,f626])).
fof(f9842,plain,(
  $false | ($spl16 | $spl96 | $spl116)),
  inference(subsumption_resolution,[],[f9841,f840])).
fof(f9841,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl96 | $spl116)),
  inference(backward_demodulation,[],[f9660,f9673])).
fof(f9823,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X19) | s(bool,X19) = s(bool,g(s(bool,'T_0')))) ) | $spl123),
  inference(superposition,[],[f9797,f7068])).
fof(f9822,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl123),
  inference(superposition,[],[f9797,f7068])).
fof(f9821,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X17)) ) | $spl123),
  inference(superposition,[],[f9797,f7068])).
fof(f9820,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X16) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl123),
  inference(superposition,[],[f9797,f7068])).
fof(f9836,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl138),
  inference(cnf_transformation,[],[f9836_D])).
fof(f9836_D,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9819,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X10) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl123),
  inference(superposition,[],[f9797,f1636])).
fof(f9835,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | $spl136),
  inference(cnf_transformation,[],[f9835_D])).
fof(f9835_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f9818,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool))) ) | $spl123),
  inference(superposition,[],[f9797,f1636])).
fof(f9834,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X7,bool))) ) | $spl134),
  inference(cnf_transformation,[],[f9834_D])).
fof(f9834_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X7,bool))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9817,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'T_0')))),X7,bool))) ) | $spl123),
  inference(superposition,[],[f9797,f1636])).
fof(f9816,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X3) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl123),
  inference(superposition,[],[f9797,f233])).
fof(f9815,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl123),
  inference(superposition,[],[f9797,f233])).
fof(f9814,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl123),
  inference(superposition,[],[f9797,f233])).
fof(f9813,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl123),
  inference(superposition,[],[f9797,f233])).
fof(f9812,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl123),
  inference(superposition,[],[f9797,f134])).
fof(f9795,plain,(
  ( ! [X208,X207] : (s(fun(bool,bool),X207) = s(fun(bool,bool),X208)) ) | $spl120),
  inference(cnf_transformation,[],[f9795_D])).
fof(f9795_D,plain,(
  ( ! [X208,X207] : (s(fun(bool,bool),X207) = s(fun(bool,bool),X208)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f9797,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl123),
  inference(cnf_transformation,[],[f9797_D])).
fof(f9797_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f9763,plain,(
  ( ! [X208,X207] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X207) = s(fun(bool,bool),X208)) ) | $spl98),
  inference(duplicate_literal_removal,[],[f9748])).
fof(f9748,plain,(
  ( ! [X208,X207] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X207) = s(fun(bool,bool),X208) | s(fun(bool,bool),X207) = s(fun(bool,bool),X208)) ) | $spl98),
  inference(superposition,[],[f186,f9661])).
fof(f9671,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X34)) ) | $spl114),
  inference(cnf_transformation,[],[f9671_D])).
fof(f9671_D,plain,(
  ( ! [X34] : (s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X34)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f9661,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl98),
  inference(cnf_transformation,[],[f9661_D])).
fof(f9661_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f9663,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(cnf_transformation,[],[f9663_D])).
fof(f9663_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f9656,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'T_0') = s(bool,sK12(X41,X42,bool,bool)) | s(bool,'F_0') = s(bool,X43) | s(bool,X43) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1104,f7068])).
fof(f9673,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl116),
  inference(cnf_transformation,[],[f9673_D])).
fof(f9673_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f9674,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X35)) ) | $spl118),
  inference(cnf_transformation,[],[f9674_D])).
fof(f9674_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X35)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f9655,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,'T_0') = s(bool,sK12(X39,X40,bool,bool)) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1104,f7068])).
fof(f9654,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'T_0') = s(bool,sK12(X36,X37,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f1104,f7068])).
fof(f9653,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,sK12(X32,X33,bool,bool)) | s(bool,X34) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f1104,f7068])).
fof(f9670,plain,(
  ( ! [X26,X27] : (s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool))) ) | $spl112),
  inference(cnf_transformation,[],[f9670_D])).
fof(f9670_D,plain,(
  ( ! [X26,X27] : (s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f9666,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl105),
  inference(cnf_transformation,[],[f9666_D])).
fof(f9666_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f9652,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,sK12(X24,X25,bool,bool)) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool))) )),
  inference(superposition,[],[f1104,f1636])).
fof(f9660,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl96),
  inference(cnf_transformation,[],[f9660_D])).
fof(f9660_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9669,plain,(
  ( ! [X23,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X23,bool))) ) | $spl110),
  inference(cnf_transformation,[],[f9669_D])).
fof(f9669_D,plain,(
  ( ! [X23,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X23,bool))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f9651,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK12(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X23,bool))) )),
  inference(superposition,[],[f1104,f1636])).
fof(f9668,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'F_0')))),X19,bool))) ) | $spl108),
  inference(cnf_transformation,[],[f9668_D])).
fof(f9668_D,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'F_0')))),X19,bool))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f9667,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl106),
  inference(cnf_transformation,[],[f9667_D])).
fof(f9667_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f9650,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,sK12(X17,X18,bool,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,g(s(bool,'F_0')))),X19,bool))) )),
  inference(superposition,[],[f1104,f1636])).
fof(f9664,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl102),
  inference(cnf_transformation,[],[f9664_D])).
fof(f9664_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f9649,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK12(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1104,f233])).
fof(f9648,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK12(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1104,f233])).
fof(f9647,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK12(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1104,f233])).
fof(f9646,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK12(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1104,f233])).
fof(f9645,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1104,f134])).
fof(f9144,plain,(
  ( ! [X54,X53,X51] : (s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,X54) | s(bool,X51) = s(bool,X54)) ) | $spl94),
  inference(cnf_transformation,[],[f9144_D])).
fof(f9144_D,plain,(
  ( ! [X54,X53,X51] : (s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'F_0') = s(bool,X54) | s(bool,X51) = s(bool,X54)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f9307,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X54) | s(bool,X51) = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8690,f134])).
fof(f8690,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'F_0') != s(bool,X52) | s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X54) | s(bool,X51) = s(bool,X54)) )),
  inference(superposition,[],[f345,f7068])).
fof(f9143,plain,(
  ( ! [X54,X52,X53,X51] : (s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,X51) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8246,f134])).
fof(f8246,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,'F_0') != s(bool,X52) | s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X53),'const.semanticPrimitives.Rval_1'(s(bool,X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,X51) = s(bool,X54) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f345,f7068])).
fof(f7700,plain,(
  ( ! [X101,X99,X102] : (s(bool,X101) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X99,bool)) | s(bool,'T_0') = s(bool,X102) | s(bool,X101) = s(bool,X102)) ) | $spl92),
  inference(cnf_transformation,[],[f7700_D])).
fof(f7700_D,plain,(
  ( ! [X101,X99,X102] : (s(bool,X101) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X99,bool)) | s(bool,'T_0') = s(bool,X102) | s(bool,X101) = s(bool,X102)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7699,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,X101) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X99,bool)) | s(bool,X101) = s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | ($spl14 | $spl86)),
  inference(subsumption_resolution,[],[f7495,f261])).
fof(f7495,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X100)) | s(bool,X101) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X99,bool)) | s(bool,X101) = s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | $spl86),
  inference(superposition,[],[f626,f7099])).
fof(f7163,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X75)),X76,bool)) | s(bool,'F_0') = s(bool,X75)) ) | $spl90),
  inference(cnf_transformation,[],[f7163_D])).
fof(f7163_D,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X75)),X76,bool)) | s(bool,'F_0') = s(bool,X75)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7187,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl58),
  inference(condensation,[],[f7186])).
fof(f7186,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X161)),X162,bool)) | s(bool,'T_0') = s(bool,X163) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X161)) ) | $spl58),
  inference(subsumption_resolution,[],[f6889,f7160])).
fof(f6889,plain,(
  ( ! [X163,X161,X164,X162,X160] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X161)),X162,bool)) | s(bool,'T_0') = s(bool,X163) | s(bool,X163) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X160)),X164,bool)) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X161)) ) | $spl58),
  inference(superposition,[],[f2510,f1850])).
fof(f2510,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X21)),X22,bool))) ) | $spl58),
  inference(superposition,[],[f2465,f1561])).
fof(f7185,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X1)),X2,bool)) | s(bool,'F_0') = s(bool,X1)) ) | $spl58),
  inference(condensation,[],[f7184])).
fof(f7184,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X157)),X158,bool)) | s(bool,'T_0') = s(bool,X159) | s(bool,X156) = s(bool,X159) | s(bool,'F_0') = s(bool,X157)) ) | $spl58),
  inference(subsumption_resolution,[],[f6888,f2465])).
fof(f6888,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X157)),X158,bool)) | s(bool,'T_0') = s(bool,X159) | s(bool,X156) = s(bool,X159) | s(bool,'F_0') = s(bool,X156) | s(bool,'F_0') = s(bool,X157)) ) | $spl58),
  inference(superposition,[],[f2465,f1850])).
fof(f7162,plain,(
  ( ! [X78,X74,X77] : (s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X74))) = s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X78))) | s(bool,'F_0') = s(bool,X74)) ) | $spl88),
  inference(cnf_transformation,[],[f7162_D])).
fof(f7162_D,plain,(
  ( ! [X78,X74,X77] : (s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X74))) = s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X78))) | s(bool,'F_0') = s(bool,X74)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f7161,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X75)),X76,bool)) | s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X74))) = s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X78))) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f6870,f1819])).
fof(f6870,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X75)),X76,bool)) | s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X74))) = s('type.semanticPrimitives.result'(bool,X77),'const.semanticPrimitives.Rval_1'(s(bool,X78))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,X75)) )),
  inference(superposition,[],[f346,f1850])).
fof(f3270,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) | $spl78),
  inference(cnf_transformation,[],[f3270_D])).
fof(f3270_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7099,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'F_0') = s(bool,X1)) ) | $spl86),
  inference(cnf_transformation,[],[f7099_D])).
fof(f7099_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7115,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,X104) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X105,bool)) | s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X102)),X103,bool))) ) | ($spl68 | $spl74)),
  inference(subsumption_resolution,[],[f6825,f4044])).
fof(f4044,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X42,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X40)),X41,bool))) ) | $spl68),
  inference(superposition,[],[f3209,f1561])).
fof(f6825,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,X104) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X105,bool)) | s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X102)),X103,bool)) | s(bool,X102) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X106,bool))) ) | $spl74),
  inference(superposition,[],[f1850,f3826])).
fof(f3826,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X42,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X40)),X41,bool))) ) | $spl74),
  inference(superposition,[],[f3699,f1561])).
fof(f3699,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl74),
  inference(trivial_inequality_removal,[],[f3693])).
fof(f3693,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X8) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl74),
  inference(superposition,[],[f3266,f233])).
fof(f7073,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool)) | s(bool,X2) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X3,bool)) | s(bool,'F_0') = s(bool,X2)) )),
  inference(condensation,[],[f6821])).
fof(f6821,plain,(
  ( ! [X87,X85,X83,X86,X84] : (s(bool,X85) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X86,bool)) | s(bool,'F_0') = s(bool,X85) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X83)),X84,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X83)),X87,bool))) )),
  inference(superposition,[],[f1850,f1777])).
fof(f7114,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f7113])).
fof(f7113,plain,(
  ( ! [X80,X78,X81,X82] : (s(bool,X80) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X81,bool)) | s(bool,'F_0') = s(bool,X80) | s(bool,X78) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f6820,f6685])).
fof(f6820,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X80) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X81,bool)) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X78)),X79,bool)) | s(bool,X78) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f1850,f1636])).
fof(f7109,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X71) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X72,bool)) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X69)),X70,bool))) ) | $spl68),
  inference(subsumption_resolution,[],[f6818,f4752])).
fof(f4752,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X56) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X56)),X57,bool))) ) | $spl68),
  inference(superposition,[],[f4044,f1561])).
fof(f6818,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X71) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X72,bool)) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X69)),X70,bool)) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f1850,f1634])).
fof(f7106,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X1,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X2)),X3,bool))) )),
  inference(condensation,[],[f7105])).
fof(f7105,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X36,bool)) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool))) )),
  inference(subsumption_resolution,[],[f6810,f1892])).
fof(f6810,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X36,bool)) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X37)),X38,bool))) )),
  inference(superposition,[],[f1850,f1636])).
fof(f7101,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'F_0') = s(bool,X1)) ) | $spl58),
  inference(condensation,[],[f7100])).
fof(f7100,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X17,bool)) | s(bool,'F_0') = s(bool,X16) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl58),
  inference(subsumption_resolution,[],[f6806,f2465])).
fof(f6806,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X17,bool)) | s(bool,'F_0') = s(bool,X16) | s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f1850,f233])).
fof(f7098,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X2,bool)) | s(bool,'F_0') = s(bool,X1)) ) | $spl58),
  inference(condensation,[],[f7097])).
fof(f7097,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X4) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X5,bool)) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl58),
  inference(subsumption_resolution,[],[f6803,f2465])).
fof(f6803,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X4) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'T_0')),X5,bool)) | s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1850,f233])).
fof(f3555,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f3554,f3400])).
fof(f3400,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl8),
  inference(superposition,[],[f239,f239])).
fof(f3554,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,X39)) ) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f3474,f133])).
fof(f3474,plain,(
  ( ! [X39] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X39)) ) | ($spl8 | $spl16)),
  inference(superposition,[],[f625,f239])).
fof(f3553,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f3391,f133])).
fof(f3391,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f131,f239,f180])).
fof(f3552,plain,(
  $false | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f3397,f133])).
fof(f3397,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f239,f625])).
fof(f3398,plain,(
  $false | ($spl8 | $spl16)),
  inference(unit_resulting_resolution,[],[f133,f239,f625])).
fof(f3393,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f131,f133,f239,f180])).
fof(f3281,plain,(
  $false | ($spl10 | $spl16)),
  inference(subsumption_resolution,[],[f240,f840])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3272,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f3272_D])).
fof(f3272_D,plain,(
  ( ! [X59] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2493,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) )),
  inference(condensation,[],[f2492])).
fof(f2492,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2491])).
fof(f2491,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f2476])).
fof(f2476,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f321,f233])).
fof(f2490,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) )),
  inference(condensation,[],[f2489])).
fof(f2489,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2488])).
fof(f2488,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f2479])).
fof(f2479,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f321,f233])).
fof(f3280,plain,(
  ( ! [X28,X29] : (s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl58),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f172,f346,f345,f321,f320,f133,f303,f453,f641,f639,f637,f474,f417,f304,f173,f174,f628,f651,f957,f959,f961,f962,f964,f974,f1100,f1102,f1104,f1106,f1108,f1110,f1277,f1276,f1278,f1279,f1280,f1286,f1288,f1290,f1291,f1292,f1294,f1297,f1299,f1303,f1304,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1400,f1401,f1402,f1404,f1407,f1408,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1448,f1449,f1450,f1452,f1457,f1458,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1634,f1635,f1636,f1686,f1641,f1698,f1664,f1673,f1712,f1717,f1744,f1763,f1752,f1777,f1813,f1815,f1816,f1817,f1818,f1797,f1819,f1849,f1850,f1867,f1892,f1903,f1933,f1956,f2286,f2298,f2332,f2308,f2309,f2310,f2311,f2314,f2316,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2326,f2333,f2340,f2346,f2352,f2353,f2354,f2357,f2358,f2359,f2361,f2365,f2368,f2369,f2373,f2374,f2375,f2377,f2381,f2397,f2400,f2410,f2413,f2446,f2465,f2482,f2484,f2510,f2527,f2529,f2530,f2543,f2687,f2688,f2590,f2695,f2705,f2706,f2709,f2713,f2628,f2714,f2716,f2731,f2753,f2838,f2840,f2814,f2907,f2909,f3254,f3255,f3257,f3260,f3263,f2738,f3269,f2636,f2637,f2585,f2576,f3275,f2516,f2519,f3278,f3279])).
fof(f3279,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X29) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool)) | s(bool,'T_0') = s(bool,X29)) ) | $spl58),
  inference(subsumption_resolution,[],[f2481,f2510])).
fof(f2481,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X28),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X29) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X26)),X27,bool)) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f321,f1634])).
fof(f2637,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') = s(bool,X105) | s(bool,X105) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X103)),X104,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X103)),X106,bool))) ) | $spl58),
  inference(trivial_inequality_removal,[],[f2623])).
fof(f2623,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X105) | s(bool,X105) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X103)),X104,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X103)),X106,bool))) ) | $spl58),
  inference(superposition,[],[f2465,f1777])).
fof(f3278,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) ) | $spl58),
  inference(condensation,[],[f3277])).
fof(f3277,plain,(
  ( ! [X4,X2,X3] : (s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl58),
  inference(subsumption_resolution,[],[f2483,f2465])).
fof(f2483,plain,(
  ( ! [X4,X2,X3] : (s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X2),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f321,f1816])).
fof(f241,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2519,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl58),
  inference(condensation,[],[f2518])).
fof(f2518,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl58),
  inference(duplicate_literal_removal,[],[f2517])).
fof(f2517,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl58),
  inference(condensation,[],[f2504])).
fof(f2504,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl58),
  inference(superposition,[],[f2465,f233])).
fof(f2516,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl58),
  inference(condensation,[],[f2515])).
fof(f2515,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1)) ) | $spl58),
  inference(duplicate_literal_removal,[],[f2514])).
fof(f2514,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl58),
  inference(condensation,[],[f2507])).
fof(f2507,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl58),
  inference(superposition,[],[f2465,f233])).
fof(f3276,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X20) | s(bool,X20) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X18)),X19,bool))) ) | $spl84),
  inference(cnf_transformation,[],[f3276_D])).
fof(f3276_D,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X20) | s(bool,X20) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X18)),X19,bool))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3275,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X20) | s(bool,X20) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X18)),X19,bool))) ) | $spl58),
  inference(subsumption_resolution,[],[f2509,f2510])).
fof(f2509,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X20) | s(bool,X20) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X18)),X19,bool)) | s(bool,'F_0') = s(bool,X18)) ) | $spl58),
  inference(superposition,[],[f2465,f1634])).
fof(f3271,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X101,bool)) | s(bool,X100) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X99,bool))) ) | $spl80),
  inference(cnf_transformation,[],[f3271_D])).
fof(f3271_D,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X101,bool)) | s(bool,X100) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X99,bool))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2576,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X59),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) | s(bool,X60) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X57)),X58,bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X57)),X61,bool))) )),
  inference(superposition,[],[f321,f1777])).
fof(f2585,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X100) | s(bool,X100) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X99,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X98)),X101,bool))) ) | $spl58),
  inference(superposition,[],[f2465,f1777])).
fof(f2636,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool))) )),
  inference(equality_factoring,[],[f1777])).
fof(f2882,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2849])).
fof(f2849,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f2837,f134])).
fof(f3268,plain,(
  ( ! [X22,X20] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X22,bool)) | s(bool,'T_0') = s(bool,X20)) ) | $spl16),
  inference(subsumption_resolution,[],[f2857,f2837])).
fof(f2857,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X22,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X20)),X21,bool)) | s(bool,'T_0') = s(bool,X20)) ) | $spl16),
  inference(superposition,[],[f2837,f2543])).
fof(f2877,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X1,bool)) | s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X2,bool))) ) | $spl16),
  inference(condensation,[],[f2858])).
fof(f2858,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X25,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X23)),X24,bool)) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X23)),X26,bool))) ) | $spl16),
  inference(superposition,[],[f2837,f1777])).
fof(f2876,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2862])).
fof(f2862,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f2837,f2543])).
fof(f3265,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(subsumption_resolution,[],[f2863,f2837])).
fof(f2863,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X2)),X4,bool))) ) | $spl16),
  inference(superposition,[],[f2837,f1777])).
fof(f3264,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f2866,f2837])).
fof(f2866,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X10)),X11,bool))) ) | $spl16),
  inference(superposition,[],[f2837,f134])).
fof(f3261,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) ) | $spl72),
  inference(cnf_transformation,[],[f3261_D])).
fof(f3261_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3263,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) )),
  inference(condensation,[],[f3262])).
fof(f3262,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,X3))) = s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2892,f413])).
fof(f2892,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,X3))) = s('type.semanticPrimitives.result'(bool,X4),'const.semanticPrimitives.Rval_1'(s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f346,f233])).
fof(f3260,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X0))) = s('type.semanticPrimitives.result'(bool,X1),'const.semanticPrimitives.Rval_1'(s(bool,X2)))) )),
  inference(condensation,[],[f3259])).
fof(f3259,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X15))) = s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2895,f413])).
fof(f2895,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X15))) = s('type.semanticPrimitives.result'(bool,X16),'const.semanticPrimitives.Rval_1'(s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f346,f233])).
fof(f3258,plain,(
  ( ! [X37,X35,X33,X36] : (s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X33))) = s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X36))) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X37,bool))) ) | $spl70),
  inference(cnf_transformation,[],[f3258_D])).
fof(f3258_D,plain,(
  ( ! [X37,X35,X33,X36] : (s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X33))) = s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X36))) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X37,bool))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3257,plain,(
  ( ! [X37,X35,X33,X36] : (s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X33))) = s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X36))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X37,bool))) )),
  inference(subsumption_resolution,[],[f3256,f2695])).
fof(f3256,plain,(
  ( ! [X37,X35,X33,X36] : (s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X33))) = s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X36))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X37,bool))) )),
  inference(forward_demodulation,[],[f2899,f1561])).
fof(f2899,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') != s(bool,'F_0') | s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,X36))) = s('type.semanticPrimitives.result'(bool,X35),'const.semanticPrimitives.Rval_1'(s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X34,bool)))) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X33)),X37,bool))) )),
  inference(superposition,[],[f346,f1777])).
fof(f3255,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X38)) )),
  inference(global_subsumption,[],[f130,f129,f128,f131,f134,f178,f175,f180,f185,f186,f188,f187,f190,f189,f191,f206,f207,f233,f235,f238,f259,f262,f263,f316,f377,f413,f415,f624,f449,f626,f627,f470,f472,f636,f172,f346,f345,f321,f320,f133,f303,f453,f641,f639,f637,f474,f417,f304,f173,f174,f628,f651,f957,f959,f961,f962,f964,f974,f1100,f1102,f1104,f1106,f1108,f1110,f1277,f1276,f1278,f1279,f1280,f1286,f1288,f1290,f1291,f1292,f1294,f1297,f1299,f1303,f1304,f1305,f1307,f1371,f1372,f1395,f1396,f1397,f1399,f1400,f1401,f1402,f1404,f1407,f1408,f1409,f1411,f1436,f1437,f1443,f1444,f1445,f1447,f1448,f1449,f1450,f1452,f1457,f1458,f1459,f1461,f1561,f1581,f1582,f1583,f1587,f1634,f1635,f1636,f1686,f1641,f1698,f1664,f1673,f1712,f1717,f1744,f1763,f1752,f1777,f1813,f1815,f1816,f1817,f1818,f1797,f1819,f1849,f1850,f1867,f1892,f1903,f1933,f1956,f2286,f2298,f2332,f2308,f2309,f2310,f2311,f2314,f2316,f2317,f2318,f2319,f2320,f2321,f2322,f2323,f2326,f2333,f2340,f2346,f2352,f2353,f2354,f2357,f2358,f2359,f2361,f2365,f2368,f2369,f2373,f2374,f2375,f2377,f2381,f2397,f2400,f2410,f2413,f2446,f2482,f2484,f2527,f2529,f2530,f2543,f2687,f2688,f2590,f2695,f2705,f2706,f2709,f2713,f2628,f2714,f2716,f2731,f2753,f2838,f2840,f2814,f2907,f2909,f3254])).
fof(f3209,plain,(
  ( ! [X92,X91] : (s(bool,X92) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X91,bool)) | s(bool,'F_0') = s(bool,X92)) ) | $spl68),
  inference(cnf_transformation,[],[f3209_D])).
fof(f3209_D,plain,(
  ( ! [X92,X91] : (s(bool,X92) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X91,bool)) | s(bool,'F_0') = s(bool,X92)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3167,plain,(
  ( ! [X92,X90,X91] : (s(bool,X92) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X91,bool)) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X90)) ) | ($spl11 | $spl16)),
  inference(duplicate_literal_removal,[],[f3073])).
fof(f3073,plain,(
  ( ! [X92,X90,X91] : (s(bool,X92) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X91,bool)) | s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | ($spl11 | $spl16)),
  inference(superposition,[],[f1867,f2985])).
fof(f2985,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | ($spl11 | $spl16)),
  inference(trivial_inequality_removal,[],[f2979])).
fof(f2979,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X8) = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X7,bool)) | s(bool,'T_0') = s(bool,X8)) ) | ($spl11 | $spl16)),
  inference(superposition,[],[f2847,f233])).
fof(f2847,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X0,bool))) ) | ($spl11 | $spl16)),
  inference(unit_resulting_resolution,[],[f241,f2837])).
fof(f3202,plain,(
  ( ! [X50,X51,X49] : (s(bool,X50) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X49,bool)) | s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,X50)))) ) | $spl66),
  inference(cnf_transformation,[],[f3202_D])).
fof(f3202_D,plain,(
  ( ! [X50,X51,X49] : (s(bool,X50) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X49,bool)) | s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,X50)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3171,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X50) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X49,bool)) | s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,X50))) | s(bool,'T_0') = s(bool,X48)) ) | ($spl11 | $spl16)),
  inference(duplicate_literal_removal,[],[f3061])).
fof(f3061,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X50) != s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X49,bool)) | s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.result'(bool,X51),'const.semanticPrimitives.Rval_1'(s(bool,X50))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | ($spl11 | $spl16)),
  inference(superposition,[],[f345,f2985])).
fof(f2660,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl60),
  inference(cnf_transformation,[],[f2660_D])).
fof(f2660_D,plain,(
  ( ! [X2,X0] : (s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2664,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,'T_0') = s(bool,X8)) ) | $spl64),
  inference(cnf_transformation,[],[f2664_D])).
fof(f2664_D,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2530,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X12)),X13,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X11)),X14,bool)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1777,f233])).
fof(f2529,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X8)),X9,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X7)),X10,bool)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1777,f233])).
fof(f2661,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool))) ) | $spl62),
  inference(cnf_transformation,[],[f2661_D])).
fof(f2661_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2527,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,'F_0')),X1,bool)) | s(bool,'F_0') = s(bool,sK10('const.semanticPrimitives.Rval_1'(s(bool,X0)),X2,bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1777,f134])).
fof(f2465,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl58),
  inference(cnf_transformation,[],[f2465_D])).
fof(f2465_D,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2464,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | $spl16),
  inference(subsumption_resolution,[],[f2458,f625])).
fof(f2458,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f627,f134])).
fof(f2258,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X48),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X48,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,X49)) ) | $spl56),
  inference(cnf_transformation,[],[f2258_D])).
fof(f2258_D,plain,(
  ( ! [X48,X49] : (s(bool,X49) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X48),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X48,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,X49)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2271,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | s(bool,X53) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X51),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'F_0') = s(bool,X53) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X51,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f2210,f625])).
fof(f2210,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X52)) | s(bool,X53) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X51),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'F_0') = s(bool,X53) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X51,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X52)) ) | $spl36),
  inference(superposition,[],[f1815,f1709])).
fof(f1709,plain,(
  ( ! [X2,X1] : (s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(trivial_inequality_removal,[],[f1704])).
fof(f1704,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,X2) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(superposition,[],[f1243,f233])).
fof(f2095,plain,(
  ( ! [X87] : (s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X87,'type.semanticPrimitives.result'(sK0,sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X87),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) | $spl48),
  inference(cnf_transformation,[],[f2095_D])).
fof(f2095_D,plain,(
  ( ! [X87] : (s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X87,'type.semanticPrimitives.result'(sK0,sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X87),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2254,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl52),
  inference(cnf_transformation,[],[f2254_D])).
fof(f2254_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2174,plain,(
  ( ! [X80,X79] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X80),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X79,'type.semanticPrimitives.result'(sK0,sK3))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X80,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X79),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) | ($spl0 | $spl36)),
  inference(superposition,[],[f1550,f1709])).
fof(f1550,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X2),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X2,'type.semanticPrimitives.result'(sK0,sK3)))) ) | $spl0),
  inference(forward_demodulation,[],[f1475,f226])).
fof(f1475,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X2),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X2,'type.semanticPrimitives.result'(sK0,sK3)))) ) | $spl0),
  inference(superposition,[],[f173,f226])).
fof(f2092,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') = s(bool,sK9('F_0',X80,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,X81))) | s(bool,'T_0') = s(bool,X81)) ) | $spl46),
  inference(cnf_transformation,[],[f2092_D])).
fof(f2092_D,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') = s(bool,sK9('F_0',X80,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,X81))) | s(bool,'T_0') = s(bool,X81)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2173,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X78),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,'F_0') = s(bool,sK9('F_0',X76,bool)) | s(bool,'T_0') = s(bool,X77) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X78,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X76),s(bool,X77)))) ) | $spl36),
  inference(superposition,[],[f303,f1709])).
fof(f2091,plain,(
  ( ! [X74,X72,X73] : (s(X72,X74) = s(X72,sK9(X74,X73,X72)) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) ) | $spl44),
  inference(cnf_transformation,[],[f2091_D])).
fof(f2091_D,plain,(
  ( ! [X74,X72,X73] : (s(X72,X74) = s(X72,sK9(X74,X73,X72)) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2171,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X71),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(X68,X70) = s(X68,sK9(X70,X69,X68)) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X71,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,X70)))) ) | $spl36),
  inference(superposition,[],[f173,f1709])).
fof(f2090,plain,(
  ( ! [X68,X71,X69,X67] : (p(s(bool,happ(s(fun(X67,bool),X68),s(X67,X71)))) | s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(X67,X69) != s(X67,X71)) ) | $spl42),
  inference(cnf_transformation,[],[f2090_D])).
fof(f2090_D,plain,(
  ( ! [X68,X71,X69,X67] : (p(s(bool,happ(s(fun(X67,bool),X68),s(X67,X71)))) | s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69))) | s(X67,X69) != s(X67,X71)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2170,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X66),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | p(s(bool,happ(s(fun(X63,bool),X64),s(X63,X67)))) | s(X63,X65) != s(X63,X67) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X66,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),X64),s(X63,X65)))) ) | $spl36),
  inference(superposition,[],[f172,f1709])).
fof(f2229,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X48),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X47) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X48,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f2168])).
fof(f2168,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X48),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X47) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X48,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X47)) ) | $spl36),
  inference(superposition,[],[f1867,f1709])).
fof(f2257,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X42),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,'F_0') = s(bool,X43) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X42,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X41)) ) | $spl36),
  inference(subsumption_resolution,[],[f2166,f1867])).
fof(f2166,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X42),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X43) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X42,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X41)) ) | $spl36),
  inference(superposition,[],[f1813,f1709])).
fof(f2256,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X34),s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool))) ) | $spl54),
  inference(cnf_transformation,[],[f2256_D])).
fof(f2256_D,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X34),s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2230,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X34),s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool)) | s(bool,'T_0') = s(bool,X33) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f2163])).
fof(f2163,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X34),s('type.semanticPrimitives.error_result'(sK3),sK6)))))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool)) | s(bool,'T_0') = s(bool,X33) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X33)) ) | $spl36),
  inference(superposition,[],[f303,f1709])).
fof(f2231,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X27),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X27,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f2160])).
fof(f2160,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X27),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X27,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X26)) ) | $spl36),
  inference(superposition,[],[f235,f1709])).
fof(f2255,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X12),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | p(s(bool,X13)) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X12,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X11)) ) | $spl36),
  inference(subsumption_resolution,[],[f2155,f262])).
fof(f2155,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X12),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X12,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X11)) ) | $spl36),
  inference(superposition,[],[f180,f1709])).
fof(f2253,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl36),
  inference(condensation,[],[f2252])).
fof(f2252,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X9),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s(bool,X8) = s(bool,X10) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X9,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X8)) ) | $spl36),
  inference(subsumption_resolution,[],[f2154,f262])).
fof(f2154,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X9),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | ~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X9,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X8)) ) | $spl36),
  inference(superposition,[],[f178,f1709])).
fof(f2112,plain,(
  ( ! [X10,X8] : (s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X10),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X8) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X10,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl50),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ( ! [X10,X8] : (s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X10),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X8) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X10,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1929,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X7),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X7,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f1920,f625])).
fof(f1920,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X7),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X7,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl24),
  inference(superposition,[],[f1552,f233])).
fof(f1552,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X6) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X5),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X5,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl24),
  inference(subsumption_resolution,[],[f926,f133])).
fof(f926,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X6) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X5),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X5,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl24),
  inference(superposition,[],[f175,f767])).
fof(f1930,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X9)) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X10),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X10,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f1921,f625])).
fof(f1921,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X9)) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X10),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X10,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl24),
  inference(superposition,[],[f1552,f233])).
fof(f2086,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl38),
  inference(cnf_transformation,[],[f2086_D])).
fof(f2086_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2002,plain,(
  ( ! [X88,X87] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X88),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),sK9('const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)),X87,'type.semanticPrimitives.result'(sK0,sK3))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X88,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK0,sK3),bool),X87),s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) | ($spl0 | $spl26)),
  inference(superposition,[],[f1550,f952])).
fof(f952,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X3),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X3,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(trivial_inequality_removal,[],[f948])).
fof(f948,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X3,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,X4) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X3),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(superposition,[],[f781,f233])).
fof(f1999,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X82),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,'F_0') = s(bool,sK9('F_0',X80,bool)) | s(bool,'T_0') = s(bool,X81) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X82,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,X81)))) ) | $spl26),
  inference(superposition,[],[f303,f952])).
fof(f1997,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X75),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(X72,X74) = s(X72,sK9(X74,X73,X72)) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X75,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),X73),s(X72,X74)))) ) | $spl26),
  inference(superposition,[],[f173,f952])).
fof(f1996,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X70),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | p(s(bool,happ(s(fun(X67,bool),X68),s(X67,X71)))) | s(X67,X69) != s(X67,X71) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X70,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,X69)))) ) | $spl26),
  inference(superposition,[],[f172,f952])).
fof(f2089,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X45),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,'F_0') = s(bool,X46) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X45,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X44)) ) | $spl26),
  inference(subsumption_resolution,[],[f1992,f1867])).
fof(f1992,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X45),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,X44) = s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X45,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X44)) ) | $spl26),
  inference(superposition,[],[f1813,f952])).
fof(f2088,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X34),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool))) ) | $spl40),
  inference(cnf_transformation,[],[f2088_D])).
fof(f2088_D,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X34),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2061,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X34),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool)) | s(bool,'T_0') = s(bool,X33) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f1988])).
fof(f1988,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X34),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))) | s(bool,'F_0') = s(bool,sK9('F_0',X35,bool)) | s(bool,'T_0') = s(bool,X33) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X34,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X33)) ) | $spl26),
  inference(superposition,[],[f303,f952])).
fof(f2062,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X27),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X27,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f1985])).
fof(f1985,plain,(
  ( ! [X26,X27] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X27),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,'T_0') = s(bool,X26) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X27,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X26)) ) | $spl26),
  inference(superposition,[],[f235,f952])).
fof(f2087,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X12),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | p(s(bool,X13)) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X12,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X11)) ) | $spl26),
  inference(subsumption_resolution,[],[f1980,f262])).
fof(f1980,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X12),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | p(s(bool,X13)) | s(bool,X11) != s(bool,X13) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X12,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X11)) ) | $spl26),
  inference(superposition,[],[f180,f952])).
fof(f2085,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl26),
  inference(condensation,[],[f2084])).
fof(f2084,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X9),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s(bool,X8) = s(bool,X10) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X9,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | $spl26),
  inference(subsumption_resolution,[],[f1979,f262])).
fof(f1979,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X9),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | ~p(s(bool,X10)) | s(bool,X8) = s(bool,X10) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X9,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X8)) ) | $spl26),
  inference(superposition,[],[f178,f952])).
fof(f1546,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1535,f1507])).
fof(f1507,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(X0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)))))) ) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f1492,f189])).
fof(f1492,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) | ($spl2 | $spl6)),
  inference(unit_resulting_resolution,[],[f232,f228])).
fof(f1535,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))))) | ($spl0 | $spl6)),
  inference(superposition,[],[f226,f1442])).
fof(f1442,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),'const.evalProps.map_error_result_2'(s(fun(sK2,X1),X2),s('type.semanticPrimitives.error_result'(sK2),sK8))))) = s('type.semanticPrimitives.result'(X0,X1),'const.evalProps.map_result_3'(s(fun(sK1,X0),X3),s(fun(sK2,X1),X2),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) ) | $spl6),
  inference(superposition,[],[f207,f232])).
fof(f1545,plain,(
  $false | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1521,f1507])).
fof(f1521,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))))) | ($spl0 | $spl6)),
  inference(superposition,[],[f1442,f226])).
fof(f1503,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1497,f232])).
fof(f1497,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | ($spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f1496])).
fof(f1496,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),sK6) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | ($spl2 | $spl4)),
  inference(superposition,[],[f228,f230])).
fof(f1502,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1501,f232])).
fof(f1501,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f1500,f1312])).
fof(f1312,plain,(
  s('type.semanticPrimitives.error_result'(sK2),sK8) = s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)) | $spl6),
  inference(unit_resulting_resolution,[],[f1275,f668])).
fof(f668,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),X0))) | s('type.semanticPrimitives.error_result'(sK2),sK8) = s('type.semanticPrimitives.error_result'(sK2),X0)) ) | $spl6),
  inference(superposition,[],[f189,f232])).
fof(f1275,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | $spl6),
  inference(unit_resulting_resolution,[],[f295,f185])).
fof(f295,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rval_1'(s(sK1,X0)))) ) | $spl6),
  inference(superposition,[],[f191,f232])).
fof(f1500,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1495,f230])).
fof(f1495,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK11(sK7,sK2,sK1)))) | ($spl2 | $spl6)),
  inference(superposition,[],[f228,f1312])).
fof(f1499,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1492,f230])).
fof(f1498,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1493,f232])).
fof(f1493,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f230,f228])).
fof(f1494,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f232,f230,f228])).
fof(f1470,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f1469])).
fof(f1469,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))) ) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f1466,f230])).
fof(f1466,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)))))))) ) | ($spl1 | $spl6)),
  inference(backward_demodulation,[],[f1442,f315])).
fof(f315,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f227,f187])).
fof(f1468,plain,(
  $false | ($spl1 | $spl4 | $spl6)),
  inference(trivial_inequality_removal,[],[f1467])).
fof(f1467,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6)))))))) ) | ($spl1 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f1465,f230])).
fof(f1465,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)))))))))) ) | ($spl1 | $spl6)),
  inference(backward_demodulation,[],[f1442,f666])).
fof(f666,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(sK0,sK3),X0),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f315,f187])).
fof(f1231,plain,(
  ( ! [X15,X16] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X16),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X15),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X15,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl30),
  inference(cnf_transformation,[],[f1231_D])).
fof(f1231_D,plain,(
  ( ! [X15,X16] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X16),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X15),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X15,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.error_result'(sK3)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1255,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X46)) | s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X45),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X47),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X47,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X45,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl6 | $spl16)),
  inference(subsumption_resolution,[],[f1194,f625])).
fof(f1194,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X46)) | s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X45),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X47),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X47,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X45,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X46)) ) | ($spl4 | $spl6 | $spl16)),
  inference(superposition,[],[f773,f630])).
fof(f630,plain,(
  ( ! [X35,X36] : (s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X35,'type.semanticPrimitives.error_result'(sK3))) | s(bool,X36) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X35),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X36)) ) | $spl4),
  inference(subsumption_resolution,[],[f469,f131])).
fof(f469,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X35,'type.semanticPrimitives.error_result'(sK3))) | s(bool,X36) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X35),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X36)) ) | $spl4),
  inference(superposition,[],[f307,f233])).
fof(f307,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl4),
  inference(forward_demodulation,[],[f302,f230])).
fof(f302,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) | s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl4),
  inference(superposition,[],[f173,f230])).
fof(f773,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl6 | $spl16)),
  inference(condensation,[],[f772])).
fof(f772,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X6),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X6,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f724,f625])).
fof(f724,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X8) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X6),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X6,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X7)) ) | $spl6),
  inference(superposition,[],[f175,f629])).
fof(f629,plain,(
  ( ! [X33,X34] : (s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X33),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X33,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(subsumption_resolution,[],[f468,f131])).
fof(f468,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X33,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,X34) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X33),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f306,f233])).
fof(f306,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X0),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X0,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl6),
  inference(forward_demodulation,[],[f301,f232])).
fof(f301,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X0),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) | s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X0,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl6),
  inference(superposition,[],[f173,f232])).
fof(f1243,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X68),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X68,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl36),
  inference(cnf_transformation,[],[f1243_D])).
fof(f1243_D,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X68),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X68,'type.semanticPrimitives.error_result'(sK3)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1254,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X40),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X40,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl16)),
  inference(subsumption_resolution,[],[f1192,f625])).
fof(f1192,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X40),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X40,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f625,f630])).
fof(f1230,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) ) | $spl28),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1252,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X35),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X35,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f1190,f625])).
fof(f1190,plain,(
  ( ! [X35,X36] : (p(s(bool,X36)) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X35),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X35,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl4 | $spl14)),
  inference(superposition,[],[f261,f630])).
fof(f1234,plain,(
  ( ! [X29] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X29,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl34),
  inference(cnf_transformation,[],[f1234_D])).
fof(f1234_D,plain,(
  ( ! [X29] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X29,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1244,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X70),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X69,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X70,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl24 | $spl26)),
  inference(subsumption_resolution,[],[f1164,f781])).
fof(f1164,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X70),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X69,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X70,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X69),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f767,f630])).
fof(f1242,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X68),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X67,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X68,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f1163,f781])).
fof(f1163,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X68),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X67,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X68,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X67),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) ) | ($spl4 | $spl26)),
  inference(superposition,[],[f781,f630])).
fof(f1209,plain,(
  ( ! [X37,X35,X36] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X36),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X35) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X36,'type.semanticPrimitives.error_result'(sK3)))) ) | ($spl4 | $spl6)),
  inference(duplicate_literal_removal,[],[f1157])).
fof(f1157,plain,(
  ( ! [X37,X35,X36] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X36),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X35) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X36,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl4 | $spl6)),
  inference(superposition,[],[f629,f630])).
fof(f1215,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1144])).
fof(f1144,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f134,f630])).
fof(f1233,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X30),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X30,'type.semanticPrimitives.error_result'(sK3))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X29,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl4 | $spl24 | $spl26)),
  inference(subsumption_resolution,[],[f1119,f781])).
fof(f1119,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X30),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X30,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X29),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X29,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f630,f767])).
fof(f1232,plain,(
  ( ! [X19,X18] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X18),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X19),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X18,'type.semanticPrimitives.error_result'(sK3))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X19,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl32),
  inference(cnf_transformation,[],[f1232_D])).
fof(f1232_D,plain,(
  ( ! [X19,X18] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X18),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X19),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X18,'type.semanticPrimitives.error_result'(sK3))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X19,'type.semanticPrimitives.error_result'(sK3)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1223,plain,(
  ( ! [X19,X17,X18] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X18),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X19),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X19,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X17) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X18,'type.semanticPrimitives.error_result'(sK3)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1117])).
fof(f1117,plain,(
  ( ! [X19,X17,X18] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X18),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X19),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X19,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X17) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X18,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X17)) ) | $spl4),
  inference(superposition,[],[f630,f630])).
fof(f1224,plain,(
  ( ! [X14,X15,X16] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X16),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X15),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X14) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X15,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl4 | $spl6)),
  inference(duplicate_literal_removal,[],[f1116])).
fof(f1116,plain,(
  ( ! [X14,X15,X16] : (s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X16),s('type.semanticPrimitives.error_result'(sK3),sK6))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X15),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X14) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X15,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl4 | $spl6)),
  inference(superposition,[],[f630,f629])).
fof(f1229,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f1111])).
fof(f1111,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.error_result'(sK3),bool),X1),s('type.semanticPrimitives.error_result'(sK3),sK6))) | s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),sK9('const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.error_result'(sK3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f630,f134])).
fof(f839,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f783,f791])).
fof(f791,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f649,f649,f178])).
fof(f783,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f133,f649,f180])).
fof(f838,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f786,f131])).
fof(f786,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f241,f649,f178])).
fof(f837,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f788,f649])).
fof(f788,plain,(
  ~p(s(bool,'F_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f241,f649,f178])).
fof(f836,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f792,f131])).
fof(f792,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f241,f649,f178])).
fof(f835,plain,(
  $false | ($spl11 | $spl20)),
  inference(subsumption_resolution,[],[f794,f649])).
fof(f794,plain,(
  ~p(s(bool,'F_0')) | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f241,f649,f178])).
fof(f834,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f821,f131])).
fof(f821,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f796,f133])).
fof(f796,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f649,f235])).
fof(f833,plain,(
  $false | ($spl11 | $spl20)),
  inference(trivial_inequality_removal,[],[f823])).
fof(f823,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))) ) | ($spl11 | $spl20)),
  inference(backward_demodulation,[],[f796,f314])).
fof(f314,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))) != s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f241,f187])).
fof(f832,plain,(
  $false | ($spl11 | $spl20)),
  inference(trivial_inequality_removal,[],[f826])).
fof(f826,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0')))))) ) | ($spl11 | $spl20)),
  inference(backward_demodulation,[],[f796,f332])).
fof(f332,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'T_0'))))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.result'(bool,X0),X1),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.result'(bool,X0),'const.semanticPrimitives.Rval_1'(s(bool,'F_0')))))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f314,f187])).
fof(f797,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f241,f649,f235])).
fof(f795,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f649,f241,f649,f178])).
fof(f793,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f131,f241,f649,f178])).
fof(f789,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f649,f241,f649,f178])).
fof(f787,plain,(
  $false | ($spl11 | $spl20)),
  inference(unit_resulting_resolution,[],[f131,f241,f649,f178])).
fof(f782,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f133,f649])).
fof(f781,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl26),
  inference(cnf_transformation,[],[f781_D])).
fof(f781_D,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f780,plain,(
  ( ! [X37,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f735,f625])).
fof(f735,plain,(
  ( ! [X37,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X37),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X37,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f625,f629])).
fof(f767,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X16),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl24),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X16),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f778,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X32),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X32,'type.semanticPrimitives.result'(sK1,sK2)))) ) | ($spl6 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f733,f625])).
fof(f733,plain,(
  ( ! [X33,X32] : (p(s(bool,X33)) | s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X32),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X32,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl6 | $spl14)),
  inference(superposition,[],[f261,f629])).
fof(f754,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f697])).
fof(f697,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f134,f629])).
fof(f760,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X16),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X15)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f688])).
fof(f688,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X16),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X16,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl6),
  inference(superposition,[],[f629,f134])).
fof(f766,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X2),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X2,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl22),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ( ! [X2,X1] : (s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X2),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X2,'type.semanticPrimitives.result'(sK1,sK2)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X2),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X2,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f683])).
fof(f683,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X1),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) = s(bool,happ(s(fun('type.semanticPrimitives.result'(sK1,sK2),bool),X2),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X2,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X0) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),sK9('const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8)),X1,'type.semanticPrimitives.result'(sK1,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f629,f629])).
fof(f648,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK9(X10,X9,X8))))) ) | $spl18),
  inference(cnf_transformation,[],[f648_D])).
fof(f648_D,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK9(X10,X9,X8))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f650,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK9(X14,X13,X12))))) ) | $spl16),
  inference(subsumption_resolution,[],[f645,f625])).
fof(f645,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK9(X14,X13,X12))))) )),
  inference(superposition,[],[f174,f233])).
fof(f647,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK9(X10,X9,X8))))) ) | $spl16),
  inference(subsumption_resolution,[],[f644,f625])).
fof(f644,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,sK9(X10,X9,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f174,f233])).
fof(f625,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f625_D])).
fof(f625_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f636,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) )),
  inference(condensation,[],[f635])).
fof(f635,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f480,f262])).
fof(f480,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f180,f233])).
fof(f624,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | s(bool,'T_0') != s(bool,X1)) )),
  inference(condensation,[],[f623])).
fof(f623,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f445,f262])).
fof(f445,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,X16) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f180,f233])).
fof(f268,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f264,f131])).
fof(f264,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f260,f133])).
fof(f261,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f261_D])).
fof(f261_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f263,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f257,f235])).
fof(f257,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f175,f134])).
fof(f259,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f255,f235])).
fof(f255,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f175,f134])).
fof(f238,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f134])).
fof(f226,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | $spl0),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f232,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | $spl6),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f128,plain,(
  s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  (s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | (s('type.semanticPrimitives.result'(sK1,sK2),sK7) = s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK8))) & s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))))) & (s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | ! [X8] : (s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),X8))) | s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),X8)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f89])).
fof(f89,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) = s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) | ? [X9] : (s('type.semanticPrimitives.result'(X1,X2),X7) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X9))) & s('type.semanticPrimitives.error_result'(X3),X6) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X9))))) & (s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) != s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) | ! [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) != s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) | s('type.semanticPrimitives.error_result'(X3),X6) != s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8))))))),
  inference(rectify,[],[f88])).
fof(f88,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) = s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) | ? [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) & s('type.semanticPrimitives.error_result'(X3),X6) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8))))) & (s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) != s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) | ! [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) != s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) | s('type.semanticPrimitives.error_result'(X3),X6) != s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) = s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) <~> ? [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) & s('type.semanticPrimitives.error_result'(X3),X6) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8)))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) = s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) <=> ? [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) & s('type.semanticPrimitives.error_result'(X3),X6) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8)))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : (s('type.semanticPrimitives.result'(X0,X3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X3),X6))) = s('type.semanticPrimitives.result'(X0,X3),'const.evalProps.map_result_3'(s(fun(X1,X0),X5),s(fun(X2,X3),X4),s('type.semanticPrimitives.result'(X1,X2),X7))) <=> ? [X8] : (s('type.semanticPrimitives.result'(X1,X2),X7) = s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X8))) & s('type.semanticPrimitives.error_result'(X3),X6) = s('type.semanticPrimitives.error_result'(X3),'const.evalProps.map_error_result_2'(s(fun(X2,X3),X4),s('type.semanticPrimitives.error_result'(X2),X8)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12,X13] : (s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X12))) = s('type.semanticPrimitives.result'(X6,X9),'const.evalProps.map_result_3'(s(fun(X7,X6),X11),s(fun(X8,X9),X10),s('type.semanticPrimitives.result'(X7,X8),X13))) <=> ? [X14] : (s('type.semanticPrimitives.result'(X7,X8),X13) = s('type.semanticPrimitives.result'(X7,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),X14))) & s('type.semanticPrimitives.error_result'(X9),X12) = s('type.semanticPrimitives.error_result'(X9),'const.evalProps.map_error_result_2'(s(fun(X8,X9),X10),s('type.semanticPrimitives.error_result'(X8),X14)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12,X13] : (s('type.semanticPrimitives.result'(X6,X9),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X9),X12))) = s('type.semanticPrimitives.result'(X6,X9),'const.evalProps.map_result_3'(s(fun(X7,X6),X11),s(fun(X8,X9),X10),s('type.semanticPrimitives.result'(X7,X8),X13))) <=> ? [X14] : (s('type.semanticPrimitives.result'(X7,X8),X13) = s('type.semanticPrimitives.result'(X7,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),X14))) & s('type.semanticPrimitives.error_result'(X9),X12) = s('type.semanticPrimitives.error_result'(X9),'const.evalProps.map_error_result_2'(s(fun(X8,X9),X10),s('type.semanticPrimitives.error_result'(X8),X14)))))),
  file('thm.evalProps.map_result_Rerr','thm.evalProps.map_result_Rerr')).
fof(f230,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) | $spl4),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f129,plain,(
  s('type.semanticPrimitives.error_result'(sK3),sK6) = s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),sK8))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) = s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))),
  inference(cnf_transformation,[],[f90])).
fof(f227,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) | $spl1),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f228,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),X8))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),X8)))) ) | $spl2),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),X8))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),X8)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f130,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(sK3),sK6) != s('type.semanticPrimitives.error_result'(sK3),'const.evalProps.map_error_result_2'(s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.error_result'(sK2),X8))) | s('type.semanticPrimitives.result'(sK1,sK2),sK7) != s('type.semanticPrimitives.result'(sK1,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),X8))) | s('type.semanticPrimitives.result'(sK0,sK3),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK3),sK6))) != s('type.semanticPrimitives.result'(sK0,sK3),'const.evalProps.map_result_3'(s(fun(sK1,sK0),sK5),s(fun(sK2,sK3),sK4),s('type.semanticPrimitives.result'(sK1,sK2),sK7)))) )),
  inference(cnf_transformation,[],[f90])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 261
% Passive clauses: 3232
% Generated clauses: 27834
% Final active clauses: 131
% Final passive clauses: 960
% Input formulas: 31
% Initial clauses: 80
% 
% Duplicate literals: 1232
% Trivial inequalities: 292
% Fw subsumption resolutions: 2158
% Fw demodulations: 2884
% Bw demodulations: 1673
% Condensations: 2140
% Global subsumptions: 16
% 
% Simple tautologies: 358
% Equational tautologies: 273
% Forward subsumptions: 14677
% Fw demodulations to eq. taut.: 248
% Bw demodulations to eq. taut.: 77
% 
% Unit resulting resolution: 1458
% Factoring: 1
% Forward superposition: 6377
% Backward superposition: 6470
% Self superposition: 218
% Equality factoring: 105
% Equality resolution: 45
% 
% Split clauses: 191
% Split components: 398
% SAT solver clauses: 3767
% SAT solver unit clauses: 65
% SAT solver binary clauses: 533
% 
% Sat splits: 191
% Sat splitting refutations: 135
% 
% Memory used [KB]: 18421
% Time elapsed: 1.566 s
% ------------------------------
% ------------------------------
% Success in time 2.91 s
