Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.closProps
% (127854)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 122
% Passive clauses: 3094
% Generated clauses: 13765
% Final active clauses: 116
% Final passive clauses: 2822
% Input formulas: 16
% Initial clauses: 20
% 
% Duplicate literals: 849
% Trivial inequalities: 80
% Fw subsumption resolutions: 1431
% Fw demodulations: 1628
% Bw demodulations: 5
% 
% Simple tautologies: 232
% Equational tautologies: 20
% Forward subsumptions: 5630
% 
% Binary resolution: 192
% Forward superposition: 5237
% Backward superposition: 4010
% Self superposition: 148
% Equality factoring: 32
% Equality resolution: 19
% 
% Split clauses: 393
% Split components: 792
% SAT solver clauses: 772
% SAT solver unit clauses: 11
% SAT solver binary clauses: 98
% 
% Sat splits: 393
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 771
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 13176
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f6235,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5951,f445,f450,f64,f81,f66,f83,f5953,f5972,f67,f84,f6234,f5991,f5995,f50,f71,f68,f85,f5930,f5934,f5932])).
fof(f5932,plain,(
  s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | $spl1738),
  inference(cnf_transformation,[],[f5932_D])).
fof(f5932_D,plain,(
  s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) <=> ~$spl1738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1738])])).
fof(f5934,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun(sK0,sK3),sK7),s(sK0,sK8(sK4,sK2,sK0)))))) != s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | $spl1741),
  inference(cnf_transformation,[],[f5934_D])).
fof(f5934_D,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun(sK0,sK3),sK7),s(sK0,sK8(sK4,sK2,sK0)))))) != s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) <=> ~$spl1741),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1741])])).
fof(f5930,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun(sK0,sK3),sK7),s(sK0,sK8(sK4,sK2,sK0)))))) != s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | ($spl1 | $spl24 | $spl78)),
  inference(forward_demodulation,[],[f5722,f85])).
fof(f5722,plain,(
  s(sK1,happ(s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(sK0,sK8(sK4,sK2,sK0)))) != s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | ($spl1 | $spl78)),
  inference(superposition,[],[f71,f450])).
fof(f85,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) ) | $spl24),
  inference(cnf_transformation,[],[f85_D])).
fof(f85_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f68,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X18,X19,X20] : ! [X21,X22,X23] : s(X18,happ(s(fun(X19,X18),X21),s(X19,happ(s(fun(X20,X19),X22),s(X20,X23))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun(X19,X18),X21),s(fun(X20,X19),X22))),s(X20,X23)))),
  file('thm.closProps.closProps_76_prove_ltb','_dst_x0x1_ax')).
fof(f71,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | $spl1),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f50,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5)))))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK3),sK7),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f41])).
fof(f41,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) != s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X1,happ(s(fun(X3,X1),X6),s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X3),X7),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5))))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),X4),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun(X0,X3),X7))),s(fun('type.semanticPrimitives.error_result'(X2),X1),'_dst_x0x1_2'(s(fun(X3,X1),X6),s(fun('type.semanticPrimitives.error_result'(X2),X3),X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12,X13] : s(X7,happ(s(fun(X9,X7),X12),s(X9,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),X10),s(fun(X6,X9),X13),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11))))) = s(X7,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),X10),s(fun(X6,X7),'_dst_x0x1_2'(s(fun(X9,X7),X12),s(fun(X6,X9),X13))),s(fun('type.semanticPrimitives.error_result'(X8),X7),'_dst_x0x1_2'(s(fun(X9,X7),X12),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12,X13] : s(X7,happ(s(fun(X9,X7),X12),s(X9,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),X10),s(fun(X6,X9),X13),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11))))) = s(X7,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),X10),s(fun(X6,X7),'_dst_x0x1_2'(s(fun(X9,X7),X12),s(fun(X6,X9),X13))),s(fun('type.semanticPrimitives.error_result'(X8),X7),'_dst_x0x1_2'(s(fun(X9,X7),X12),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11)))))),
  file('thm.closProps.closProps_76_prove_ltb','thm.closProps.closProps_76_prove_ltb')).
fof(f5995,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | $spl1783),
  inference(cnf_transformation,[],[f5995_D])).
fof(f5995_D,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) <=> ~$spl1783),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1783])])).
fof(f5991,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))))) != s(sK1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun(sK0,sK3),sK7))),s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))))) | ($spl1 | $spl1774)),
  inference(backward_demodulation,[],[f5972,f71])).
fof(f6234,plain,(
  $false | ($spl24 | $spl1774 | $spl1783)),
  inference(subsumption_resolution,[],[f6122,f85])).
fof(f6122,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))))) != s(sK1,happ(s(fun('type.semanticPrimitives.error_result'(sK2),sK1),'_dst_x0x1_2'(s(fun(sK3,sK1),sK6),s(fun('type.semanticPrimitives.error_result'(sK2),sK3),sK5))),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | ($spl1774 | $spl1783)),
  inference(superposition,[],[f5995,f5972])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))) ) | $spl22),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f67,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun('type.semanticPrimitives.error_result'(X2),X1),X5),s('type.semanticPrimitives.error_result'(X2),X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X3))),s(fun(X0,X1),X4),s(fun('type.semanticPrimitives.error_result'(X2),X1),X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X9,X8] : ! [X15,X13,X11] : s(X9,happ(s(fun('type.semanticPrimitives.error_result'(X8),X9),X11),s('type.semanticPrimitives.error_result'(X8),X15))) = s(X9,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),X15))),s(fun(X6,X9),X13),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11)))),
  file('thm.closProps.closProps_76_prove_ltb','thm.semanticPrimitives.result_case_def|split|1')).
fof(f5972,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3)))) ) | $spl1774),
  inference(cnf_transformation,[],[f5972_D])).
fof(f5972_D,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3)))) ) <=> ~$spl1774),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1774])])).
fof(f5953,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3),s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) = s(X2,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(sK0,sK2),sK4),s(fun(sK0,X2),X4),s(fun('type.semanticPrimitives.error_result'(sK2),X2),X3)))) ) | ($spl22 | $spl1738)),
  inference(superposition,[],[f84,f5932])).
fof(f83,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))) ) | $spl20),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f66,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X2,X0),'const.semanticPrimitives.Rval_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun('type.semanticPrimitives.error_result'(X0),X1),X5)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X9,X6] : ! [X15,X13,X11] : s(X9,happ(s(fun(X6,X9),X13),s(X6,X15))) = s(X9,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X6,X8),'const.semanticPrimitives.Rval_1'(s(X6,X15))),s(fun(X6,X9),X13),s(fun('type.semanticPrimitives.error_result'(X8),X9),X11)))),
  file('thm.closProps.closProps_76_prove_ltb','thm.semanticPrimitives.result_case_def|split|0')).
fof(f81,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK8(X2,X1,X0))))) ) | $spl16),
  inference(cnf_transformation,[],[f81_D])).
fof(f81_D,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK8(X2,X1,X0))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f64,plain,(
  ( ! [X2,X0,X1] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK8(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : (s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,sK8(X2,X1,X0)))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X4))) | ? [X3] : s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X3))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X8] : ! [X16] : (? [X15] : s('type.semanticPrimitives.result'(X6,X8),X16) = s('type.semanticPrimitives.result'(X6,X8),'const.semanticPrimitives.Rval_1'(s(X6,X15))) | ? [X17] : s('type.semanticPrimitives.result'(X6,X8),X16) = s('type.semanticPrimitives.result'(X6,X8),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X8),X17))))),
  file('thm.closProps.closProps_76_prove_ltb','thm.semanticPrimitives.result_nchotomy')).
fof(f450,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X3,happ(s(fun(X0,X3),X4),s(X0,sK8(X2,X1,X0)))) = s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X2),s(fun(X0,X3),X4),s(fun('type.semanticPrimitives.error_result'(X1),X3),X5))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0))))) ) | $spl78),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X3,happ(s(fun(X0,X3),X4),s(X0,sK8(X2,X1,X0)))) = s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X2),s(fun(X0,X3),X4),s(fun('type.semanticPrimitives.error_result'(X1),X3),X5))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f445,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X3,happ(s(fun(X0,X3),X4),s(X0,sK8(X2,X1,X0)))) = s(X3,'const.semanticPrimitives.result_CASE_3'(s('type.semanticPrimitives.result'(X0,X1),X2),s(fun(X0,X3),X4),s(fun('type.semanticPrimitives.error_result'(X1),X3),X5))) | s('type.semanticPrimitives.result'(X0,X1),X2) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),sK9(X2,X1,X0))))) ) | ($spl16 | $spl20)),
  inference(superposition,[],[f83,f81])).
fof(f5951,plain,(
  s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | ($spl78 | $spl1741)),
  inference(trivial_inequality_removal,[],[f5950])).
fof(f5950,plain,(
  s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun(sK0,sK3),sK7),s(sK0,sK8(sK4,sK2,sK0)))))) != s(sK1,happ(s(fun(sK3,sK1),sK6),s(sK3,happ(s(fun(sK0,sK3),sK7),s(sK0,sK8(sK4,sK2,sK0)))))) | s('type.semanticPrimitives.result'(sK0,sK2),sK4) = s('type.semanticPrimitives.result'(sK0,sK2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(sK2),sK9(sK4,sK2,sK0)))) | ($spl78 | $spl1741)),
  inference(superposition,[],[f5934,f450])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 78
% Passive clauses: 892
% Generated clauses: 6044
% Final active clauses: 69
% Final passive clauses: 801
% Input formulas: 16
% Initial clauses: 20
% 
% Duplicate literals: 330
% Trivial inequalities: 54
% Fw subsumption resolutions: 250
% Fw demodulations: 654
% Bw demodulations: 18
% Condensations: 311
% 
% Simple tautologies: 106
% Equational tautologies: 19
% Forward subsumptions: 2226
% 
% Binary resolution: 54
% Forward superposition: 1939
% Backward superposition: 1394
% Self superposition: 85
% Equality factoring: 30
% Equality resolution: 10
% 
% Split clauses: 25
% Split components: 55
% SAT solver clauses: 1222
% SAT solver unit clauses: 28
% SAT solver binary clauses: 41
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 25
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1220
% TWLsolver calls for satisfiability: 132
% 
% Memory used [KB]: 9466
% Time elapsed: 0.390 s
% ------------------------------
% ------------------------------
% Success in time 1.021 s
