Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.word_instProof
% (22096)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 882
% Passive clauses: 6205
% Generated clauses: 24150
% Final active clauses: 879
% Final passive clauses: 5022
% Input formulas: 69
% Initial clauses: 232
% 
% Duplicate literals: 634
% Trivial inequalities: 302
% Fw subsumption resolutions: 2038
% Fw demodulations: 1598
% Bw demodulations: 2
% 
% Simple tautologies: 392
% Equational tautologies: 57
% Forward subsumptions: 12383
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 6612
% Factoring: 1431
% Forward superposition: 5638
% Backward superposition: 5310
% Self superposition: 49
% Equality factoring: 26
% Equality resolution: 51
% 
% Split clauses: 536
% Split components: 1087
% SAT solver clauses: 837
% SAT solver unit clauses: 18
% SAT solver binary clauses: 149
% 
% Sat splits: 536
% 
% TWLsolver clauses: 837
% TWLsolver calls for satisfiability: 440
% 
% Memory used [KB]: 12920
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.word_instProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.word_instProof
% SZS output start Proof for thm.word_instProof
fof(f14738,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14696,f334,f592,f14511,f5734,f593,f595,f447,f640,f14501,f14502,f392,f614,f837,f844,f442,f643,f14356,f14357,f871,f877,f353,f602,f393,f615,f1370,f1384,f357,f604,f733,f735,f333,f590,f410,f626,f2099,f2122,f908,f918,f443,f642,f13414,f13415,f411,f627,f2971,f3000,f404,f622])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))) ) | $spl34),
  inference(cnf_transformation,[],[f622_D])).
fof(f622_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X12,X13] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),X13))),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X13),s('type.list.list'(X6),X12))))))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.sorting.PERM_APPEND')).
fof(f3000,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))) ) | $spl1100),
  inference(cnf_transformation,[],[f3000_D])).
fof(f3000_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))) ) <=> ~$spl1100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1100])])).
fof(f2971,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1)))))) ) | ($spl10 | $spl34)),
  inference(resolution,[],[f622,f602])).
fof(f627,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4)))) ) | $spl44),
  inference(cnf_transformation,[],[f627_D])).
fof(f627_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f411,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) | ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) | s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) => s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2)))) => s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X25,X19,X12,X13] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),X25),s('type.list.list'(X6),X19)))) => s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X19),s('type.list.list'(X6),X12))),s('type.list.list'(X6),X13))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X25),s('type.list.list'(X6),X12))),s('type.list.list'(X6),X13))))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.sorting.PERM_REWR')).
fof(f13415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4)))) ) | $spl3702),
  inference(cnf_transformation,[],[f13415_D])).
fof(f13415_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4)))) ) <=> ~$spl3702),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3702])])).
fof(f13414,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4)))) ) | ($spl2 | $spl42 | $spl44 | $spl1100)),
  inference(forward_demodulation,[],[f13413,f626])).
fof(f13413,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4)))) ) | ($spl2 | $spl42 | $spl44 | $spl1100)),
  inference(forward_demodulation,[],[f13412,f626])).
fof(f13412,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4)))) ) | ($spl2 | $spl44 | $spl1100)),
  inference(subsumption_resolution,[],[f13214,f592])).
fof(f13214,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3))),s('type.list.list'(X0),X4)))) ) | ($spl44 | $spl1100)),
  inference(superposition,[],[f627,f3000])).
fof(f642,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl64),
  inference(cnf_transformation,[],[f642_D])).
fof(f642_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f443,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f279])).
fof(f279,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f918,plain,(
  ( ! [X4,X3] : (~sP1('F_0',X3) | ~p(s(bool,X4)) | ~sP1(X4,X3)) ) | $spl346),
  inference(cnf_transformation,[],[f918_D])).
fof(f918_D,plain,(
  ( ! [X4,X3] : (~sP1('F_0',X3) | ~p(s(bool,X4)) | ~sP1(X4,X3)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f908,plain,(
  ( ! [X4,X3] : (~sP1('F_0',X3) | ~p(s(bool,X4)) | ~sP1(X4,X3)) ) | ($spl64 | $spl290)),
  inference(resolution,[],[f844,f642])).
fof(f2122,plain,(
  ( ! [X6,X7] : (~sP1('const.combin.I_1'(s(bool,X6)),X7) | ~sP1('F_0',X7) | ~p(s(bool,X6))) ) | $spl770),
  inference(cnf_transformation,[],[f2122_D])).
fof(f2122_D,plain,(
  ( ! [X6,X7] : (~sP1('const.combin.I_1'(s(bool,X6)),X7) | ~sP1('F_0',X7) | ~p(s(bool,X6))) ) <=> ~$spl770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl770])])).
fof(f2099,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP1('F_0',X7) | ~sP1('const.combin.I_1'(s(bool,X6)),X7)) ) | ($spl18 | $spl346)),
  inference(superposition,[],[f918,f614])).
fof(f626,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) ) | $spl42),
  inference(cnf_transformation,[],[f626_D])).
fof(f626_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f410,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X12,X13,X26] : s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X13),s('type.list.list'(X6),X26))))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),X13))),s('type.list.list'(X6),X26)))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.list.APPEND_ASSOC')).
fof(f590,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK14))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) | $spl1),
  inference(cnf_transformation,[],[f590_D])).
fof(f590_D,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK14))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f333,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK14))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))))),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK15))),s('type.list.list'(sK13),sK14))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f178])).
fof(f178,plain,(
  ? [X0,X1,X2,X3] : ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X1))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8))),s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8))),s('type.list.list'(X6),X7))),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))))))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.word_instProof.word_instProof_2_prove_ltb')).
fof(f735,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) | $spl233),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f733,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f590,f626])).
fof(f604,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f357,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X18] : (s(bool,'F_0') = s(bool,X18) <=> ~p(s(bool,X18)))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f1384,plain,(
  s(bool,'F_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))) | $spl502),
  inference(cnf_transformation,[],[f1384_D])).
fof(f1384_D,plain,(
  s(bool,'F_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f1370,plain,(
  s(bool,'F_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))) | ($spl14 | $spl233)),
  inference(resolution,[],[f735,f604])).
fof(f615,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))) ) | $spl20),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f393,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X1))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X11] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),X11),s('type.list.list'(X6),X11))))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.sorting.PERM_REFL')).
fof(f602,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X18] : (s(bool,'T_0') = s(bool,X18) <=> p(s(bool,X18)))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f877,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X10),X11),s('type.list.list'(X10),X11)))) ) | $spl322),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X10),X11),s('type.list.list'(X10),X11)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f871,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X10),X11),s('type.list.list'(X10),X11)))) ) | ($spl10 | $spl20)),
  inference(resolution,[],[f615,f602])).
fof(f14357,plain,(
  ( ! [X29] : (~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29)) ) | $spl4176),
  inference(cnf_transformation,[],[f14357_D])).
fof(f14357_D,plain,(
  ( ! [X29] : (~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29)) ) <=> ~$spl4176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4176])])).
fof(f14356,plain,(
  ( ! [X29] : (~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29)) ) | ($spl2 | $spl322 | $spl502 | $spl770 | $spl3702)),
  inference(subsumption_resolution,[],[f14355,f592])).
fof(f14355,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | ~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29)) ) | ($spl322 | $spl502 | $spl770 | $spl3702)),
  inference(forward_demodulation,[],[f14354,f877])).
fof(f14354,plain,(
  ( ! [X29] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14)))))))) | ~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29)) ) | ($spl502 | $spl770 | $spl3702)),
  inference(forward_demodulation,[],[f14329,f13415])).
fof(f14329,plain,(
  ( ! [X29] : (~sP1('const.combin.I_1'(s(bool,'F_0')),X29) | ~sP1('F_0',X29) | ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),sK14))))),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK16),s('type.list.list'(sK13),sK14))))))))) ) | ($spl502 | $spl770)),
  inference(superposition,[],[f2122,f1384])).
fof(f643,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f643_D])).
fof(f643_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f442,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f280])).
fof(f844,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP1('F_0',X0)) ) | $spl290),
  inference(cnf_transformation,[],[f844_D])).
fof(f844_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP1('F_0',X0)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f837,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP1('F_0',X0)) ) | ($spl5 | $spl66)),
  inference(resolution,[],[f643,f595])).
fof(f614,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) ) | $spl18),
  inference(cnf_transformation,[],[f614_D])).
fof(f614_D,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f392,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X27] : s(X6,X27) = s(X6,'const.combin.I_1'(s(X6,X27)))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.combin.I_THM')).
fof(f14502,plain,(
  ( ! [X1] : (~sP1('F_0',X1)) ) | $spl4196),
  inference(cnf_transformation,[],[f14502_D])).
fof(f14502_D,plain,(
  ( ! [X1] : (~sP1('F_0',X1)) ) <=> ~$spl4196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4196])])).
fof(f14501,plain,(
  ( ! [X1] : (~sP1('F_0',X1)) ) | ($spl5 | $spl18 | $spl60 | $spl290 | $spl4176)),
  inference(subsumption_resolution,[],[f14500,f595])).
fof(f14500,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~sP1('F_0',X1)) ) | ($spl18 | $spl60 | $spl290 | $spl4176)),
  inference(forward_demodulation,[],[f14499,f614])).
fof(f14499,plain,(
  ( ! [X1] : (~sP1('F_0',X1) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | ($spl60 | $spl290 | $spl4176)),
  inference(subsumption_resolution,[],[f14492,f844])).
fof(f14492,plain,(
  ( ! [X1] : (~sP1('F_0',X1) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | ~p(s(bool,X1))) ) | ($spl60 | $spl4176)),
  inference(resolution,[],[f14357,f640])).
fof(f640,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl60),
  inference(cnf_transformation,[],[f640_D])).
fof(f640_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f447,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f595,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f595_D])).
fof(f595_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f593,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f337,f334])).
fof(f337,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f5734,plain,(
  ( ! [X59] : (~p(s(bool,X59))) ) | $spl1878),
  inference(cnf_transformation,[],[f5734_D])).
fof(f5734_D,plain,(
  ( ! [X59] : (~p(s(bool,X59))) ) <=> ~$spl1878),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1878])])).
fof(f14511,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl5 | $spl60 | $spl4196)),
  inference(subsumption_resolution,[],[f14510,f595])).
fof(f14510,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,X1))) ) | ($spl60 | $spl4196)),
  inference(resolution,[],[f14502,f640])).
fof(f592,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f334,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,axiom,(
  p(s(bool,'T_0'))),
  file('thm.word_instProof.word_instProof_2_prove_ltb','thm.bool.TRUTH')).
fof(f14696,plain,(
  $false | ($spl2 | $spl1878)),
  inference(resolution,[],[f5734,f592])).
% SZS output end Proof for thm.word_instProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 460
% Passive clauses: 2096
% Generated clauses: 14326
% Final active clauses: 444
% Final passive clauses: 1612
% Input formulas: 69
% Initial clauses: 232
% 
% Duplicate literals: 322
% Trivial inequalities: 57
% Fw subsumption resolutions: 1098
% Fw demodulations: 790
% Bw demodulations: 30
% Condensations: 611
% Global subsumptions: 19
% 
% Simple tautologies: 324
% Equational tautologies: 29
% Forward subsumptions: 6260
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 3188
% Factoring: 337
% Forward superposition: 2544
% Backward superposition: 2865
% Self superposition: 53
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 2758
% SAT solver unit clauses: 241
% SAT solver binary clauses: 84
% SAT solver learnt clauses: 39
% SAT solver learnt literals: 12
% 
% Sat splits: 48
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 2679
% TWLsolver calls for satisfiability: 1959
% 
% Memory used [KB]: 13432
% Time elapsed: 0.570 s
% ------------------------------
% ------------------------------
% Success in time 1.21 s
