Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (144507)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 3481
% Generated clauses: 20934
% Final active clauses: 124
% Final passive clauses: 3139
% Input formulas: 68
% Initial clauses: 98
% 
% Duplicate literals: 639
% Trivial inequalities: 317
% Fw subsumption resolutions: 1972
% Fw demodulations: 3430
% Bw demodulations: 13
% 
% Simple tautologies: 301
% Equational tautologies: 52
% Forward subsumptions: 10316
% 
% Binary resolution: 170
% Factoring: 1
% Forward superposition: 6371
% Backward superposition: 7506
% Self superposition: 214
% Equality factoring: 111
% Equality resolution: 35
% 
% Split clauses: 157
% Split components: 316
% SAT solver clauses: 431
% SAT solver unit clauses: 8
% SAT solver binary clauses: 68
% 
% Sat splits: 157
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 430
% TWLsolver calls for satisfiability: 131
% 
% Memory used [KB]: 18805
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f9835,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9834,f9742,f9738,f9743,f239,f339,f307,f385,f8909,f9267,f224,f326,f308,f386,f9830,f9523,f9711,f9831,f9832,f9731,f9732,f313,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) ) | $spl108),
  inference(cnf_transformation,[],[f390_D])).
fof(f390_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X45,X46,X47] : (p(s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),'=_0'),s(X45,X46))),s(X45,X47)))) <=> s(X45,X46) = s(X45,X47))),
  file('thm.wordSem.result_distinct|split|8',arity_0_HL_EQ)).
fof(f9732,plain,(
  ( ! [X146,X152,X151,X149] : (s(X146,X149) != s(X146,X152) | p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) | $spl1170),
  inference(cnf_transformation,[],[f9732_D])).
fof(f9732_D,plain,(
  ( ! [X146,X152,X151,X149] : (s(X146,X149) != s(X146,X152) | p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) <=> ~$spl1170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1170])])).
fof(f9731,plain,(
  ( ! [X146,X152,X151,X149] : (s(X146,X149) != s(X146,X152) | p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) | ($spl98 | $spl108 | $spl1086)),
  inference(forward_demodulation,[],[f9561,f385])).
fof(f9561,plain,(
  ( ! [X146,X152,X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151)))) | s(X146,X152) != s(X146,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X146)),X147),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X146)),X148),s(X146,X149),s(X146,X150),s(X146,X151)))) ) | ($spl108 | $spl1086)),
  inference(superposition,[],[f390,f9267])).
fof(f9832,plain,(
  ( ! [X146,X152,X151] : (p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) | $spl1234),
  inference(cnf_transformation,[],[f9832_D])).
fof(f9832_D,plain,(
  ( ! [X146,X152,X151] : (p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) <=> ~$spl1234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1234])])).
fof(f9831,plain,(
  ( ! [X146,X152,X151] : (p(s(bool,happ(s(fun(X146,bool),happ(s(fun(X146,fun(X146,bool)),'=_0'),s(X146,X152))),s(X146,X151))))) ) | ($spl1146 | $spl1170)),
  inference(subsumption_resolution,[],[f9732,f9711])).
fof(f9711,plain,(
  ( ! [X6,X11,X9] : (s(X6,X9) = s(X6,X11)) ) | $spl1146),
  inference(cnf_transformation,[],[f9711_D])).
fof(f9711_D,plain,(
  ( ! [X6,X11,X9] : (s(X6,X9) = s(X6,X11)) ) <=> ~$spl1146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1146])])).
fof(f9523,plain,(
  ( ! [X6,X11,X9] : (s(X6,X9) = s(X6,X11)) ) | ($spl98 | $spl1086)),
  inference(superposition,[],[f9267,f385])).
fof(f9830,plain,(
  $false | ($spl1146 | $spl1180)),
  inference(subsumption_resolution,[],[f9743,f9711])).
fof(f386,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) ) | $spl100),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f308,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X16] : ! [X23,X34,X35,X36,X37] : s(X16,X37) = s(X16,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X16)),X23),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X16)),X34),s(X16,X35),s(X16,X36),s(X16,X37)))),
  file('thm.wordSem.result_distinct|split|8','thm.wordSem.result_case_def|split|4')).
fof(f326,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Error_0') | $spl0),
  inference(cnf_transformation,[],[f326_D])).
fof(f326_D,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Error_0') <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f224,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Error_0')),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Error_0')),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f185])).
fof(f185,plain,(
  ? [X0] : s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),'const.wordSem.Error_0')),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(X0),'const.wordSem.Error_0')),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ~s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),'const.wordSem.Error_0')),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),'const.wordSem.Error_0')),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),'const.wordSem.Error_0')),
  file('thm.wordSem.result_distinct|split|8','thm.wordSem.result_distinct|split|8')).
fof(f9267,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X1)))) ) | $spl1086),
  inference(cnf_transformation,[],[f9267_D])).
fof(f9267_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X1)))) ) <=> ~$spl1086),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1086])])).
fof(f8909,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X1)))) ) | ($spl0 | $spl100)),
  inference(superposition,[],[f386,f326])).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) ) | $spl98),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f307,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X16] : ! [X23,X34,X35,X36,X37] : s(X16,X35) = s(X16,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X16)),X23),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X16)),X34),s(X16,X35),s(X16,X36),s(X16,X37)))),
  file('thm.wordSem.result_distinct|split|8','thm.wordSem.result_case_def|split|2')).
fof(f339,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f239,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X26] : (s(bool,'F_0') = s(bool,X26) <=> ~p(s(bool,X26)))),
  file('thm.wordSem.result_distinct|split|8','thm.bool.EQ_CLAUSES|split|3')).
fof(f9743,plain,(
  ( ! [X222] : (s(bool,'F_0') != s(bool,X222)) ) | $spl1180),
  inference(cnf_transformation,[],[f9743_D])).
fof(f9743_D,plain,(
  ( ! [X222] : (s(bool,'F_0') != s(bool,X222)) ) <=> ~$spl1180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1180])])).
fof(f9738,plain,(
  ( ! [X207] : (~p(s(bool,X207))) ) | $spl1176),
  inference(cnf_transformation,[],[f9738_D])).
fof(f9738_D,plain,(
  ( ! [X207] : (~p(s(bool,X207))) ) <=> ~$spl1176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1176])])).
fof(f9742,plain,(
  ( ! [X222,X220] : (~p(s(bool,X220)) | s(bool,'F_0') != s(bool,X222)) ) | ($spl14 | $spl98 | $spl1086)),
  inference(forward_demodulation,[],[f9572,f385])).
fof(f9572,plain,(
  ( ! [X218,X222,X220,X219,X221] : (s(bool,'F_0') != s(bool,X222) | ~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X218),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X219),s(bool,X220),s(bool,X221),s(bool,X222))))) ) | ($spl14 | $spl1086)),
  inference(superposition,[],[f339,f9267])).
fof(f9834,plain,(
  $false | ($spl1176 | $spl1234)),
  inference(subsumption_resolution,[],[f9832,f9738])).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 78
% Passive clauses: 592
% Generated clauses: 9502
% Final active clauses: 76
% Final passive clauses: 502
% Input formulas: 68
% Initial clauses: 98
% 
% Duplicate literals: 471
% Trivial inequalities: 116
% Fw subsumption resolutions: 649
% Fw demodulations: 930
% Bw demodulations: 12
% Condensations: 1006
% Global subsumptions: 8
% 
% Simple tautologies: 181
% Equational tautologies: 46
% Forward subsumptions: 4354
% 
% Binary resolution: 84
% Factoring: 1
% Forward superposition: 2606
% Backward superposition: 2620
% Self superposition: 143
% Equality factoring: 96
% Equality resolution: 19
% 
% Split clauses: 35
% Split components: 70
% SAT solver clauses: 1251
% SAT solver unit clauses: 151
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 1
% 
% Sat splits: 35
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1217
% TWLsolver calls for satisfiability: 308
% 
% Memory used [KB]: 10234
% Time elapsed: 0.468 s
% ------------------------------
% ------------------------------
% Success in time 1.09 s
