Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.ml_translator
% (3956)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 49
% Passive clauses: 2286
% Generated clauses: 10818
% Final active clauses: 47
% Final passive clauses: 2204
% Input formulas: 14
% Initial clauses: 16
% 
% Duplicate literals: 762
% Trivial inequalities: 100
% Fw subsumption resolutions: 1179
% Fw demodulations: 344
% 
% Simple tautologies: 195
% Equational tautologies: 100
% Forward subsumptions: 4496
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 25
% Forward superposition: 4111
% Backward superposition: 3553
% Self superposition: 372
% Equality factoring: 283
% Equality resolution: 7
% 
% Split clauses: 462
% Split components: 929
% SAT solver clauses: 1372
% SAT solver unit clauses: 243
% SAT solver binary clauses: 447
% 
% Sat splits: 462
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1368
% TWLsolver calls for satisfiability: 36
% 
% Memory used [KB]: 7036
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ml_translator
% (3980)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 36
% Passive clauses: 434
% Generated clauses: 5562
% Final active clauses: 33
% Final passive clauses: 394
% Input formulas: 14
% Initial clauses: 16
% 
% Duplicate literals: 757
% Trivial inequalities: 32
% Fw subsumption resolutions: 113
% Fw demodulations: 107
% Bw demodulations: 3
% Condensations: 1230
% 
% Simple tautologies: 66
% Equational tautologies: 49
% Forward subsumptions: 1754
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 12
% Forward superposition: 1415
% Backward superposition: 1145
% Self superposition: 204
% Equality factoring: 154
% 
% Split clauses: 65
% Split components: 131
% SAT solver clauses: 1045
% SAT solver unit clauses: 22
% SAT solver binary clauses: 78
% 
% Sat splits: 65
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1045
% TWLsolver calls for satisfiability: 52
% 
% Memory used [KB]: 7931
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ml_translator
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_translator
% SZS output start Proof for thm.ml_translator
fof(f91499,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f91498])).
fof(f91498,plain,(
  s(sK2,happ(s(fun(sK3,sK2),'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6))),s(sK3,sK5))) != s(sK2,happ(s(fun(sK3,sK2),'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6))),s(sK3,sK5)))),
  inference(superposition,[],[f44,f49278])).
fof(f49278,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X19,X16),'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X17),X22),s(fun(X18,fun(X17,fun(X19,X16))),X20))),s(X19,X21))) = s(X16,'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X17),X22),s(fun(X18,fun(X17,X16)),'_dst_x11x00_2'(s(fun(X18,fun(X17,fun(X19,X16))),X20),s(X19,X21)))))) )),
  inference(forward_demodulation,[],[f49064,f1681])).
fof(f1681,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X3,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X2),s(fun(X0,fun(X1,X3)),X4))) = s(X3,happ(s(fun(X1,X3),happ(s(fun(X0,fun(X1,X3)),X4),s(X0,sK7(X2,X1,X0)))),s(X1,sK8(X2,X1,X0))))) )),
  inference(superposition,[],[f58,f53])).
fof(f53,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK7(X2,X1,X0)),s(X1,sK8(X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : ! [X2] : ? [X3,X4] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : ! [X10] : ? [X14,X15] : s('type.pair.prod'(X6,X7),X10) = s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X14),s(X7,X15)))),
  file('thm.ml_translator.ml_translator_438_prove_ltb','thm.pair.ABS_PAIR_THM')).
fof(f58,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X9] : ! [X11,X10,X12] : s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X7,X9),'const.pair.,_2'(s(X7,X10),s(X9,X11))),s(fun(X7,fun(X9,X6)),X12))) = s(X6,happ(s(fun(X9,X6),happ(s(fun(X7,fun(X9,X6)),X12),s(X7,X10))),s(X9,X11)))),
  file('thm.ml_translator.ml_translator_438_prove_ltb','thm.pair.pair_case_thm')).
fof(f49064,plain,(
  ( ! [X21,X19,X17,X22,X20,X18,X16] : (s(X16,'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X17),X22),s(fun(X18,fun(X17,X16)),'_dst_x11x00_2'(s(fun(X18,fun(X17,fun(X19,X16))),X20),s(X19,X21))))) = s(X16,happ(s(fun(X19,X16),happ(s(fun(X17,fun(X19,X16)),happ(s(fun(X18,fun(X17,fun(X19,X16))),X20),s(X18,sK7(X22,X17,X18)))),s(X17,sK8(X22,X17,X18)))),s(X19,X21)))) )),
  inference(superposition,[],[f3595,f1681])).
fof(f3595,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),happ(s(fun(X3,fun(X0,X1)),'_dst_x11x00_2'(s(fun(X3,fun(X0,fun(X2,X1))),X4),s(X2,X6))),s(X3,X5))),s(X0,X7))) = s(X1,happ(s(fun(X2,X1),happ(s(fun(X0,fun(X2,X1)),happ(s(fun(X3,fun(X0,fun(X2,X1))),X4),s(X3,X5))),s(X0,X7))),s(X2,X6)))) )),
  inference(superposition,[],[f57,f59])).
fof(f59,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x11x00_2'(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X3,X5))),s(X2,X6))) = s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X3,X1)),happ(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X2,X6))),s(X3,X5)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x11x00_2'(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X3,X5))),s(X2,X6))) = s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X3,X1)),happ(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X2,X6))),s(X3,X5)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x11x00_2'(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X3,X5))),s(X2,X6))) = s(fun(X0,X1),'_dst_x1x0_2'(s(fun(X0,fun(X3,X1)),happ(s(fun(X2,fun(X0,fun(X3,X1))),X4),s(X2,X6))),s(X3,X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X18,X16,X17,X22] : ! [X19,X20,X21] : s(fun(X18,X16),happ(s(fun(X17,fun(X18,X16)),'_dst_x11x00_2'(s(fun(X17,fun(X18,fun(X22,X16))),X19),s(X22,X20))),s(X17,X21))) = s(fun(X18,X16),'_dst_x1x0_2'(s(fun(X18,fun(X22,X16)),happ(s(fun(X17,fun(X18,fun(X22,X16))),X19),s(X17,X21))),s(X22,X20)))),
  file('thm.ml_translator.ml_translator_438_prove_ltb','_dst_x11x00_ax')).
fof(f57,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),'_dst_x1x0_2'(s(fun(X17,fun(X18,X16)),X19),s(X18,X20))),s(X17,X21))) = s(X16,happ(s(fun(X18,X16),happ(s(fun(X17,fun(X18,X16)),X19),s(X17,X21))),s(X18,X20)))),
  file('thm.ml_translator.ml_translator_438_prove_ltb','_dst_x1x0_ax')).
fof(f44,plain,(
  s(sK2,happ(s(fun(sK3,sK2),'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6))),s(sK3,sK5))) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,sK2)),'_dst_x11x00_2'(s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6),s(sK3,sK5)))))),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  s(sK2,happ(s(fun(sK3,sK2),'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6))),s(sK3,sK5))) != s(sK2,'const.pair.pair_CASE_2'(s('type.pair.prod'(sK0,sK1),sK4),s(fun(sK0,fun(sK1,sK2)),'_dst_x11x00_2'(s(fun(sK0,fun(sK1,fun(sK3,sK2))),sK6),s(sK3,sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f36])).
fof(f36,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : s(X2,happ(s(fun(X3,X2),'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,fun(X3,X2))),X6))),s(X3,X5))) != s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),'_dst_x11x00_2'(s(fun(X0,fun(X1,fun(X3,X2))),X6),s(X3,X5)))))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : s(X2,happ(s(fun(X3,X2),'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,fun(X3,X2))),X6))),s(X3,X5))) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),'_dst_x11x00_2'(s(fun(X0,fun(X1,fun(X3,X2))),X6),s(X3,X5)))))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(X2,happ(s(fun(X3,X2),'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,fun(X3,X2))),X6))),s(X3,X5))) = s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X4),s(fun(X0,fun(X1,X2)),'_dst_x11x00_2'(s(fun(X0,fun(X1,fun(X3,X2))),X6),s(X3,X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(X8,happ(s(fun(X9,X8),'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X10),s(fun(X6,fun(X7,fun(X9,X8))),X12))),s(X9,X11))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X10),s(fun(X6,fun(X7,X8)),'_dst_x11x00_2'(s(fun(X6,fun(X7,fun(X9,X8))),X12),s(X9,X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : s(X8,happ(s(fun(X9,X8),'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X10),s(fun(X6,fun(X7,fun(X9,X8))),X12))),s(X9,X11))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X10),s(fun(X6,fun(X7,X8)),'_dst_x11x00_2'(s(fun(X6,fun(X7,fun(X9,X8))),X12),s(X9,X11)))))),
  file('thm.ml_translator.ml_translator_438_prove_ltb','thm.ml_translator.ml_translator_438_prove_ltb')).
% SZS output end Proof for thm.ml_translator
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 287
% Passive clauses: 8590
% Generated clauses: 101971
% Final active clauses: 98
% Final passive clauses: 673
% Input formulas: 14
% Initial clauses: 16
% Discarded non-redundant clauses: 10917
% 
% Duplicate literals: 6483
% Trivial inequalities: 595
% Fw subsumption resolutions: 4789
% Fw demodulations: 7228
% Bw demodulations: 4938
% Condensations: 11947
% Global subsumptions: 69
% 
% Simple tautologies: 848
% Equational tautologies: 1961
% Forward subsumptions: 58198
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 292
% 
% Unit resulting resolution: 4959
% Forward superposition: 28838
% Backward superposition: 16539
% Self superposition: 572
% Equality factoring: 318
% Equality resolution: 33
% 
% Split clauses: 261
% Split components: 525
% SAT solver clauses: 9284
% SAT solver unit clauses: 155
% SAT solver binary clauses: 813
% 
% Sat splits: 261
% Sat splitting refutations: 101
% 
% Memory used [KB]: 55265
% Time elapsed: 9.484 s
% ------------------------------
% ------------------------------
% Success in time 10.84 s
