Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.word_instProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.word_instProof
% SZS output start Proof for thm.word_instProof
fof(f5249,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5242,f113,f157,f131,f158,f156,f802,f403])).
fof(f403,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X21),'const.list.APPEND_2'(s('type.list.list'(X21),X22),s('type.list.list'(X21),X23))),s('type.list.list'(X21),'const.list.APPEND_2'(s('type.list.list'(X21),X23),s('type.list.list'(X21),X22)))))) ) | $spl8),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X21),'const.list.APPEND_2'(s('type.list.list'(X21),X22),s('type.list.list'(X21),X23))),s('type.list.list'(X21),'const.list.APPEND_2'(s('type.list.list'(X21),X23),s('type.list.list'(X21),X22)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f802,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),X5),s('type.list.list'(X4),X6))),s('type.list.list'(X4),'const.list.APPEND_2'(s('type.list.list'(X4),X6),s('type.list.list'(X4),X5)))))) )),
  inference(resolution,[],[f148,f121])).
fof(f121,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f148,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X12,X13] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),X13))),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X13),s('type.list.list'(X6),X12))))))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.sorting.PERM_APPEND')).
fof(f156,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f158,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f131,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f157,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f113,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb',aHL_FALSITY)).
fof(f5242,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f5241,f158])).
fof(f5241,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f5240,f403])).
fof(f5240,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK1),s('type.list.list'(sK0),sK3))),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1)))))) | $spl1),
  inference(forward_demodulation,[],[f5239,f149])).
fof(f149,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2,X3] : s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X2))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X17,X13,X12] : s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),X13))) = s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X12))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X13)))))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.sorting.PERM_CONS_IFF')).
fof(f5239,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK1),s('type.list.list'(sK0),sK3))))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1)))))))) | $spl1),
  inference(subsumption_resolution,[],[f5210,f157])).
fof(f5210,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK1),s('type.list.list'(sK0),sK3))))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1))))))))),
  inference(superposition,[],[f955,f1062])).
fof(f1062,plain,(
  s(bool,'F_0') = s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1)))))))),
  inference(resolution,[],[f106,f125])).
fof(f125,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f106,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1))))))))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.APPEND_2'(s('type.list.list'(sK0),sK3),s('type.list.list'(sK0),sK1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f70])).
fof(f70,plain,(
  ? [X0,X1,X2,X3] : ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2,X3] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0] : ! [X1,X2,X3] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X3),s('type.list.list'(X0),X1))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X7))))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X7))))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X7))))))))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.word_instProof.word_instProof_|48|_prove_ltb')).
fof(f955,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X6),s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s('type.list.list'(X3),X7)))) | ~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),'const.list.APPEND_2'(s('type.list.list'(X3),X5),s('type.list.list'(X3),X6))))),s('type.list.list'(X3),X7))))) )),
  inference(superposition,[],[f151,f150])).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X2)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))))) = s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X12,X13,X14] : s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X14),s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X12),s('type.list.list'(X6),X13))))) = s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X14),s('type.list.list'(X6),X12))),s('type.list.list'(X6),X13)))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.list.APPEND|split|1')).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) | p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) | p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) => p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X1),s('type.list.list'(X0),X2))),s('type.list.list'(X0),X3)))) => p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X1))),s('type.list.list'(X0),X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X9),s('type.list.list'(X6),X8))),s('type.list.list'(X6),X7)))) => p(s(bool,'const.sorting.PERM_2'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X8),s('type.list.list'(X6),X9))),s('type.list.list'(X6),X7)))))),
  file('thm.word_instProof.word_instProof_|48|_prove_ltb','thm.sorting.APPEND_PERM_SYM')).
% SZS output end Proof for thm.word_instProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 95
% Passive clauses: 551
% Generated clauses: 5030
% Final active clauses: 81
% Final passive clauses: 362
% Input formulas: 28
% Initial clauses: 49
% 
% Duplicate literals: 343
% Trivial inequalities: 55
% Fw subsumption resolutions: 610
% Fw demodulations: 621
% Bw demodulations: 13
% 
% Simple tautologies: 171
% Equational tautologies: 17
% Forward subsumptions: 2478
% 
% Binary resolution: 139
% Forward superposition: 2077
% Backward superposition: 987
% Self superposition: 47
% Equality factoring: 28
% Equality resolution: 14
% 
% Split clauses: 131
% Split components: 269
% SAT solver clauses: 277
% SAT solver unit clauses: 9
% SAT solver binary clauses: 63
% SAT solver learnt clauses: 1
% 
% Sat splits: 131
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 272
% TWLsolver calls for satisfiability: 96
% 
% Memory used [KB]: 3070
% Time elapsed: 0.116 s
% ------------------------------
% ------------------------------
% Success in time 0.201 s
