Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.conSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.conSem
% SZS output start Proof for thm.conSem
fof(f7557,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f130,f198,f196,f194,f192,f190,f3844,f126,f191,f127,f128,f195,f129,f189,f7545,f7548,f7543,f193,f7546,f197])).
fof(f197,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) = s('type.ffi.ffi_state'(sK0),sK8) | $spl8),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) = s('type.ffi.ffi_state'(sK0),sK8) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f7546,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) = s('type.ffi.ffi_state'(sK0),sK8) | $spl0),
  inference(forward_demodulation,[],[f7521,f169])).
fof(f169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),'const.conSem.state_ffi_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),'const.conSem.state_ffi_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),'const.conSem.state_ffi_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X32,X42,X43,X31] : s('type.ffi.ffi_state'(X6),X31) = s('type.ffi.ffi_state'(X6),'const.conSem.state_ffi_1'(s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X42),s('type.list.list'('type.option.option'('type.conSem.v')),X43),s('type.ffi.ffi_state'(X6),X31)))))),
  file('thm.conSem.state_|49||49|','thm.conSem.state_ffi')).
fof(f7521,plain,(
  s('type.ffi.ffi_state'(sK0),sK8) = s('type.ffi.ffi_state'(sK0),'const.conSem.state_ffi_1'(s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))))) | $spl0),
  inference(superposition,[],[f169,f189])).
fof(f193,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | $spl4),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f7543,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | $spl0),
  inference(forward_demodulation,[],[f7519,f171])).
fof(f171,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.conSem.state_refs_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.conSem.state_refs_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.conSem.state_refs_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X32,X42,X43,X31] : s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X42) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.conSem.state_refs_1'(s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X42),s('type.list.list'('type.option.option'('type.conSem.v')),X43),s('type.ffi.ffi_state'(X6),X31)))))),
  file('thm.conSem.state_|49||49|','thm.conSem.state_refs')).
fof(f7519,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),'const.conSem.state_refs_1'(s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))))) | $spl0),
  inference(superposition,[],[f171,f189])).
fof(f7548,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f7547,f192])).
fof(f7547,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | $spl0),
  inference(forward_demodulation,[],[f7522,f168])).
fof(f168,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num',X1) = s('type.num.num','const.conSem.state_clock_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.conSem.state_clock_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.conSem.state_clock_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X32,X42,X43,X31] : s('type.num.num',X32) = s('type.num.num','const.conSem.state_clock_1'(s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X42),s('type.list.list'('type.option.option'('type.conSem.v')),X43),s('type.ffi.ffi_state'(X6),X31)))))),
  file('thm.conSem.state_|49||49|','thm.conSem.state_clock')).
fof(f7522,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.conSem.state_clock_1'(s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))))) | $spl0),
  inference(superposition,[],[f168,f189])).
fof(f7545,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f7544,f196])).
fof(f7544,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) = s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | $spl0),
  inference(forward_demodulation,[],[f7520,f170])).
fof(f170,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.conSem.state_globals_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.conSem.state_globals_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.conSem.state_globals_1'(s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X32,X42,X43,X31] : s('type.list.list'('type.option.option'('type.conSem.v')),X43) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.conSem.state_globals_1'(s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X42),s('type.list.list'('type.option.option'('type.conSem.v')),X43),s('type.ffi.ffi_state'(X6),X31)))))),
  file('thm.conSem.state_|49||49|','thm.conSem.state_globals')).
fof(f7520,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK7) = s('type.list.list'('type.option.option'('type.conSem.v')),'const.conSem.state_globals_1'(s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))))) | $spl0),
  inference(superposition,[],[f170,f189])).
fof(f189,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) | $spl0),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f129,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) = s('type.ffi.ffi_state'(sK0),sK8) | s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8)))),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  (s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) | (s('type.num.num',sK1) = s('type.num.num',sK5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) & s('type.list.list'('type.option.option'('type.conSem.v')),sK3) = s('type.list.list'('type.option.option'('type.conSem.v')),sK7) & s('type.ffi.ffi_state'(sK0),sK4) = s('type.ffi.ffi_state'(sK0),sK8))) & (s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.list.list'('type.option.option'('type.conSem.v')),sK3) != s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | s('type.ffi.ffi_state'(sK0),sK4) != s('type.ffi.ffi_state'(sK0),sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f108])).
fof(f108,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) = s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),X7) & s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),X8))) & (s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) != s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) | s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) | s('type.list.list'('type.option.option'('type.conSem.v')),X3) != s('type.list.list'('type.option.option'('type.conSem.v')),X7) | s('type.ffi.ffi_state'(X0),X4) != s('type.ffi.ffi_state'(X0),X8)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) = s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),X7) & s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),X8))) & (s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) != s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) | (s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) | s('type.list.list'('type.option.option'('type.conSem.v')),X3) != s('type.list.list'('type.option.option'('type.conSem.v')),X7) | s('type.ffi.ffi_state'(X0),X4) != s('type.ffi.ffi_state'(X0),X8))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) = s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) <~> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),X7) & s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),X8)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) = s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),X7) & s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),X8)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2),s('type.list.list'('type.option.option'('type.conSem.v')),X3),s('type.ffi.ffi_state'(X0),X4))) = s('type.conSem.state'(X0),'const.conSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6),s('type.list.list'('type.option.option'('type.conSem.v')),X7),s('type.ffi.ffi_state'(X0),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X6) & s('type.list.list'('type.option.option'('type.conSem.v')),X3) = s('type.list.list'('type.option.option'('type.conSem.v')),X7) & s('type.ffi.ffi_state'(X0),X4) = s('type.ffi.ffi_state'(X0),X8)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8),s('type.list.list'('type.option.option'('type.conSem.v')),X9),s('type.ffi.ffi_state'(X6),X10))) = s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12),s('type.list.list'('type.option.option'('type.conSem.v')),X13),s('type.ffi.ffi_state'(X6),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12) & s('type.list.list'('type.option.option'('type.conSem.v')),X9) = s('type.list.list'('type.option.option'('type.conSem.v')),X13) & s('type.ffi.ffi_state'(X6),X10) = s('type.ffi.ffi_state'(X6),X14)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8),s('type.list.list'('type.option.option'('type.conSem.v')),X9),s('type.ffi.ffi_state'(X6),X10))) = s('type.conSem.state'(X6),'const.conSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12),s('type.list.list'('type.option.option'('type.conSem.v')),X13),s('type.ffi.ffi_state'(X6),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),X12) & s('type.list.list'('type.option.option'('type.conSem.v')),X9) = s('type.list.list'('type.option.option'('type.conSem.v')),X13) & s('type.ffi.ffi_state'(X6),X10) = s('type.ffi.ffi_state'(X6),X14)))),
  file('thm.conSem.state_|49||49|','thm.conSem.state_|49||49|')).
fof(f195,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) = s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | $spl6),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) = s('type.list.list'('type.option.option'('type.conSem.v')),sK7) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f128,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) = s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8)))),
  inference(cnf_transformation,[],[f109])).
fof(f127,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8)))),
  inference(cnf_transformation,[],[f109])).
fof(f191,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | $spl2),
  inference(cnf_transformation,[],[f191_D])).
fof(f191_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f126,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) = s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8)))),
  inference(cnf_transformation,[],[f109])).
fof(f3844,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f3843])).
fof(f3843,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3842,f191])).
fof(f3842,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3841,f193])).
fof(f3841,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) | ($spl1 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3840,f195])).
fof(f3840,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK4))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f190,f197])).
fof(f190,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) | $spl1),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f192,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) | $spl3),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f194,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | $spl5),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f196,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) != s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | $spl7),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  s('type.list.list'('type.option.option'('type.conSem.v')),sK3) != s('type.list.list'('type.option.option'('type.conSem.v')),sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f198,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) != s('type.ffi.ffi_state'(sK0),sK8) | $spl9),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) != s('type.ffi.ffi_state'(sK0),sK8) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f130,plain,(
  s('type.ffi.ffi_state'(sK0),sK4) != s('type.ffi.ffi_state'(sK0),sK8) | s('type.list.list'('type.option.option'('type.conSem.v')),sK3) != s('type.list.list'('type.option.option'('type.conSem.v')),sK7) | s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK2),s('type.list.list'('type.option.option'('type.conSem.v')),sK3),s('type.ffi.ffi_state'(sK0),sK4))) != s('type.conSem.state'(sK0),'const.conSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.conSem.v')),sK6),s('type.list.list'('type.option.option'('type.conSem.v')),sK7),s('type.ffi.ffi_state'(sK0),sK8)))),
  inference(cnf_transformation,[],[f109])).
% SZS output end Proof for thm.conSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 90
% Passive clauses: 1413
% Generated clauses: 7262
% Final active clauses: 80
% Final passive clauses: 1187
% Input formulas: 37
% Initial clauses: 62
% 
% Duplicate literals: 334
% Trivial inequalities: 177
% Fw subsumption resolutions: 1003
% Fw demodulations: 491
% 
% Simple tautologies: 197
% Equational tautologies: 49
% Forward subsumptions: 3269
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 171
% Factoring: 1
% Forward superposition: 2284
% Backward superposition: 2369
% Self superposition: 157
% Equality factoring: 116
% Equality resolution: 23
% 
% Split clauses: 202
% Split components: 415
% SAT solver clauses: 444
% SAT solver unit clauses: 30
% SAT solver binary clauses: 106
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 3
% 
% Sat splits: 202
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 441
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 4861
% Time elapsed: 0.273 s
% ------------------------------
% ------------------------------
% Success in time 0.291 s
