Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (8485)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 3853
% Generated clauses: 20623
% Final active clauses: 143
% Final passive clauses: 3436
% Input formulas: 68
% Initial clauses: 103
% 
% Duplicate literals: 587
% Trivial inequalities: 345
% Fw subsumption resolutions: 1643
% Fw demodulations: 2829
% Bw demodulations: 19
% 
% Simple tautologies: 258
% Equational tautologies: 44
% Forward subsumptions: 10227
% 
% Binary resolution: 543
% Factoring: 1
% Forward superposition: 6906
% Backward superposition: 7290
% Self superposition: 167
% Equality factoring: 72
% Equality resolution: 31
% 
% Split clauses: 259
% Split components: 520
% SAT solver clauses: 711
% SAT solver unit clauses: 9
% SAT solver binary clauses: 144
% 
% Sat splits: 259
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 710
% TWLsolver calls for satisfiability: 186
% 
% Memory used [KB]: 19701
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (8505)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 1434
% Generated clauses: 17144
% Final active clauses: 120
% Final passive clauses: 783
% Input formulas: 68
% Initial clauses: 103
% 
% Duplicate literals: 519
% Trivial inequalities: 189
% Fw subsumption resolutions: 1004
% Fw demodulations: 2072
% Bw demodulations: 425
% Condensations: 1129
% Global subsumptions: 8
% 
% Simple tautologies: 228
% Equational tautologies: 45
% Forward subsumptions: 8623
% Bw demodulations to eq. taut.: 102
% 
% Binary resolution: 397
% Factoring: 1
% Forward superposition: 5460
% Backward superposition: 4159
% Self superposition: 135
% Equality factoring: 82
% Equality resolution: 21
% 
% Split clauses: 114
% Split components: 228
% SAT solver clauses: 1975
% SAT solver unit clauses: 152
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 16
% 
% Sat splits: 114
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1943
% TWLsolver calls for satisfiability: 420
% 
% Memory used [KB]: 17398
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f89149,plain,(
  $false),
  inference(subsumption_resolution,[],[f89148,f89147])).
fof(f89147,plain,(
  ( ! [X662] : (~p(s(bool,X662))) )),
  inference(forward_demodulation,[],[f89146,f316])).
fof(f316,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X19] : ! [X21,X35,X36,X37,X38] : s(X19,X36) = s(X19,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X19)),X21),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X19)),X35),s(X19,X36),s(X19,X37),s(X19,X38)))),
  file('thm.wordSem.result_distinct|split|1','thm.wordSem.result_case_def|split|2')).
fof(f89146,plain,(
  ( ! [X660,X662,X664,X661,X663] : (~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X660),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X661),s(bool,X662),s(bool,X663),s(bool,X664))))) )),
  inference(subsumption_resolution,[],[f88593,f88486])).
fof(f88486,plain,(
  ( ! [X6,X7,X9] : (s(X6,X9) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X6)),X7),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f50423,f316])).
fof(f50423,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(X0,X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f318,f229])).
fof(f229,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f182])).
fof(f182,plain,(
  ? [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.wordSem.result_distinct|split|1','thm.wordSem.result_distinct|split|1')).
fof(f318,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X19,X6] : ! [X8,X7,X21,X35,X36,X37,X38] : s(X19,happ(s(fun('type.wordSem.word_loc'(X6),X19),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X19)),X21),s('type.wordSem.word_loc'(X6),X8))),s('type.wordSem.word_loc'(X6),X7))) = s(X19,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X19)),X21),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X19)),X35),s(X19,X36),s(X19,X37),s(X19,X38)))),
  file('thm.wordSem.result_distinct|split|1','thm.wordSem.result_case_def|split|0')).
fof(f88593,plain,(
  ( ! [X660,X662,X664,X661,X663] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X660),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X660),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X661),s(bool,X662),s(bool,X663),s(bool,X664))))) )),
  inference(superposition,[],[f244,f50423])).
fof(f244,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.wordSem.result_distinct|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f89148,plain,(
  ( ! [X718,X723] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X718),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X723))))) )),
  inference(subsumption_resolution,[],[f88603,f89147])).
fof(f88603,plain,(
  ( ! [X720,X722,X718,X721,X723,X719] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X718),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X723)))) | p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X718),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X719),s(bool,X720),s(bool,X721),s(bool,X722))))) )),
  inference(superposition,[],[f305,f50423])).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X46,X47] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X46))),s(bool,X47)))) <=> (p(s(bool,X46)) => p(s(bool,X47))))),
  file('thm.wordSem.result_distinct|split|1','==>_HL_0')).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 269
% Passive clauses: 3755
% Generated clauses: 89925
% Final active clauses: 170
% Final passive clauses: 1664
% Input formulas: 68
% Initial clauses: 103
% Discarded non-redundant clauses: 1362
% 
% Duplicate literals: 1618
% Trivial inequalities: 555
% Fw subsumption resolutions: 6755
% Fw demodulations: 20702
% Bw demodulations: 395
% Condensations: 4905
% Global subsumptions: 18
% 
% Simple tautologies: 503
% Equational tautologies: 235
% Forward subsumptions: 48035
% Fw demodulations to eq. taut.: 2281
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 6462
% Factoring: 1
% Forward superposition: 25111
% Backward superposition: 21059
% Self superposition: 478
% Equality factoring: 208
% Equality resolution: 58
% 
% Split clauses: 90
% Split components: 180
% SAT solver clauses: 4008
% SAT solver unit clauses: 825
% SAT solver binary clauses: 1288
% 
% Sat splits: 90
% Sat splitting refutations: 53
% 
% Memory used [KB]: 78292
% Time elapsed: 4.406 s
% ------------------------------
% ------------------------------
% Success in time 5.78 s
