Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.dec_to_exhProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.dec_to_exhProof
% SZS output start Proof for thm.dec_to_exhProof
fof(f377,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f366,f327,f77,f128,f126,f155,f291,f80,f130,f78,f134,f76,f127,f125])).
fof(f125,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) | $spl0),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f127,plain,(
  s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | $spl2),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f76,plain,(
  s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))))),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) | s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) | s('type.exhSem.v',sK0) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f47])).
fof(f47,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X0)))) | s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))) & (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X0)))) | s('type.exhSem.v',X0) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X0)))) <~> s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X0)))) <=> s('type.exhSem.v',X0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))),s('type.exhSem.v',X6)))) <=> s('type.exhSem.v',X6) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))),s('type.exhSem.v',X6)))) <=> s('type.exhSem.v',X6) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|0','thm.dec_to_exhProof.v_rel_lit_loc|split|0')).
fof(f134,plain,(
  p(s(bool,'T_0')) | $spl10),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f78,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|0',aHL_TRUTH)).
fof(f130,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f80,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|0',aHL_FALSITY)).
fof(f291,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))) | $spl16),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f155,plain,(
  s(bool,'T_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f154,f130])).
fof(f154,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))) | $spl0),
  inference(superposition,[],[f125,f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|0',aHL_BOOL_CASES)).
fof(f126,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f128,plain,(
  s('type.exhSem.v',sK0) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | $spl3),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s('type.exhSem.v',sK0) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f77,plain,(
  s('type.exhSem.v',sK0) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | ~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),sK2),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))))),
  inference(cnf_transformation,[],[f48])).
fof(f327,plain,(
  $false | ($spl3 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f326,f128])).
fof(f326,plain,(
  s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | ($spl10 | $spl16)),
  inference(subsumption_resolution,[],[f320,f134])).
fof(f320,plain,(
  ~p(s(bool,'T_0')) | s('type.exhSem.v',sK0) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',sK1))) | $spl16),
  inference(superposition,[],[f121,f291])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X2)))) | s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X2)))) | s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X2))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X2)))) <=> s('type.exhSem.v',X2) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X7,X9] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X8),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X7))),s('type.exhSem.v',X9)))) <=> s('type.exhSem.v',X9) = s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.dec_to_exhProof.v_rel_lit_loc|split|0','thm.dec_to_exhProof.dec_to_exhProof_|49||48|_prove_ltb|split|0')).
fof(f366,plain,(
  $false | ($spl1 | $spl2)),
  inference(resolution,[],[f355,f126])).
fof(f355,plain,(
  ( ! [X0] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',sK1))),s('type.exhSem.v',sK0))))) ) | $spl2),
  inference(superposition,[],[f351,f127])).
fof(f351,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))))))) )),
  inference(equality_resolution,[],[f122])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (s('type.exhSem.v',X2) != s('type.exhSem.v','const.exhSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.dec_to_exhProof.v_rel_3'(s('type.finite_map.fmap'('type.ast.id'('type.list.list'('type.string.char')),'type.sptree.spt'('type.num.num')),X0),s('type.conSem.v','const.conSem.Litv_1'(s('type.ast.lit',X1))),s('type.exhSem.v',X2))))) )),
  inference(cnf_transformation,[],[f75])).
% SZS output end Proof for thm.dec_to_exhProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 18
% Passive clauses: 44
% Generated clauses: 284
% Final active clauses: 14
% Final passive clauses: 17
% Input formulas: 21
% Initial clauses: 47
% 
% Duplicate literals: 43
% Trivial inequalities: 4
% Fw subsumption resolutions: 16
% Fw demodulations: 3
% Bw demodulations: 1
% 
% Simple tautologies: 46
% Equational tautologies: 15
% Forward subsumptions: 88
% 
% Binary resolution: 11
% Forward superposition: 55
% Backward superposition: 40
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 33
% SAT solver unit clauses: 4
% SAT solver binary clauses: 15
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 19
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 28
% TWLsolver calls for satisfiability: 12
% 
% Memory used [KB]: 767
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
% Success in time 0.02 s
