Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (9117)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 3756
% Generated clauses: 23974
% Final active clauses: 146
% Final passive clauses: 3356
% Input formulas: 69
% Initial clauses: 89
% 
% Duplicate literals: 609
% Trivial inequalities: 349
% Fw subsumption resolutions: 1972
% Fw demodulations: 4687
% Bw demodulations: 23
% 
% Simple tautologies: 258
% Equational tautologies: 50
% Forward subsumptions: 11756
% 
% Binary resolution: 356
% Factoring: 1
% Forward superposition: 7110
% Backward superposition: 8383
% Self superposition: 200
% Equality factoring: 81
% Equality resolution: 38
% 
% Split clauses: 209
% Split components: 418
% SAT solver clauses: 610
% SAT solver unit clauses: 75
% SAT solver binary clauses: 285
% 
% Sat splits: 209
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 608
% TWLsolver calls for satisfiability: 138
% 
% Memory used [KB]: 19445
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (9139)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 158
% Passive clauses: 1102
% Generated clauses: 17617
% Final active clauses: 134
% Final passive clauses: 785
% Input formulas: 69
% Initial clauses: 89
% 
% Duplicate literals: 403
% Trivial inequalities: 195
% Fw subsumption resolutions: 1039
% Fw demodulations: 3410
% Bw demodulations: 178
% Condensations: 781
% 
% Simple tautologies: 160
% Equational tautologies: 44
% Forward subsumptions: 8354
% 
% Binary resolution: 237
% Factoring: 1
% Forward superposition: 4751
% Backward superposition: 5105
% Self superposition: 164
% Equality factoring: 69
% Equality resolution: 33
% 
% Split clauses: 74
% Split components: 148
% SAT solver clauses: 2160
% SAT solver unit clauses: 133
% SAT solver binary clauses: 51
% 
% Sat splits: 74
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 2160
% TWLsolver calls for satisfiability: 362
% 
% Memory used [KB]: 17654
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f65088,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1440,f1457,f1456,f1441,f1460,f1459,f3573,f3590,f3574,f3593,f3594,f3596,f3604,f3606,f3608,f3612,f3614,f3616,f3618,f3620,f3622,f3626,f3628,f3630,f3632,f3634,f3636,f3638,f3640,f3643,f3645,f3647,f3649,f3651,f3653,f3655,f3688,f3746,f3747,f3748,f3749,f3750,f3751,f3753,f3755,f3757,f3759,f3761,f3763,f3765,f3767,f3768,f3769,f3770,f3771,f3772,f3773,f3774,f3775,f3776,f3777,f3778,f3779,f3780,f3781,f3782,f9071,f9073,f9074,f9076,f13916,f13918,f13884,f13942,f13920,f13885,f13946,f13944,f14924,f14928,f14926,f14925,f50281,f3591,f50282,f64773,f64781,f64785,f64787,f64806,f64814,f64820,f64834,f64836,f64837,f64839,f64840,f64841,f64842,f64844,f64846,f64845,f64847,f64848,f64849,f64850,f64853,f64854,f64855,f64856,f64858,f64859,f64860,f64861,f64862,f64851,f64863,f64865,f64866,f64869,f64870,f64835,f64875,f64876,f64877,f64879,f64880,f64881,f64885,f64889,f64878,f64894,f64895,f64899,f64843,f64900,f64903,f64906,f64907,f64909,f64913,f64922,f64923,f64941,f64783,f64782,f64952,f64960,f64775,f64774,f64964,f64965,f64838,f64998,f65000,f65001,f65002,f65003,f65004,f65005,f65006,f65007,f65008,f65009,f65010,f64999,f65021,f65022,f65023,f65025,f65049,f65053,f65087])).
fof(f65087,plain,(
  $false | ($spl36 | $spl48)),
  inference(subsumption_resolution,[],[f65086,f331])).
fof(f331,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f231,f234])).
fof(f234,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X29] : (s(bool,'T_0') = s(bool,X29) <=> p(s(bool,X29)))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f231,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.stackSem.result_distinct|split|0',aHL_FALSITY)).
fof(f65086,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl36 | $spl48)),
  inference(forward_demodulation,[],[f65085,f242])).
fof(f242,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X17] : ! [X30] : s(X17,X30) = s(X17,'const.combin.I_1'(s(X17,X30)))),
  file('thm.stackSem.result_distinct|split|0','thm.combin.I_THM')).
fof(f65085,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl36 | $spl48)),
  inference(forward_demodulation,[],[f64775,f64839])).
fof(f65053,plain,(
  $false | ($spl34 | $spl50)),
  inference(subsumption_resolution,[],[f64841,f65046])).
fof(f65046,plain,(
  ( ! [X241,X242,X247] : (s(X241,X247) = s(X241,happ(s(fun('type.wordSem.word_loc'(sK0),X241),X242),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl34),
  inference(subsumption_resolution,[],[f65028,f229])).
fof(f229,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f38])).
fof(f38,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.TRUTH')).
fof(f65028,plain,(
  ( ! [X241,X242,X247] : (~p(s(bool,'T_0')) | s(X241,X247) = s(X241,happ(s(fun('type.wordSem.word_loc'(sK0),X241),X242),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl34),
  inference(backward_demodulation,[],[f64774,f64830])).
fof(f64830,plain,(
  ( ! [X243,X241,X242,X247] : (s(X241,X247) = s(X241,happ(s(fun('type.wordSem.word_loc'(sK0),X241),X242),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X241,fun(X241,bool)),'=_0'),s(X241,X247),s(fun('type.wordSem.word_loc'(sK0),X241),X243))),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64829,f305])).
fof(f305,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X3),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X17,X6] : ! [X8,X18,X31,X32,X33,X34] : s(X17,happ(s(fun('type.wordSem.word_loc'(X6),X17),X18),s('type.wordSem.word_loc'(X6),X8))) = s(X17,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X8))),s(fun('type.wordSem.word_loc'(X6),X17),X18),s(fun('type.wordSem.word_loc'(X6),X17),X31),s(fun('type.wordSem.word_loc'(X6),X17),X32),s(X17,X33),s(X17,X34)))),
  file('thm.stackSem.result_distinct|split|0','thm.stackSem.result_case_def|split|0')).
fof(f64829,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X241,fun(X241,bool)),'=_0'),s(X241,X247),s(fun('type.wordSem.word_loc'(sK0),X241),X243))),s('type.wordSem.word_loc'(sK0),sK1)))) | s(X241,X247) = s(X241,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),X241),X242),s(fun('type.wordSem.word_loc'(sK0),X241),X243),s(fun('type.wordSem.word_loc'(sK0),X241),X244),s(X241,X245),s(X241,X246)))) )),
  inference(forward_demodulation,[],[f64493,f313])).
fof(f313,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X11,X12,X37,X13] : ! [X14,X15,X38,X16] : s(X11,happ(s(fun(X13,X11),'_dst_x0x0x1_3'(s(fun(X12,fun(X37,X11)),X14),s(X12,X15),s(fun(X13,X37),X38))),s(X13,X16))) = s(X11,happ(s(fun(X37,X11),happ(s(fun(X12,fun(X37,X11)),X14),s(X12,X15))),s(X37,happ(s(fun(X13,X37),X38),s(X13,X16)))))),
  file('thm.stackSem.result_distinct|split|0','_dst_x0x0x1_ax')).
fof(f64493,plain,(
  ( ! [X243,X241,X246,X244,X242,X247,X245] : (~p(s(bool,happ(s(fun(X241,bool),happ(s(fun(X241,fun(X241,bool)),'=_0'),s(X241,X247))),s(X241,happ(s(fun('type.wordSem.word_loc'(sK0),X241),X243),s('type.wordSem.word_loc'(sK0),sK1)))))) | s(X241,X247) = s(X241,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),X241),X242),s(fun('type.wordSem.word_loc'(sK0),X241),X243),s(fun('type.wordSem.word_loc'(sK0),X241),X244),s(X241,X245),s(X241,X246)))) )),
  inference(superposition,[],[f306,f30531])).
fof(f30531,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f304,f228])).
fof(f228,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) = s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) = s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2] : s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2] : s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) != s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.stackSem.result'(X0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X0),X2))) = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X8))) = s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.stackSem.result'(X6),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(X6),X8))) = s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.stackSem.result_distinct|split|0','thm.stackSem.result_distinct|split|0')).
fof(f304,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X17,X6] : ! [X8,X18,X31,X32,X33,X34] : s(X17,happ(s(fun('type.wordSem.word_loc'(X6),X17),X31),s('type.wordSem.word_loc'(X6),X8))) = s(X17,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X8))),s(fun('type.wordSem.word_loc'(X6),X17),X18),s(fun('type.wordSem.word_loc'(X6),X17),X31),s(fun('type.wordSem.word_loc'(X6),X17),X32),s(X17,X33),s(X17,X34)))),
  file('thm.stackSem.result_distinct|split|0','thm.stackSem.result_case_def|split|1')).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X40,X41,X42] : (p(s(bool,happ(s(fun(X40,bool),happ(s(fun(X40,fun(X40,bool)),'=_0'),s(X40,X41))),s(X40,X42)))) <=> s(X40,X41) = s(X40,X42))),
  file('thm.stackSem.result_distinct|split|0',arity_0_HL_EQ)).
fof(f65049,plain,(
  $false | ($spl34 | $spl46)),
  inference(subsumption_resolution,[],[f65038,f229])).
fof(f65038,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl46)),
  inference(backward_demodulation,[],[f64774,f64838])).
fof(f65025,plain,(
  $false | ($spl40 | $spl52)),
  inference(subsumption_resolution,[],[f65024,f1295])).
fof(f1295,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f331,f1221])).
fof(f1221,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1183,f231])).
fof(f1183,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f234,f1073])).
fof(f1073,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1072,f331])).
fof(f1072,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f916,f243])).
fof(f243,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X17] : ! [X30] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X17,X30)))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.DATATYPE_TAG_THM')).
fof(f916,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(X11,X12))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f787,f233])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.stackSem.result_distinct|split|0',aHL_BOOL_CASES)).
fof(f787,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,'const.bool.DATATYPE_1'(s(X42,X43))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X41)))) )),
  inference(superposition,[],[f767,f242])).
fof(f767,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.DATATYPE_1'(s(X4,X5))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f757,f229])).
fof(f757,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,'const.bool.DATATYPE_1'(s(X4,X5))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f586,f243])).
fof(f586,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f585,f242])).
fof(f585,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f573,f229])).
fof(f573,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.combin.I_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f280,f364])).
fof(f364,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f233,f242])).
fof(f280,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f65024,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ($spl40 | $spl52)),
  inference(forward_demodulation,[],[f64783,f64843])).
fof(f65023,plain,(
  $false | ($spl42 | $spl46)),
  inference(subsumption_resolution,[],[f64835,f64838])).
fof(f65022,plain,(
  $false | ($spl52 | $spl60)),
  inference(subsumption_resolution,[],[f65020,f231])).
fof(f65020,plain,(
  p(s(bool,'F_0')) | ($spl52 | $spl60)),
  inference(backward_demodulation,[],[f64843,f64851])).
fof(f65021,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f65019,f231])).
fof(f65019,plain,(
  p(s(bool,'F_0')) | $spl52),
  inference(backward_demodulation,[],[f64843,f64769])).
fof(f64769,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))),s(fun('type.wordSem.word_loc'(sK0),X0),X2))),s('type.wordSem.word_loc'(sK0),sK2))))) )),
  inference(forward_demodulation,[],[f64768,f313])).
fof(f64768,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))))),s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X2),s('type.wordSem.word_loc'(sK0),sK2))))))) )),
  inference(forward_demodulation,[],[f63955,f305])).
fof(f63955,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))))),s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(X0,X4),s(X0,X5))))))) )),
  inference(unit_resulting_resolution,[],[f30531,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f227])).
fof(f64999,plain,(
  ( ! [X1752] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1752),s('type.wordSem.word_loc'(sK0),sK1)))))) ) | $spl94),
  inference(cnf_transformation,[],[f64999_D])).
fof(f64999_D,plain,(
  ( ! [X1752] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1752),s('type.wordSem.word_loc'(sK0),sK1)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f65010,plain,(
  ( ! [X1784,X1779] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1779),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),X1784)))) ) | $spl106),
  inference(cnf_transformation,[],[f65010_D])).
fof(f65010_D,plain,(
  ( ! [X1784,X1779] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1779),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),X1784)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f65009,plain,(
  ( ! [X1780,X1784,X1779] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1779),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),X1784))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1780),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64700,f305])).
fof(f64700,plain,(
  ( ! [X1780,X1782,X1784,X1779,X1781,X1783] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1780),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1779),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1780),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1781),s('type.wordSem.word_loc'(sK0),X1782),s('type.wordSem.word_loc'(sK0),X1783))),s('type.wordSem.word_loc'(sK0),X1784)))) )),
  inference(superposition,[],[f1909,f30531])).
fof(f1909,plain,(
  ( ! [X8,X9] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X8))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.wordSem.word_loc'(sK0),X8),s('type.wordSem.word_loc'(sK0),X9)))) )),
  inference(superposition,[],[f1894,f255])).
fof(f255,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X17] : ! [X9,X10] : s(X17,X9) = s(X17,'const.bool.COND_3'(s(bool,'T_0'),s(X17,X9),s(X17,X10)))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f1894,plain,(
  ( ! [X0] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X0))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),X0)) )),
  inference(superposition,[],[f269,f228])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1))) != s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X2))) | s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1))) != s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X2))) | s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2)) & (s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),X2) | s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X2)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X2))) <=> s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : (s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1))) = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X2))) <=> s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),X2))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X8,X7] : (s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7))) = s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X8))) <=> s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),X8))),
  file('thm.stackSem.result_distinct|split|0','thm.stackSem.result_|49||49||split|1')).
fof(f65008,plain,(
  ( ! [X1774] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.combin.I_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1774),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl104),
  inference(cnf_transformation,[],[f65008_D])).
fof(f65008_D,plain,(
  ( ! [X1774] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.combin.I_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1774),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f65007,plain,(
  ( ! [X1775,X1774] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.combin.I_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1774),s('type.wordSem.word_loc'(sK0),sK2))))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1775),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64699,f305])).
fof(f64699,plain,(
  ( ! [X1775,X1776,X1778,X1774,X1777] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1775),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.combin.I_1'(s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1774),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1775),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1776),s('type.wordSem.word_loc'(sK0),X1777),s('type.wordSem.word_loc'(sK0),X1778)))))) )),
  inference(superposition,[],[f1908,f30531])).
fof(f1908,plain,(
  ( ! [X7] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X7))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.combin.I_1'(s('type.wordSem.word_loc'(sK0),X7)))) )),
  inference(superposition,[],[f1894,f242])).
fof(f65006,plain,(
  ( ! [X1773,X1768] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1768),s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X1773),s('type.num.num','const.num.|48|_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f65006_D])).
fof(f65006_D,plain,(
  ( ! [X1773,X1768] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1768),s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X1773),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f65005,plain,(
  ( ! [X1769,X1773,X1768] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1768),s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X1773),s('type.num.num','const.num.|48|_0'))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1769),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64698,f305])).
fof(f64698,plain,(
  ( ! [X1769,X1771,X1773,X1768,X1770,X1772] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1769),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1768),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1769),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1770),s('type.wordSem.word_loc'(sK0),X1771),s('type.wordSem.word_loc'(sK0),X1772))),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X1773),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1907,f30531])).
fof(f1907,plain,(
  ( ! [X6,X5] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X5))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),'const.ind_type.FCONS_3'(s('type.wordSem.word_loc'(sK0),X5),s(fun('type.num.num','type.wordSem.word_loc'(sK0)),X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1894,f256])).
fof(f256,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X17] : ! [X8,X18] : s(X17,X8) = s(X17,'const.ind_type.FCONS_3'(s(X17,X8),s(fun('type.num.num',X17),X18),s('type.num.num','const.num.|48|_0')))),
  file('thm.stackSem.result_distinct|split|0','thm.ind_type.FCONS|split|0')).
fof(f65004,plain,(
  ( ! [X1761,X1767,X1766] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun(X1766,'type.wordSem.word_loc'(sK0)),'_dst_x0_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1761),s('type.wordSem.word_loc'(sK0),sK2))))),s(X1766,X1767)))) ) | $spl100),
  inference(cnf_transformation,[],[f65004_D])).
fof(f65004_D,plain,(
  ( ! [X1761,X1767,X1766] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun(X1766,'type.wordSem.word_loc'(sK0)),'_dst_x0_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1761),s('type.wordSem.word_loc'(sK0),sK2))))),s(X1766,X1767)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f65003,plain,(
  ( ! [X1761,X1767,X1762,X1766] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun(X1766,'type.wordSem.word_loc'(sK0)),'_dst_x0_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1761),s('type.wordSem.word_loc'(sK0),sK2))))),s(X1766,X1767))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1762),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64697,f305])).
fof(f64697,plain,(
  ( ! [X1761,X1763,X1765,X1767,X1762,X1764,X1766] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1762),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun(X1766,'type.wordSem.word_loc'(sK0)),'_dst_x0_1'(s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1761),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1762),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1763),s('type.wordSem.word_loc'(sK0),X1764),s('type.wordSem.word_loc'(sK0),X1765))))),s(X1766,X1767)))) )),
  inference(superposition,[],[f1906,f30531])).
fof(f1906,plain,(
  ( ! [X4,X2,X3] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X3))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun(X2,'type.wordSem.word_loc'(sK0)),'_dst_x0_1'(s('type.wordSem.word_loc'(sK0),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f1894,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X13,X11] : ! [X14,X16] : s(X11,X14) = s(X11,happ(s(fun(X13,X11),'_dst_x0_1'(s(X11,X14))),s(X13,X16)))),
  file('thm.stackSem.result_distinct|split|0','_dst_x0_ax')).
fof(f65002,plain,(
  ( ! [X1756] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'const.combin.I_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1756),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl98),
  inference(cnf_transformation,[],[f65002_D])).
fof(f65002_D,plain,(
  ( ! [X1756] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'const.combin.I_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1756),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f65001,plain,(
  ( ! [X1756,X1757] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'const.combin.I_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1756),s('type.wordSem.word_loc'(sK0),sK2))))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1757),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64696,f305])).
fof(f64696,plain,(
  ( ! [X1756,X1758,X1757,X1759,X1760] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1757),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'const.combin.I_0'),s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1756),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1757),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1758),s('type.wordSem.word_loc'(sK0),X1759),s('type.wordSem.word_loc'(sK0),X1760)))))) )),
  inference(superposition,[],[f1905,f30531])).
fof(f1905,plain,(
  ( ! [X1] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X1))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'const.combin.I_0'),s('type.wordSem.word_loc'(sK0),X1)))) )),
  inference(superposition,[],[f1894,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f247,f242])).
fof(f247,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X17] : ! [X20] : s(X17,'const.combin.I_1'(s(X17,X20))) = s(X17,happ(s(fun(X17,X17),'const.combin.I_0'),s(X17,X20)))),
  file('thm.stackSem.result_distinct|split|0','ARITY_0_1_const.combin.I')).
fof(f65000,plain,(
  ( ! [X1751] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'_i_0_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1751),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl96),
  inference(cnf_transformation,[],[f65000_D])).
fof(f65000_D,plain,(
  ( ! [X1751] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'_i_0_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1751),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f64998,plain,(
  ( ! [X1751,X1752] : (s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'_i_0_0'),s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1751),s('type.wordSem.word_loc'(sK0),sK2))))) | s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1752),s('type.wordSem.word_loc'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f64695,f305])).
fof(f64695,plain,(
  ( ! [X1754,X1751,X1753,X1755,X1752] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1752),s('type.wordSem.word_loc'(sK0),sK1))))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'_i_0_0'),s('type.wordSem.word_loc'(sK0),'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1751),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1752),s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),X1753),s('type.wordSem.word_loc'(sK0),X1754),s('type.wordSem.word_loc'(sK0),X1755)))))) )),
  inference(superposition,[],[f1904,f30531])).
fof(f1904,plain,(
  ( ! [X0] : (s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))) != s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),X0))) | s('type.wordSem.word_loc'(sK0),sK1) = s('type.wordSem.word_loc'(sK0),happ(s(fun('type.wordSem.word_loc'(sK0),'type.wordSem.word_loc'(sK0)),'_i_0_0'),s('type.wordSem.word_loc'(sK0),X0)))) )),
  inference(superposition,[],[f1894,f244])).
fof(f244,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X13] : ! [X15] : s(X13,X15) = s(X13,happ(s(fun(X13,X13),'_i_0_0'),s(X13,X15)))),
  file('thm.stackSem.result_distinct|split|0','_i0_ax')).
fof(f64838,plain,(
  ( ! [X460] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X460),s('type.wordSem.word_loc'(sK0),sK1))))) ) | $spl46),
  inference(cnf_transformation,[],[f64838_D])).
fof(f64838_D,plain,(
  ( ! [X460] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X460),s('type.wordSem.word_loc'(sK0),sK1))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f64965,plain,(
  ( ! [X973,X968] : (p(s(bool,happ(s(fun(fun(X973,bool),bool),'!_1'),s(fun(X973,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X968),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) | $spl92),
  inference(cnf_transformation,[],[f64965_D])).
fof(f64965_D,plain,(
  ( ! [X973,X968] : (p(s(bool,happ(s(fun(fun(X973,bool),bool),'!_1'),s(fun(X973,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X968),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f64964,plain,(
  ( ! [X969,X973,X968] : (p(s(bool,happ(s(fun(fun(X973,bool),bool),'!_1'),s(fun(X973,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X968),s('type.wordSem.word_loc'(sK0),sK2)))))))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X969),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64605,f305])).
fof(f64605,plain,(
  ( ! [X971,X969,X973,X970,X968,X972] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X969),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X973,bool),bool),'!_1'),s(fun(X973,bool),'_dst_x0_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X968),s(fun('type.wordSem.word_loc'(sK0),bool),X969),s(fun('type.wordSem.word_loc'(sK0),bool),X970),s(bool,X971),s(bool,X972))))))))) )),
  inference(superposition,[],[f13886,f30531])).
fof(f13886,plain,(
  ( ! [X74,X75] : (~p(s(bool,X75)) | p(s(bool,happ(s(fun(fun(X74,bool),bool),'!_1'),s(fun(X74,bool),'_dst_x0_1'(s(bool,X75))))))) )),
  inference(superposition,[],[f282,f292])).
fof(f282,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X40,X43] : (p(s(bool,happ(s(fun(fun(X40,bool),bool),'!_1'),s(fun(X40,bool),X43)))) <=> ! [X41] : p(s(bool,happ(s(fun(X40,bool),X43),s(X40,X41)))))),
  file('thm.stackSem.result_distinct|split|0','!_HL_QUANT')).
fof(f64774,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X55),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl34),
  inference(cnf_transformation,[],[f64774_D])).
fof(f64774_D,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X55),s('type.wordSem.word_loc'(sK0),sK1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f64775,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X54),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl36),
  inference(cnf_transformation,[],[f64775_D])).
fof(f64775_D,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X54),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f64960,plain,(
  ( ! [X947,X948] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X947),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X948),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl16),
  inference(forward_demodulation,[],[f64959,f305])).
fof(f64959,plain,(
  ( ! [X947,X951,X949,X950,X948] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X948),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X947),s(fun('type.wordSem.word_loc'(sK0),bool),X948),s(fun('type.wordSem.word_loc'(sK0),bool),X949),s(bool,X950),s(bool,X951)))))) ) | $spl16),
  inference(forward_demodulation,[],[f64601,f9144])).
fof(f9144,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f331,f9076,f586])).
fof(f64601,plain,(
  ( ! [X947,X951,X949,X950,X948] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X948),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X947),s(fun('type.wordSem.word_loc'(sK0),bool),X948),s(fun('type.wordSem.word_loc'(sK0),bool),X949),s(bool,X950),s(bool,X951)))))) ) | $spl16),
  inference(superposition,[],[f10753,f30531])).
fof(f10753,plain,(
  ( ! [X54] : (s(bool,X54) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X54)))) ) | $spl16),
  inference(subsumption_resolution,[],[f10752,f331])).
fof(f10752,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X54) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X54)))) ) | $spl16),
  inference(forward_demodulation,[],[f10494,f243])).
fof(f10494,plain,(
  ( ! [X54,X52,X53] : (s(bool,X54) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(X52,X53))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X54)))) ) | $spl16),
  inference(superposition,[],[f9727,f787])).
fof(f9727,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X11)) ) | $spl16),
  inference(subsumption_resolution,[],[f9649,f229])).
fof(f9649,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X11)) ) | $spl16),
  inference(superposition,[],[f586,f9144])).
fof(f64952,plain,(
  ( ! [X926,X925] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X925),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X926),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl14),
  inference(forward_demodulation,[],[f64951,f305])).
fof(f64951,plain,(
  ( ! [X926,X928,X927,X925,X929] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X926),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X925),s(fun('type.wordSem.word_loc'(sK0),bool),X926),s(fun('type.wordSem.word_loc'(sK0),bool),X927),s(bool,X928),s(bool,X929)))))) ) | $spl14),
  inference(forward_demodulation,[],[f64597,f9087])).
fof(f9087,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f331,f9073,f586])).
fof(f64597,plain,(
  ( ! [X926,X928,X927,X925,X929] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X926),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X925),s(fun('type.wordSem.word_loc'(sK0),bool),X926),s(fun('type.wordSem.word_loc'(sK0),bool),X927),s(bool,X928),s(bool,X929)))))) ) | $spl14),
  inference(superposition,[],[f10334,f30531])).
fof(f10334,plain,(
  ( ! [X53] : (s(bool,X53) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X53)))) ) | $spl14),
  inference(subsumption_resolution,[],[f10333,f331])).
fof(f10333,plain,(
  ( ! [X53] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X53)))) ) | $spl14),
  inference(forward_demodulation,[],[f10079,f243])).
fof(f10079,plain,(
  ( ! [X52,X53,X51] : (s(bool,X53) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(X51,X52))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X53)))) ) | $spl14),
  inference(superposition,[],[f9591,f787])).
fof(f9591,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X11)) ) | $spl14),
  inference(subsumption_resolution,[],[f9513,f229])).
fof(f9513,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X11)) ) | $spl14),
  inference(superposition,[],[f586,f9087])).
fof(f64782,plain,(
  ( ! [X108] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X108),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl38),
  inference(cnf_transformation,[],[f64782_D])).
fof(f64782_D,plain,(
  ( ! [X108] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X108),s('type.wordSem.word_loc'(sK0),sK1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f64783,plain,(
  ( ! [X107] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X107),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl40),
  inference(cnf_transformation,[],[f64783_D])).
fof(f64783_D,plain,(
  ( ! [X107] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X107),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f64941,plain,(
  ( ! [X888,X889] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X888),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X889),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64940,f305])).
fof(f64940,plain,(
  ( ! [X890,X888,X892,X891,X889] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X889),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X888),s(fun('type.wordSem.word_loc'(sK0),bool),X889),s(fun('type.wordSem.word_loc'(sK0),bool),X890),s(bool,X891),s(bool,X892)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64590,f242])).
fof(f64590,plain,(
  ( ! [X890,X888,X892,X891,X889] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X889),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X888),s(fun('type.wordSem.word_loc'(sK0),bool),X889),s(fun('type.wordSem.word_loc'(sK0),bool),X890),s(bool,X891),s(bool,X892)))))) ) | $spl2),
  inference(superposition,[],[f6686,f30531])).
fof(f6686,plain,(
  ( ! [X62] : (s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62)))) ) | $spl2),
  inference(subsumption_resolution,[],[f6475,f1468])).
fof(f1468,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK3(X0,'_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1457,f234])).
fof(f6475,plain,(
  ( ! [X61,X62] : (s(bool,X62) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK3(X61,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X62)))) ) | $spl2),
  inference(superposition,[],[f6081,f1987])).
fof(f1987,plain,(
  ( ! [X57,X56] : (s(bool,X56) = s(bool,sK3(X57,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X56)))) ) | $spl2),
  inference(superposition,[],[f1550,f242])).
fof(f1550,plain,(
  ( ! [X12,X13] : (s(bool,X13) = s(bool,sK3(X12,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X13)) ) | $spl2),
  inference(subsumption_resolution,[],[f1518,f231])).
fof(f1518,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,sK3(X12,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X13)) ) | $spl2),
  inference(superposition,[],[f439,f1467])).
fof(f1467,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK3(X0,'_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1457,f239])).
fof(f239,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X29] : (s(bool,'F_0') = s(bool,X29) <=> ~p(s(bool,X29)))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f439,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f427,f231])).
fof(f427,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f277,f233])).
fof(f277,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f6081,plain,(
  ( ! [X185] : (s(bool,X185) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(trivial_inequality_removal,[],[f5825])).
fof(f5825,plain,(
  ( ! [X185] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X185) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f1295,f361])).
fof(f361,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f233,f233])).
fof(f64923,plain,(
  ( ! [X822,X827] : (p(s(bool,happ(s(fun(fun(X827,bool),bool),'?_1'),s(fun(X827,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X822),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) | $spl90),
  inference(cnf_transformation,[],[f64923_D])).
fof(f64923_D,plain,(
  ( ! [X822,X827] : (p(s(bool,happ(s(fun(fun(X827,bool),bool),'?_1'),s(fun(X827,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X822),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f64922,plain,(
  ( ! [X823,X822,X827] : (p(s(bool,happ(s(fun(fun(X827,bool),bool),'?_1'),s(fun(X827,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X822),s('type.wordSem.word_loc'(sK0),sK2)))))))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X823),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64579,f305])).
fof(f64579,plain,(
  ( ! [X825,X823,X826,X824,X822,X827] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X823),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X827,bool),bool),'?_1'),s(fun(X827,bool),'_dst_x0_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X822),s(fun('type.wordSem.word_loc'(sK0),bool),X823),s(fun('type.wordSem.word_loc'(sK0),bool),X824),s(bool,X825),s(bool,X826))))))))) )),
  inference(superposition,[],[f3575,f30531])).
fof(f3575,plain,(
  ( ! [X30,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(fun(X29,bool),bool),'?_1'),s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))))))) )),
  inference(superposition,[],[f284,f292])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X40,X43] : (p(s(bool,happ(s(fun(fun(X40,bool),bool),'?_1'),s(fun(X40,bool),X43)))) <=> ? [X41] : p(s(bool,happ(s(fun(X40,bool),X43),s(X40,X41)))))),
  file('thm.stackSem.result_distinct|split|0','?_HL_QUANT')).
fof(f64913,plain,(
  ( ! [X784,X785] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X784),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X785),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64912,f305])).
fof(f64912,plain,(
  ( ! [X786,X784,X788,X787,X785] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X785),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X784),s(fun('type.wordSem.word_loc'(sK0),bool),X785),s(fun('type.wordSem.word_loc'(sK0),bool),X786),s(bool,X787),s(bool,X788)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64573,f1623])).
fof(f1623,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl2),
  inference(unit_resulting_resolution,[],[f1564,f235])).
fof(f235,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f1564,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl2),
  inference(unit_resulting_resolution,[],[f1468,f261])).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK3(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK3(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X17] : ! [X28,X29] : (! [X30] : (s(X17,X29) = s(X17,X30) => p(s(bool,happ(s(fun(X17,bool),X28),s(X17,X30))))) => p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X28)))))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.JRH_INDUCT_UTIL')).
fof(f64573,plain,(
  ( ! [X786,X784,X788,X787,X785] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X785),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X784),s(fun('type.wordSem.word_loc'(sK0),bool),X785),s(fun('type.wordSem.word_loc'(sK0),bool),X786),s(bool,X787),s(bool,X788)))))) ) | $spl2),
  inference(superposition,[],[f2846,f30531])).
fof(f2846,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl2),
  inference(subsumption_resolution,[],[f2845,f331])).
fof(f2845,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X17) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl2),
  inference(forward_demodulation,[],[f2700,f243])).
fof(f2700,plain,(
  ( ! [X17,X15,X16] : (s(bool,X17) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(X15,X16))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl2),
  inference(superposition,[],[f1732,f787])).
fof(f1732,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl2),
  inference(subsumption_resolution,[],[f1707,f229])).
fof(f1707,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl2),
  inference(superposition,[],[f586,f1623])).
fof(f64909,plain,(
  ( ! [X775,X774] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X774),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X775),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl4),
  inference(forward_demodulation,[],[f64908,f305])).
fof(f64908,plain,(
  ( ! [X777,X775,X778,X776,X774] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X775),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X774),s(fun('type.wordSem.word_loc'(sK0),bool),X775),s(fun('type.wordSem.word_loc'(sK0),bool),X776),s(bool,X777),s(bool,X778)))))) ) | $spl4),
  inference(forward_demodulation,[],[f64571,f1485])).
fof(f1485,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | $spl4),
  inference(unit_resulting_resolution,[],[f1459,f235])).
fof(f64571,plain,(
  ( ! [X777,X775,X778,X776,X774] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X775),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X774),s(fun('type.wordSem.word_loc'(sK0),bool),X775),s(fun('type.wordSem.word_loc'(sK0),bool),X776),s(bool,X777),s(bool,X778)))))) ) | $spl4),
  inference(superposition,[],[f2634,f30531])).
fof(f2634,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl4),
  inference(subsumption_resolution,[],[f2633,f331])).
fof(f2633,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl4),
  inference(forward_demodulation,[],[f2494,f243])).
fof(f2494,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,'const.bool.DATATYPE_1'(s(X13,X14))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X15)))) ) | $spl4),
  inference(superposition,[],[f1672,f787])).
fof(f1672,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl4),
  inference(subsumption_resolution,[],[f1647,f229])).
fof(f1647,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) | s(bool,'F_0') = s(bool,X5)) ) | $spl4),
  inference(superposition,[],[f586,f1485])).
fof(f64907,plain,(
  ( ! [X770] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X770),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl88),
  inference(cnf_transformation,[],[f64907_D])).
fof(f64907_D,plain,(
  ( ! [X770] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X770),s('type.wordSem.word_loc'(sK0),sK1)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f64906,plain,(
  ( ! [X769,X770] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X769),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X770),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64570,f305])).
fof(f64570,plain,(
  ( ! [X771,X769,X773,X770,X772] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X770),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X769),s(fun('type.wordSem.word_loc'(sK0),bool),X770),s(fun('type.wordSem.word_loc'(sK0),bool),X771),s(bool,X772),s(bool,X773)))))) )),
  inference(superposition,[],[f2197,f30531])).
fof(f2197,plain,(
  ( ! [X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X77)))) )),
  inference(forward_demodulation,[],[f2163,f243])).
fof(f2163,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,'const.combin.I_1'(s(bool,X77))) = s(bool,'const.bool.DATATYPE_1'(s(X78,X79)))) )),
  inference(superposition,[],[f1402,f242])).
fof(f1402,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | s(bool,X24) = s(bool,'const.bool.DATATYPE_1'(s(X25,X26)))) )),
  inference(superposition,[],[f1151,f242])).
fof(f1151,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.DATATYPE_1'(s(X8,X9)))) )),
  inference(subsumption_resolution,[],[f1037,f231])).
fof(f1037,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.bool.DATATYPE_1'(s(X8,X9)))) )),
  inference(superposition,[],[f234,f787])).
fof(f64903,plain,(
  ( ! [X757,X758] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X757),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X758),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64902,f305])).
fof(f64902,plain,(
  ( ! [X761,X757,X759,X760,X758] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X758),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X757),s(fun('type.wordSem.word_loc'(sK0),bool),X758),s(fun('type.wordSem.word_loc'(sK0),bool),X759),s(bool,X760),s(bool,X761)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64568,f1467])).
fof(f64568,plain,(
  ( ! [X761,X757,X759,X760,X762,X758] : (s(bool,sK3(X762,'_i_0_0',bool)) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X758),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X757),s(fun('type.wordSem.word_loc'(sK0),bool),X758),s(fun('type.wordSem.word_loc'(sK0),bool),X759),s(bool,X760),s(bool,X761)))))) ) | $spl2),
  inference(superposition,[],[f1987,f30531])).
fof(f64900,plain,(
  ( ! [X750,X746] : (p(s(bool,'|dollar|exists_1'(s(fun(X750,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X746),s('type.wordSem.word_loc'(sK0),sK1))))))))) ) | $spl86),
  inference(cnf_transformation,[],[f64900_D])).
fof(f64900_D,plain,(
  ( ! [X750,X746] : (p(s(bool,'|dollar|exists_1'(s(fun(X750,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X746),s('type.wordSem.word_loc'(sK0),sK1))))))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f64843,plain,(
  ( ! [X469] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X469),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl52),
  inference(cnf_transformation,[],[f64843_D])).
fof(f64843_D,plain,(
  ( ! [X469] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X469),s('type.wordSem.word_loc'(sK0),sK2)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f64899,plain,(
  ( ! [X745,X750,X746] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X745),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,'|dollar|exists_1'(s(fun(X750,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X746),s('type.wordSem.word_loc'(sK0),sK1))))))))) )),
  inference(forward_demodulation,[],[f64566,f305])).
fof(f64566,plain,(
  ( ! [X749,X745,X747,X748,X750,X746] : (p(s(bool,'|dollar|exists_1'(s(fun(X750,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X746),s('type.wordSem.word_loc'(sK0),sK1)))))))) | s(bool,'F_0') = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X745),s(fun('type.wordSem.word_loc'(sK0),bool),X746),s(fun('type.wordSem.word_loc'(sK0),bool),X747),s(bool,X748),s(bool,X749)))) )),
  inference(superposition,[],[f1761,f30531])).
fof(f1761,plain,(
  ( ! [X0,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1760,f242])).
fof(f1760,plain,(
  ( ! [X0,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(subsumption_resolution,[],[f1759,f229])).
fof(f1759,plain,(
  ( ! [X0,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(forward_demodulation,[],[f1739,f243])).
fof(f1739,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X1,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(superposition,[],[f1442,f787])).
fof(f1442,plain,(
  ( ! [X15,X16] : (~p(s(bool,X16)) | p(s(bool,'|dollar|exists_1'(s(fun(X15,bool),'_dst_x0_1'(s(bool,X16))))))) )),
  inference(superposition,[],[f262,f292])).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK3(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f208])).
fof(f64895,plain,(
  ( ! [X728,X723] : (p(s(bool,'|dollar|exists_1'(s(fun(X728,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X723),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) | $spl84),
  inference(cnf_transformation,[],[f64895_D])).
fof(f64895_D,plain,(
  ( ! [X728,X723] : (p(s(bool,'|dollar|exists_1'(s(fun(X728,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X723),s('type.wordSem.word_loc'(sK0),sK2))))))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f64894,plain,(
  ( ! [X728,X724,X723] : (p(s(bool,'|dollar|exists_1'(s(fun(X728,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X723),s('type.wordSem.word_loc'(sK0),sK2)))))))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X724),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64562,f305])).
fof(f64562,plain,(
  ( ! [X728,X724,X726,X725,X727,X723] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X724),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X728,bool),'_dst_x0_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X723),s(fun('type.wordSem.word_loc'(sK0),bool),X724),s(fun('type.wordSem.word_loc'(sK0),bool),X725),s(bool,X726),s(bool,X727))))))))) )),
  inference(superposition,[],[f1442,f30531])).
fof(f64878,plain,(
  ( ! [X634] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X634),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl78),
  inference(cnf_transformation,[],[f64878_D])).
fof(f64878_D,plain,(
  ( ! [X634] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X634),s('type.wordSem.word_loc'(sK0),sK1)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f64889,plain,(
  ( ! [X699,X700] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X699),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X700),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64558,f305])).
fof(f64558,plain,(
  ( ! [X699,X700,X702,X701,X703] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X700),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X699),s(fun('type.wordSem.word_loc'(sK0),bool),X700),s(fun('type.wordSem.word_loc'(sK0),bool),X701),s(bool,X702),s(bool,X703)))))) )),
  inference(superposition,[],[f1111,f30531])).
fof(f1111,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) )),
  inference(subsumption_resolution,[],[f1110,f229])).
fof(f1110,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) )),
  inference(forward_demodulation,[],[f978,f243])).
fof(f978,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X10,X11)))) | s(bool,'F_0') != s(bool,X9) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,X9)))) )),
  inference(superposition,[],[f238,f787])).
fof(f238,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f205])).
fof(f64885,plain,(
  ( ! [X656,X657] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X656),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X657),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64884,f305])).
fof(f64884,plain,(
  ( ! [X660,X656,X658,X657,X659] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X657),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X656),s(fun('type.wordSem.word_loc'(sK0),bool),X657),s(fun('type.wordSem.word_loc'(sK0),bool),X658),s(bool,X659),s(bool,X660)))))) )),
  inference(forward_demodulation,[],[f64551,f243])).
fof(f64551,plain,(
  ( ! [X660,X662,X656,X658,X661,X657,X659] : (s(bool,'const.bool.DATATYPE_1'(s(X661,X662))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X657),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X656),s(fun('type.wordSem.word_loc'(sK0),bool),X657),s(fun('type.wordSem.word_loc'(sK0),bool),X658),s(bool,X659),s(bool,X660)))))) )),
  inference(superposition,[],[f787,f30531])).
fof(f64881,plain,(
  ( ! [X638] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X638),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl82),
  inference(cnf_transformation,[],[f64881_D])).
fof(f64881_D,plain,(
  ( ! [X638] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X638),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f64880,plain,(
  ( ! [X639,X638] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X638),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X639),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64548,f305])).
fof(f64548,plain,(
  ( ! [X639,X640,X642,X638,X641] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X639),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X638),s(fun('type.wordSem.word_loc'(sK0),bool),X639),s(fun('type.wordSem.word_loc'(sK0),bool),X640),s(bool,X641),s(bool,X642)))))) )),
  inference(superposition,[],[f609,f30531])).
fof(f609,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X8)))) )),
  inference(forward_demodulation,[],[f596,f242])).
fof(f596,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X8)))) )),
  inference(superposition,[],[f584,f318])).
fof(f584,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f569,f229])).
fof(f569,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f238,f364])).
fof(f64879,plain,(
  ( ! [X633] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X633),s('type.wordSem.word_loc'(sK0),sK2)))))) ) | $spl80),
  inference(cnf_transformation,[],[f64879_D])).
fof(f64879_D,plain,(
  ( ! [X633] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X633),s('type.wordSem.word_loc'(sK0),sK2)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f64877,plain,(
  ( ! [X633,X634] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X633),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X634),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64547,f305])).
fof(f64547,plain,(
  ( ! [X633,X635,X636,X634,X637] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X634),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X633),s(fun('type.wordSem.word_loc'(sK0),bool),X634),s(fun('type.wordSem.word_loc'(sK0),bool),X635),s(bool,X636),s(bool,X637)))))) )),
  inference(superposition,[],[f608,f30531])).
fof(f608,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X7)))) )),
  inference(forward_demodulation,[],[f595,f242])).
fof(f595,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X7)))) )),
  inference(superposition,[],[f584,f244])).
fof(f64876,plain,(
  ( ! [X627,X632] : (s(bool,X632) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X627),s('type.wordSem.word_loc'(sK0),sK2))) | s(bool,'F_0') = s(bool,X632)) ) | $spl76),
  inference(cnf_transformation,[],[f64876_D])).
fof(f64876_D,plain,(
  ( ! [X627,X632] : (s(bool,X632) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X627),s('type.wordSem.word_loc'(sK0),sK2))) | s(bool,'F_0') = s(bool,X632)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f64875,plain,(
  ( ! [X627,X632,X628] : (s(bool,X632) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X627),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X628),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'F_0') = s(bool,X632)) )),
  inference(forward_demodulation,[],[f64546,f305])).
fof(f64546,plain,(
  ( ! [X629,X631,X627,X632,X628,X630] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X628),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,X632) = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X627),s(fun('type.wordSem.word_loc'(sK0),bool),X628),s(fun('type.wordSem.word_loc'(sK0),bool),X629),s(bool,X630),s(bool,X631))) | s(bool,'F_0') = s(bool,X632)) )),
  inference(superposition,[],[f586,f30531])).
fof(f64835,plain,(
  ( ! [X455] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X455),s('type.wordSem.word_loc'(sK0),sK1))))) ) | $spl42),
  inference(cnf_transformation,[],[f64835_D])).
fof(f64835_D,plain,(
  ( ! [X455] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X455),s('type.wordSem.word_loc'(sK0),sK1))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f64870,plain,(
  ( ! [X594,X589] : (s(bool,X594) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X589),s('type.wordSem.word_loc'(sK0),sK2))) | s(bool,'T_0') = s(bool,X594)) ) | $spl74),
  inference(cnf_transformation,[],[f64870_D])).
fof(f64870_D,plain,(
  ( ! [X594,X589] : (s(bool,X594) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X589),s('type.wordSem.word_loc'(sK0),sK2))) | s(bool,'T_0') = s(bool,X594)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f64869,plain,(
  ( ! [X594,X590,X589] : (s(bool,X594) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X589),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X590),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'T_0') = s(bool,X594)) )),
  inference(forward_demodulation,[],[f64539,f305])).
fof(f64539,plain,(
  ( ! [X592,X594,X590,X593,X589,X591] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X590),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,X594) = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X589),s(fun('type.wordSem.word_loc'(sK0),bool),X590),s(fun('type.wordSem.word_loc'(sK0),bool),X591),s(bool,X592),s(bool,X593))) | s(bool,'T_0') = s(bool,X594)) )),
  inference(superposition,[],[f439,f30531])).
fof(f64866,plain,(
  ( ! [X549,X554] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X554),s(fun('type.wordSem.word_loc'(sK0),bool),X549))),s('type.wordSem.word_loc'(sK0),sK2))))) ) | $spl72),
  inference(cnf_transformation,[],[f64866_D])).
fof(f64866_D,plain,(
  ( ! [X549,X554] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X554),s(fun('type.wordSem.word_loc'(sK0),bool),X549))),s('type.wordSem.word_loc'(sK0),sK2))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f64865,plain,(
  ( ! [X549,X554,X550] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X554),s(fun('type.wordSem.word_loc'(sK0),bool),X549))),s('type.wordSem.word_loc'(sK0),sK2)))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X550),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64864,f313])).
fof(f64864,plain,(
  ( ! [X549,X554,X550] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X554))),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X549),s('type.wordSem.word_loc'(sK0),sK2)))))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X550),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64532,f305])).
fof(f64532,plain,(
  ( ! [X553,X549,X551,X552,X554,X550] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X550),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X554))),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X549),s(fun('type.wordSem.word_loc'(sK0),bool),X550),s(fun('type.wordSem.word_loc'(sK0),bool),X551),s(bool,X552),s(bool,X553))))))) )),
  inference(superposition,[],[f290,f30531])).
fof(f290,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) => p(s(bool,X42))))),
  file('thm.stackSem.result_distinct|split|0','==>_HL_0')).
fof(f64863,plain,(
  ( ! [X542,X538] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X538),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X542))))) ) | $spl70),
  inference(cnf_transformation,[],[f64863_D])).
fof(f64863_D,plain,(
  ( ! [X542,X538] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X538),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X542))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f64851,plain,(
  ( ! [X507] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X507),s('type.wordSem.word_loc'(sK0),sK2))))) ) | $spl60),
  inference(cnf_transformation,[],[f64851_D])).
fof(f64851_D,plain,(
  ( ! [X507] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X507),s('type.wordSem.word_loc'(sK0),sK2))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f64862,plain,(
  ( ! [X542,X538,X537] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X537),s('type.wordSem.word_loc'(sK0),sK2)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X538),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X542))))) )),
  inference(forward_demodulation,[],[f64530,f305])).
fof(f64530,plain,(
  ( ! [X540,X542,X538,X541,X537,X539] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X538),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X542)))) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X537),s(fun('type.wordSem.word_loc'(sK0),bool),X538),s(fun('type.wordSem.word_loc'(sK0),bool),X539),s(bool,X540),s(bool,X541))))) )),
  inference(superposition,[],[f289,f30531])).
fof(f289,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f224])).
fof(f64861,plain,(
  ( ! [X536,X531] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X531),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X536)))) | p(s(bool,X536))) ) | $spl68),
  inference(cnf_transformation,[],[f64861_D])).
fof(f64861_D,plain,(
  ( ! [X536,X531] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X531),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X536)))) | p(s(bool,X536))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f64860,plain,(
  ( ! [X536,X532,X531] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X531),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X536)))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X532),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X536))) )),
  inference(forward_demodulation,[],[f64529,f305])).
fof(f64529,plain,(
  ( ! [X536,X532,X534,X533,X535,X531] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X532),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X536)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X531),s(fun('type.wordSem.word_loc'(sK0),bool),X532),s(fun('type.wordSem.word_loc'(sK0),bool),X533),s(bool,X534),s(bool,X535))))),s(bool,X536))))) )),
  inference(superposition,[],[f288,f30531])).
fof(f288,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f64859,plain,(
  ( ! [X530,X525] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X530),s(fun('type.wordSem.word_loc'(sK0),bool),X525))),s('type.wordSem.word_loc'(sK0),sK2))))) ) | $spl66),
  inference(cnf_transformation,[],[f64859_D])).
fof(f64859_D,plain,(
  ( ! [X530,X525] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X530),s(fun('type.wordSem.word_loc'(sK0),bool),X525))),s('type.wordSem.word_loc'(sK0),sK2))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f64858,plain,(
  ( ! [X530,X526,X525] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X530),s(fun('type.wordSem.word_loc'(sK0),bool),X525))),s('type.wordSem.word_loc'(sK0),sK2)))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X526),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64857,f313])).
fof(f64857,plain,(
  ( ! [X530,X526,X525] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X530))),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X525),s('type.wordSem.word_loc'(sK0),sK2)))))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X526),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64528,f305])).
fof(f64528,plain,(
  ( ! [X528,X530,X526,X529,X525,X527] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X526),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X530))),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X525),s(fun('type.wordSem.word_loc'(sK0),bool),X526),s(fun('type.wordSem.word_loc'(sK0),bool),X527),s(bool,X528),s(bool,X529))))))) )),
  inference(superposition,[],[f287,f30531])).
fof(f287,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) | p(s(bool,X42))))),
  file('thm.stackSem.result_distinct|split|0','|bslash||slash|_HL_0')).
fof(f64856,plain,(
  ( ! [X524,X519] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X519),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X524))))) ) | $spl64),
  inference(cnf_transformation,[],[f64856_D])).
fof(f64856_D,plain,(
  ( ! [X524,X519] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X519),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X524))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f64855,plain,(
  ( ! [X524,X520,X519] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X519),s('type.wordSem.word_loc'(sK0),sK2))))),s(bool,X524)))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X520),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64527,f305])).
fof(f64527,plain,(
  ( ! [X524,X520,X522,X521,X523,X519] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X520),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X519),s(fun('type.wordSem.word_loc'(sK0),bool),X520),s(fun('type.wordSem.word_loc'(sK0),bool),X521),s(bool,X522),s(bool,X523))))),s(bool,X524))))) )),
  inference(superposition,[],[f286,f30531])).
fof(f286,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f64854,plain,(
  ( ! [X518,X514] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X518),s(fun('type.wordSem.word_loc'(sK0),bool),X514))),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X518))) ) | $spl62),
  inference(cnf_transformation,[],[f64854_D])).
fof(f64854_D,plain,(
  ( ! [X518,X514] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X518),s(fun('type.wordSem.word_loc'(sK0),bool),X514))),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X518))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f64853,plain,(
  ( ! [X518,X514,X513] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X513),s('type.wordSem.word_loc'(sK0),sK2)))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X518),s(fun('type.wordSem.word_loc'(sK0),bool),X514))),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X518))) )),
  inference(forward_demodulation,[],[f64852,f305])).
fof(f64852,plain,(
  ( ! [X516,X518,X514,X517,X513,X515] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X518),s(fun('type.wordSem.word_loc'(sK0),bool),X514))),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X518)) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X513),s(fun('type.wordSem.word_loc'(sK0),bool),X514),s(fun('type.wordSem.word_loc'(sK0),bool),X515),s(bool,X516),s(bool,X517))))) )),
  inference(forward_demodulation,[],[f64526,f313])).
fof(f64526,plain,(
  ( ! [X516,X518,X514,X517,X513,X515] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X518))),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X514),s('type.wordSem.word_loc'(sK0),sK1)))))) | p(s(bool,X518)) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X513),s(fun('type.wordSem.word_loc'(sK0),bool),X514),s(fun('type.wordSem.word_loc'(sK0),bool),X515),s(bool,X516),s(bool,X517))))) )),
  inference(superposition,[],[f285,f30531])).
fof(f285,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f64850,plain,(
  ( ! [X512,X508] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X508),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X512)))) | p(s(bool,X512))) ) | $spl58),
  inference(cnf_transformation,[],[f64850_D])).
fof(f64850_D,plain,(
  ( ! [X512,X508] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X508),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X512)))) | p(s(bool,X512))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f64849,plain,(
  ( ! [X507,X512,X508] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X507),s('type.wordSem.word_loc'(sK0),sK2)))) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X508),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X512)))) | p(s(bool,X512))) )),
  inference(forward_demodulation,[],[f64525,f305])).
fof(f64525,plain,(
  ( ! [X509,X511,X507,X512,X508,X510] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X508),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X512)))) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X507),s(fun('type.wordSem.word_loc'(sK0),bool),X508),s(fun('type.wordSem.word_loc'(sK0),bool),X509),s(bool,X510),s(bool,X511)))) | p(s(bool,X512))) )),
  inference(superposition,[],[f285,f30531])).
fof(f64848,plain,(
  ( ! [X506,X501] : (s(bool,X506) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X501),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,X506))) ) | $spl56),
  inference(cnf_transformation,[],[f64848_D])).
fof(f64848_D,plain,(
  ( ! [X506,X501] : (s(bool,X506) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X501),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,X506))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f64847,plain,(
  ( ! [X502,X506,X501] : (s(bool,X506) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X501),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X502),s('type.wordSem.word_loc'(sK0),sK1)))) | ~p(s(bool,X506))) )),
  inference(forward_demodulation,[],[f64524,f305])).
fof(f64524,plain,(
  ( ! [X505,X502,X504,X506,X501,X503] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X502),s('type.wordSem.word_loc'(sK0),sK1)))) | ~p(s(bool,X506)) | s(bool,X506) = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X501),s(fun('type.wordSem.word_loc'(sK0),bool),X502),s(fun('type.wordSem.word_loc'(sK0),bool),X503),s(bool,X504),s(bool,X505)))) )),
  inference(superposition,[],[f280,f30531])).
fof(f64845,plain,(
  ( ! [X494,X489] : (s(bool,X494) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X489),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,X494))) ) | $spl54),
  inference(cnf_transformation,[],[f64845_D])).
fof(f64845_D,plain,(
  ( ! [X494,X489] : (s(bool,X494) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X489),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,X494))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f64846,plain,(
  ( ! [X500,X496,X495] : (s(bool,X500) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X495),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X496),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X500))) )),
  inference(forward_demodulation,[],[f64523,f305])).
fof(f64523,plain,(
  ( ! [X499,X500,X496,X498,X495,X497] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X496),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X500)) | s(bool,X500) = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X495),s(fun('type.wordSem.word_loc'(sK0),bool),X496),s(fun('type.wordSem.word_loc'(sK0),bool),X497),s(bool,X498),s(bool,X499)))) )),
  inference(superposition,[],[f277,f30531])).
fof(f64844,plain,(
  ( ! [X494,X490,X489] : (s(bool,X494) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X489),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X490),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X494))) )),
  inference(forward_demodulation,[],[f64522,f305])).
fof(f64522,plain,(
  ( ! [X492,X494,X490,X493,X489,X491] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X490),s('type.wordSem.word_loc'(sK0),sK1)))) | p(s(bool,X494)) | s(bool,X494) = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X489),s(fun('type.wordSem.word_loc'(sK0),bool),X490),s(fun('type.wordSem.word_loc'(sK0),bool),X491),s(bool,X492),s(bool,X493)))) )),
  inference(superposition,[],[f277,f30531])).
fof(f64842,plain,(
  ( ! [X469,X470] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X469),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X470),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64520,f305])).
fof(f64520,plain,(
  ( ! [X472,X469,X471,X473,X470] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X470),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'F_0') = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X469),s(fun('type.wordSem.word_loc'(sK0),bool),X470),s(fun('type.wordSem.word_loc'(sK0),bool),X471),s(bool,X472),s(bool,X473)))) )),
  inference(superposition,[],[f239,f30531])).
fof(f64841,plain,(
  ( ! [X464] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X464),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl50),
  inference(cnf_transformation,[],[f64841_D])).
fof(f64841_D,plain,(
  ( ! [X464] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X464),s('type.wordSem.word_loc'(sK0),sK2)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f64840,plain,(
  ( ! [X465,X464] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X464),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X465),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64519,f305])).
fof(f64519,plain,(
  ( ! [X466,X465,X467,X468,X464] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X465),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'F_0') != s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X464),s(fun('type.wordSem.word_loc'(sK0),bool),X465),s(fun('type.wordSem.word_loc'(sK0),bool),X466),s(bool,X467),s(bool,X468)))) )),
  inference(superposition,[],[f238,f30531])).
fof(f64839,plain,(
  ( ! [X459] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X459),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl48),
  inference(cnf_transformation,[],[f64839_D])).
fof(f64839_D,plain,(
  ( ! [X459] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X459),s('type.wordSem.word_loc'(sK0),sK2)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f64837,plain,(
  ( ! [X459,X460] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X459),s('type.wordSem.word_loc'(sK0),sK2))) | ~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X460),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64518,f305])).
fof(f64518,plain,(
  ( ! [X461,X463,X459,X460,X462] : (~p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X460),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'T_0') = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X459),s(fun('type.wordSem.word_loc'(sK0),bool),X460),s(fun('type.wordSem.word_loc'(sK0),bool),X461),s(bool,X462),s(bool,X463)))) )),
  inference(superposition,[],[f235,f30531])).
fof(f64836,plain,(
  ( ! [X454] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X454),s('type.wordSem.word_loc'(sK0),sK2)))) ) | $spl44),
  inference(cnf_transformation,[],[f64836_D])).
fof(f64836_D,plain,(
  ( ! [X454] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X454),s('type.wordSem.word_loc'(sK0),sK2)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f64834,plain,(
  ( ! [X455,X454] : (s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X454),s('type.wordSem.word_loc'(sK0),sK2))) | p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X455),s('type.wordSem.word_loc'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f64517,f305])).
fof(f64517,plain,(
  ( ! [X455,X457,X454,X456,X458] : (p(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X455),s('type.wordSem.word_loc'(sK0),sK1)))) | s(bool,'T_0') != s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X454),s(fun('type.wordSem.word_loc'(sK0),bool),X455),s(fun('type.wordSem.word_loc'(sK0),bool),X456),s(bool,X457),s(bool,X458)))) )),
  inference(superposition,[],[f234,f30531])).
fof(f64820,plain,(
  ( ! [X227,X226] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X226),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X227),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl16),
  inference(forward_demodulation,[],[f64819,f305])).
fof(f64819,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X227),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X226),s(fun('type.wordSem.word_loc'(sK0),bool),X227),s(fun('type.wordSem.word_loc'(sK0),bool),X228),s(bool,X229),s(bool,X230)))))) ) | $spl16),
  inference(forward_demodulation,[],[f64458,f9144])).
fof(f64458,plain,(
  ( ! [X229,X227,X230,X228,X226] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X227),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X226),s(fun('type.wordSem.word_loc'(sK0),bool),X227),s(fun('type.wordSem.word_loc'(sK0),bool),X228),s(bool,X229),s(bool,X230)))))) ) | $spl16),
  inference(superposition,[],[f30531,f10753])).
fof(f64814,plain,(
  ( ! [X210,X209] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X209),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X210),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl14),
  inference(forward_demodulation,[],[f64813,f305])).
fof(f64813,plain,(
  ( ! [X212,X210,X213,X211,X209] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X210),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X209),s(fun('type.wordSem.word_loc'(sK0),bool),X210),s(fun('type.wordSem.word_loc'(sK0),bool),X211),s(bool,X212),s(bool,X213)))))) ) | $spl14),
  inference(forward_demodulation,[],[f64455,f9087])).
fof(f64455,plain,(
  ( ! [X212,X210,X213,X211,X209] : (s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X210),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X209),s(fun('type.wordSem.word_loc'(sK0),bool),X210),s(fun('type.wordSem.word_loc'(sK0),bool),X211),s(bool,X212),s(bool,X213)))))) ) | $spl14),
  inference(superposition,[],[f30531,f10334])).
fof(f64806,plain,(
  ( ! [X182,X183] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X182),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X183),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64805,f305])).
fof(f64805,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X183),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X182),s(fun('type.wordSem.word_loc'(sK0),bool),X183),s(fun('type.wordSem.word_loc'(sK0),bool),X184),s(bool,X185),s(bool,X186)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64450,f242])).
fof(f64450,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'const.combin.I_1'(s(bool,'F_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X183),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X182),s(fun('type.wordSem.word_loc'(sK0),bool),X183),s(fun('type.wordSem.word_loc'(sK0),bool),X184),s(bool,X185),s(bool,X186)))))) ) | $spl2),
  inference(superposition,[],[f30531,f6686])).
fof(f64787,plain,(
  ( ! [X118,X119] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X118),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X119),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64786,f305])).
fof(f64786,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X119),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X118),s(fun('type.wordSem.word_loc'(sK0),bool),X119),s(fun('type.wordSem.word_loc'(sK0),bool),X120),s(bool,X121),s(bool,X122)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64440,f1623])).
fof(f64440,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X119),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X118),s(fun('type.wordSem.word_loc'(sK0),bool),X119),s(fun('type.wordSem.word_loc'(sK0),bool),X120),s(bool,X121),s(bool,X122)))))) ) | $spl2),
  inference(superposition,[],[f30531,f2846])).
fof(f64785,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X113),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X114),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl4),
  inference(forward_demodulation,[],[f64784,f305])).
fof(f64784,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X114),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X113),s(fun('type.wordSem.word_loc'(sK0),bool),X114),s(fun('type.wordSem.word_loc'(sK0),bool),X115),s(bool,X116),s(bool,X117)))))) ) | $spl4),
  inference(forward_demodulation,[],[f64439,f1485])).
fof(f64439,plain,(
  ( ! [X116,X114,X117,X115,X113] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X114),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X113),s(fun('type.wordSem.word_loc'(sK0),bool),X114),s(fun('type.wordSem.word_loc'(sK0),bool),X115),s(bool,X116),s(bool,X117)))))) ) | $spl4),
  inference(superposition,[],[f30531,f2634])).
fof(f64781,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X107),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X108),s('type.wordSem.word_loc'(sK0),sK1)))) ) | $spl2),
  inference(forward_demodulation,[],[f64780,f305])).
fof(f64780,plain,(
  ( ! [X111,X109,X107,X110,X108] : (s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X108),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X107),s(fun('type.wordSem.word_loc'(sK0),bool),X108),s(fun('type.wordSem.word_loc'(sK0),bool),X109),s(bool,X110),s(bool,X111)))))) ) | $spl2),
  inference(forward_demodulation,[],[f64438,f1467])).
fof(f64438,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,sK3(X112,'_i_0_0',bool)) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X108),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X107),s(fun('type.wordSem.word_loc'(sK0),bool),X108),s(fun('type.wordSem.word_loc'(sK0),bool),X109),s(bool,X110),s(bool,X111)))))) ) | $spl2),
  inference(superposition,[],[f30531,f1987])).
fof(f64773,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X54),s('type.wordSem.word_loc'(sK0),sK2))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X55),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f64772,f305])).
fof(f64772,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X55),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X54),s(fun('type.wordSem.word_loc'(sK0),bool),X55),s(fun('type.wordSem.word_loc'(sK0),bool),X56),s(bool,X57),s(bool,X58)))))) )),
  inference(forward_demodulation,[],[f64429,f243])).
fof(f64429,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55] : (s(bool,'const.bool.DATATYPE_1'(s(X59,X60))) = s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X55),s('type.wordSem.word_loc'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Result_1'(s('type.wordSem.word_loc'(sK0),sK2))),s(fun('type.wordSem.word_loc'(sK0),bool),X54),s(fun('type.wordSem.word_loc'(sK0),bool),X55),s(fun('type.wordSem.word_loc'(sK0),bool),X56),s(bool,X57),s(bool,X58)))))) )),
  inference(superposition,[],[f30531,f787])).
fof(f50282,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) | $spl32),
  inference(cnf_transformation,[],[f50282_D])).
fof(f50282_D,plain,(
  ( ! [X4,X2,X0,X5,X3] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3591,plain,(
  ( ! [X27] : (~p(s(bool,X27))) ) | $spl10),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  ( ! [X27] : (~p(s(bool,X27))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f50281,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5)))) )),
  inference(subsumption_resolution,[],[f50193,f238])).
fof(f50193,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | s(X2,happ(s(fun(X0,X2),X3),s(X0,X4))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(fun(X0,X2),fun(X0,X2)),X5),s(fun(X0,X2),'const.relation.RESTRICT_3'(s(fun(X0,X2),X3),s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(X0,X4))))),s(X0,X4))) | s(fun(X0,X2),X3) != s(fun(X0,X2),'const.relation.WFREC_2'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'),s(fun(fun(X0,X2),fun(X0,X2)),X5))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f298,f769])).
fof(f769,plain,(
  ( ! [X10,X9] : (s(bool,X10) = s(bool,'const.relation.WF_1'(s(fun(X9,fun(X9,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f759,f229])).
fof(f759,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | s(bool,X10) = s(bool,'const.relation.WF_1'(s(fun(X9,fun(X9,bool)),'const.relation.EMPTY_REL_0'))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f586,f340])).
fof(f340,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0')))) )),
  inference(unit_resulting_resolution,[],[f232,f235])).
fof(f232,plain,(
  ( ! [X0] : (p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.relation.EMPTY_REL_0'))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X17] : p(s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),'const.relation.EMPTY_REL_0'))))),
  file('thm.stackSem.result_distinct|split|0','thm.relation.WF_EMPTY_REL')).
fof(f298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))) | s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | ~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) != s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) | (~p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) | ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(ennf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (s(fun(X1,X0),X4) = s(fun(X1,X0),'const.relation.WFREC_2'(s(fun(X1,fun(X1,bool)),X3),s(fun(fun(X1,X0),fun(X1,X0)),X2))) => (p(s(bool,'const.relation.WF_1'(s(fun(X1,fun(X1,bool)),X3)))) => ! [X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,X0),fun(X1,X0)),X2),s(fun(X1,X0),'const.relation.RESTRICT_3'(s(fun(X1,X0),X4),s(fun(X1,fun(X1,bool)),X3),s(X1,X5))))),s(X1,X5)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X23,X17] : ! [X60,X61,X18] : (s(fun(X17,X23),X18) = s(fun(X17,X23),'const.relation.WFREC_2'(s(fun(X17,fun(X17,bool)),X61),s(fun(fun(X17,X23),fun(X17,X23)),X60))) => (p(s(bool,'const.relation.WF_1'(s(fun(X17,fun(X17,bool)),X61)))) => ! [X30] : s(X23,happ(s(fun(X17,X23),X18),s(X17,X30))) = s(X23,happ(s(fun(X17,X23),happ(s(fun(fun(X17,X23),fun(X17,X23)),X60),s(fun(X17,X23),'const.relation.RESTRICT_3'(s(fun(X17,X23),X18),s(fun(X17,fun(X17,bool)),X61),s(X17,X30))))),s(X17,X30)))))),
  file('thm.stackSem.result_distinct|split|0','thm.relation.WFREC_COROLLARY')).
fof(f14925,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) ) | $spl28),
  inference(cnf_transformation,[],[f14925_D])).
fof(f14925_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f14926,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f14926_D])).
fof(f14926_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f14928,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) )),
  inference(condensation,[],[f14927])).
fof(f14927,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X77)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X77)))) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f14854,f439])).
fof(f14854,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X77)))) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f288,f361])).
fof(f14924,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X1))))) )),
  inference(condensation,[],[f14923])).
fof(f14923,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X74)))) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f14853,f439])).
fof(f14853,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X74)))) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f288,f361])).
fof(f13944,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'const.combin.I_0')))) | $spl24),
  inference(cnf_transformation,[],[f13944_D])).
fof(f13944_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f13946,plain,(
  ~p(s(bool,sK6('const.combin.I_0',bool))) | $spl27),
  inference(cnf_transformation,[],[f13946_D])).
fof(f13946_D,plain,(
  ~p(s(bool,sK6('const.combin.I_0',bool))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f13885,plain,(
  ~p(s(bool,sK6('const.combin.I_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'const.combin.I_0'))))),
  inference(superposition,[],[f282,f318])).
fof(f13920,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl20),
  inference(cnf_transformation,[],[f13920_D])).
fof(f13920_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f13942,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | $spl23),
  inference(cnf_transformation,[],[f13942_D])).
fof(f13942_D,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f13884,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f282,f244])).
fof(f13918,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl18),
  inference(cnf_transformation,[],[f13918_D])).
fof(f13918_D,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f13916,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK6('_i_0_0',bool))),
  inference(subsumption_resolution,[],[f13859,f229])).
fof(f13859,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK6('_i_0_0',bool))),
  inference(superposition,[],[f282,f409])).
fof(f409,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f244,f233])).
fof(f9076,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl16),
  inference(cnf_transformation,[],[f9076_D])).
fof(f9076_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f9074,plain,(
  ( ! [X72] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X72))) )),
  inference(forward_demodulation,[],[f9040,f318])).
fof(f9040,plain,(
  ( ! [X72] : (~p(s(bool,X72)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f260,f318])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X17] : ! [X28,X30] : (p(s(bool,happ(s(fun(X17,bool),X28),s(X17,X30)))) => p(s(bool,happ(s(fun(X17,bool),X28),s(X17,'const.min.@_1'(s(fun(X17,bool),X28)))))))),
  file('thm.stackSem.result_distinct|split|0','thm.bool.SELECT_AX')).
fof(f9073,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl14),
  inference(cnf_transformation,[],[f9073_D])).
fof(f9073_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f9071,plain,(
  ( ! [X71] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X71))) )),
  inference(forward_demodulation,[],[f9039,f244])).
fof(f9039,plain,(
  ( ! [X71] : (~p(s(bool,X71)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f260,f244])).
fof(f3782,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3595,f3591])).
fof(f3595,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f289])).
fof(f3781,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3601,f1295])).
fof(f3601,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1295,f3591,f439])).
fof(f3780,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3602,f1468])).
fof(f3602,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK3(X0,'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f3591,f439])).
fof(f3779,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3605,f3744])).
fof(f3744,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f3705,f231])).
fof(f3705,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2)) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f307])).
fof(f3644,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f1761])).
fof(f3605,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1295,f3591,f439])).
fof(f3778,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3607,f3744])).
fof(f3607,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,sK3(X1,'_i_0_0',bool))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f3591,f439])).
fof(f3777,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3610,f3744])).
fof(f3610,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f3591,f277])).
fof(f3776,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3617,f231])).
fof(f3617,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f331,f3591,f277])).
fof(f3775,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3619,f3591])).
fof(f3619,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1295,f3591,f277])).
fof(f3774,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3621,f3591])).
fof(f3621,plain,(
  ( ! [X0] : (p(s(bool,sK3(X0,'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f3591,f277])).
fof(f3773,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3624,f3744])).
fof(f3624,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f3591,f277])).
fof(f3772,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3631,f231])).
fof(f3631,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f331,f3591,f277])).
fof(f3771,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3633,f3591])).
fof(f3633,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1295,f3591,f277])).
fof(f3770,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3635,f3591])).
fof(f3635,plain,(
  ( ! [X0] : (p(s(bool,sK3(X0,'_i_0_0',bool)))) ) | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f3591,f277])).
fof(f3769,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3642,f3744])).
fof(f3642,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,sK3(X1,X2,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f261])).
fof(f3768,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3738,f231])).
fof(f3738,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f3644,f321])).
fof(f321,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),sK5(X0)))))) )),
  inference(forward_demodulation,[],[f320,f273])).
fof(f273,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X17] : ! [X20,X21,X22] : s('type.ind_type.recspace'(X17),'const.ind_type.CONSTR_3'(s('type.num.num',X20),s(X17,X21),s(fun('type.num.num','type.ind_type.recspace'(X17)),X22))) = s('type.ind_type.recspace'(X17),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X17)),'type.ind_type.recspace'(X17)),happ(s(fun(X17,fun(fun('type.num.num','type.ind_type.recspace'(X17)),'type.ind_type.recspace'(X17))),'const.ind_type.CONSTR_1'(s('type.num.num',X20))),s(X17,X21))),s(fun('type.num.num','type.ind_type.recspace'(X17)),X22)))),
  file('thm.stackSem.result_distinct|split|0','ARITY_1_3_const.ind_type.CONSTR')).
fof(f320,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),sK5(X0)))))) )),
  inference(forward_demodulation,[],[f276,f273])).
fof(f276,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),sK5(X0)))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),sK5(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f137])).
fof(f137,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.wordSem.word_loc'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),happ(s(fun('type.wordSem.word_loc'(X0),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'type.ind_type.recspace'('type.wordSem.word_loc'(X0)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.wordSem.word_loc'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X0)),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X0),'type.ind_type.recspace'('type.wordSem.word_loc'(X0))),X1))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ? [X36] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.wordSem.word_loc'(X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X6),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.wordSem.word_loc'(X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X6),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1_2'(s(fun(fun('type.wordSem.word_loc'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.wordSem.word_loc'(X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_i_0_0'),s(fun('type.wordSem.word_loc'(X6),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x1x0_2'(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),happ(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.wordSem.word_loc'(X6),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'=_0'),s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),happ(s(fun('type.wordSem.word_loc'(X6),fun(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'type.ind_type.recspace'('type.wordSem.word_loc'(X6)))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.wordSem.word_loc'(X6),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_dst_x0_1'(s('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),fun(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.wordSem.word_loc'(X6)),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),'_i_0_0'))))))),s(fun('type.stackSem.result'(X6),'type.ind_type.recspace'('type.wordSem.word_loc'(X6))),X36))))),
  file('thm.stackSem.result_distinct|split|0','thm.stackSem.result_TY_DEF')).
fof(f3767,plain,(
  $false | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3766,f3744])).
fof(f3766,plain,(
  ( ! [X112,X115,X113] : (s(X112,X115) = s(X112,sK3(X115,X113,X112))) ) | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f3734,f1295])).
fof(f3734,plain,(
  ( ! [X112,X115,X113] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(X112,X115) = s(X112,sK3(X115,X113,X112))) ) | ($spl2 | $spl10)),
  inference(backward_demodulation,[],[f3644,f3434])).
fof(f3434,plain,(
  ( ! [X112,X115,X113] : (s(X112,X115) = s(X112,sK3(X115,X113,X112)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(X112,bool),X113)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f3433,f231])).
fof(f3433,plain,(
  ( ! [X112,X115,X113] : (p(s(bool,'F_0')) | s(X112,X115) = s(X112,sK3(X115,X113,X112)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(X112,bool),X113)))))) ) | $spl2),
  inference(forward_demodulation,[],[f3202,f1467])).
fof(f3202,plain,(
  ( ! [X114,X112,X115,X113] : (p(s(bool,sK3(X114,'_i_0_0',bool))) | s(X112,X115) = s(X112,sK3(X115,X113,X112)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'|dollar|exists_1'(s(fun(X112,bool),X113)))))) ) | $spl2),
  inference(superposition,[],[f261,f1987])).
fof(f3765,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3764,f3744])).
fof(f3764,plain,(
  ( ! [X21,X19,X20,X18] : (s(X18,X21) = s(X18,sK3(X21,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X18,bool),X19),s(fun(X18,bool),X20)),X18))) ) | $spl10),
  inference(subsumption_resolution,[],[f3733,f231])).
fof(f3733,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'F_0')) | s(X18,X21) = s(X18,sK3(X21,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X18,bool),X19),s(fun(X18,bool),X20)),X18))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f749])).
fof(f749,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X18,bool),X19)))) | s(X18,X21) = s(X18,sK3(X21,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X18,bool),X19),s(fun(X18,bool),X20)),X18))) )),
  inference(superposition,[],[f261,f255])).
fof(f3763,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3762,f3744])).
fof(f3762,plain,(
  ( ! [X17,X15,X16] : (s(X15,X17) = s(X15,sK3(X17,'const.combin.I_1'(s(fun(X15,bool),X16)),X15))) ) | $spl10),
  inference(subsumption_resolution,[],[f3732,f231])).
fof(f3732,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'F_0')) | s(X15,X17) = s(X15,sK3(X17,'const.combin.I_1'(s(fun(X15,bool),X16)),X15))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f748])).
fof(f748,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16)))) | s(X15,X17) = s(X15,sK3(X17,'const.combin.I_1'(s(fun(X15,bool),X16)),X15))) )),
  inference(superposition,[],[f261,f242])).
fof(f3761,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3760,f3744])).
fof(f3760,plain,(
  ( ! [X14,X12,X13,X11] : (s(X11,X14) = s(X11,sK3(X14,'const.ind_type.FCONS_3'(s(fun(X11,bool),X12),s(fun('type.num.num',fun(X11,bool)),X13),s('type.num.num','const.num.|48|_0')),X11))) ) | $spl10),
  inference(subsumption_resolution,[],[f3731,f231])).
fof(f3731,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'F_0')) | s(X11,X14) = s(X11,sK3(X14,'const.ind_type.FCONS_3'(s(fun(X11,bool),X12),s(fun('type.num.num',fun(X11,bool)),X13),s('type.num.num','const.num.|48|_0')),X11))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f747])).
fof(f747,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'|dollar|exists_1'(s(fun(X11,bool),X12)))) | s(X11,X14) = s(X11,sK3(X14,'const.ind_type.FCONS_3'(s(fun(X11,bool),X12),s(fun('type.num.num',fun(X11,bool)),X13),s('type.num.num','const.num.|48|_0')),X11))) )),
  inference(superposition,[],[f261,f256])).
fof(f3759,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3758,f3744])).
fof(f3758,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X6,X10) = s(X6,sK3(X10,happ(s(fun(X7,fun(X6,bool)),'_dst_x0_1'(s(fun(X6,bool),X8))),s(X7,X9)),X6))) ) | $spl10),
  inference(subsumption_resolution,[],[f3730,f231])).
fof(f3730,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'F_0')) | s(X6,X10) = s(X6,sK3(X10,happ(s(fun(X7,fun(X6,bool)),'_dst_x0_1'(s(fun(X6,bool),X8))),s(X7,X9)),X6))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f746])).
fof(f746,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X8)))) | s(X6,X10) = s(X6,sK3(X10,happ(s(fun(X7,fun(X6,bool)),'_dst_x0_1'(s(fun(X6,bool),X8))),s(X7,X9)),X6))) )),
  inference(superposition,[],[f261,f292])).
fof(f3757,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3756,f3744])).
fof(f3756,plain,(
  ( ! [X4,X5,X3] : (s(X3,X5) = s(X3,sK3(X5,happ(s(fun(fun(X3,bool),fun(X3,bool)),'const.combin.I_0'),s(fun(X3,bool),X4)),X3))) ) | $spl10),
  inference(subsumption_resolution,[],[f3729,f231])).
fof(f3729,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | s(X3,X5) = s(X3,sK3(X5,happ(s(fun(fun(X3,bool),fun(X3,bool)),'const.combin.I_0'),s(fun(X3,bool),X4)),X3))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f745])).
fof(f745,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) | s(X3,X5) = s(X3,sK3(X5,happ(s(fun(fun(X3,bool),fun(X3,bool)),'const.combin.I_0'),s(fun(X3,bool),X4)),X3))) )),
  inference(superposition,[],[f261,f318])).
fof(f3755,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3754,f3744])).
fof(f3754,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK3(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f3728,f231])).
fof(f3728,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK3(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f744])).
fof(f744,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK3(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f261,f244])).
fof(f3753,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3752,f3744])).
fof(f3752,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK3(X2,X1,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f3727,f231])).
fof(f3727,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK3(X2,X1,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f261])).
fof(f3751,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3722,f231])).
fof(f3722,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f3644,f2904])).
fof(f2904,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f229,f290])).
fof(f3750,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3721,f231])).
fof(f3721,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f3644,f2332])).
fof(f2332,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f229,f287])).
fof(f3749,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3720,f231])).
fof(f3720,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f3644,f2411])).
fof(f2411,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f231,f289])).
fof(f3748,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3719,f231])).
fof(f3719,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(backward_demodulation,[],[f3644,f2199])).
fof(f2199,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f229,f286])).
fof(f3747,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3741,f231])).
fof(f3741,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(condensation,[],[f3717])).
fof(f3717,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f289])).
fof(f3746,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3745,f3744])).
fof(f3745,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,X13)) ) | $spl10),
  inference(subsumption_resolution,[],[f3715,f231])).
fof(f3715,plain,(
  ( ! [X13] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X13)) ) | $spl10),
  inference(backward_demodulation,[],[f3644,f2469])).
fof(f2469,plain,(
  ( ! [X14,X13] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,X14)))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f2416,f231])).
fof(f2416,plain,(
  ( ! [X14,X13] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,X14)))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f289,f233])).
fof(f3688,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3654,f3591])).
fof(f3654,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f289])).
fof(f3655,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f231,f3591,f289])).
fof(f3653,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f229,f3591,f290])).
fof(f3651,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f229,f3591,f286])).
fof(f3649,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f229,f3591,f287])).
fof(f3647,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f229,f3591,f1442])).
fof(f3645,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f331,f3591,f1761])).
fof(f3643,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f3591,f261])).
fof(f3640,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f243,f3591,f234])).
fof(f3638,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f331,f3591,f239])).
fof(f3636,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f3591,f1468,f3591,f277])).
fof(f3634,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f1295,f3591,f277])).
fof(f3632,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f231,f331,f3591,f277])).
fof(f3630,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f3591,f1468,f3591,f277])).
fof(f3628,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f1295,f3591,f277])).
fof(f3626,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f331,f3591,f277])).
fof(f3622,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f3591,f1468,f3591,f277])).
fof(f3620,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f1295,f3591,f277])).
fof(f3618,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f231,f331,f3591,f277])).
fof(f3616,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f3591,f1468,f3591,f277])).
fof(f3614,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f1295,f3591,f277])).
fof(f3612,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f331,f3591,f277])).
fof(f3608,plain,(
  $false | ($spl2 | $spl10)),
  inference(unit_resulting_resolution,[],[f1468,f1468,f3591,f439])).
fof(f3606,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1295,f1295,f3591,f439])).
fof(f3604,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f331,f331,f3591,f439])).
fof(f3596,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3591,f3591,f289])).
fof(f3594,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f229,f3591])).
fof(f3593,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'const.combin.I_0')))) | $spl12),
  inference(cnf_transformation,[],[f3593_D])).
fof(f3593_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3574,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f284,f318])).
fof(f3590,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl8),
  inference(cnf_transformation,[],[f3590_D])).
fof(f3590_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3573,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f284,f244])).
fof(f1459,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl4),
  inference(cnf_transformation,[],[f1459_D])).
fof(f1459_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1460,plain,(
  ( ! [X14] : (~p(s(bool,sK3(X14,'const.combin.I_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1460_D])).
fof(f1460_D,plain,(
  ( ! [X14] : (~p(s(bool,sK3(X14,'const.combin.I_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1441,plain,(
  ( ! [X14] : (~p(s(bool,sK3(X14,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f262,f318])).
fof(f1456,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl0),
  inference(cnf_transformation,[],[f1456_D])).
fof(f1456_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1457,plain,(
  ( ! [X13] : (~p(s(bool,sK3(X13,'_i_0_0',bool)))) ) | $spl2),
  inference(cnf_transformation,[],[f1457_D])).
fof(f1457_D,plain,(
  ( ! [X13] : (~p(s(bool,sK3(X13,'_i_0_0',bool)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1440,plain,(
  ( ! [X13] : (~p(s(bool,sK3(X13,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f262,f244])).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 233
% Passive clauses: 2802
% Generated clauses: 64643
% Final active clauses: 187
% Final passive clauses: 2395
% Input formulas: 69
% Initial clauses: 89
% 
% Duplicate literals: 1158
% Trivial inequalities: 422
% Fw subsumption resolutions: 4497
% Fw demodulations: 16493
% Bw demodulations: 207
% Condensations: 1730
% Global subsumptions: 1
% 
% Simple tautologies: 302
% Equational tautologies: 186
% Forward subsumptions: 36795
% Fw demodulations to eq. taut.: 138
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 1546
% Factoring: 1
% Forward superposition: 19364
% Backward superposition: 18522
% Self superposition: 361
% Equality factoring: 119
% Equality resolution: 19
% 
% Split clauses: 56
% Split components: 112
% SAT solver clauses: 2980
% SAT solver unit clauses: 514
% SAT solver binary clauses: 1662
% 
% Sat splits: 56
% Sat splitting refutations: 63
% 
% Memory used [KB]: 75862
% Time elapsed: 2.703 s
% ------------------------------
% ------------------------------
% Success in time 4.03 s
