Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (101170)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1096
% Passive clauses: 6722
% Generated clauses: 42800
% Final active clauses: 1029
% Final passive clauses: 5258
% Input formulas: 516
% Initial clauses: 875
% 
% Duplicate literals: 159
% Trivial inequalities: 61
% Fw subsumption resolutions: 1861
% Fw demodulations: 3322
% Bw demodulations: 66
% 
% Simple tautologies: 300
% Equational tautologies: 26
% Forward subsumptions: 30199
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 5524
% Factoring: 495
% Forward superposition: 17505
% Backward superposition: 12693
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 389
% Split components: 780
% SAT solver clauses: 520
% SAT solver unit clauses: 22
% SAT solver binary clauses: 99
% 
% Sat splits: 389
% 
% TWLsolver clauses: 520
% TWLsolver calls for satisfiability: 354
% 
% Memory used [KB]: 22899
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (101197)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 620
% Passive clauses: 3531
% Generated clauses: 26802
% Final active clauses: 581
% Final passive clauses: 2710
% Input formulas: 516
% Initial clauses: 875
% 
% Duplicate literals: 109
% Trivial inequalities: 29
% Fw subsumption resolutions: 915
% Fw demodulations: 2880
% Bw demodulations: 241
% Condensations: 90
% Global subsumptions: 62
% 
% Simple tautologies: 318
% Equational tautologies: 16
% Forward subsumptions: 14606
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 2466
% Factoring: 201
% Forward superposition: 9662
% Backward superposition: 5597
% Self superposition: 95
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 5333
% SAT solver unit clauses: 1301
% SAT solver binary clauses: 279
% SAT solver learnt clauses: 124
% SAT solver learnt literals: 16
% 
% Sat splits: 13
% 
% TWLsolver clauses: 5085
% TWLsolver calls for satisfiability: 5364
% 
% Memory used [KB]: 24306
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% (101224)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1591
% Passive clauses: 13524
% Generated clauses: 341013
% Final active clauses: 1472
% Final passive clauses: 3407
% Input formulas: 516
% Initial clauses: 875
% Discarded non-redundant clauses: 187368
% 
% Duplicate literals: 9
% Trivial inequalities: 89
% Fw subsumption resolutions: 1502
% Fw demodulations: 43257
% Bw demodulations: 1648
% Condensations: 68
% 
% Simple tautologies: 165
% Equational tautologies: 63
% Forward subsumptions: 96598
% Fw demodulations to eq. taut.: 18352
% Bw demodulations to eq. taut.: 50
% 
% Unit resulting resolution: 207765
% Factoring: 1
% Forward superposition: 44038
% Backward superposition: 40484
% Self superposition: 223
% Equality factoring: 2
% Equality resolution: 37
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 14675
% SAT solver unit clauses: 5396
% SAT solver binary clauses: 6915
% 
% Sat splits: 18
% Sat splitting refutations: 1113
% 
% Memory used [KB]: 191126
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f8251,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1010,f1435,f1434,f1454,f1467,f1455,f1456,f1457,f1468,f1458,f1469,f1470,f1737,f1472,f1742,f1740,f1473,f1739,f1741,f1474,f1744,f1476,f1747,f1478,f1749,f1479,f1751,f1481,f1752,f1482,f1483,f1484,f1485,f1486,f1488,f1755,f1489,f1589,f1753,f1600,f1601,f1764,f1602,f1603,f1765,f1604,f1606,f1607,f1767,f1763,f1608,f1609,f1769,f1762,f1626,f1627,f1773,f1771,f1628,f1629,f1776,f1781,f1782,f1783,f1784,f1785,f1786,f1788,f1888,f2151,f2064,f2233,f2234,f2235,f2252,f2253,f2254,f2255,f2257,f2258,f2259,f2260,f2261,f2263,f2264,f2364,f2393,f2394,f2395,f2396,f2397,f2398,f2400,f2500,f2539,f2540,f2541,f2542,f2543,f2544,f2546,f2646,f1475,f2686,f1471,f2688,f2741,f2752,f2742,f2743,f2744,f2753,f2745,f2754,f1466,f2788,f2789,f2790,f2791,f2792,f2676,f2793,f2795,f2895,f2675,f2942,f2943,f2981,f2964,f2965,f2983,f2998,f2999,f3000,f3001,f3002,f3251,f3003,f3252,f3005,f3105,f3250,f3106,f3253,f3110,f3111,f3112,f3113,f3114,f3115,f3116,f3117,f3118,f3125,f3126,f3263,f3128,f3266,f3129,f3267,f3130,f3268,f3131,f3269,f3132,f3133,f3270,f3136,f3137,f3138,f3272,f3139,f3271,f3140,f3141,f3142,f3274,f3143,f3273,f3144,f3145,f3146,f3275,f3147,f3276,f3150,f3265,f3278,f3246,f3245,f3260,f3293,f3294,f3295,f3296,f3297,f3299,f3300,f3400,f3432,f3433,f3434,f3435,f3436,f3438,f3439,f3539,f3617,f3618,f3619,f3620,f3621,f3622,f3624,f3724,f3763,f3798,f3799,f3800,f3801,f3802,f3804,f3577,f3805,f3905,f3576,f4040,f4039,f4038,f4043,f4079,f4080,f2985,f4184,f4183,f4187,f4247,f4248,f4249,f4250,f4251,f4253,f4528,f4254,f4354,f4360,f4361,f4532,f4362,f4535,f4363,f4536,f4364,f4537,f4365,f4538,f4366,f4367,f4539,f4368,f4369,f4540,f4541,f4371,f4372,f4373,f4374,f4375,f4376,f4377,f4543,f4378,f4542,f4379,f4380,f4381,f4545,f4382,f4544,f4383,f4384,f4385,f4546,f4386,f4547,f4387,f4391,f4524,f4523,f4522,f4534,f4569,f4570,f4571,f4572,f4573,f1756,f1914,f4574,f4576,f4676,f1913,f4870,f4872,f4885,f4919,f4943,f4997,f4999,f5000,f5002,f5004,f5104,f5158,f5160,f5161,f5162,f5165,f5265,f2387,f5685,f5691,f5693,f5694,f5696,f5698,f5798,f5129,f5831,f6138,f6139,f6140,f6141,f6142,f6143,f6144,f6145,f6146,f6147,f6148,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6164,f6167,f6170,f6173,f6208,f6209,f6210,f6211,f6212,f6215,f6216,f6217,f6218,f6219,f6220,f6221,f6222,f6223,f6224,f6225,f6226,f6227,f6228,f6229,f6230,f6231,f6232,f6233,f6234,f6235,f6236,f6237,f6238,f6239,f6240,f6241,f6243,f6245,f6246,f6247,f6248,f6249,f6250,f6251,f6252,f6253,f6254,f6255,f6256,f6257,f6258,f6259,f6260,f6261,f6262,f6263,f6264,f6265,f6266,f6267,f6268,f6269,f6270,f6271,f6272,f6273,f6274,f6276,f6278,f6279,f6280,f6281,f6282,f6283,f6284,f6285,f6286,f6287,f6288,f6289,f6290,f6291,f6292,f6293,f6294,f6295,f6296,f6297,f6298,f6299,f6301,f6302,f6304,f6305,f6306,f6307,f6308,f6309,f6310,f6311,f6312,f6313,f6314,f6315,f6316,f6317,f6318,f6319,f6320,f6321,f6322,f6323,f6324,f6325,f6326,f6327,f6328,f6329,f6330,f6640,f6642,f6646,f6647,f6339,f6340,f6341,f6342,f6343,f6344,f6345,f6346,f6392,f6393,f6394,f6395,f6729,f6751,f6752,f6770,f6753,f6754,f6771,f6755,f6757,f6758,f6773,f6769,f6759,f6760,f6775,f6768,f6338,f6648,f6334,f6333,f6840,f6501,f6841,f6504,f6842,f6505,f6843,f6506,f6844,f6507,f6845,f6508,f6846,f6509,f6847,f6510,f6848,f6511,f6849,f6512,f6850,f6513,f6851,f6514,f6852,f6515,f6853,f6516,f6854,f6517,f6855,f6518,f6856,f6519,f6857,f6520,f6858,f6521,f6859,f6522,f6860,f6523,f6861,f6524,f6862,f6525,f6863,f6526,f6866,f6527,f6869,f6528,f6872,f6529,f6875,f6530,f6910,f6911,f6912,f6913,f6914,f6555,f6917,f6556,f6918,f6557,f6919,f6558,f6920,f6559,f6921,f6560,f6922,f6561,f6923,f6562,f6924,f6563,f6925,f6564,f6926,f6565,f6927,f6566,f6928,f6567,f6929,f6568,f6930,f6569,f6931,f6570,f6932,f6571,f6933,f6572,f6934,f6573,f6935,f6574,f6936,f6552,f6937,f6551,f6938,f6553,f6939,f6554,f6940,f6941,f6942,f6943,f6945,f6579,f6947,f6580,f6948,f6581,f6949,f6582,f6950,f6583,f6951,f6584,f6952,f6585,f6953,f6586,f6954,f6587,f6955,f6588,f6956,f6589,f6957,f6590,f6958,f6591,f6959,f6592,f6960,f6593,f6961,f6594,f6962,f6595,f6963,f6596,f6964,f6597,f6965,f6598,f6966,f6575,f6967,f6576,f6968,f6577,f6969,f6578,f6970,f6971,f6972,f6973,f6974,f6603,f6975,f6604,f6976,f6605,f6978,f6606,f6980,f6607,f6981,f6608,f6982,f6609,f6983,f6610,f6984,f6611,f6985,f6612,f6986,f6613,f6987,f6614,f6988,f6615,f6989,f6616,f6990,f6617,f6991,f6618,f6992,f6619,f6993,f6599,f6994,f6600,f6995,f6601,f6996,f6602,f6997,f6620,f6998,f6621,f6999,f6622,f7000,f7001,f7003,f6625,f7004,f6626,f7006,f7007,f7008,f7009,f7010,f6627,f7011,f6628,f7012,f6629,f7013,f6630,f7014,f6624,f7015,f6623,f7016,f7017,f7018,f7019,f7020,f7021,f7022,f7023,f7024,f7025,f6638,f7026,f6639,f7027,f6636,f7028,f6637,f7029,f6635,f7030,f6634,f7031,f6633,f7032,f6632,f7223,f7224,f7035,f7036,f7225,f6641,f7226,f6643,f7227,f6644,f7040,f6645,f7041,f6649,f7042,f6650,f7043,f6651,f7044,f6652,f7045,f6653,f7046,f6654,f7047,f6655,f7048,f6656,f7094,f6685,f7095,f6502,f6686,f7096,f6687,f7097,f6631,f7371,f7373,f7374,f7375,f7378,f7478,f4527,f7623,f7662,f7901,f7824,f1777,f7823,f7912,f1774,f7784,f7684,f7681,f7680,f7679,f7677,f7913,f7914,f7916,f7917,f1754,f7918,f1758,f7921,f1757,f8021,f8196,f1759,f8166,f6503,f8207,f8250])).
fof(f8250,plain,(
  $false | $spl0),
  inference(trivial_inequality_removal,[],[f8236])).
fof(f8236,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),sK16) | $spl0),
  inference(superposition,[],[f1434,f1477])).
fof(f1477,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK31(sK25,sK24,sK23,sK22,sK20,sK19,sK18,sK17,sK16,'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)),sK13))))),
  inference(resolution,[],[f1009,f1080])).
fof(f1080,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK31(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))))) )),
  inference(cnf_transformation,[],[f839])).
fof(f839,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK31(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f751])).
fof(f751,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | ? [X11] : s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X11))))),
  inference(ennf_transformation,[],[f555])).
fof(f555,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) => ? [X11] : s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X11))))),
  inference(flattening,[],[f554])).
fof(f554,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) => ? [X11] : s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X11))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X8,X9,X10,X11,X12,X13,X15,X16,X17,X18] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))),s('type.wordSem.word_loc'(X6),X9))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))) => ? [X7] : s('type.wordSem.word_loc'(X6),X9) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X7))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_Number_IMP_Word')).
fof(f1009,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))))))),
  inference(cnf_transformation,[],[f810])).
fof(f810,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) & ! [X13,X14] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X13))) | s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17,sK18,sK19,sK20,sK21,sK22,sK23,sK24,sK25])],[f729])).
fof(f729,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X3))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X2)))))))) & ! [X13,X14] : (s('type.wordSem.word_loc'(X0),X3) != s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) | s('type.wordSem.word_loc'(X0),X1) != s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X3))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X2)))))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X3) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14)))))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X3))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X2)))))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X3) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))),s('type.wordSem.word_loc'(X6),X9))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X14))),s('type.wordSem.word_loc'(X6),X7))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))))) => ? [X19,X20] : (s('type.wordSem.word_loc'(X6),X9) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X19))) & s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X20)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))),s('type.wordSem.word_loc'(X6),X9))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X14))),s('type.wordSem.word_loc'(X6),X7))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))))) => ? [X19,X20] : (s('type.wordSem.word_loc'(X6),X9) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X19))) & s('type.wordSem.word_loc'(X6),X7) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X20)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2')).
fof(f8207,plain,(
  sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),sK25) | $spl452),
  inference(cnf_transformation,[],[f8207_D])).
fof(f8207_D,plain,(
  sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),sK25) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f6503,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | $spl170),
  inference(cnf_transformation,[],[f6503_D])).
fof(f6503_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f8166,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),sK25) | ($spl20 | $spl42)),
  inference(superposition,[],[f5238,f1747])).
fof(f5238,plain,(
  ( ! [X163] : (p(s(bool,X163)) | sP11(X163,sK25)) ) | $spl42),
  inference(resolution,[],[f1763,f1173])).
fof(f1173,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f880])).
fof(f880,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1)))),
  inference(rectify,[],[f879])).
fof(f879,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP11(X1,X0)))),
  inference(flattening,[],[f878])).
fof(f878,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1759,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl40),
  inference(cnf_transformation,[],[f1759_D])).
fof(f1759_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8196,plain,(
  ( ! [X2,X0] : (sP11(X0,sK25) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl450),
  inference(cnf_transformation,[],[f8196_D])).
fof(f8196_D,plain,(
  ( ! [X2,X0] : (sP11(X0,sK25) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f8021,plain,(
  ( ! [X298,X297,X299] : (sP11(X297,sK25) | p(s(bool,X298)) | p(s(bool,X299)) | s(bool,X297) != s(bool,X299)) ) | $spl42),
  inference(resolution,[],[f5238,f1424])).
fof(f1424,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f1015])).
fof(f1015,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f812,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f811])).
fof(f811,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f731])).
fof(f731,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f730])).
fof(f730,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X30,X40,X31,X41] : ((s(bool,X30) = s(bool,X40) & (p(s(bool,X40)) => s(bool,X31) = s(bool,X41))) => ((p(s(bool,X30)) => p(s(bool,X31))) <=> (p(s(bool,X40)) => p(s(bool,X41)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CONG')).
fof(f1757,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) | $spl36),
  inference(cnf_transformation,[],[f1757_D])).
fof(f1757_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f7921,plain,(
  ( ! [X30,X28,X31,X29] : (sP11(X28,sK25) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl42),
  inference(resolution,[],[f5238,f1018])).
fof(f1018,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f1758,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | $spl38),
  inference(cnf_transformation,[],[f1758_D])).
fof(f1758_D,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7918,plain,(
  ( ! [X19,X17,X18,X16] : (sP11(X16,sK25) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl42),
  inference(resolution,[],[f5238,f1017])).
fof(f1017,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f1754,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f7917,plain,(
  ( ! [X14,X12,X15,X13] : (sP11(X12,sK25) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f5238,f1016])).
fof(f1016,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f7916,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(X8,sK25) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl42),
  inference(resolution,[],[f5238,f1014])).
fof(f1014,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f7914,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,sK25) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl42),
  inference(resolution,[],[f5238,f1012])).
fof(f1012,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f812])).
fof(f7913,plain,(
  ( ! [X192] : (p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) ) | $spl448),
  inference(cnf_transformation,[],[f7913_D])).
fof(f7913_D,plain,(
  ( ! [X192] : (p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f7677,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) )),
  inference(resolution,[],[f1480,f1012])).
fof(f1480,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))))),
  inference(resolution,[],[f1009,f1347])).
fof(f1347,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),X2),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))))) )),
  inference(cnf_transformation,[],[f785])).
fof(f785,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),X2),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))),
  inference(ennf_transformation,[],[f666])).
fof(f666,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),X2),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))),
  inference(flattening,[],[f665])).
fof(f665,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),X2),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X8,X9,X10,X11,X12,X13,X16,X17,X18] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),X9),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_tail')).
fof(f7679,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f1480,f1014])).
fof(f7680,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) )),
  inference(resolution,[],[f1480,f1016])).
fof(f7681,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) )),
  inference(resolution,[],[f1480,f1017])).
fof(f7684,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f1480,f1018])).
fof(f7784,plain,(
  ( ! [X191,X192] : (p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) )),
  inference(resolution,[],[f1480,f1424])).
fof(f1774,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) | $spl56),
  inference(cnf_transformation,[],[f1774_D])).
fof(f1774_D,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7912,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) | $spl446),
  inference(cnf_transformation,[],[f7912_D])).
fof(f7912_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f7823,plain,(
  ( ! [X8] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) )),
  inference(superposition,[],[f1480,f1206])).
fof(f1206,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f895])).
fof(f895,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f894])).
fof(f894,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f197])).
fof(f197,axiom,(
  ! [X76,X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X76))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X13)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X76) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_LT_CALCULATE|split|2')).
fof(f1777,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) ) | $spl60),
  inference(cnf_transformation,[],[f1777_D])).
fof(f1777_D,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7824,plain,(
  ( ! [X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) )),
  inference(superposition,[],[f1480,f1207])).
fof(f1207,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f895])).
fof(f7901,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f7672,f1435])).
fof(f7672,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK31(sK25,sK24,sK23,sK21,sK20,sK19,sK18,sK17,sK14,sK15,sK13))))),
  inference(resolution,[],[f1480,f1080])).
fof(f7662,plain,(
  sP11(sK25,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl444),
  inference(cnf_transformation,[],[f7662_D])).
fof(f7662_D,plain,(
  sP11(sK25,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f7623,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(sK25,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ($spl20 | $spl42)),
  inference(superposition,[],[f5237,f1747])).
fof(f5237,plain,(
  ( ! [X162] : (p(s(bool,X162)) | sP11(sK25,X162)) ) | $spl42),
  inference(resolution,[],[f1763,f1170])).
fof(f1170,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f880])).
fof(f4527,plain,(
  ( ! [X2,X0] : (sP11(sK25,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl134),
  inference(cnf_transformation,[],[f4527_D])).
fof(f4527_D,plain,(
  ( ! [X2,X0] : (sP11(sK25,X0) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7478,plain,(
  ( ! [X298,X297,X299] : (sP11(sK25,X297) | p(s(bool,X298)) | p(s(bool,X299)) | s(bool,X297) != s(bool,X299)) ) | $spl42),
  inference(resolution,[],[f5237,f1424])).
fof(f7378,plain,(
  ( ! [X30,X28,X31,X29] : (sP11(sK25,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl42),
  inference(resolution,[],[f5237,f1018])).
fof(f7375,plain,(
  ( ! [X19,X17,X18,X16] : (sP11(sK25,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl42),
  inference(resolution,[],[f5237,f1017])).
fof(f7374,plain,(
  ( ! [X14,X12,X15,X13] : (sP11(sK25,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f5237,f1016])).
fof(f7373,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(sK25,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl42),
  inference(resolution,[],[f5237,f1014])).
fof(f7371,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(sK25,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl42),
  inference(resolution,[],[f5237,f1012])).
fof(f6631,plain,(
  ( ! [X699,X698] : (~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X698,X699)) ) | $spl386),
  inference(cnf_transformation,[],[f6631_D])).
fof(f6631_D,plain,(
  ( ! [X699,X698] : (~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X698,X699)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f7097,plain,(
  ( ! [X783,X782] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X782,X783)) ) | $spl20),
  inference(superposition,[],[f1428,f1747])).
fof(f1428,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1145])).
fof(f1145,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f868])).
fof(f868,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | (p(s(bool,X2)) & p(s(bool,X0))) | (p(s(bool,X1)) & p(s(bool,X0)))) & (((~p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f867])).
fof(f867,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP10(X0,X2,X1)))),
  inference(flattening,[],[f866])).
fof(f866,plain,(
  ! [X0,X2,X1] : ((~sP10(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0))))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP10(X0,X2,X1)))),
  inference(nnf_transformation,[],[f804])).
fof(f804,plain,(
  ! [X0,X2,X1] : (sP10(X0,X2,X1) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f6687,plain,(
  ( ! [X781,X780] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X780,X781) | ~sP8(X780,X781,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl432),
  inference(cnf_transformation,[],[f6687_D])).
fof(f6687_D,plain,(
  ( ! [X781,X780] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X780,X781) | ~sP8(X780,X781,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f7096,plain,(
  ( ! [X781,X780] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X780,X781) | ~sP8(X780,X781,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1427,f1747])).
fof(f1427,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1125])).
fof(f1125,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f859])).
fof(f859,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f858])).
fof(f858,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f857])).
fof(f857,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f6686,plain,(
  ( ! [X779,X778] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779)) ) | $spl430),
  inference(cnf_transformation,[],[f6686_D])).
fof(f6686_D,plain,(
  ( ! [X779,X778] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779)) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f6502,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | $spl171),
  inference(cnf_transformation,[],[f6502_D])).
fof(f6502_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f7095,plain,(
  ( ! [X779,X778] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779)) ) | $spl20),
  inference(superposition,[],[f1426,f1747])).
fof(f1426,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1105])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f851])).
fof(f851,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f850])).
fof(f850,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f849])).
fof(f849,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f799])).
fof(f799,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f6685,plain,(
  ( ! [X777,X776] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777)) ) | $spl428),
  inference(cnf_transformation,[],[f6685_D])).
fof(f6685_D,plain,(
  ( ! [X777,X776] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f7094,plain,(
  ( ! [X777,X776] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777)) ) | $spl20),
  inference(superposition,[],[f1425,f1747])).
fof(f1425,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1088])).
fof(f1088,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f843])).
fof(f843,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f842])).
fof(f842,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f841])).
fof(f841,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f796])).
fof(f796,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f6656,plain,(
  ( ! [X745,X747,X746] : (p(s(bool,happ(s(fun(X745,bool),X746),s(X745,sK43('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X746,X745))))) | ~p(s(bool,happ(s(fun(X745,bool),X746),s(X745,X747))))) ) | $spl426),
  inference(cnf_transformation,[],[f6656_D])).
fof(f6656_D,plain,(
  ( ! [X745,X747,X746] : (p(s(bool,happ(s(fun(X745,bool),X746),s(X745,sK43('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X746,X745))))) | ~p(s(bool,happ(s(fun(X745,bool),X746),s(X745,X747))))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f7048,plain,(
  ( ! [X745,X747,X746] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X745,bool),X746),s(X745,sK43('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X746,X745))))) | ~p(s(bool,happ(s(fun(X745,bool),X746),s(X745,X747))))) ) | $spl20),
  inference(superposition,[],[f1238,f1747])).
fof(f1238,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f920])).
fof(f920,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK42(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK42,sK43])],[f919])).
fof(f919,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ? [X5] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f918])).
fof(f918,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X2)) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(flattening,[],[f917])).
fof(f917,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X2)) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ? [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | (p(s(bool,X2)) & ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) & ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(flattening,[],[f609])).
fof(f609,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X2)) & ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X2)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))))),
  inference(rectify,[],[f316])).
fof(f316,axiom,(
  ! [X6] : ! [X38,X77] : ((p(s(bool,X77)) & ? [X30] : p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30))))) <=> ? [X30] : (p(s(bool,X77)) & p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30))))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.PULL_EXISTS|split|2')).
fof(f6655,plain,(
  ( ! [X743,X744,X742] : (p(s(bool,happ(s(fun(X742,bool),X743),s(X742,sK42('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X743,X742))))) | ~p(s(bool,happ(s(fun(X742,bool),X743),s(X742,X744))))) ) | $spl424),
  inference(cnf_transformation,[],[f6655_D])).
fof(f6655_D,plain,(
  ( ! [X743,X744,X742] : (p(s(bool,happ(s(fun(X742,bool),X743),s(X742,sK42('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X743,X742))))) | ~p(s(bool,happ(s(fun(X742,bool),X743),s(X742,X744))))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f7047,plain,(
  ( ! [X743,X744,X742] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X742,bool),X743),s(X742,sK42('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X743,X742))))) | ~p(s(bool,happ(s(fun(X742,bool),X743),s(X742,X744))))) ) | $spl20),
  inference(superposition,[],[f1236,f1747])).
fof(f1236,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK42(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f920])).
fof(f6654,plain,(
  ( ! [X741,X739,X740] : (p(s(bool,happ(s(fun(X739,bool),X740),s(X739,sK41('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X740,X739))))) | ~p(s(bool,happ(s(fun(X739,bool),X740),s(X739,X741))))) ) | $spl422),
  inference(cnf_transformation,[],[f6654_D])).
fof(f6654_D,plain,(
  ( ! [X741,X739,X740] : (p(s(bool,happ(s(fun(X739,bool),X740),s(X739,sK41('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X740,X739))))) | ~p(s(bool,happ(s(fun(X739,bool),X740),s(X739,X741))))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f7046,plain,(
  ( ! [X741,X739,X740] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X739,bool),X740),s(X739,sK41('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X740,X739))))) | ~p(s(bool,happ(s(fun(X739,bool),X740),s(X739,X741))))) ) | $spl20),
  inference(superposition,[],[f1233,f1747])).
fof(f1233,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f916,plain,(
  ! [X0,X1,X2] : ((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40,sK41])],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2] : ((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f914])).
fof(f914,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f913])).
fof(f913,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f608])).
fof(f608,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f607])).
fof(f607,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f315])).
fof(f315,axiom,(
  ! [X6] : ! [X38,X77] : ((? [X30] : p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30)))) & p(s(bool,X77))) <=> ? [X30] : (p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30)))) & p(s(bool,X77))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.PULL_EXISTS|split|1')).
fof(f6653,plain,(
  ( ! [X736,X738,X737] : (p(s(bool,happ(s(fun(X736,bool),X737),s(X736,sK40('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X737,X736))))) | ~p(s(bool,happ(s(fun(X736,bool),X737),s(X736,X738))))) ) | $spl420),
  inference(cnf_transformation,[],[f6653_D])).
fof(f6653_D,plain,(
  ( ! [X736,X738,X737] : (p(s(bool,happ(s(fun(X736,bool),X737),s(X736,sK40('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X737,X736))))) | ~p(s(bool,happ(s(fun(X736,bool),X737),s(X736,X738))))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f7045,plain,(
  ( ! [X736,X738,X737] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X736,bool),X737),s(X736,sK40('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X737,X736))))) | ~p(s(bool,happ(s(fun(X736,bool),X737),s(X736,X738))))) ) | $spl20),
  inference(superposition,[],[f1231,f1747])).
fof(f1231,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f916])).
fof(f6652,plain,(
  ( ! [X734,X733,X735] : (~p(s(bool,happ(s(fun(X733,bool),X734),s(X733,X735)))) | p(s(bool,happ(s(fun(X733,bool),X734),s(X733,sK39('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X734,X733)))))) ) | $spl418),
  inference(cnf_transformation,[],[f6652_D])).
fof(f6652_D,plain,(
  ( ! [X734,X733,X735] : (~p(s(bool,happ(s(fun(X733,bool),X734),s(X733,X735)))) | p(s(bool,happ(s(fun(X733,bool),X734),s(X733,sK39('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X734,X733)))))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f7044,plain,(
  ( ! [X734,X733,X735] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,happ(s(fun(X733,bool),X734),s(X733,X735)))) | p(s(bool,happ(s(fun(X733,bool),X734),s(X733,sK39('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X734,X733)))))) ) | $spl20),
  inference(superposition,[],[f1229,f1747])).
fof(f1229,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK39(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f912])).
fof(f912,plain,(
  ! [X0,X1,X2] : (((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK39(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK38,sK39])],[f911])).
fof(f911,plain,(
  ! [X0,X1,X2] : (((? [X6] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f910])).
fof(f910,plain,(
  ! [X0,X1,X2] : (((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f909])).
fof(f909,plain,(
  ! [X0,X1,X2] : (((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f772])).
fof(f772,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f606])).
fof(f606,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(rectify,[],[f314])).
fof(f314,axiom,(
  ! [X6] : ! [X38,X77] : ((? [X30] : p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30)))) => p(s(bool,X77))) <=> ! [X30] : (p(s(bool,happ(s(fun(X6,bool),X38),s(X6,X30)))) => p(s(bool,X77))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.PULL_EXISTS|split|0')).
fof(f6651,plain,(
  ( ! [X732,X730,X731] : (~p(s(bool,happ(s(fun(X730,bool),X731),s(X730,X732)))) | p(s(bool,happ(s(fun(X730,bool),X731),s(X730,sK38('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X731,X730)))))) ) | $spl416),
  inference(cnf_transformation,[],[f6651_D])).
fof(f6651_D,plain,(
  ( ! [X732,X730,X731] : (~p(s(bool,happ(s(fun(X730,bool),X731),s(X730,X732)))) | p(s(bool,happ(s(fun(X730,bool),X731),s(X730,sK38('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X731,X730)))))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f7043,plain,(
  ( ! [X732,X730,X731] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,happ(s(fun(X730,bool),X731),s(X730,X732)))) | p(s(bool,happ(s(fun(X730,bool),X731),s(X730,sK38('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X731,X730)))))) ) | $spl20),
  inference(superposition,[],[f1227,f1747])).
fof(f1227,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f912])).
fof(f6650,plain,(
  ( ! [X728,X729,X727] : (p(s(bool,happ(s(fun(X727,bool),X728),s(X727,sK37(X728,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X727))))) | ~p(s(bool,happ(s(fun(X727,bool),X728),s(X727,X729))))) ) | $spl414),
  inference(cnf_transformation,[],[f6650_D])).
fof(f6650_D,plain,(
  ( ! [X728,X729,X727] : (p(s(bool,happ(s(fun(X727,bool),X728),s(X727,sK37(X728,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X727))))) | ~p(s(bool,happ(s(fun(X727,bool),X728),s(X727,X729))))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f7042,plain,(
  ( ! [X728,X729,X727] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X727,bool),X728),s(X727,sK37(X728,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X727))))) | ~p(s(bool,happ(s(fun(X727,bool),X728),s(X727,X729))))) ) | $spl20),
  inference(superposition,[],[f1226,f1747])).
fof(f1226,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f908])).
fof(f908,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36,sK37])],[f907])).
fof(f907,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f906])).
fof(f906,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f905])).
fof(f905,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f504])).
fof(f504,axiom,(
  ! [X6] : ! [X38,X77] : (? [X30] : (p(s(bool,X38)) & p(s(bool,happ(s(fun(X6,bool),X77),s(X6,X30))))) <=> (p(s(bool,X38)) & ? [X30] : p(s(bool,happ(s(fun(X6,bool),X77),s(X6,X30))))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f6649,plain,(
  ( ! [X724,X726,X725] : (p(s(bool,happ(s(fun(X724,bool),X725),s(X724,sK36(X725,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X724))))) | ~p(s(bool,happ(s(fun(X724,bool),X725),s(X724,X726))))) ) | $spl412),
  inference(cnf_transformation,[],[f6649_D])).
fof(f6649_D,plain,(
  ( ! [X724,X726,X725] : (p(s(bool,happ(s(fun(X724,bool),X725),s(X724,sK36(X725,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X724))))) | ~p(s(bool,happ(s(fun(X724,bool),X725),s(X724,X726))))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f7041,plain,(
  ( ! [X724,X726,X725] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X724,bool),X725),s(X724,sK36(X725,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X724))))) | ~p(s(bool,happ(s(fun(X724,bool),X725),s(X724,X726))))) ) | $spl20),
  inference(superposition,[],[f1224,f1747])).
fof(f1224,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f908])).
fof(f6645,plain,(
  ( ! [X719] : (sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X719)) ) | $spl410),
  inference(cnf_transformation,[],[f6645_D])).
fof(f6645_D,plain,(
  ( ! [X719] : (sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X719)) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f7040,plain,(
  ( ! [X723] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X723)) ) | $spl20),
  inference(superposition,[],[f1191,f1747])).
fof(f1191,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f887])).
fof(f887,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f886])).
fof(f886,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f809])).
fof(f809,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f764,f808])).
fof(f808,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f764,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f763])).
fof(f763,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f382])).
fof(f382,axiom,(
  ! [X49,X6] : ((~(p(s(bool,X6)) | p(s(bool,X49))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X49)) => p(s(bool,'F_0')))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_9_prove_ltb')).
fof(f6644,plain,(
  ( ! [X718] : (sP12(X718,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl408),
  inference(cnf_transformation,[],[f6644_D])).
fof(f6644_D,plain,(
  ( ! [X718] : (sP12(X718,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f7227,plain,(
  ( ! [X722] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X722,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f7039,f1250])).
fof(f1250,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f931])).
fof(f931,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f930])).
fof(f930,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f774])).
fof(f774,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X39] : ((p(s(bool,X39)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CLAUSES|split|4')).
fof(f7039,plain,(
  ( ! [X722] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X722,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,'F_0'))) ) | $spl20),
  inference(superposition,[],[f1190,f1747])).
fof(f1190,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | p(s(bool,'F_0')) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f887])).
fof(f6643,plain,(
  ( ! [X717] : (p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) | $spl406),
  inference(cnf_transformation,[],[f6643_D])).
fof(f6643_D,plain,(
  ( ! [X717] : (p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f7226,plain,(
  ( ! [X721] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X721)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X721)) ) | $spl20),
  inference(subsumption_resolution,[],[f7038,f1270])).
fof(f1270,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f945])).
fof(f945,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f944])).
fof(f944,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X39] : ((p(s(bool,'F_0')) | p(s(bool,X39))) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.OR_CLAUSES|split|2')).
fof(f7038,plain,(
  ( ! [X721] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X721)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X721)) ) | $spl20),
  inference(superposition,[],[f1188,f1747])).
fof(f1188,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f887])).
fof(f6641,plain,(
  ( ! [X716] : (p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl404),
  inference(cnf_transformation,[],[f6641_D])).
fof(f6641_D,plain,(
  ( ! [X716] : (p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f7225,plain,(
  ( ! [X720] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X720)) | ~sP12(X720,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f7037,f1270])).
fof(f7037,plain,(
  ( ! [X720] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X720)) | ~sP12(X720,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1188,f1747])).
fof(f7036,plain,(
  ( ! [X719] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X719)) ) | $spl20),
  inference(superposition,[],[f1186,f1747])).
fof(f1186,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f885])).
fof(f885,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f884])).
fof(f884,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f808])).
fof(f7035,plain,(
  ( ! [X718] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X718,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1185,f1747])).
fof(f1185,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f885])).
fof(f7224,plain,(
  ( ! [X717] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) | $spl20),
  inference(subsumption_resolution,[],[f7034,f1270])).
fof(f7034,plain,(
  ( ! [X717] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) | $spl20),
  inference(superposition,[],[f1184,f1747])).
fof(f1184,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f885])).
fof(f7223,plain,(
  ( ! [X716] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f7033,f1270])).
fof(f7033,plain,(
  ( ! [X716] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1184,f1747])).
fof(f6632,plain,(
  ( ! [X700] : (p(s(bool,X700)) | ~sP11(X700,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl388),
  inference(cnf_transformation,[],[f6632_D])).
fof(f6632_D,plain,(
  ( ! [X700] : (p(s(bool,X700)) | ~sP11(X700,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f7032,plain,(
  ( ! [X715] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X715)) | ~sP11(X715,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1177,f1747])).
fof(f1177,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X1,X0)) )),
  inference(cnf_transformation,[],[f881])).
fof(f881,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP11(X1,X0)) & (~sP11(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f807])).
fof(f807,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP11(X1,X0))),
  inference(definition_folding,[],[f577,f806])).
fof(f577,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f396])).
fof(f396,axiom,(
  ! [X88,X78] : ((p(s(bool,X78)) <=> ~p(s(bool,X88))) <=> ((p(s(bool,X78)) | p(s(bool,X88))) & (~p(s(bool,X88)) | ~p(s(bool,X78)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_29_prove_ltb')).
fof(f6633,plain,(
  ( ! [X701] : (p(s(bool,X701)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X701)) ) | $spl390),
  inference(cnf_transformation,[],[f6633_D])).
fof(f6633_D,plain,(
  ( ! [X701] : (p(s(bool,X701)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X701)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f7031,plain,(
  ( ! [X714] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X714)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X714)) ) | $spl20),
  inference(superposition,[],[f1177,f1747])).
fof(f6634,plain,(
  ( ! [X702] : (~p(s(bool,X702)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X702)) ) | $spl392),
  inference(cnf_transformation,[],[f6634_D])).
fof(f6634_D,plain,(
  ( ! [X702] : (~p(s(bool,X702)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X702)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f7030,plain,(
  ( ! [X713] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X713)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X713)) ) | $spl20),
  inference(superposition,[],[f1176,f1747])).
fof(f1176,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X1,X0)) )),
  inference(cnf_transformation,[],[f881])).
fof(f6635,plain,(
  ( ! [X703] : (~p(s(bool,X703)) | ~sP11(X703,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl394),
  inference(cnf_transformation,[],[f6635_D])).
fof(f6635_D,plain,(
  ( ! [X703] : (~p(s(bool,X703)) | ~sP11(X703,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f7029,plain,(
  ( ! [X712] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X712)) | ~sP11(X712,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1176,f1747])).
fof(f6637,plain,(
  ( ! [X705] : (sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X705) | p(s(bool,X705))) ) | $spl398),
  inference(cnf_transformation,[],[f6637_D])).
fof(f6637_D,plain,(
  ( ! [X705] : (sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X705) | p(s(bool,X705))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f7028,plain,(
  ( ! [X711] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X711) | p(s(bool,X711))) ) | $spl20),
  inference(superposition,[],[f1175,f1747])).
fof(f1175,plain,(
  ( ! [X0,X1] : (sP11(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f881])).
fof(f6636,plain,(
  ( ! [X704] : (sP11(X704,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X704))) ) | $spl396),
  inference(cnf_transformation,[],[f6636_D])).
fof(f6636_D,plain,(
  ( ! [X704] : (sP11(X704,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X704))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f7027,plain,(
  ( ! [X710] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X710,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X710))) ) | $spl20),
  inference(superposition,[],[f1175,f1747])).
fof(f6639,plain,(
  ( ! [X707] : (sP11(X707,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X707))) ) | $spl402),
  inference(cnf_transformation,[],[f6639_D])).
fof(f6639_D,plain,(
  ( ! [X707] : (sP11(X707,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X707))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f7026,plain,(
  ( ! [X709] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X709,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X709))) ) | $spl20),
  inference(superposition,[],[f1174,f1747])).
fof(f1174,plain,(
  ( ! [X0,X1] : (sP11(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f881])).
fof(f6638,plain,(
  ( ! [X706] : (sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X706) | ~p(s(bool,X706))) ) | $spl400),
  inference(cnf_transformation,[],[f6638_D])).
fof(f6638_D,plain,(
  ( ! [X706] : (sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X706) | ~p(s(bool,X706))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f7025,plain,(
  ( ! [X708] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X708) | ~p(s(bool,X708))) ) | $spl20),
  inference(superposition,[],[f1174,f1747])).
fof(f7024,plain,(
  ( ! [X707] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X707,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X707))) ) | $spl20),
  inference(superposition,[],[f1173,f1747])).
fof(f7023,plain,(
  ( ! [X706] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X706) | ~p(s(bool,X706))) ) | $spl20),
  inference(superposition,[],[f1173,f1747])).
fof(f7022,plain,(
  ( ! [X705] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X705) | p(s(bool,X705))) ) | $spl20),
  inference(superposition,[],[f1170,f1747])).
fof(f7021,plain,(
  ( ! [X704] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X704,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X704))) ) | $spl20),
  inference(superposition,[],[f1170,f1747])).
fof(f7020,plain,(
  ( ! [X703] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X703)) | ~sP11(X703,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1169,f1747])).
fof(f1169,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f880])).
fof(f7019,plain,(
  ( ! [X702] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X702)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X702)) ) | $spl20),
  inference(superposition,[],[f1169,f1747])).
fof(f7018,plain,(
  ( ! [X701] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X701)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X701)) ) | $spl20),
  inference(superposition,[],[f1168,f1747])).
fof(f1168,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f880])).
fof(f7017,plain,(
  ( ! [X700] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X700)) | ~sP11(X700,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1168,f1747])).
fof(f7016,plain,(
  ( ! [X699,X698] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X698,X699)) ) | $spl20),
  inference(superposition,[],[f1151,f1747])).
fof(f1151,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP10(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f870])).
fof(f870,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0)) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(flattening,[],[f869])).
fof(f869,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0))) | sP10(X0,X2,X1)) & (~sP10(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f805])).
fof(f805,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> sP10(X0,X2,X1))),
  inference(definition_folding,[],[f568,f804])).
fof(f568,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  inference(rectify,[],[f338])).
fof(f338,axiom,(
  ! [X6,X49,X65] : (((p(s(bool,X49)) | p(s(bool,X65))) & p(s(bool,X6))) <=> ((p(s(bool,X49)) & p(s(bool,X6))) | (p(s(bool,X65)) & p(s(bool,X6)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.RIGHT_AND_OVER_OR')).
fof(f6623,plain,(
  ( ! [X666,X667] : (p(s(bool,X666)) | ~sP10(X667,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X666)) ) | $spl370),
  inference(cnf_transformation,[],[f6623_D])).
fof(f6623_D,plain,(
  ( ! [X666,X667] : (p(s(bool,X666)) | ~sP10(X667,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X666)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f7015,plain,(
  ( ! [X697,X696] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X696)) | ~sP10(X697,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X696)) ) | $spl20),
  inference(superposition,[],[f1150,f1747])).
fof(f1150,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP10(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f870])).
fof(f6624,plain,(
  ( ! [X668,X669] : (p(s(bool,X668)) | ~sP10(X669,X668,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl372),
  inference(cnf_transformation,[],[f6624_D])).
fof(f6624_D,plain,(
  ( ! [X668,X669] : (p(s(bool,X668)) | ~sP10(X669,X668,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f7014,plain,(
  ( ! [X695,X694] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X694)) | ~sP10(X695,X694,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1150,f1747])).
fof(f6630,plain,(
  ( ! [X685,X684] : (sP10(X684,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X685) | ~p(s(bool,X684))) ) | $spl384),
  inference(cnf_transformation,[],[f6630_D])).
fof(f6630_D,plain,(
  ( ! [X685,X684] : (sP10(X684,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X685) | ~p(s(bool,X684))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f7013,plain,(
  ( ! [X693,X692] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X692,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X693) | ~p(s(bool,X692))) ) | $spl20),
  inference(superposition,[],[f1149,f1747])).
fof(f1149,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X2,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f870])).
fof(f6629,plain,(
  ( ! [X683,X682] : (sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X682,X683) | ~p(s(bool,X682))) ) | $spl382),
  inference(cnf_transformation,[],[f6629_D])).
fof(f6629_D,plain,(
  ( ! [X683,X682] : (sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X682,X683) | ~p(s(bool,X682))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f7012,plain,(
  ( ! [X691,X690] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X690,X691) | ~p(s(bool,X690))) ) | $spl20),
  inference(superposition,[],[f1149,f1747])).
fof(f6628,plain,(
  ( ! [X681,X680] : (sP10(X680,X681,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X680))) ) | $spl380),
  inference(cnf_transformation,[],[f6628_D])).
fof(f6628_D,plain,(
  ( ! [X681,X680] : (sP10(X680,X681,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X680))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f7011,plain,(
  ( ! [X689,X688] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X688,X689,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X688))) ) | $spl20),
  inference(superposition,[],[f1148,f1747])).
fof(f1148,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X2,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f870])).
fof(f6627,plain,(
  ( ! [X679,X678] : (sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X678,X679) | ~p(s(bool,X679))) ) | $spl378),
  inference(cnf_transformation,[],[f6627_D])).
fof(f6627_D,plain,(
  ( ! [X679,X678] : (sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X678,X679) | ~p(s(bool,X679))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f7010,plain,(
  ( ! [X687,X686] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X686,X687) | ~p(s(bool,X687))) ) | $spl20),
  inference(superposition,[],[f1148,f1747])).
fof(f7009,plain,(
  ( ! [X685,X684] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X684,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X685) | ~p(s(bool,X684))) ) | $spl20),
  inference(superposition,[],[f1147,f1747])).
fof(f1147,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f868])).
fof(f7008,plain,(
  ( ! [X683,X682] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X682,X683) | ~p(s(bool,X682))) ) | $spl20),
  inference(superposition,[],[f1147,f1747])).
fof(f7007,plain,(
  ( ! [X681,X680] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X680,X681,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X680))) ) | $spl20),
  inference(superposition,[],[f1146,f1747])).
fof(f1146,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f868])).
fof(f7006,plain,(
  ( ! [X679,X678] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X678,X679) | ~p(s(bool,X679))) ) | $spl20),
  inference(superposition,[],[f1146,f1747])).
fof(f6626,plain,(
  ( ! [X675,X674] : (p(s(bool,X674)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X675,X674)) ) | $spl376),
  inference(cnf_transformation,[],[f6626_D])).
fof(f6626_D,plain,(
  ( ! [X675,X674] : (p(s(bool,X674)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X675,X674)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f7004,plain,(
  ( ! [X675,X674] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X674)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X675,X674)) ) | $spl20),
  inference(superposition,[],[f1144,f1747])).
fof(f1144,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f868])).
fof(f6625,plain,(
  ( ! [X673,X672] : (p(s(bool,X672)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X672,X673)) ) | $spl374),
  inference(cnf_transformation,[],[f6625_D])).
fof(f6625_D,plain,(
  ( ! [X673,X672] : (p(s(bool,X672)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X672,X673)) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f7003,plain,(
  ( ! [X673,X672] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X672)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X672,X673)) ) | $spl20),
  inference(superposition,[],[f1143,f1747])).
fof(f1143,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f868])).
fof(f7001,plain,(
  ( ! [X668,X669] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X668)) | ~sP10(X669,X668,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1142,f1747])).
fof(f1142,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f868])).
fof(f7000,plain,(
  ( ! [X666,X667] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X666)) | ~sP10(X667,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X666)) ) | $spl20),
  inference(superposition,[],[f1142,f1747])).
fof(f6622,plain,(
  ( ! [X664,X665] : (p(s(bool,X664)) | ~sP9(X664,X665,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X665))) ) | $spl368),
  inference(cnf_transformation,[],[f6622_D])).
fof(f6622_D,plain,(
  ( ! [X664,X665] : (p(s(bool,X664)) | ~sP9(X664,X665,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X665))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f6999,plain,(
  ( ! [X664,X665] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X664)) | ~p(s(bool,X665)) | ~sP9(X664,X665,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f1135,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f863])).
fof(f863,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f862])).
fof(f862,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f803])).
fof(f803,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f566,f802,f801])).
fof(f801,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f566,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f565])).
fof(f565,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f392])).
fof(f392,axiom,(
  ! [X89,X88,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X88)) & p(s(bool,X89)))) <=> ((p(s(bool,X78)) | ~p(s(bool,X88)) | ~p(s(bool,X89))) & (p(s(bool,X88)) | ~p(s(bool,X78))) & (p(s(bool,X89)) | ~p(s(bool,X78)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_23_prove_ltb')).
fof(f6621,plain,(
  ( ! [X662,X663] : (p(s(bool,X662)) | ~sP9(X662,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X663) | ~p(s(bool,X663))) ) | $spl366),
  inference(cnf_transformation,[],[f6621_D])).
fof(f6621_D,plain,(
  ( ! [X662,X663] : (p(s(bool,X662)) | ~sP9(X662,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X663) | ~p(s(bool,X663))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f6998,plain,(
  ( ! [X662,X663] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X662)) | ~p(s(bool,X663)) | ~sP9(X662,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X663)) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f6620,plain,(
  ( ! [X660,X661] : (~p(s(bool,X660)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X660,X661) | ~p(s(bool,X661))) ) | $spl364),
  inference(cnf_transformation,[],[f6620_D])).
fof(f6620_D,plain,(
  ( ! [X660,X661] : (~p(s(bool,X660)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X660,X661) | ~p(s(bool,X661))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f6997,plain,(
  ( ! [X660,X661] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X660)) | ~p(s(bool,X661)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X660,X661)) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f6602,plain,(
  ( ! [X613,X612] : (p(s(bool,X612)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X612,X613)) ) | $spl328),
  inference(cnf_transformation,[],[f6602_D])).
fof(f6602_D,plain,(
  ( ! [X613,X612] : (p(s(bool,X612)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X612,X613)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f6996,plain,(
  ( ! [X658,X659] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X658)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X658,X659)) ) | $spl20),
  inference(superposition,[],[f1134,f1747])).
fof(f1134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f863])).
fof(f6601,plain,(
  ( ! [X611,X610] : (~p(s(bool,X610)) | ~sP9(X610,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X611)) ) | $spl326),
  inference(cnf_transformation,[],[f6601_D])).
fof(f6601_D,plain,(
  ( ! [X611,X610] : (~p(s(bool,X610)) | ~sP9(X610,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X611)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f6995,plain,(
  ( ! [X656,X657] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X656)) | ~sP9(X656,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X657)) ) | $spl20),
  inference(superposition,[],[f1134,f1747])).
fof(f6600,plain,(
  ( ! [X609,X608] : (p(s(bool,X608)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X609,X608)) ) | $spl324),
  inference(cnf_transformation,[],[f6600_D])).
fof(f6600_D,plain,(
  ( ! [X609,X608] : (p(s(bool,X608)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X609,X608)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f6994,plain,(
  ( ! [X654,X655] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X654)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X655,X654)) ) | $spl20),
  inference(superposition,[],[f1133,f1747])).
fof(f1133,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f863])).
fof(f6599,plain,(
  ( ! [X606,X607] : (~p(s(bool,X606)) | ~sP9(X606,X607,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl322),
  inference(cnf_transformation,[],[f6599_D])).
fof(f6599_D,plain,(
  ( ! [X606,X607] : (~p(s(bool,X606)) | ~sP9(X606,X607,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f6993,plain,(
  ( ! [X652,X653] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X652)) | ~sP9(X652,X653,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1133,f1747])).
fof(f6619,plain,(
  ( ! [X650,X651] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X650,X651) | ~p(s(bool,X651)) | ~p(s(bool,X650))) ) | $spl362),
  inference(cnf_transformation,[],[f6619_D])).
fof(f6619_D,plain,(
  ( ! [X650,X651] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X650,X651) | ~p(s(bool,X651)) | ~p(s(bool,X650))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f6992,plain,(
  ( ! [X650,X651] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X650,X651) | ~p(s(bool,X650)) | ~p(s(bool,X651))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f863])).
fof(f6618,plain,(
  ( ! [X648,X649] : (sP9(X648,X649,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X648)) | ~p(s(bool,X649))) ) | $spl360),
  inference(cnf_transformation,[],[f6618_D])).
fof(f6618_D,plain,(
  ( ! [X648,X649] : (sP9(X648,X649,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X648)) | ~p(s(bool,X649))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f6991,plain,(
  ( ! [X648,X649] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X648,X649,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X649)) | ~p(s(bool,X648))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f6617,plain,(
  ( ! [X646,X647] : (sP9(X646,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X647) | ~p(s(bool,X646)) | ~p(s(bool,X647))) ) | $spl358),
  inference(cnf_transformation,[],[f6617_D])).
fof(f6617_D,plain,(
  ( ! [X646,X647] : (sP9(X646,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X647) | ~p(s(bool,X646)) | ~p(s(bool,X647))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f6990,plain,(
  ( ! [X646,X647] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X646,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X647) | ~p(s(bool,X647)) | ~p(s(bool,X646))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f6616,plain,(
  ( ! [X644,X645] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X644,X645) | p(s(bool,X644))) ) | $spl356),
  inference(cnf_transformation,[],[f6616_D])).
fof(f6616_D,plain,(
  ( ! [X644,X645] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X644,X645) | p(s(bool,X644))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f6989,plain,(
  ( ! [X644,X645] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X644,X645) | p(s(bool,X644))) ) | $spl20),
  inference(superposition,[],[f1131,f1747])).
fof(f1131,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f863])).
fof(f6615,plain,(
  ( ! [X642,X643] : (sP9(X642,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X643) | p(s(bool,X642))) ) | $spl354),
  inference(cnf_transformation,[],[f6615_D])).
fof(f6615_D,plain,(
  ( ! [X642,X643] : (sP9(X642,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X643) | p(s(bool,X642))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f6988,plain,(
  ( ! [X642,X643] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X642,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X643) | p(s(bool,X642))) ) | $spl20),
  inference(superposition,[],[f1131,f1747])).
fof(f6614,plain,(
  ( ! [X640,X641] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X640,X641) | p(s(bool,X641))) ) | $spl352),
  inference(cnf_transformation,[],[f6614_D])).
fof(f6614_D,plain,(
  ( ! [X640,X641] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X640,X641) | p(s(bool,X641))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f6987,plain,(
  ( ! [X640,X641] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X640,X641) | p(s(bool,X641))) ) | $spl20),
  inference(superposition,[],[f1130,f1747])).
fof(f1130,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f863])).
fof(f6613,plain,(
  ( ! [X639,X638] : (sP9(X638,X639,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X638))) ) | $spl350),
  inference(cnf_transformation,[],[f6613_D])).
fof(f6613_D,plain,(
  ( ! [X639,X638] : (sP9(X638,X639,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X638))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f6986,plain,(
  ( ! [X639,X638] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X638,X639,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X638))) ) | $spl20),
  inference(superposition,[],[f1130,f1747])).
fof(f6612,plain,(
  ( ! [X636,X637] : (sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X636,X637)) ) | $spl348),
  inference(cnf_transformation,[],[f6612_D])).
fof(f6612_D,plain,(
  ( ! [X636,X637] : (sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X636,X637)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f6985,plain,(
  ( ! [X636,X637] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X636,X637)) ) | $spl20),
  inference(superposition,[],[f1129,f1747])).
fof(f1129,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f861])).
fof(f861,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(flattening,[],[f860])).
fof(f860,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(nnf_transformation,[],[f801])).
fof(f6611,plain,(
  ( ! [X635,X634] : (sP8(X634,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X635)) ) | $spl346),
  inference(cnf_transformation,[],[f6611_D])).
fof(f6611_D,plain,(
  ( ! [X635,X634] : (sP8(X634,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X635)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f6984,plain,(
  ( ! [X635,X634] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8(X634,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X635)) ) | $spl20),
  inference(superposition,[],[f1128,f1747])).
fof(f1128,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f861])).
fof(f6610,plain,(
  ( ! [X633,X632] : (sP8(X632,X633,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl344),
  inference(cnf_transformation,[],[f6610_D])).
fof(f6610_D,plain,(
  ( ! [X633,X632] : (sP8(X632,X633,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f6983,plain,(
  ( ! [X633,X632] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8(X632,X633,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1127,f1747])).
fof(f1127,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f861])).
fof(f6609,plain,(
  ( ! [X631,X630] : (~p(s(bool,X630)) | ~sP8(X630,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X631) | p(s(bool,X631))) ) | $spl342),
  inference(cnf_transformation,[],[f6609_D])).
fof(f6609_D,plain,(
  ( ! [X631,X630] : (~p(s(bool,X630)) | ~sP8(X630,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X631) | p(s(bool,X631))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f6982,plain,(
  ( ! [X631,X630] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X630)) | p(s(bool,X631)) | ~sP8(X630,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X631)) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f861])).
fof(f6608,plain,(
  ( ! [X629,X628] : (~p(s(bool,X628)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X628,X629) | p(s(bool,X629))) ) | $spl340),
  inference(cnf_transformation,[],[f6608_D])).
fof(f6608_D,plain,(
  ( ! [X629,X628] : (~p(s(bool,X628)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X628,X629) | p(s(bool,X629))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f6981,plain,(
  ( ! [X629,X628] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X628)) | p(s(bool,X629)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X628,X629)) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f6607,plain,(
  ( ! [X627,X626] : (~p(s(bool,X626)) | ~sP8(X626,X627,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X627))) ) | $spl338),
  inference(cnf_transformation,[],[f6607_D])).
fof(f6607_D,plain,(
  ( ! [X627,X626] : (~p(s(bool,X626)) | ~sP8(X626,X627,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X627))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f6980,plain,(
  ( ! [X627,X626] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X626)) | ~p(s(bool,X627)) | ~sP8(X626,X627,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f6606,plain,(
  ( ! [X623,X622] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X622,X623) | ~sP8(X622,X623,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X623))) ) | $spl336),
  inference(cnf_transformation,[],[f6606_D])).
fof(f6606_D,plain,(
  ( ! [X623,X622] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X622,X623) | ~sP8(X622,X623,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X623))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f6978,plain,(
  ( ! [X623,X622] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X622,X623) | ~p(s(bool,X623)) | ~sP8(X622,X623,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1124,f1747])).
fof(f1124,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f859])).
fof(f6605,plain,(
  ( ! [X619,X618] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X618,X619) | ~sP8(X618,X619,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X618))) ) | $spl334),
  inference(cnf_transformation,[],[f6605_D])).
fof(f6605_D,plain,(
  ( ! [X619,X618] : (sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X618,X619) | ~sP8(X618,X619,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X618))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f6976,plain,(
  ( ! [X619,X618] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X618,X619) | ~p(s(bool,X618)) | ~sP8(X618,X619,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1123,f1747])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f859])).
fof(f6604,plain,(
  ( ! [X617,X616] : (sP9(X616,X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP8(X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X616) | ~p(s(bool,X617))) ) | $spl332),
  inference(cnf_transformation,[],[f6604_D])).
fof(f6604_D,plain,(
  ( ! [X617,X616] : (sP9(X616,X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP8(X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X616) | ~p(s(bool,X617))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f6975,plain,(
  ( ! [X617,X616] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X616,X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X617)) | ~sP8(X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X616)) ) | $spl20),
  inference(superposition,[],[f1122,f1747])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f859])).
fof(f6603,plain,(
  ( ! [X615,X614] : (sP9(X614,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615,X614) | ~p(s(bool,X615))) ) | $spl330),
  inference(cnf_transformation,[],[f6603_D])).
fof(f6603_D,plain,(
  ( ! [X615,X614] : (sP9(X614,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615,X614) | ~p(s(bool,X615))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f6974,plain,(
  ( ! [X615,X614] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X614,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615) | ~p(s(bool,X615)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615,X614)) ) | $spl20),
  inference(superposition,[],[f1122,f1747])).
fof(f6973,plain,(
  ( ! [X613,X612] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X612)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X612,X613)) ) | $spl20),
  inference(superposition,[],[f1121,f1747])).
fof(f1121,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f859])).
fof(f6972,plain,(
  ( ! [X611,X610] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X610)) | ~sP9(X610,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X611)) ) | $spl20),
  inference(superposition,[],[f1121,f1747])).
fof(f6971,plain,(
  ( ! [X609,X608] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X608)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X609,X608)) ) | $spl20),
  inference(superposition,[],[f1120,f1747])).
fof(f1120,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f859])).
fof(f6970,plain,(
  ( ! [X606,X607] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X606)) | ~sP9(X606,X607,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1120,f1747])).
fof(f6578,plain,(
  ( ! [X553,X552] : (p(s(bool,X552)) | ~sP7(X552,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X553)) ) | $spl280),
  inference(cnf_transformation,[],[f6578_D])).
fof(f6578_D,plain,(
  ( ! [X553,X552] : (p(s(bool,X552)) | ~sP7(X552,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X553)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f6969,plain,(
  ( ! [X605,X604] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X604)) | ~sP7(X604,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X605)) ) | $spl20),
  inference(superposition,[],[f1118,f1747])).
fof(f1118,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f856])).
fof(f856,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f855])).
fof(f855,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f800])).
fof(f800,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f564,f799,f798])).
fof(f798,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f564,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f393])).
fof(f393,axiom,(
  ! [X89,X88,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X88)) | p(s(bool,X89)))) <=> ((p(s(bool,X78)) | ~p(s(bool,X88))) & (p(s(bool,X78)) | ~p(s(bool,X89))) & (p(s(bool,X88)) | p(s(bool,X89)) | ~p(s(bool,X78)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_25_prove_ltb')).
fof(f6577,plain,(
  ( ! [X551,X550] : (~p(s(bool,X550)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X550,X551)) ) | $spl278),
  inference(cnf_transformation,[],[f6577_D])).
fof(f6577_D,plain,(
  ( ! [X551,X550] : (~p(s(bool,X550)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X550,X551)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f6968,plain,(
  ( ! [X602,X603] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X602)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X602,X603)) ) | $spl20),
  inference(superposition,[],[f1118,f1747])).
fof(f6576,plain,(
  ( ! [X549,X548] : (p(s(bool,X548)) | ~sP7(X548,X549,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl276),
  inference(cnf_transformation,[],[f6576_D])).
fof(f6576_D,plain,(
  ( ! [X549,X548] : (p(s(bool,X548)) | ~sP7(X548,X549,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f6967,plain,(
  ( ! [X600,X601] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X600)) | ~sP7(X600,X601,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1117,f1747])).
fof(f1117,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X1)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f856])).
fof(f6575,plain,(
  ( ! [X546,X547] : (~p(s(bool,X546)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X547,X546)) ) | $spl274),
  inference(cnf_transformation,[],[f6575_D])).
fof(f6575_D,plain,(
  ( ! [X546,X547] : (~p(s(bool,X546)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X547,X546)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f6966,plain,(
  ( ! [X598,X599] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X598)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X599,X598)) ) | $spl20),
  inference(superposition,[],[f1117,f1747])).
fof(f6598,plain,(
  ( ! [X596,X597] : (p(s(bool,X596)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X596,X597) | p(s(bool,X597))) ) | $spl320),
  inference(cnf_transformation,[],[f6598_D])).
fof(f6598_D,plain,(
  ( ! [X596,X597] : (p(s(bool,X596)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X596,X597) | p(s(bool,X597))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f6965,plain,(
  ( ! [X596,X597] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X596)) | p(s(bool,X597)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X596,X597)) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f1116,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f856])).
fof(f6597,plain,(
  ( ! [X594,X595] : (p(s(bool,X594)) | ~p(s(bool,X595)) | ~sP7(X595,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X594)) ) | $spl318),
  inference(cnf_transformation,[],[f6597_D])).
fof(f6597_D,plain,(
  ( ! [X594,X595] : (p(s(bool,X594)) | ~p(s(bool,X595)) | ~sP7(X595,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X594)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f6964,plain,(
  ( ! [X594,X595] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X594)) | ~p(s(bool,X595)) | ~sP7(X595,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X594)) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f6596,plain,(
  ( ! [X592,X593] : (p(s(bool,X592)) | ~p(s(bool,X593)) | ~sP7(X593,X592,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl316),
  inference(cnf_transformation,[],[f6596_D])).
fof(f6596_D,plain,(
  ( ! [X592,X593] : (p(s(bool,X592)) | ~p(s(bool,X593)) | ~sP7(X593,X592,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f6963,plain,(
  ( ! [X592,X593] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X592)) | ~p(s(bool,X593)) | ~sP7(X593,X592,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f6595,plain,(
  ( ! [X590,X591] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X590,X591) | ~p(s(bool,X590))) ) | $spl314),
  inference(cnf_transformation,[],[f6595_D])).
fof(f6595_D,plain,(
  ( ! [X590,X591] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X590,X591) | ~p(s(bool,X590))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f6962,plain,(
  ( ! [X590,X591] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X590,X591) | ~p(s(bool,X590))) ) | $spl20),
  inference(superposition,[],[f1115,f1747])).
fof(f1115,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f856])).
fof(f6594,plain,(
  ( ! [X588,X589] : (sP7(X588,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X589) | ~p(s(bool,X588))) ) | $spl312),
  inference(cnf_transformation,[],[f6594_D])).
fof(f6594_D,plain,(
  ( ! [X588,X589] : (sP7(X588,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X589) | ~p(s(bool,X588))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f6961,plain,(
  ( ! [X588,X589] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X588,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X589) | ~p(s(bool,X588))) ) | $spl20),
  inference(superposition,[],[f1115,f1747])).
fof(f6593,plain,(
  ( ! [X586,X587] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X586,X587) | ~p(s(bool,X587))) ) | $spl310),
  inference(cnf_transformation,[],[f6593_D])).
fof(f6593_D,plain,(
  ( ! [X586,X587] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X586,X587) | ~p(s(bool,X587))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f6960,plain,(
  ( ! [X586,X587] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X586,X587) | ~p(s(bool,X587))) ) | $spl20),
  inference(superposition,[],[f1114,f1747])).
fof(f1114,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f856])).
fof(f6592,plain,(
  ( ! [X584,X585] : (sP7(X584,X585,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X584))) ) | $spl308),
  inference(cnf_transformation,[],[f6592_D])).
fof(f6592_D,plain,(
  ( ! [X584,X585] : (sP7(X584,X585,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X584))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f6959,plain,(
  ( ! [X584,X585] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X584,X585,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X584))) ) | $spl20),
  inference(superposition,[],[f1114,f1747])).
fof(f6591,plain,(
  ( ! [X582,X583] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X582,X583) | p(s(bool,X583)) | p(s(bool,X582))) ) | $spl306),
  inference(cnf_transformation,[],[f6591_D])).
fof(f6591_D,plain,(
  ( ! [X582,X583] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X582,X583) | p(s(bool,X583)) | p(s(bool,X582))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f6958,plain,(
  ( ! [X582,X583] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X582,X583) | p(s(bool,X582)) | p(s(bool,X583))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f1113,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f856])).
fof(f6590,plain,(
  ( ! [X580,X581] : (sP7(X580,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X581) | p(s(bool,X580)) | p(s(bool,X581))) ) | $spl304),
  inference(cnf_transformation,[],[f6590_D])).
fof(f6590_D,plain,(
  ( ! [X580,X581] : (sP7(X580,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X581) | p(s(bool,X580)) | p(s(bool,X581))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f6957,plain,(
  ( ! [X580,X581] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X580,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X581) | p(s(bool,X581)) | p(s(bool,X580))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f6589,plain,(
  ( ! [X578,X579] : (sP7(X578,X579,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X578)) | p(s(bool,X579))) ) | $spl302),
  inference(cnf_transformation,[],[f6589_D])).
fof(f6589_D,plain,(
  ( ! [X578,X579] : (sP7(X578,X579,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X578)) | p(s(bool,X579))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f6956,plain,(
  ( ! [X578,X579] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X578,X579,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X579)) | p(s(bool,X578))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f6588,plain,(
  ( ! [X576,X577] : (sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X576,X577)) ) | $spl300),
  inference(cnf_transformation,[],[f6588_D])).
fof(f6588_D,plain,(
  ( ! [X576,X577] : (sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X576,X577)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f6955,plain,(
  ( ! [X576,X577] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X576,X577)) ) | $spl20),
  inference(superposition,[],[f1112,f1747])).
fof(f1112,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f854])).
fof(f854,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f853])).
fof(f853,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f852])).
fof(f852,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f798])).
fof(f6587,plain,(
  ( ! [X575,X574] : (sP6(X574,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X575)) ) | $spl298),
  inference(cnf_transformation,[],[f6587_D])).
fof(f6587_D,plain,(
  ( ! [X575,X574] : (sP6(X574,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X575)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f6954,plain,(
  ( ! [X575,X574] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6(X574,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X575)) ) | $spl20),
  inference(superposition,[],[f1111,f1747])).
fof(f1111,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f854])).
fof(f6586,plain,(
  ( ! [X573,X572] : (sP6(X572,X573,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl296),
  inference(cnf_transformation,[],[f6586_D])).
fof(f6586_D,plain,(
  ( ! [X573,X572] : (sP6(X572,X573,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f6953,plain,(
  ( ! [X573,X572] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6(X572,X573,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1110,f1747])).
fof(f1110,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f854])).
fof(f6585,plain,(
  ( ! [X571,X570] : (p(s(bool,X570)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X570,X571) | p(s(bool,X571))) ) | $spl294),
  inference(cnf_transformation,[],[f6585_D])).
fof(f6585_D,plain,(
  ( ! [X571,X570] : (p(s(bool,X570)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X570,X571) | p(s(bool,X571))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f6952,plain,(
  ( ! [X571,X570] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X570)) | p(s(bool,X571)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X570,X571)) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f1109,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f854])).
fof(f6584,plain,(
  ( ! [X569,X568] : (~p(s(bool,X568)) | ~sP6(X568,X569,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X569))) ) | $spl292),
  inference(cnf_transformation,[],[f6584_D])).
fof(f6584_D,plain,(
  ( ! [X569,X568] : (~p(s(bool,X568)) | ~sP6(X568,X569,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X569))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f6951,plain,(
  ( ! [X569,X568] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X568)) | p(s(bool,X569)) | ~sP6(X568,X569,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f6583,plain,(
  ( ! [X567,X566] : (~p(s(bool,X566)) | ~sP6(X566,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X567) | p(s(bool,X567))) ) | $spl290),
  inference(cnf_transformation,[],[f6583_D])).
fof(f6583_D,plain,(
  ( ! [X567,X566] : (~p(s(bool,X566)) | ~sP6(X566,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X567) | p(s(bool,X567))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f6950,plain,(
  ( ! [X567,X566] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X566)) | p(s(bool,X567)) | ~sP6(X566,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X567)) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f6582,plain,(
  ( ! [X565,X564] : (sP7(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X565)) | ~sP6(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl288),
  inference(cnf_transformation,[],[f6582_D])).
fof(f6582_D,plain,(
  ( ! [X565,X564] : (sP7(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X565)) | ~sP6(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f6949,plain,(
  ( ! [X565,X564] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP6(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X565))) ) | $spl20),
  inference(superposition,[],[f1108,f1747])).
fof(f1108,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f851])).
fof(f6581,plain,(
  ( ! [X563,X562] : (sP7(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | p(s(bool,X563)) | ~sP6(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563)) ) | $spl286),
  inference(cnf_transformation,[],[f6581_D])).
fof(f6581_D,plain,(
  ( ! [X563,X562] : (sP7(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | p(s(bool,X563)) | ~sP6(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f6948,plain,(
  ( ! [X563,X562] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | ~sP6(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | p(s(bool,X563))) ) | $spl20),
  inference(superposition,[],[f1108,f1747])).
fof(f6580,plain,(
  ( ! [X561,X560] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | p(s(bool,X560)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561)) ) | $spl284),
  inference(cnf_transformation,[],[f6580_D])).
fof(f6580_D,plain,(
  ( ! [X561,X560] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | p(s(bool,X560)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561)) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f6947,plain,(
  ( ! [X561,X560] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | p(s(bool,X560))) ) | $spl20),
  inference(superposition,[],[f1107,f1747])).
fof(f1107,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f851])).
fof(f6579,plain,(
  ( ! [X557,X556] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | p(s(bool,X557)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557)) ) | $spl282),
  inference(cnf_transformation,[],[f6579_D])).
fof(f6579_D,plain,(
  ( ! [X557,X556] : (sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | p(s(bool,X557)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f6945,plain,(
  ( ! [X557,X556] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | p(s(bool,X557))) ) | $spl20),
  inference(superposition,[],[f1106,f1747])).
fof(f1106,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f851])).
fof(f6943,plain,(
  ( ! [X553,X552] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X552)) | ~sP7(X552,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X553)) ) | $spl20),
  inference(superposition,[],[f1103,f1747])).
fof(f1103,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f851])).
fof(f6942,plain,(
  ( ! [X551,X550] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X550)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X550,X551)) ) | $spl20),
  inference(superposition,[],[f1103,f1747])).
fof(f6941,plain,(
  ( ! [X549,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X548)) | ~sP7(X548,X549,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1102,f1747])).
fof(f1102,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f851])).
fof(f6940,plain,(
  ( ! [X546,X547] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X546)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X547,X546)) ) | $spl20),
  inference(superposition,[],[f1102,f1747])).
fof(f6554,plain,(
  ( ! [X492,X493] : (p(s(bool,X492)) | ~sP5(X492,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X493)) ) | $spl232),
  inference(cnf_transformation,[],[f6554_D])).
fof(f6554_D,plain,(
  ( ! [X492,X493] : (p(s(bool,X492)) | ~sP5(X492,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X493)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f6939,plain,(
  ( ! [X545,X544] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X544)) | ~sP5(X544,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X545)) ) | $spl20),
  inference(superposition,[],[f1101,f1747])).
fof(f1101,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f848])).
fof(f848,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f847])).
fof(f847,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f797])).
fof(f797,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f756,f796,f795])).
fof(f795,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f756,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f562])).
fof(f562,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f561])).
fof(f561,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f394])).
fof(f394,axiom,(
  ! [X89,X88,X78] : ((p(s(bool,X78)) <=> (p(s(bool,X88)) => p(s(bool,X89)))) <=> ((p(s(bool,X78)) | p(s(bool,X88))) & (p(s(bool,X78)) | ~p(s(bool,X89))) & (~p(s(bool,X88)) | p(s(bool,X89)) | ~p(s(bool,X78)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_27_prove_ltb')).
fof(f6553,plain,(
  ( ! [X490,X491] : (~p(s(bool,X490)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X490,X491)) ) | $spl230),
  inference(cnf_transformation,[],[f6553_D])).
fof(f6553_D,plain,(
  ( ! [X490,X491] : (~p(s(bool,X490)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X490,X491)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f6938,plain,(
  ( ! [X542,X543] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X542)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X542,X543)) ) | $spl20),
  inference(superposition,[],[f1101,f1747])).
fof(f6551,plain,(
  ( ! [X486,X487] : (p(s(bool,X486)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X487,X486)) ) | $spl226),
  inference(cnf_transformation,[],[f6551_D])).
fof(f6551_D,plain,(
  ( ! [X486,X487] : (p(s(bool,X486)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X487,X486)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f6937,plain,(
  ( ! [X540,X541] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X540)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X541,X540)) ) | $spl20),
  inference(superposition,[],[f1100,f1747])).
fof(f1100,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f848])).
fof(f6552,plain,(
  ( ! [X488,X489] : (p(s(bool,X488)) | ~sP5(X488,X489,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl228),
  inference(cnf_transformation,[],[f6552_D])).
fof(f6552_D,plain,(
  ( ! [X488,X489] : (p(s(bool,X488)) | ~sP5(X488,X489,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f6936,plain,(
  ( ! [X538,X539] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X538)) | ~sP5(X538,X539,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1100,f1747])).
fof(f6574,plain,(
  ( ! [X536,X537] : (p(s(bool,X536)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X536,X537) | ~p(s(bool,X537))) ) | $spl272),
  inference(cnf_transformation,[],[f6574_D])).
fof(f6574_D,plain,(
  ( ! [X536,X537] : (p(s(bool,X536)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X536,X537) | ~p(s(bool,X537))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f6935,plain,(
  ( ! [X536,X537] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X536)) | ~p(s(bool,X537)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X536,X537)) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f1099,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f848])).
fof(f6573,plain,(
  ( ! [X534,X535] : (p(s(bool,X534)) | ~p(s(bool,X535)) | ~sP5(X535,X534,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl270),
  inference(cnf_transformation,[],[f6573_D])).
fof(f6573_D,plain,(
  ( ! [X534,X535] : (p(s(bool,X534)) | ~p(s(bool,X535)) | ~sP5(X535,X534,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f6934,plain,(
  ( ! [X534,X535] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X534)) | ~p(s(bool,X535)) | ~sP5(X535,X534,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f6572,plain,(
  ( ! [X532,X533] : (~p(s(bool,X532)) | ~p(s(bool,X533)) | ~sP5(X533,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X532)) ) | $spl268),
  inference(cnf_transformation,[],[f6572_D])).
fof(f6572_D,plain,(
  ( ! [X532,X533] : (~p(s(bool,X532)) | ~p(s(bool,X533)) | ~sP5(X533,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X532)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f6933,plain,(
  ( ! [X532,X533] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X532)) | ~p(s(bool,X533)) | ~sP5(X533,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X532)) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f6571,plain,(
  ( ! [X530,X531] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X530,X531) | ~p(s(bool,X530))) ) | $spl266),
  inference(cnf_transformation,[],[f6571_D])).
fof(f6571_D,plain,(
  ( ! [X530,X531] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X530,X531) | ~p(s(bool,X530))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f6932,plain,(
  ( ! [X530,X531] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X530,X531) | ~p(s(bool,X530))) ) | $spl20),
  inference(superposition,[],[f1098,f1747])).
fof(f1098,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f848])).
fof(f6570,plain,(
  ( ! [X528,X529] : (sP5(X528,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X529) | ~p(s(bool,X528))) ) | $spl264),
  inference(cnf_transformation,[],[f6570_D])).
fof(f6570_D,plain,(
  ( ! [X528,X529] : (sP5(X528,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X529) | ~p(s(bool,X528))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f6931,plain,(
  ( ! [X528,X529] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X528,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X529) | ~p(s(bool,X528))) ) | $spl20),
  inference(superposition,[],[f1098,f1747])).
fof(f6569,plain,(
  ( ! [X526,X527] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X526,X527) | p(s(bool,X527))) ) | $spl262),
  inference(cnf_transformation,[],[f6569_D])).
fof(f6569_D,plain,(
  ( ! [X526,X527] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X526,X527) | p(s(bool,X527))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f6930,plain,(
  ( ! [X526,X527] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X526,X527) | p(s(bool,X527))) ) | $spl20),
  inference(superposition,[],[f1097,f1747])).
fof(f1097,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f848])).
fof(f6568,plain,(
  ( ! [X524,X525] : (sP5(X524,X525,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X524))) ) | $spl260),
  inference(cnf_transformation,[],[f6568_D])).
fof(f6568_D,plain,(
  ( ! [X524,X525] : (sP5(X524,X525,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X524))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f6929,plain,(
  ( ! [X524,X525] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X524,X525,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X524))) ) | $spl20),
  inference(superposition,[],[f1097,f1747])).
fof(f6567,plain,(
  ( ! [X522,X523] : (sP5(X522,X523,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X522)) | p(s(bool,X523))) ) | $spl258),
  inference(cnf_transformation,[],[f6567_D])).
fof(f6567_D,plain,(
  ( ! [X522,X523] : (sP5(X522,X523,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X522)) | p(s(bool,X523))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f6928,plain,(
  ( ! [X522,X523] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X522,X523,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X523)) | p(s(bool,X522))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f1096,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f848])).
fof(f6566,plain,(
  ( ! [X520,X521] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X520,X521) | ~p(s(bool,X521)) | p(s(bool,X520))) ) | $spl256),
  inference(cnf_transformation,[],[f6566_D])).
fof(f6566_D,plain,(
  ( ! [X520,X521] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X520,X521) | ~p(s(bool,X521)) | p(s(bool,X520))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f6927,plain,(
  ( ! [X520,X521] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X520,X521) | p(s(bool,X520)) | ~p(s(bool,X521))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f6565,plain,(
  ( ! [X518,X519] : (sP5(X518,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X519) | p(s(bool,X518)) | ~p(s(bool,X519))) ) | $spl254),
  inference(cnf_transformation,[],[f6565_D])).
fof(f6565_D,plain,(
  ( ! [X518,X519] : (sP5(X518,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X519) | p(s(bool,X518)) | ~p(s(bool,X519))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f6926,plain,(
  ( ! [X518,X519] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X518,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X519) | ~p(s(bool,X519)) | p(s(bool,X518))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f6564,plain,(
  ( ! [X516,X517] : (sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X516,X517)) ) | $spl252),
  inference(cnf_transformation,[],[f6564_D])).
fof(f6564_D,plain,(
  ( ! [X516,X517] : (sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X516,X517)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f6925,plain,(
  ( ! [X516,X517] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X516,X517)) ) | $spl20),
  inference(superposition,[],[f1095,f1747])).
fof(f1095,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f846])).
fof(f846,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f845])).
fof(f845,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f844])).
fof(f844,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f795])).
fof(f6563,plain,(
  ( ! [X514,X515] : (sP4(X514,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X515)) ) | $spl250),
  inference(cnf_transformation,[],[f6563_D])).
fof(f6563_D,plain,(
  ( ! [X514,X515] : (sP4(X514,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X515)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f6924,plain,(
  ( ! [X514,X515] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4(X514,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X515)) ) | $spl20),
  inference(superposition,[],[f1094,f1747])).
fof(f1094,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f846])).
fof(f6562,plain,(
  ( ! [X512,X513] : (sP4(X512,X513,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl248),
  inference(cnf_transformation,[],[f6562_D])).
fof(f6562_D,plain,(
  ( ! [X512,X513] : (sP4(X512,X513,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f6923,plain,(
  ( ! [X512,X513] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4(X512,X513,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1093,f1747])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f846])).
fof(f6561,plain,(
  ( ! [X511,X510] : (~p(s(bool,X510)) | ~sP4(X510,X511,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X511))) ) | $spl246),
  inference(cnf_transformation,[],[f6561_D])).
fof(f6561_D,plain,(
  ( ! [X511,X510] : (~p(s(bool,X510)) | ~sP4(X510,X511,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X511))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f6922,plain,(
  ( ! [X511,X510] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X510)) | p(s(bool,X511)) | ~sP4(X510,X511,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f1092,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f846])).
fof(f6560,plain,(
  ( ! [X509,X508] : (p(s(bool,X508)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X508,X509) | ~p(s(bool,X509))) ) | $spl244),
  inference(cnf_transformation,[],[f6560_D])).
fof(f6560_D,plain,(
  ( ! [X509,X508] : (p(s(bool,X508)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X508,X509) | ~p(s(bool,X509))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f6921,plain,(
  ( ! [X509,X508] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X508)) | ~p(s(bool,X509)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X508,X509)) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f6559,plain,(
  ( ! [X507,X506] : (~p(s(bool,X506)) | ~sP4(X506,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X507) | ~p(s(bool,X507))) ) | $spl242),
  inference(cnf_transformation,[],[f6559_D])).
fof(f6559_D,plain,(
  ( ! [X507,X506] : (~p(s(bool,X506)) | ~sP4(X506,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X507) | ~p(s(bool,X507))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f6920,plain,(
  ( ! [X507,X506] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X506)) | ~p(s(bool,X507)) | ~sP4(X506,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X507)) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f6558,plain,(
  ( ! [X505,X504] : (sP5(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X505)) | ~sP4(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl240),
  inference(cnf_transformation,[],[f6558_D])).
fof(f6558_D,plain,(
  ( ! [X505,X504] : (sP5(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X505)) | ~sP4(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f6919,plain,(
  ( ! [X505,X504] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP4(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X505))) ) | $spl20),
  inference(superposition,[],[f1091,f1747])).
fof(f1091,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f843])).
fof(f6557,plain,(
  ( ! [X502,X503] : (sP5(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~p(s(bool,X503)) | ~sP4(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503)) ) | $spl238),
  inference(cnf_transformation,[],[f6557_D])).
fof(f6557_D,plain,(
  ( ! [X502,X503] : (sP5(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~p(s(bool,X503)) | ~sP4(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f6918,plain,(
  ( ! [X502,X503] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~sP4(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~p(s(bool,X503))) ) | $spl20),
  inference(superposition,[],[f1091,f1747])).
fof(f6556,plain,(
  ( ! [X500,X501] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | p(s(bool,X500)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501)) ) | $spl236),
  inference(cnf_transformation,[],[f6556_D])).
fof(f6556_D,plain,(
  ( ! [X500,X501] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | p(s(bool,X500)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f6917,plain,(
  ( ! [X500,X501] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | p(s(bool,X500))) ) | $spl20),
  inference(superposition,[],[f1090,f1747])).
fof(f1090,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f843])).
fof(f6555,plain,(
  ( ! [X494,X495] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~p(s(bool,X495)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495)) ) | $spl234),
  inference(cnf_transformation,[],[f6555_D])).
fof(f6555_D,plain,(
  ( ! [X494,X495] : (sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~p(s(bool,X495)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f6914,plain,(
  ( ! [X494,X495] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~p(s(bool,X495))) ) | $spl20),
  inference(superposition,[],[f1089,f1747])).
fof(f1089,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f843])).
fof(f6913,plain,(
  ( ! [X492,X493] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X492)) | ~sP5(X492,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X493)) ) | $spl20),
  inference(superposition,[],[f1086,f1747])).
fof(f1086,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f843])).
fof(f6912,plain,(
  ( ! [X490,X491] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X490)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X490,X491)) ) | $spl20),
  inference(superposition,[],[f1086,f1747])).
fof(f6911,plain,(
  ( ! [X488,X489] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X488)) | ~sP5(X488,X489,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1085,f1747])).
fof(f1085,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f843])).
fof(f6910,plain,(
  ( ! [X486,X487] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X486)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X487,X486)) ) | $spl20),
  inference(superposition,[],[f1085,f1747])).
fof(f6530,plain,(
  ( ! [X173,X172] : (s(bool,X172) != s(bool,X173) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X173,X172)) ) | $spl224),
  inference(cnf_transformation,[],[f6530_D])).
fof(f6530_D,plain,(
  ( ! [X173,X172] : (s(bool,X172) != s(bool,X173) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X173,X172)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f6875,plain,(
  ( ! [X173,X172] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | s(bool,X172) != s(bool,X173) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X173,X172)) ) | $spl20),
  inference(superposition,[],[f1041,f1747])).
fof(f1041,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f824])).
fof(f824,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f794])).
fof(f794,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f522,f793,f792,f791,f790])).
fof(f790,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f791,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f792,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f793,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f522,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f389])).
fof(f389,axiom,(
  ! [X89,X88,X78] : ((p(s(bool,X78)) <=> s(bool,X88) = s(bool,X89)) <=> ((p(s(bool,X78)) | p(s(bool,X88)) | p(s(bool,X89))) & (p(s(bool,X78)) | ~p(s(bool,X89)) | ~p(s(bool,X88))) & (p(s(bool,X88)) | ~p(s(bool,X89)) | ~p(s(bool,X78))) & (p(s(bool,X89)) | ~p(s(bool,X88)) | ~p(s(bool,X78)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.sat.sat_2|49|_prove_ltb')).
fof(f6529,plain,(
  ( ! [X167,X166] : (s(bool,X166) = s(bool,X167) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X167,X166)) ) | $spl222),
  inference(cnf_transformation,[],[f6529_D])).
fof(f6529_D,plain,(
  ( ! [X167,X166] : (s(bool,X166) = s(bool,X167) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X167,X166)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f6872,plain,(
  ( ! [X167,X166] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | s(bool,X166) = s(bool,X167) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X167,X166)) ) | $spl20),
  inference(superposition,[],[f1040,f1747])).
fof(f1040,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f824])).
fof(f6528,plain,(
  ( ! [X161,X160] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X160,X161) | s(bool,X160) != s(bool,X161)) ) | $spl220),
  inference(cnf_transformation,[],[f6528_D])).
fof(f6528_D,plain,(
  ( ! [X161,X160] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X160,X161) | s(bool,X160) != s(bool,X161)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f6869,plain,(
  ( ! [X161,X160] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X160,X161) | s(bool,X160) != s(bool,X161)) ) | $spl20),
  inference(superposition,[],[f1039,f1747])).
fof(f1039,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f824])).
fof(f6527,plain,(
  ( ! [X154,X155] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X154,X155) | s(bool,X154) = s(bool,X155)) ) | $spl218),
  inference(cnf_transformation,[],[f6527_D])).
fof(f6527_D,plain,(
  ( ! [X154,X155] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X154,X155) | s(bool,X154) = s(bool,X155)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f6866,plain,(
  ( ! [X154,X155] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X154,X155) | s(bool,X154) = s(bool,X155)) ) | $spl20),
  inference(superposition,[],[f1038,f1747])).
fof(f1038,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f824])).
fof(f6526,plain,(
  ( ! [X149,X148] : (sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X148,X149)) ) | $spl216),
  inference(cnf_transformation,[],[f6526_D])).
fof(f6526_D,plain,(
  ( ! [X149,X148] : (sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X148,X149)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f6863,plain,(
  ( ! [X149,X148] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X148,X149)) ) | $spl20),
  inference(superposition,[],[f1037,f1747])).
fof(f1037,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f823])).
fof(f823,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f822])).
fof(f822,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f790])).
fof(f6525,plain,(
  ( ! [X146,X147] : (sP0(X146,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X147)) ) | $spl214),
  inference(cnf_transformation,[],[f6525_D])).
fof(f6525_D,plain,(
  ( ! [X146,X147] : (sP0(X146,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X147)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f6862,plain,(
  ( ! [X146,X147] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0(X146,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X147)) ) | $spl20),
  inference(superposition,[],[f1036,f1747])).
fof(f1036,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f823])).
fof(f6524,plain,(
  ( ! [X144,X145] : (sP0(X144,X145,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl212),
  inference(cnf_transformation,[],[f6524_D])).
fof(f6524_D,plain,(
  ( ! [X144,X145] : (sP0(X144,X145,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f6861,plain,(
  ( ! [X144,X145] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0(X144,X145,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1035,f1747])).
fof(f1035,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f823])).
fof(f6523,plain,(
  ( ! [X142,X143] : (p(s(bool,X142)) | ~sP0(X142,X143,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X143))) ) | $spl210),
  inference(cnf_transformation,[],[f6523_D])).
fof(f6523_D,plain,(
  ( ! [X142,X143] : (p(s(bool,X142)) | ~sP0(X142,X143,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X143))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f6860,plain,(
  ( ! [X142,X143] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X142)) | p(s(bool,X143)) | ~sP0(X142,X143,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f1034,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f823])).
fof(f6522,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X140,X141) | p(s(bool,X141))) ) | $spl208),
  inference(cnf_transformation,[],[f6522_D])).
fof(f6522_D,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X140,X141) | p(s(bool,X141))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f6859,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X140)) | p(s(bool,X141)) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X140,X141)) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f6521,plain,(
  ( ! [X138,X139] : (p(s(bool,X138)) | ~sP0(X138,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X139) | p(s(bool,X139))) ) | $spl206),
  inference(cnf_transformation,[],[f6521_D])).
fof(f6521_D,plain,(
  ( ! [X138,X139] : (p(s(bool,X138)) | ~sP0(X138,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X139) | p(s(bool,X139))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f6858,plain,(
  ( ! [X138,X139] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X138)) | p(s(bool,X139)) | ~sP0(X138,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X139)) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f6520,plain,(
  ( ! [X136,X137] : (sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X136,X137)) ) | $spl204),
  inference(cnf_transformation,[],[f6520_D])).
fof(f6520_D,plain,(
  ( ! [X136,X137] : (sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X136,X137)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f6857,plain,(
  ( ! [X136,X137] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X136,X137)) ) | $spl20),
  inference(superposition,[],[f1033,f1747])).
fof(f1033,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f821])).
fof(f821,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f820])).
fof(f820,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f819])).
fof(f819,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f791])).
fof(f6519,plain,(
  ( ! [X134,X135] : (sP1(X134,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X135)) ) | $spl202),
  inference(cnf_transformation,[],[f6519_D])).
fof(f6519_D,plain,(
  ( ! [X134,X135] : (sP1(X134,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X135)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f6856,plain,(
  ( ! [X134,X135] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1(X134,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X135)) ) | $spl20),
  inference(superposition,[],[f1032,f1747])).
fof(f1032,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f821])).
fof(f6518,plain,(
  ( ! [X132,X133] : (sP1(X132,X133,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl200),
  inference(cnf_transformation,[],[f6518_D])).
fof(f6518_D,plain,(
  ( ! [X132,X133] : (sP1(X132,X133,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f6855,plain,(
  ( ! [X132,X133] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1(X132,X133,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1031,f1747])).
fof(f1031,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f821])).
fof(f6517,plain,(
  ( ! [X130,X131] : (~p(s(bool,X130)) | ~sP1(X130,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X131) | p(s(bool,X131))) ) | $spl198),
  inference(cnf_transformation,[],[f6517_D])).
fof(f6517_D,plain,(
  ( ! [X130,X131] : (~p(s(bool,X130)) | ~sP1(X130,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X131) | p(s(bool,X131))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f6854,plain,(
  ( ! [X130,X131] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X130)) | p(s(bool,X131)) | ~sP1(X130,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X131)) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f1030,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f821])).
fof(f6516,plain,(
  ( ! [X128,X129] : (~p(s(bool,X128)) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X128,X129) | p(s(bool,X129))) ) | $spl196),
  inference(cnf_transformation,[],[f6516_D])).
fof(f6516_D,plain,(
  ( ! [X128,X129] : (~p(s(bool,X128)) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X128,X129) | p(s(bool,X129))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f6853,plain,(
  ( ! [X128,X129] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X128)) | p(s(bool,X129)) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X128,X129)) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f6515,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | ~sP1(X126,X127,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X127))) ) | $spl194),
  inference(cnf_transformation,[],[f6515_D])).
fof(f6515_D,plain,(
  ( ! [X127,X126] : (~p(s(bool,X126)) | ~sP1(X126,X127,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X127))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f6852,plain,(
  ( ! [X127,X126] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | ~sP1(X126,X127,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f6514,plain,(
  ( ! [X125,X124] : (sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X124,X125)) ) | $spl192),
  inference(cnf_transformation,[],[f6514_D])).
fof(f6514_D,plain,(
  ( ! [X125,X124] : (sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X124,X125)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f6851,plain,(
  ( ! [X125,X124] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X124,X125)) ) | $spl20),
  inference(superposition,[],[f1029,f1747])).
fof(f1029,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f818])).
fof(f818,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f817])).
fof(f817,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f816])).
fof(f816,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f792])).
fof(f6513,plain,(
  ( ! [X123,X122] : (sP2(X122,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X123)) ) | $spl190),
  inference(cnf_transformation,[],[f6513_D])).
fof(f6513_D,plain,(
  ( ! [X123,X122] : (sP2(X122,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X123)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f6850,plain,(
  ( ! [X123,X122] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2(X122,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X123)) ) | $spl20),
  inference(superposition,[],[f1028,f1747])).
fof(f1028,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f818])).
fof(f6512,plain,(
  ( ! [X121,X120] : (sP2(X120,X121,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl188),
  inference(cnf_transformation,[],[f6512_D])).
fof(f6512_D,plain,(
  ( ! [X121,X120] : (sP2(X120,X121,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f6849,plain,(
  ( ! [X121,X120] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2(X120,X121,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1027,f1747])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f818])).
fof(f6511,plain,(
  ( ! [X118,X119] : (~p(s(bool,X118)) | ~sP2(X118,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X119) | p(s(bool,X119))) ) | $spl186),
  inference(cnf_transformation,[],[f6511_D])).
fof(f6511_D,plain,(
  ( ! [X118,X119] : (~p(s(bool,X118)) | ~sP2(X118,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X119) | p(s(bool,X119))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f6848,plain,(
  ( ! [X118,X119] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X118)) | p(s(bool,X119)) | ~sP2(X118,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X119)) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f1026,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f818])).
fof(f6510,plain,(
  ( ! [X116,X117] : (~p(s(bool,X116)) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X116,X117) | p(s(bool,X117))) ) | $spl184),
  inference(cnf_transformation,[],[f6510_D])).
fof(f6510_D,plain,(
  ( ! [X116,X117] : (~p(s(bool,X116)) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X116,X117) | p(s(bool,X117))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f6847,plain,(
  ( ! [X116,X117] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X116)) | p(s(bool,X117)) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X116,X117)) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f6509,plain,(
  ( ! [X114,X115] : (~p(s(bool,X114)) | ~sP2(X114,X115,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X115))) ) | $spl182),
  inference(cnf_transformation,[],[f6509_D])).
fof(f6509_D,plain,(
  ( ! [X114,X115] : (~p(s(bool,X114)) | ~sP2(X114,X115,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X115))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6846,plain,(
  ( ! [X114,X115] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | ~sP2(X114,X115,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f6508,plain,(
  ( ! [X112,X113] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X112,X113) | ~sP0(X113,X112,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP1(X112,X113,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X113,X112)) ) | $spl180),
  inference(cnf_transformation,[],[f6508_D])).
fof(f6508_D,plain,(
  ( ! [X112,X113] : (sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X112,X113) | ~sP0(X113,X112,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP1(X112,X113,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X113,X112)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6845,plain,(
  ( ! [X112,X113] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X112,X113) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X113,X112) | ~sP1(X112,X113,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP0(X113,X112,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1025,f1747])).
fof(f1025,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f815])).
fof(f815,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f814])).
fof(f814,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f813])).
fof(f813,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f793])).
fof(f6507,plain,(
  ( ! [X111,X110] : (sP3(X110,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111) | ~sP0(X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X110) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111,X110) | ~sP2(X110,X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl178),
  inference(cnf_transformation,[],[f6507_D])).
fof(f6507_D,plain,(
  ( ! [X111,X110] : (sP3(X110,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111) | ~sP0(X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X110) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111,X110) | ~sP2(X110,X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6844,plain,(
  ( ! [X111,X110] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3(X110,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111) | ~sP2(X110,X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111,X110) | ~sP0(X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X110)) ) | $spl20),
  inference(superposition,[],[f1024,f1747])).
fof(f1024,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f815])).
fof(f6506,plain,(
  ( ! [X109,X108] : (sP3(X108,X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109,X108) | ~sP1(X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X108) | ~sP2(X108,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109)) ) | $spl176),
  inference(cnf_transformation,[],[f6506_D])).
fof(f6506_D,plain,(
  ( ! [X109,X108] : (sP3(X108,X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109,X108) | ~sP1(X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X108) | ~sP2(X108,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6843,plain,(
  ( ! [X109,X108] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3(X108,X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP2(X108,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109) | ~sP1(X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X108) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109,X108)) ) | $spl20),
  inference(superposition,[],[f1023,f1747])).
fof(f1023,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f815])).
fof(f6505,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | ~sP3(X106,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X107) | p(s(bool,X107))) ) | $spl174),
  inference(cnf_transformation,[],[f6505_D])).
fof(f6505_D,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | ~sP3(X106,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X107) | p(s(bool,X107))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6842,plain,(
  ( ! [X107,X106] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3(X106,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X107)) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f1022,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f815])).
fof(f6504,plain,(
  ( ! [X105,X104] : (~p(s(bool,X104)) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X104,X105) | p(s(bool,X105))) ) | $spl172),
  inference(cnf_transformation,[],[f6504_D])).
fof(f6504_D,plain,(
  ( ! [X105,X104] : (~p(s(bool,X104)) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X104,X105) | p(s(bool,X105))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f6841,plain,(
  ( ! [X105,X104] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X104)) | p(s(bool,X105)) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X104,X105)) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f6501,plain,(
  ( ! [X103,X102] : (~p(s(bool,X102)) | ~sP3(X102,X103,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X103))) ) | $spl168),
  inference(cnf_transformation,[],[f6501_D])).
fof(f6501_D,plain,(
  ( ! [X103,X102] : (~p(s(bool,X102)) | ~sP3(X102,X103,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X103))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f6840,plain,(
  ( ! [X103,X102] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | ~sP3(X102,X103,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f6333,plain,(
  ( ! [X718] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X718,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1185,f1747])).
fof(f6334,plain,(
  ( ! [X719] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X719)) ) | $spl20),
  inference(superposition,[],[f1186,f1747])).
fof(f6648,plain,(
  ( ! [X722] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X722,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f6337,f1250])).
fof(f6337,plain,(
  ( ! [X722] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12(X722,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,'F_0'))) ) | $spl20),
  inference(superposition,[],[f1190,f1747])).
fof(f6338,plain,(
  ( ! [X723] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X723)) ) | $spl20),
  inference(superposition,[],[f1191,f1747])).
fof(f6768,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl435),
  inference(cnf_transformation,[],[f6768_D])).
fof(f6768_D,plain,(
  ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) <=> ~$spl435),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl435])])).
fof(f6775,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | $spl443),
  inference(cnf_transformation,[],[f6775_D])).
fof(f6775_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) <=> ~$spl443),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl443])])).
fof(f6760,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl21),
  inference(superposition,[],[f1746,f1328])).
fof(f1328,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f977])).
fof(f977,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f653])).
fof(f653,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X39] : (s(bool,'T_0') = s(bool,X39) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.EQ_CLAUSES|split|0')).
fof(f1746,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl21),
  inference(cnf_transformation,[],[f1746_D])).
fof(f1746_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) != s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f6759,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl21),
  inference(superposition,[],[f1746,f1326])).
fof(f1326,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f976])).
fof(f976,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X39] : (s(bool,'T_0') = s(bool,X39) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.EQ_CLAUSES|split|1')).
fof(f6769,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl434),
  inference(cnf_transformation,[],[f6769_D])).
fof(f6769_D,plain,(
  p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f6773,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | $spl441),
  inference(cnf_transformation,[],[f6773_D])).
fof(f6773_D,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) <=> ~$spl441),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl441])])).
fof(f6758,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl21),
  inference(superposition,[],[f1746,f1324])).
fof(f1324,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f975])).
fof(f975,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f651])).
fof(f651,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f650])).
fof(f650,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X39] : (s(bool,'F_0') = s(bool,X39) <=> ~p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.EQ_CLAUSES|split|3')).
fof(f6757,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl21),
  inference(superposition,[],[f1746,f1322])).
fof(f1322,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f974])).
fof(f974,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f649])).
fof(f649,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f648])).
fof(f648,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X39] : (s(bool,'F_0') = s(bool,X39) <=> ~p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.EQ_CLAUSES|split|2')).
fof(f6755,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | $spl21),
  inference(superposition,[],[f1746,f1252])).
fof(f1252,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f933])).
fof(f933,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f932])).
fof(f932,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f775])).
fof(f775,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X39] : ((p(s(bool,X39)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X39))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_F_EQ_F')).
fof(f6771,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,X10))) ) | $spl438),
  inference(cnf_transformation,[],[f6771_D])).
fof(f6771_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,X10))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f6754,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,X11)) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))))) ) | $spl21),
  inference(superposition,[],[f1746,f1053])).
fof(f1053,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f733])).
fof(f733,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f732])).
fof(f732,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f527])).
fof(f527,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X27,X28] : ((p(s(bool,X27)) => p(s(bool,X28))) => ((p(s(bool,X28)) => p(s(bool,X27))) => s(bool,X27) = s(bool,X28)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_ANTISYM_AX')).
fof(f6753,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | ~p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | ~p(s(bool,X10))) ) | $spl21),
  inference(superposition,[],[f1746,f1053])).
fof(f6770,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,X8))) ) | $spl436),
  inference(cnf_transformation,[],[f6770_D])).
fof(f6770_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,X8))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f6752,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) | p(s(bool,X9))) ) | $spl21),
  inference(superposition,[],[f1746,f1050])).
fof(f1050,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f733])).
fof(f6751,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) | p(s(bool,X8)) | p(s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))))) ) | $spl21),
  inference(superposition,[],[f1746,f1050])).
fof(f6729,plain,(
  $false | ($spl42 | $spl94)),
  inference(subsumption_resolution,[],[f5150,f1763])).
fof(f5150,plain,(
  ~p(s(bool,sK25)) | $spl94),
  inference(subsumption_resolution,[],[f5146,f1429])).
fof(f1429,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f1245])).
fof(f1245,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f928])).
fof(f928,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f616])).
fof(f616,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f615])).
fof(f615,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X30] : (s(X6,X30) = s(X6,X30) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.REFL_CLAUSE')).
fof(f5146,plain,(
  ~p(s(bool,sK25)) | ~p(s(bool,'T_0')) | $spl94),
  inference(resolution,[],[f3260,f1176])).
fof(f6395,plain,(
  ( ! [X783,X782] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X782,X783)) ) | $spl20),
  inference(superposition,[],[f1428,f1747])).
fof(f6394,plain,(
  ( ! [X781,X780] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X780,X781) | ~sP8(X780,X781,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1427,f1747])).
fof(f6393,plain,(
  ( ! [X779,X778] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X778,X779)) ) | $spl20),
  inference(superposition,[],[f1426,f1747])).
fof(f6392,plain,(
  ( ! [X777,X776] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X776,X777)) ) | $spl20),
  inference(superposition,[],[f1425,f1747])).
fof(f6346,plain,(
  ( ! [X745,X747,X746] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X745,bool),X746),s(X745,sK43('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X746,X745))))) | ~p(s(bool,happ(s(fun(X745,bool),X746),s(X745,X747))))) ) | $spl20),
  inference(superposition,[],[f1238,f1747])).
fof(f6345,plain,(
  ( ! [X743,X744,X742] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X742,bool),X743),s(X742,sK42('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X743,X742))))) | ~p(s(bool,happ(s(fun(X742,bool),X743),s(X742,X744))))) ) | $spl20),
  inference(superposition,[],[f1236,f1747])).
fof(f6344,plain,(
  ( ! [X741,X739,X740] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X739,bool),X740),s(X739,sK41('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X740,X739))))) | ~p(s(bool,happ(s(fun(X739,bool),X740),s(X739,X741))))) ) | $spl20),
  inference(superposition,[],[f1233,f1747])).
fof(f6343,plain,(
  ( ! [X736,X738,X737] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X736,bool),X737),s(X736,sK40('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X737,X736))))) | ~p(s(bool,happ(s(fun(X736,bool),X737),s(X736,X738))))) ) | $spl20),
  inference(superposition,[],[f1231,f1747])).
fof(f6342,plain,(
  ( ! [X734,X733,X735] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,happ(s(fun(X733,bool),X734),s(X733,X735)))) | p(s(bool,happ(s(fun(X733,bool),X734),s(X733,sK39('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X734,X733)))))) ) | $spl20),
  inference(superposition,[],[f1229,f1747])).
fof(f6341,plain,(
  ( ! [X732,X730,X731] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,happ(s(fun(X730,bool),X731),s(X730,X732)))) | p(s(bool,happ(s(fun(X730,bool),X731),s(X730,sK38('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X731,X730)))))) ) | $spl20),
  inference(superposition,[],[f1227,f1747])).
fof(f6340,plain,(
  ( ! [X728,X729,X727] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X727,bool),X728),s(X727,sK37(X728,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X727))))) | ~p(s(bool,happ(s(fun(X727,bool),X728),s(X727,X729))))) ) | $spl20),
  inference(superposition,[],[f1226,f1747])).
fof(f6339,plain,(
  ( ! [X724,X726,X725] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,happ(s(fun(X724,bool),X725),s(X724,sK36(X725,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X724))))) | ~p(s(bool,happ(s(fun(X724,bool),X725),s(X724,X726))))) ) | $spl20),
  inference(superposition,[],[f1224,f1747])).
fof(f6647,plain,(
  ( ! [X721] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X721)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X721)) ) | $spl20),
  inference(subsumption_resolution,[],[f6336,f1270])).
fof(f6336,plain,(
  ( ! [X721] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X721)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X721)) ) | $spl20),
  inference(superposition,[],[f1188,f1747])).
fof(f6646,plain,(
  ( ! [X720] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X720)) | ~sP12(X720,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f6335,f1270])).
fof(f6335,plain,(
  ( ! [X720] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X720)) | ~sP12(X720,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1188,f1747])).
fof(f6642,plain,(
  ( ! [X717] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) | $spl20),
  inference(subsumption_resolution,[],[f6332,f1270])).
fof(f6332,plain,(
  ( ! [X717] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X717)) | ~sP12('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X717)) ) | $spl20),
  inference(superposition,[],[f1184,f1747])).
fof(f6640,plain,(
  ( ! [X716] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(subsumption_resolution,[],[f6331,f1270])).
fof(f6331,plain,(
  ( ! [X716] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,'F_0')) | p(s(bool,X716)) | ~sP12(X716,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1184,f1747])).
fof(f6330,plain,(
  ( ! [X715] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X715)) | ~sP11(X715,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1177,f1747])).
fof(f6329,plain,(
  ( ! [X714] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X714)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X714)) ) | $spl20),
  inference(superposition,[],[f1177,f1747])).
fof(f6328,plain,(
  ( ! [X713] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X713)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X713)) ) | $spl20),
  inference(superposition,[],[f1176,f1747])).
fof(f6327,plain,(
  ( ! [X712] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X712)) | ~sP11(X712,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1176,f1747])).
fof(f6326,plain,(
  ( ! [X711] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X711) | p(s(bool,X711))) ) | $spl20),
  inference(superposition,[],[f1175,f1747])).
fof(f6325,plain,(
  ( ! [X710] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X710,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X710))) ) | $spl20),
  inference(superposition,[],[f1175,f1747])).
fof(f6324,plain,(
  ( ! [X709] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X709,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X709))) ) | $spl20),
  inference(superposition,[],[f1174,f1747])).
fof(f6323,plain,(
  ( ! [X708] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X708) | ~p(s(bool,X708))) ) | $spl20),
  inference(superposition,[],[f1174,f1747])).
fof(f6322,plain,(
  ( ! [X707] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X707,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X707))) ) | $spl20),
  inference(superposition,[],[f1173,f1747])).
fof(f6321,plain,(
  ( ! [X706] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X706) | ~p(s(bool,X706))) ) | $spl20),
  inference(superposition,[],[f1173,f1747])).
fof(f6320,plain,(
  ( ! [X705] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X705) | p(s(bool,X705))) ) | $spl20),
  inference(superposition,[],[f1170,f1747])).
fof(f6319,plain,(
  ( ! [X704] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP11(X704,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X704))) ) | $spl20),
  inference(superposition,[],[f1170,f1747])).
fof(f6318,plain,(
  ( ! [X703] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X703)) | ~sP11(X703,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1169,f1747])).
fof(f6317,plain,(
  ( ! [X702] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X702)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X702)) ) | $spl20),
  inference(superposition,[],[f1169,f1747])).
fof(f6316,plain,(
  ( ! [X701] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X701)) | ~sP11('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X701)) ) | $spl20),
  inference(superposition,[],[f1168,f1747])).
fof(f6315,plain,(
  ( ! [X700] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X700)) | ~sP11(X700,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1168,f1747])).
fof(f6314,plain,(
  ( ! [X699,X698] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X698,X699)) ) | $spl20),
  inference(superposition,[],[f1151,f1747])).
fof(f6313,plain,(
  ( ! [X697,X696] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X696)) | ~sP10(X697,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X696)) ) | $spl20),
  inference(superposition,[],[f1150,f1747])).
fof(f6312,plain,(
  ( ! [X695,X694] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X694)) | ~sP10(X695,X694,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1150,f1747])).
fof(f6311,plain,(
  ( ! [X693,X692] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X692,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X693) | ~p(s(bool,X692))) ) | $spl20),
  inference(superposition,[],[f1149,f1747])).
fof(f6310,plain,(
  ( ! [X691,X690] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X690,X691) | ~p(s(bool,X690))) ) | $spl20),
  inference(superposition,[],[f1149,f1747])).
fof(f6309,plain,(
  ( ! [X689,X688] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X688,X689,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X688))) ) | $spl20),
  inference(superposition,[],[f1148,f1747])).
fof(f6308,plain,(
  ( ! [X687,X686] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X686,X687) | ~p(s(bool,X687))) ) | $spl20),
  inference(superposition,[],[f1148,f1747])).
fof(f6307,plain,(
  ( ! [X685,X684] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X684,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X685) | ~p(s(bool,X684))) ) | $spl20),
  inference(superposition,[],[f1147,f1747])).
fof(f6306,plain,(
  ( ! [X683,X682] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X682,X683) | ~p(s(bool,X682))) ) | $spl20),
  inference(superposition,[],[f1147,f1747])).
fof(f6305,plain,(
  ( ! [X681,X680] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10(X680,X681,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X680))) ) | $spl20),
  inference(superposition,[],[f1146,f1747])).
fof(f6304,plain,(
  ( ! [X679,X678] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X678,X679) | ~p(s(bool,X679))) ) | $spl20),
  inference(superposition,[],[f1146,f1747])).
fof(f6302,plain,(
  ( ! [X675,X674] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X674)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X675,X674)) ) | $spl20),
  inference(superposition,[],[f1144,f1747])).
fof(f6301,plain,(
  ( ! [X673,X672] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X672)) | ~sP10('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X672,X673)) ) | $spl20),
  inference(superposition,[],[f1143,f1747])).
fof(f6299,plain,(
  ( ! [X668,X669] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X668)) | ~sP10(X669,X668,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1142,f1747])).
fof(f6298,plain,(
  ( ! [X666,X667] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X666)) | ~sP10(X667,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X666)) ) | $spl20),
  inference(superposition,[],[f1142,f1747])).
fof(f6297,plain,(
  ( ! [X664,X665] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X664)) | ~p(s(bool,X665)) | ~sP9(X664,X665,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f6296,plain,(
  ( ! [X662,X663] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X662)) | ~p(s(bool,X663)) | ~sP9(X662,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X663)) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f6295,plain,(
  ( ! [X660,X661] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X660)) | ~p(s(bool,X661)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X660,X661)) ) | $spl20),
  inference(superposition,[],[f1135,f1747])).
fof(f6294,plain,(
  ( ! [X658,X659] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X658)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X658,X659)) ) | $spl20),
  inference(superposition,[],[f1134,f1747])).
fof(f6293,plain,(
  ( ! [X656,X657] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X656)) | ~sP9(X656,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X657)) ) | $spl20),
  inference(superposition,[],[f1134,f1747])).
fof(f6292,plain,(
  ( ! [X654,X655] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X654)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X655,X654)) ) | $spl20),
  inference(superposition,[],[f1133,f1747])).
fof(f6291,plain,(
  ( ! [X652,X653] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X652)) | ~sP9(X652,X653,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1133,f1747])).
fof(f6290,plain,(
  ( ! [X650,X651] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X650,X651) | ~p(s(bool,X650)) | ~p(s(bool,X651))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f6289,plain,(
  ( ! [X648,X649] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X648,X649,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X649)) | ~p(s(bool,X648))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f6288,plain,(
  ( ! [X646,X647] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X646,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X647) | ~p(s(bool,X647)) | ~p(s(bool,X646))) ) | $spl20),
  inference(superposition,[],[f1132,f1747])).
fof(f6287,plain,(
  ( ! [X644,X645] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X644,X645) | p(s(bool,X644))) ) | $spl20),
  inference(superposition,[],[f1131,f1747])).
fof(f6286,plain,(
  ( ! [X642,X643] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X642,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X643) | p(s(bool,X642))) ) | $spl20),
  inference(superposition,[],[f1131,f1747])).
fof(f6285,plain,(
  ( ! [X640,X641] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X640,X641) | p(s(bool,X641))) ) | $spl20),
  inference(superposition,[],[f1130,f1747])).
fof(f6284,plain,(
  ( ! [X639,X638] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X638,X639,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X638))) ) | $spl20),
  inference(superposition,[],[f1130,f1747])).
fof(f6283,plain,(
  ( ! [X636,X637] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X636,X637)) ) | $spl20),
  inference(superposition,[],[f1129,f1747])).
fof(f6282,plain,(
  ( ! [X635,X634] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8(X634,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X635)) ) | $spl20),
  inference(superposition,[],[f1128,f1747])).
fof(f6281,plain,(
  ( ! [X633,X632] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP8(X632,X633,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1127,f1747])).
fof(f6280,plain,(
  ( ! [X631,X630] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X630)) | p(s(bool,X631)) | ~sP8(X630,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X631)) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f6279,plain,(
  ( ! [X629,X628] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X628)) | p(s(bool,X629)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X628,X629)) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f6278,plain,(
  ( ! [X627,X626] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X626)) | ~p(s(bool,X627)) | ~sP8(X626,X627,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1126,f1747])).
fof(f6276,plain,(
  ( ! [X623,X622] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X622,X623) | ~p(s(bool,X623)) | ~sP8(X622,X623,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1124,f1747])).
fof(f6274,plain,(
  ( ! [X619,X618] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X618,X619) | ~p(s(bool,X618)) | ~sP8(X618,X619,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1123,f1747])).
fof(f6273,plain,(
  ( ! [X617,X616] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X616,X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X617)) | ~sP8(X617,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X616)) ) | $spl20),
  inference(superposition,[],[f1122,f1747])).
fof(f6272,plain,(
  ( ! [X615,X614] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP9(X614,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615) | ~p(s(bool,X615)) | ~sP8('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X615,X614)) ) | $spl20),
  inference(superposition,[],[f1122,f1747])).
fof(f6271,plain,(
  ( ! [X613,X612] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X612)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X612,X613)) ) | $spl20),
  inference(superposition,[],[f1121,f1747])).
fof(f6270,plain,(
  ( ! [X611,X610] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X610)) | ~sP9(X610,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X611)) ) | $spl20),
  inference(superposition,[],[f1121,f1747])).
fof(f6269,plain,(
  ( ! [X609,X608] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X608)) | ~sP9('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X609,X608)) ) | $spl20),
  inference(superposition,[],[f1120,f1747])).
fof(f6268,plain,(
  ( ! [X606,X607] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X606)) | ~sP9(X606,X607,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1120,f1747])).
fof(f6267,plain,(
  ( ! [X605,X604] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X604)) | ~sP7(X604,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X605)) ) | $spl20),
  inference(superposition,[],[f1118,f1747])).
fof(f6266,plain,(
  ( ! [X602,X603] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X602)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X602,X603)) ) | $spl20),
  inference(superposition,[],[f1118,f1747])).
fof(f6265,plain,(
  ( ! [X600,X601] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X600)) | ~sP7(X600,X601,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1117,f1747])).
fof(f6264,plain,(
  ( ! [X598,X599] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X598)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X599,X598)) ) | $spl20),
  inference(superposition,[],[f1117,f1747])).
fof(f6263,plain,(
  ( ! [X596,X597] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X596)) | p(s(bool,X597)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X596,X597)) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f6262,plain,(
  ( ! [X594,X595] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X594)) | ~p(s(bool,X595)) | ~sP7(X595,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X594)) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f6261,plain,(
  ( ! [X592,X593] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X592)) | ~p(s(bool,X593)) | ~sP7(X593,X592,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1116,f1747])).
fof(f6260,plain,(
  ( ! [X590,X591] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X590,X591) | ~p(s(bool,X590))) ) | $spl20),
  inference(superposition,[],[f1115,f1747])).
fof(f6259,plain,(
  ( ! [X588,X589] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X588,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X589) | ~p(s(bool,X588))) ) | $spl20),
  inference(superposition,[],[f1115,f1747])).
fof(f6258,plain,(
  ( ! [X586,X587] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X586,X587) | ~p(s(bool,X587))) ) | $spl20),
  inference(superposition,[],[f1114,f1747])).
fof(f6257,plain,(
  ( ! [X584,X585] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X584,X585,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X584))) ) | $spl20),
  inference(superposition,[],[f1114,f1747])).
fof(f6256,plain,(
  ( ! [X582,X583] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X582,X583) | p(s(bool,X582)) | p(s(bool,X583))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f6255,plain,(
  ( ! [X580,X581] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X580,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X581) | p(s(bool,X581)) | p(s(bool,X580))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f6254,plain,(
  ( ! [X578,X579] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X578,X579,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X579)) | p(s(bool,X578))) ) | $spl20),
  inference(superposition,[],[f1113,f1747])).
fof(f6253,plain,(
  ( ! [X576,X577] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X576,X577)) ) | $spl20),
  inference(superposition,[],[f1112,f1747])).
fof(f6252,plain,(
  ( ! [X575,X574] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6(X574,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X575)) ) | $spl20),
  inference(superposition,[],[f1111,f1747])).
fof(f6251,plain,(
  ( ! [X573,X572] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP6(X572,X573,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1110,f1747])).
fof(f6250,plain,(
  ( ! [X571,X570] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X570)) | p(s(bool,X571)) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X570,X571)) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f6249,plain,(
  ( ! [X569,X568] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X568)) | p(s(bool,X569)) | ~sP6(X568,X569,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f6248,plain,(
  ( ! [X567,X566] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X566)) | p(s(bool,X567)) | ~sP6(X566,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X567)) ) | $spl20),
  inference(superposition,[],[f1109,f1747])).
fof(f6247,plain,(
  ( ! [X565,X564] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP6(X564,X565,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X565))) ) | $spl20),
  inference(superposition,[],[f1108,f1747])).
fof(f6246,plain,(
  ( ! [X563,X562] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | ~sP6(X562,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X563) | p(s(bool,X563))) ) | $spl20),
  inference(superposition,[],[f1108,f1747])).
fof(f6245,plain,(
  ( ! [X561,X560] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X560,X561) | p(s(bool,X560))) ) | $spl20),
  inference(superposition,[],[f1107,f1747])).
fof(f6243,plain,(
  ( ! [X557,X556] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | ~sP6('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X556,X557) | p(s(bool,X557))) ) | $spl20),
  inference(superposition,[],[f1106,f1747])).
fof(f6241,plain,(
  ( ! [X553,X552] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X552)) | ~sP7(X552,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X553)) ) | $spl20),
  inference(superposition,[],[f1103,f1747])).
fof(f6240,plain,(
  ( ! [X551,X550] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X550)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X550,X551)) ) | $spl20),
  inference(superposition,[],[f1103,f1747])).
fof(f6239,plain,(
  ( ! [X549,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X548)) | ~sP7(X548,X549,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1102,f1747])).
fof(f6238,plain,(
  ( ! [X546,X547] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X546)) | ~sP7('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X547,X546)) ) | $spl20),
  inference(superposition,[],[f1102,f1747])).
fof(f6237,plain,(
  ( ! [X545,X544] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X544)) | ~sP5(X544,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X545)) ) | $spl20),
  inference(superposition,[],[f1101,f1747])).
fof(f6236,plain,(
  ( ! [X542,X543] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X542)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X542,X543)) ) | $spl20),
  inference(superposition,[],[f1101,f1747])).
fof(f6235,plain,(
  ( ! [X540,X541] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X540)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X541,X540)) ) | $spl20),
  inference(superposition,[],[f1100,f1747])).
fof(f6234,plain,(
  ( ! [X538,X539] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X538)) | ~sP5(X538,X539,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1100,f1747])).
fof(f6233,plain,(
  ( ! [X536,X537] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X536)) | ~p(s(bool,X537)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X536,X537)) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f6232,plain,(
  ( ! [X534,X535] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X534)) | ~p(s(bool,X535)) | ~sP5(X535,X534,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f6231,plain,(
  ( ! [X532,X533] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X532)) | ~p(s(bool,X533)) | ~sP5(X533,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X532)) ) | $spl20),
  inference(superposition,[],[f1099,f1747])).
fof(f6230,plain,(
  ( ! [X530,X531] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X530,X531) | ~p(s(bool,X530))) ) | $spl20),
  inference(superposition,[],[f1098,f1747])).
fof(f6229,plain,(
  ( ! [X528,X529] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X528,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X529) | ~p(s(bool,X528))) ) | $spl20),
  inference(superposition,[],[f1098,f1747])).
fof(f6228,plain,(
  ( ! [X526,X527] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X526,X527) | p(s(bool,X527))) ) | $spl20),
  inference(superposition,[],[f1097,f1747])).
fof(f6227,plain,(
  ( ! [X524,X525] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X524,X525,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,X524))) ) | $spl20),
  inference(superposition,[],[f1097,f1747])).
fof(f6226,plain,(
  ( ! [X522,X523] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X522,X523,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X523)) | p(s(bool,X522))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f6225,plain,(
  ( ! [X520,X521] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X520,X521) | p(s(bool,X520)) | ~p(s(bool,X521))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f6224,plain,(
  ( ! [X518,X519] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X518,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X519) | ~p(s(bool,X519)) | p(s(bool,X518))) ) | $spl20),
  inference(superposition,[],[f1096,f1747])).
fof(f6223,plain,(
  ( ! [X516,X517] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X516,X517)) ) | $spl20),
  inference(superposition,[],[f1095,f1747])).
fof(f6222,plain,(
  ( ! [X514,X515] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4(X514,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X515)) ) | $spl20),
  inference(superposition,[],[f1094,f1747])).
fof(f6221,plain,(
  ( ! [X512,X513] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP4(X512,X513,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1093,f1747])).
fof(f6220,plain,(
  ( ! [X511,X510] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X510)) | p(s(bool,X511)) | ~sP4(X510,X511,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f6219,plain,(
  ( ! [X509,X508] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X508)) | ~p(s(bool,X509)) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X508,X509)) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f6218,plain,(
  ( ! [X507,X506] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X506)) | ~p(s(bool,X507)) | ~sP4(X506,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X507)) ) | $spl20),
  inference(superposition,[],[f1092,f1747])).
fof(f6217,plain,(
  ( ! [X505,X504] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP4(X504,X505,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | p(s(bool,X505))) ) | $spl20),
  inference(superposition,[],[f1091,f1747])).
fof(f6216,plain,(
  ( ! [X502,X503] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~sP4(X502,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X503) | ~p(s(bool,X503))) ) | $spl20),
  inference(superposition,[],[f1091,f1747])).
fof(f6215,plain,(
  ( ! [X500,X501] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X500,X501) | p(s(bool,X500))) ) | $spl20),
  inference(superposition,[],[f1090,f1747])).
fof(f6212,plain,(
  ( ! [X494,X495] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~sP4('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X494,X495) | ~p(s(bool,X495))) ) | $spl20),
  inference(superposition,[],[f1089,f1747])).
fof(f6211,plain,(
  ( ! [X492,X493] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X492)) | ~sP5(X492,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X493)) ) | $spl20),
  inference(superposition,[],[f1086,f1747])).
fof(f6210,plain,(
  ( ! [X490,X491] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X490)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X490,X491)) ) | $spl20),
  inference(superposition,[],[f1086,f1747])).
fof(f6209,plain,(
  ( ! [X488,X489] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X488)) | ~sP5(X488,X489,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1085,f1747])).
fof(f6208,plain,(
  ( ! [X486,X487] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X486)) | ~sP5('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X487,X486)) ) | $spl20),
  inference(superposition,[],[f1085,f1747])).
fof(f6173,plain,(
  ( ! [X173,X172] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | s(bool,X172) != s(bool,X173) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X173,X172)) ) | $spl20),
  inference(superposition,[],[f1041,f1747])).
fof(f6170,plain,(
  ( ! [X167,X166] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | s(bool,X166) = s(bool,X167) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X167,X166)) ) | $spl20),
  inference(superposition,[],[f1040,f1747])).
fof(f6167,plain,(
  ( ! [X161,X160] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X160,X161) | s(bool,X160) != s(bool,X161)) ) | $spl20),
  inference(superposition,[],[f1039,f1747])).
fof(f6164,plain,(
  ( ! [X154,X155] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X154,X155) | s(bool,X154) = s(bool,X155)) ) | $spl20),
  inference(superposition,[],[f1038,f1747])).
fof(f6161,plain,(
  ( ! [X149,X148] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X148,X149)) ) | $spl20),
  inference(superposition,[],[f1037,f1747])).
fof(f6160,plain,(
  ( ! [X146,X147] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0(X146,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X147)) ) | $spl20),
  inference(superposition,[],[f1036,f1747])).
fof(f6159,plain,(
  ( ! [X144,X145] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP0(X144,X145,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1035,f1747])).
fof(f6158,plain,(
  ( ! [X142,X143] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X142)) | p(s(bool,X143)) | ~sP0(X142,X143,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f6157,plain,(
  ( ! [X140,X141] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X140)) | p(s(bool,X141)) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X140,X141)) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f6156,plain,(
  ( ! [X138,X139] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | p(s(bool,X138)) | p(s(bool,X139)) | ~sP0(X138,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X139)) ) | $spl20),
  inference(superposition,[],[f1034,f1747])).
fof(f6155,plain,(
  ( ! [X136,X137] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X136,X137)) ) | $spl20),
  inference(superposition,[],[f1033,f1747])).
fof(f6154,plain,(
  ( ! [X134,X135] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1(X134,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X135)) ) | $spl20),
  inference(superposition,[],[f1032,f1747])).
fof(f6153,plain,(
  ( ! [X132,X133] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP1(X132,X133,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1031,f1747])).
fof(f6152,plain,(
  ( ! [X130,X131] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X130)) | p(s(bool,X131)) | ~sP1(X130,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X131)) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f6151,plain,(
  ( ! [X128,X129] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X128)) | p(s(bool,X129)) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X128,X129)) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f6150,plain,(
  ( ! [X127,X126] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | ~sP1(X126,X127,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1030,f1747])).
fof(f6149,plain,(
  ( ! [X125,X124] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X124,X125)) ) | $spl20),
  inference(superposition,[],[f1029,f1747])).
fof(f6148,plain,(
  ( ! [X123,X122] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2(X122,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X123)) ) | $spl20),
  inference(superposition,[],[f1028,f1747])).
fof(f6147,plain,(
  ( ! [X121,X120] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP2(X120,X121,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1027,f1747])).
fof(f6146,plain,(
  ( ! [X118,X119] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X118)) | p(s(bool,X119)) | ~sP2(X118,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X119)) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f6145,plain,(
  ( ! [X116,X117] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X116)) | p(s(bool,X117)) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X116,X117)) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f6144,plain,(
  ( ! [X114,X115] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X114)) | ~p(s(bool,X115)) | ~sP2(X114,X115,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1026,f1747])).
fof(f6143,plain,(
  ( ! [X112,X113] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X112,X113) | ~sP2('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X113,X112) | ~sP1(X112,X113,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP0(X113,X112,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1025,f1747])).
fof(f6142,plain,(
  ( ! [X111,X110] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3(X110,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111) | ~sP2(X110,X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP1('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X111,X110) | ~sP0(X111,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X110)) ) | $spl20),
  inference(superposition,[],[f1024,f1747])).
fof(f6141,plain,(
  ( ! [X109,X108] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | sP3(X108,X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~sP2(X108,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109) | ~sP1(X109,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X108) | ~sP0('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X109,X108)) ) | $spl20),
  inference(superposition,[],[f1023,f1747])).
fof(f6140,plain,(
  ( ! [X107,X106] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X106)) | p(s(bool,X107)) | ~sP3(X106,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X107)) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f6139,plain,(
  ( ! [X105,X104] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X104)) | p(s(bool,X105)) | ~sP3('const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))),X104,X105)) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f6138,plain,(
  ( ! [X103,X102] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21)))) | ~p(s(bool,X102)) | ~p(s(bool,X103)) | ~sP3(X102,X103,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13))))) ) | $spl20),
  inference(superposition,[],[f1022,f1747])).
fof(f5831,plain,(
  $false | ($spl42 | $spl99)),
  inference(subsumption_resolution,[],[f4877,f1763])).
fof(f4877,plain,(
  ~p(s(bool,sK25)) | $spl99),
  inference(subsumption_resolution,[],[f4873,f1250])).
fof(f4873,plain,(
  ~p(s(bool,sK25)) | p(s(bool,'F_0')) | $spl99),
  inference(resolution,[],[f3264,f1174])).
fof(f3264,plain,(
  ~sP11('F_0',sK25) | $spl99),
  inference(cnf_transformation,[],[f3264_D])).
fof(f3264_D,plain,(
  ~sP11('F_0',sK25) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f5129,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK46) != s(bool,X1)) ) | $spl166),
  inference(cnf_transformation,[],[f5129_D])).
fof(f5129_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK46) != s(bool,X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f5798,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,sK46) != s(bool,X181)) ) | $spl123),
  inference(resolution,[],[f3278,f1424])).
fof(f5698,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK46) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl123),
  inference(resolution,[],[f3278,f1018])).
fof(f5696,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK46) != s(bool,X16)) ) | $spl123),
  inference(resolution,[],[f3278,f1017])).
fof(f5694,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK46) != s(bool,X10)) ) | $spl123),
  inference(resolution,[],[f3278,f1016])).
fof(f5693,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK46) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl123),
  inference(resolution,[],[f3278,f1014])).
fof(f5691,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK46) != s(bool,X1)) ) | $spl123),
  inference(resolution,[],[f3278,f1012])).
fof(f5685,plain,(
  $false | $spl122),
  inference(subsumption_resolution,[],[f5554,f1336])).
fof(f1336,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f661])).
fof(f661,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2',aHL_FALSITY)).
fof(f5554,plain,(
  p(s(bool,'F_0')) | $spl122),
  inference(resolution,[],[f3277,f1344])).
fof(f1344,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,sK46))),
  inference(cnf_transformation,[],[f983])).
fof(f983,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK46)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46])],[f982])).
fof(f982,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f981])).
fof(f981,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f664])).
fof(f664,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f317])).
fof(f317,axiom,(
  p(s(bool,'F_0')) <=> ! [X39] : p(s(bool,X39))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.F_DEF')).
fof(f3277,plain,(
  p(s(bool,sK46)) | $spl122),
  inference(cnf_transformation,[],[f3277_D])).
fof(f3277_D,plain,(
  p(s(bool,sK46)) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2387,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK25) != s(bool,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f2387_D])).
fof(f2387_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,sK25) != s(bool,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5265,plain,(
  ( ! [X191,X190] : (p(s(bool,X190)) | p(s(bool,X191)) | s(bool,sK25) != s(bool,X191)) ) | $spl42),
  inference(resolution,[],[f1763,f1424])).
fof(f5165,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl42),
  inference(resolution,[],[f1763,f1018])).
fof(f5162,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK25) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f1763,f1017])).
fof(f5161,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl42),
  inference(resolution,[],[f1763,f1016])).
fof(f5160,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl42),
  inference(resolution,[],[f1763,f1014])).
fof(f5158,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl42),
  inference(resolution,[],[f1763,f1012])).
fof(f5104,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,sK46) != s(bool,X181)) ) | $spl123),
  inference(resolution,[],[f3278,f1424])).
fof(f5004,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK46) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl123),
  inference(resolution,[],[f3278,f1018])).
fof(f5002,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK46) != s(bool,X16)) ) | $spl123),
  inference(resolution,[],[f3278,f1017])).
fof(f5000,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK46) != s(bool,X10)) ) | $spl123),
  inference(resolution,[],[f3278,f1016])).
fof(f4999,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK46) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl123),
  inference(resolution,[],[f3278,f1014])).
fof(f4997,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK46) != s(bool,X1)) ) | $spl123),
  inference(resolution,[],[f3278,f1012])).
fof(f4943,plain,(
  $false | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f1744,f1434])).
fof(f4919,plain,(
  $false | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f1737,f1434])).
fof(f4885,plain,(
  $false | ($spl0 | $spl22)),
  inference(subsumption_resolution,[],[f1749,f1434])).
fof(f4872,plain,(
  $false | ($spl43 | $spl98)),
  inference(subsumption_resolution,[],[f4871,f1762])).
fof(f4871,plain,(
  p(s(bool,sK25)) | $spl98),
  inference(subsumption_resolution,[],[f4867,f1343])).
fof(f1343,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f983])).
fof(f4867,plain,(
  p(s(bool,sK25)) | p(s(bool,'F_0')) | $spl98),
  inference(resolution,[],[f3265,f1168])).
fof(f4870,plain,(
  $false | ($spl43 | $spl98)),
  inference(subsumption_resolution,[],[f4869,f1762])).
fof(f4869,plain,(
  p(s(bool,sK25)) | $spl98),
  inference(subsumption_resolution,[],[f4866,f1343])).
fof(f4866,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK25)) | $spl98),
  inference(resolution,[],[f3265,f1177])).
fof(f1913,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK25) != s(bool,X1)) ) | $spl62),
  inference(cnf_transformation,[],[f1913_D])).
fof(f1913_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,sK25) != s(bool,X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4676,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,sK25) != s(bool,X181)) ) | $spl43),
  inference(resolution,[],[f1762,f1424])).
fof(f4576,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl43),
  inference(resolution,[],[f1762,f1018])).
fof(f4574,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK25) != s(bool,X16)) ) | $spl43),
  inference(resolution,[],[f1762,f1017])).
fof(f1914,plain,(
  ( ! [X5] : (s(bool,sK25) != s(bool,X5)) ) | $spl64),
  inference(cnf_transformation,[],[f1914_D])).
fof(f1914_D,plain,(
  ( ! [X5] : (s(bool,sK25) != s(bool,X5)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1756,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl34),
  inference(cnf_transformation,[],[f1756_D])).
fof(f1756_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4573,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,sK25) != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f4572,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f4571,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl43),
  inference(resolution,[],[f1762,f1014])).
fof(f4570,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK25) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f4569,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f4534,plain,(
  sP11(sK25,'F_0') | $spl140),
  inference(cnf_transformation,[],[f4534_D])).
fof(f4534_D,plain,(
  sP11(sK25,'F_0') <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f4522,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP11(sK25,'F_0')) ) | $spl42),
  inference(duplicate_literal_removal,[],[f4495])).
fof(f4495,plain,(
  ( ! [X30] : (p(s(bool,X30)) | sP11(sK25,'F_0') | p(s(bool,X30))) ) | $spl42),
  inference(superposition,[],[f3372,f1252])).
fof(f3372,plain,(
  ( ! [X162] : (p(s(bool,X162)) | sP11(sK25,X162)) ) | $spl42),
  inference(resolution,[],[f1763,f1170])).
fof(f4523,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP11(sK25,'F_0')) ) | $spl42),
  inference(duplicate_literal_removal,[],[f4494])).
fof(f4494,plain,(
  ( ! [X29] : (p(s(bool,X29)) | sP11(sK25,'F_0') | p(s(bool,X29))) ) | $spl42),
  inference(superposition,[],[f3372,f1322])).
fof(f4524,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP11(sK25,'F_0')) ) | $spl42),
  inference(duplicate_literal_removal,[],[f4493])).
fof(f4493,plain,(
  ( ! [X28] : (p(s(bool,X28)) | sP11(sK25,'F_0') | p(s(bool,X28))) ) | $spl42),
  inference(superposition,[],[f3372,f1324])).
fof(f4391,plain,(
  ( ! [X343] : (sP11(sK25,'F_0') | p(s(bool,X343))) ) | $spl42),
  inference(resolution,[],[f3372,f1343])).
fof(f4387,plain,(
  ( ! [X342] : (sP11(sK25,'F_0') | p(s(bool,X342))) ) | $spl42),
  inference(resolution,[],[f3372,f1334])).
fof(f1334,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f784])).
fof(f784,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f658])).
fof(f658,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X39] : (p(s(bool,'F_0')) => p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.FALSITY')).
fof(f4547,plain,(
  ( ! [X339] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X339)))) ) | $spl164),
  inference(cnf_transformation,[],[f4547_D])).
fof(f4547_D,plain,(
  ( ! [X339] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X339)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4386,plain,(
  ( ! [X341] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X341)))) ) | $spl42),
  inference(resolution,[],[f3372,f1320])).
fof(f1320,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f973])).
fof(f973,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f180])).
fof(f180,axiom,(
  ! [X76] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|2')).
fof(f4546,plain,(
  ( ! [X338] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X338)))) ) | $spl162),
  inference(cnf_transformation,[],[f4546_D])).
fof(f4546_D,plain,(
  ( ! [X338] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X338)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f4385,plain,(
  ( ! [X340] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X340)))) ) | $spl42),
  inference(resolution,[],[f3372,f1318])).
fof(f1318,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f972])).
fof(f972,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f646])).
fof(f646,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X76] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|0')).
fof(f4384,plain,(
  ( ! [X339] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X339)))) ) | $spl42),
  inference(resolution,[],[f3372,f1316])).
fof(f1316,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f971])).
fof(f971,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f141])).
fof(f141,axiom,(
  ! [X76] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|3')).
fof(f4383,plain,(
  ( ! [X338] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X338)))) ) | $spl42),
  inference(resolution,[],[f3372,f1314])).
fof(f1314,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f970])).
fof(f970,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f159])).
fof(f159,axiom,(
  ! [X76] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|1')).
fof(f4544,plain,(
  ( ! [X334] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X334)))))))) ) | $spl158),
  inference(cnf_transformation,[],[f4544_D])).
fof(f4544_D,plain,(
  ( ! [X334] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X334)))))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4382,plain,(
  ( ! [X337] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X337)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1312])).
fof(f1312,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f969])).
fof(f969,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f431])).
fof(f431,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|5')).
fof(f4545,plain,(
  ( ! [X335] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X335)))))))) ) | $spl160),
  inference(cnf_transformation,[],[f4545_D])).
fof(f4545_D,plain,(
  ( ! [X335] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X335)))))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f4381,plain,(
  ( ! [X336] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X336)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1310])).
fof(f1310,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f968])).
fof(f968,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f390])).
fof(f390,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|2')).
fof(f4380,plain,(
  ( ! [X335] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X335)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1308])).
fof(f1308,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f967])).
fof(f967,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f377])).
fof(f377,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|6')).
fof(f4379,plain,(
  ( ! [X334] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X334)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1306])).
fof(f1306,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f966])).
fof(f966,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f640])).
fof(f640,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f386])).
fof(f386,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|1')).
fof(f4542,plain,(
  ( ! [X330] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))))))) ) | $spl154),
  inference(cnf_transformation,[],[f4542_D])).
fof(f4542_D,plain,(
  ( ! [X330] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4378,plain,(
  ( ! [X333] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X333)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1304])).
fof(f1304,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f965])).
fof(f965,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f378])).
fof(f378,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|7')).
fof(f4543,plain,(
  ( ! [X331] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X331)))))))))) ) | $spl156),
  inference(cnf_transformation,[],[f4543_D])).
fof(f4543_D,plain,(
  ( ! [X331] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X331)))))))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f4377,plain,(
  ( ! [X332] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X332)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1302])).
fof(f1302,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f964])).
fof(f964,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f638])).
fof(f638,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f422])).
fof(f422,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|4')).
fof(f4376,plain,(
  ( ! [X331] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X331)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1300])).
fof(f1300,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f963])).
fof(f963,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f381])).
fof(f381,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|8')).
fof(f4375,plain,(
  ( ! [X330] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1298])).
fof(f1298,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f962])).
fof(f962,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f404])).
fof(f404,axiom,(
  ! [X76] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|3')).
fof(f4374,plain,(
  ( ! [X329] : (sP11(sK25,'F_0') | p(s(bool,X329))) ) | $spl42),
  inference(resolution,[],[f3372,f1295])).
fof(f1295,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f961])).
fof(f961,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f960])).
fof(f960,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f635])).
fof(f635,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X39] : ((p(s(bool,X39)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.AND_CLAUSES|split|3')).
fof(f4373,plain,(
  ( ! [X328] : (sP11(sK25,'F_0') | p(s(bool,X328))) ) | $spl42),
  inference(resolution,[],[f3372,f1284])).
fof(f1284,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f953])).
fof(f953,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f952])).
fof(f952,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X39] : ((p(s(bool,'F_0')) & p(s(bool,X39))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.AND_CLAUSES|split|2')).
fof(f4372,plain,(
  ( ! [X327] : (sP11(sK25,'F_0') | p(s(bool,X327))) ) | $spl42),
  inference(resolution,[],[f3372,f1280])).
fof(f1280,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f951])).
fof(f951,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f950])).
fof(f950,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f630])).
fof(f630,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X39] : ((p(s(bool,X39)) | p(s(bool,'F_0'))) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.OR_CLAUSES|split|3')).
fof(f4371,plain,(
  ( ! [X326] : (sP11(sK25,'F_0') | p(s(bool,X326))) ) | $spl42),
  inference(resolution,[],[f3372,f1270])).
fof(f4541,plain,(
  ( ! [X325] : (sP11(sK25,'F_0') | p(s(bool,X325))) ) | $spl42),
  inference(subsumption_resolution,[],[f4370,f1429])).
fof(f4370,plain,(
  ( ! [X325] : (sP11(sK25,'F_0') | p(s(bool,X325)) | ~p(s(bool,'T_0'))) ) | $spl42),
  inference(resolution,[],[f3372,f1257])).
fof(f1257,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f935])).
fof(f935,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f934])).
fof(f934,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f776])).
fof(f776,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X39] : ((p(s(bool,'F_0')) => p(s(bool,X39))) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CLAUSES|split|2')).
fof(f4540,plain,(
  ( ! [X324] : (s(bool,'F_0') = s(bool,X324)) ) | $spl152),
  inference(cnf_transformation,[],[f4540_D])).
fof(f4540_D,plain,(
  ( ! [X324] : (s(bool,'F_0') = s(bool,X324)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f4369,plain,(
  ( ! [X324] : (sP11(sK25,'F_0') | s(bool,'F_0') = s(bool,X324)) ) | $spl42),
  inference(resolution,[],[f3372,f1253])).
fof(f1253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f933])).
fof(f4368,plain,(
  ( ! [X323] : (sP11(sK25,'F_0') | ~p(s(bool,X323))) ) | $spl42),
  inference(resolution,[],[f3372,f1250])).
fof(f4539,plain,(
  ( ! [X320,X319] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X319))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X320)))) ) | $spl150),
  inference(cnf_transformation,[],[f4539_D])).
fof(f4539_D,plain,(
  ( ! [X320,X319] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X319))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X320)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4367,plain,(
  ( ! [X321,X322] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X321))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X322)))) ) | $spl42),
  inference(resolution,[],[f3372,f1211])).
fof(f1211,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f897])).
fof(f897,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f169])).
fof(f169,axiom,(
  ! [X76,X13] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|4')).
fof(f4366,plain,(
  ( ! [X320,X319] : (sP11(sK25,'F_0') | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X319))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X320)))) ) | $spl42),
  inference(resolution,[],[f3372,f1209])).
fof(f1209,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f896])).
fof(f896,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f590])).
fof(f590,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f165])).
fof(f165,axiom,(
  ! [X76,X13] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_eq|split|5')).
fof(f4538,plain,(
  ( ! [X318,X317] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X317))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X318)))))))))) ) | $spl148),
  inference(cnf_transformation,[],[f4538_D])).
fof(f4538_D,plain,(
  ( ! [X318,X317] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X317))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X318)))))))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f4365,plain,(
  ( ! [X318,X317] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X317))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X318)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1204])).
fof(f1204,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f893])).
fof(f893,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f375])).
fof(f375,axiom,(
  ! [X76,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|13')).
fof(f4537,plain,(
  ( ! [X315,X316] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X315))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316)))))))) ) | $spl146),
  inference(cnf_transformation,[],[f4537_D])).
fof(f4537_D,plain,(
  ( ! [X315,X316] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X315))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316)))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f4364,plain,(
  ( ! [X315,X316] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X315))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1202])).
fof(f1202,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f892])).
fof(f892,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))),
  inference(nnf_transformation,[],[f586])).
fof(f586,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f374])).
fof(f374,axiom,(
  ! [X76,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|10')).
fof(f4536,plain,(
  ( ! [X313,X314] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X313))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X314)))))))))) ) | $spl144),
  inference(cnf_transformation,[],[f4536_D])).
fof(f4536_D,plain,(
  ( ! [X313,X314] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X313))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X314)))))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4363,plain,(
  ( ! [X313,X314] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X313))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X314)))))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1200])).
fof(f1200,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f891])).
fof(f891,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))))),
  inference(nnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f373])).
fof(f373,axiom,(
  ! [X76,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X76))))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|12')).
fof(f4535,plain,(
  ( ! [X312,X311] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X311))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X312)))))))) ) | $spl142),
  inference(cnf_transformation,[],[f4535_D])).
fof(f4535_D,plain,(
  ( ! [X312,X311] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X311))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X312)))))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f4362,plain,(
  ( ! [X312,X311] : (sP11(sK25,'F_0') | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X311))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X312)))))))) ) | $spl42),
  inference(resolution,[],[f3372,f1198])).
fof(f1198,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f890])).
fof(f890,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))),
  inference(nnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f372])).
fof(f372,axiom,(
  ! [X76,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X76))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))))) <=> p(s(bool,'F_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_EQ_REDUCE|split|11')).
fof(f4532,plain,(
  ( ! [X308,X307] : (sP12(X307,X308)) ) | $spl138),
  inference(cnf_transformation,[],[f4532_D])).
fof(f4532_D,plain,(
  ( ! [X308,X307] : (sP12(X307,X308)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f4361,plain,(
  ( ! [X310,X309] : (sP11(sK25,'F_0') | sP12(X309,X310)) ) | $spl42),
  inference(resolution,[],[f3372,f1192])).
fof(f1192,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f887])).
fof(f4360,plain,(
  ( ! [X308,X307] : (sP11(sK25,'F_0') | sP12(X307,X308)) ) | $spl42),
  inference(resolution,[],[f3372,f1187])).
fof(f1187,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f885])).
fof(f4354,plain,(
  ( ! [X298,X297,X299] : (sP11(sK25,X297) | p(s(bool,X298)) | p(s(bool,X299)) | s(bool,X297) != s(bool,X299)) ) | $spl42),
  inference(resolution,[],[f3372,f1424])).
fof(f4254,plain,(
  ( ! [X30,X28,X31,X29] : (sP11(sK25,X28) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl42),
  inference(resolution,[],[f3372,f1018])).
fof(f4528,plain,(
  ( ! [X4,X7] : (sP11(sK25,X4) | s(bool,X4) != s(bool,X7)) ) | $spl136),
  inference(cnf_transformation,[],[f4528_D])).
fof(f4528_D,plain,(
  ( ! [X4,X7] : (sP11(sK25,X4) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f4253,plain,(
  ( ! [X26,X24,X27,X25] : (sP11(sK25,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27)) ) | $spl42),
  inference(resolution,[],[f3372,f1018])).
fof(f4251,plain,(
  ( ! [X19,X17,X18,X16] : (sP11(sK25,X16) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl42),
  inference(resolution,[],[f3372,f1017])).
fof(f4250,plain,(
  ( ! [X14,X12,X15,X13] : (sP11(sK25,X12) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f3372,f1016])).
fof(f4249,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(sK25,X8) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl42),
  inference(resolution,[],[f3372,f1014])).
fof(f4248,plain,(
  ( ! [X6,X4,X7,X5] : (sP11(sK25,X4) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X4) != s(bool,X7)) ) | $spl42),
  inference(resolution,[],[f3372,f1014])).
fof(f4247,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(sK25,X0) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl42),
  inference(resolution,[],[f3372,f1012])).
fof(f4187,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) | $spl132),
  inference(cnf_transformation,[],[f4187_D])).
fof(f4187_D,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4183,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f4104])).
fof(f4104,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X38))) ) | $spl36),
  inference(superposition,[],[f1757,f1328])).
fof(f4184,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f4103])).
fof(f4103,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | ~p(s(bool,X36))) ) | $spl36),
  inference(superposition,[],[f1757,f1326])).
fof(f2985,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl87),
  inference(cnf_transformation,[],[f2985_D])).
fof(f2985_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f4080,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) ) | $spl53),
  inference(superposition,[],[f2978,f1207])).
fof(f2978,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK18) | $spl53),
  inference(trivial_inequality_removal,[],[f2953])).
fof(f2953,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK18) | $spl53),
  inference(superposition,[],[f1771,f1058])).
fof(f1058,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f831])).
fof(f831,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X76] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X76))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_distrib|split|18')).
fof(f4079,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) | $spl53),
  inference(superposition,[],[f2978,f1206])).
fof(f4043,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl130),
  inference(cnf_transformation,[],[f4043_D])).
fof(f4043_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4038,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f3958])).
fof(f3958,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl30),
  inference(superposition,[],[f1754,f1324])).
fof(f4039,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f3957])).
fof(f3957,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl30),
  inference(superposition,[],[f1754,f1322])).
fof(f4040,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f3955])).
fof(f3955,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl30),
  inference(superposition,[],[f1754,f1252])).
fof(f3576,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl124),
  inference(cnf_transformation,[],[f3576_D])).
fof(f3576_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f3905,plain,(
  ( ! [X198,X199] : (p(s(bool,X198)) | p(s(bool,X199)) | s(bool,X199) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1424])).
fof(f3805,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X30) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f1751,f1018])).
fof(f3577,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl126),
  inference(cnf_transformation,[],[f3577_D])).
fof(f3577_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f3804,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1018])).
fof(f3802,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1017])).
fof(f3801,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1016])).
fof(f3800,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f1751,f1014])).
fof(f3799,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1014])).
fof(f3798,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1012])).
fof(f3763,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl128),
  inference(cnf_transformation,[],[f3763_D])).
fof(f3763_D,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f3724,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | ~p(s(bool,X190)) | s(bool,X190) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl25),
  inference(resolution,[],[f1750,f1424])).
fof(f1750,plain,(
  ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) | $spl25),
  inference(cnf_transformation,[],[f1750_D])).
fof(f1750_D,plain,(
  ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f3624,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X30) != s(bool,X32)) ) | $spl25),
  inference(resolution,[],[f1750,f1018])).
fof(f3622,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl25),
  inference(resolution,[],[f1750,f1017])).
fof(f3621,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X23) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X21) != s(bool,X22)) ) | $spl25),
  inference(resolution,[],[f1750,f1016])).
fof(f3620,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl25),
  inference(resolution,[],[f1750,f1016])).
fof(f3619,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X15) != s(bool,X17)) ) | $spl25),
  inference(resolution,[],[f1750,f1014])).
fof(f3618,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X12) != s(bool,X13)) ) | $spl25),
  inference(resolution,[],[f1750,f1012])).
fof(f3617,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl25),
  inference(resolution,[],[f1750,f1012])).
fof(f3539,plain,(
  ( ! [X198,X199] : (p(s(bool,X198)) | p(s(bool,X199)) | s(bool,X199) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1424])).
fof(f3439,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X30) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f1751,f1018])).
fof(f3438,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X27) | s(bool,X28) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1018])).
fof(f3436,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1017])).
fof(f3435,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1016])).
fof(f3434,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22))) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f1751,f1014])).
fof(f3433,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X12) | s(bool,X13) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1014])).
fof(f3432,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) ) | $spl24),
  inference(resolution,[],[f1751,f1012])).
fof(f3400,plain,(
  ( ! [X191,X190] : (p(s(bool,X190)) | p(s(bool,X191)) | s(bool,sK25) != s(bool,X191)) ) | $spl42),
  inference(resolution,[],[f1763,f1424])).
fof(f3300,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl42),
  inference(resolution,[],[f1763,f1018])).
fof(f3299,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK25) != s(bool,X20)) ) | $spl42),
  inference(resolution,[],[f1763,f1018])).
fof(f3297,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK25) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f1763,f1017])).
fof(f3296,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl42),
  inference(resolution,[],[f1763,f1016])).
fof(f3295,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl42),
  inference(resolution,[],[f1763,f1014])).
fof(f3294,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK25) != s(bool,X5)) ) | $spl42),
  inference(resolution,[],[f1763,f1014])).
fof(f3293,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl42),
  inference(resolution,[],[f1763,f1012])).
fof(f3260,plain,(
  sP11('T_0',sK25) | $spl94),
  inference(cnf_transformation,[],[f3260_D])).
fof(f3260_D,plain,(
  sP11('T_0',sK25) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f3245,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP11('T_0',sK25)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f3215])).
fof(f3215,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | sP11('T_0',sK25) | ~p(s(bool,X27))) ) | $spl43),
  inference(superposition,[],[f2473,f1326])).
fof(f2473,plain,(
  ( ! [X164] : (~p(s(bool,X164)) | sP11(X164,sK25)) ) | $spl43),
  inference(resolution,[],[f1762,f1170])).
fof(f3246,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP11('T_0',sK25)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f3214])).
fof(f3214,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | sP11('T_0',sK25) | ~p(s(bool,X26))) ) | $spl43),
  inference(superposition,[],[f2473,f1328])).
fof(f3278,plain,(
  ~p(s(bool,sK46)) | $spl123),
  inference(cnf_transformation,[],[f3278_D])).
fof(f3278_D,plain,(
  ~p(s(bool,sK46)) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f3265,plain,(
  sP11('F_0',sK25) | $spl98),
  inference(cnf_transformation,[],[f3265_D])).
fof(f3265_D,plain,(
  sP11('F_0',sK25) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3150,plain,(
  sP11('F_0',sK25) | ~p(s(bool,sK46)) | $spl43),
  inference(resolution,[],[f2473,f1344])).
fof(f3276,plain,(
  ( ! [X336] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X336)))) ) | $spl120),
  inference(cnf_transformation,[],[f3276_D])).
fof(f3276_D,plain,(
  ( ! [X336] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X336)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f3147,plain,(
  ( ! [X338] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X338)))) ) | $spl43),
  inference(resolution,[],[f2473,f1319])).
fof(f1319,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f973])).
fof(f3275,plain,(
  ( ! [X335] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X335)))) ) | $spl118),
  inference(cnf_transformation,[],[f3275_D])).
fof(f3275_D,plain,(
  ( ! [X335] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X335)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3146,plain,(
  ( ! [X337] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X337)))) ) | $spl43),
  inference(resolution,[],[f2473,f1317])).
fof(f1317,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f972])).
fof(f3145,plain,(
  ( ! [X336] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X336)))) ) | $spl43),
  inference(resolution,[],[f2473,f1315])).
fof(f1315,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f971])).
fof(f3144,plain,(
  ( ! [X335] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X335)))) ) | $spl43),
  inference(resolution,[],[f2473,f1313])).
fof(f1313,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f970])).
fof(f3273,plain,(
  ( ! [X331] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X331)))))))) ) | $spl114),
  inference(cnf_transformation,[],[f3273_D])).
fof(f3273_D,plain,(
  ( ! [X331] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X331)))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f3143,plain,(
  ( ! [X334] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X334)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1311])).
fof(f1311,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f969])).
fof(f3274,plain,(
  ( ! [X332] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X332)))))))) ) | $spl116),
  inference(cnf_transformation,[],[f3274_D])).
fof(f3274_D,plain,(
  ( ! [X332] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X332)))))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f3142,plain,(
  ( ! [X333] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X333)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1309])).
fof(f1309,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f968])).
fof(f3141,plain,(
  ( ! [X332] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X332)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1307])).
fof(f1307,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f967])).
fof(f3140,plain,(
  ( ! [X331] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X331)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1305])).
fof(f1305,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f966])).
fof(f3271,plain,(
  ( ! [X327] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X327)))))))))) ) | $spl110),
  inference(cnf_transformation,[],[f3271_D])).
fof(f3271_D,plain,(
  ( ! [X327] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X327)))))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f3139,plain,(
  ( ! [X330] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X330)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1303])).
fof(f1303,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f965])).
fof(f3272,plain,(
  ( ! [X328] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X328)))))))))) ) | $spl112),
  inference(cnf_transformation,[],[f3272_D])).
fof(f3272_D,plain,(
  ( ! [X328] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X328)))))))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f3138,plain,(
  ( ! [X329] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X329)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1301])).
fof(f1301,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f964])).
fof(f3137,plain,(
  ( ! [X328] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X328)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1299])).
fof(f1299,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f963])).
fof(f3136,plain,(
  ( ! [X327] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X327)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1297])).
fof(f1297,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f962])).
fof(f3270,plain,(
  ( ! [X323,X322] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X322))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X323)))) ) | $spl108),
  inference(cnf_transformation,[],[f3270_D])).
fof(f3270_D,plain,(
  ( ! [X323,X322] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X322))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X323)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f3133,plain,(
  ( ! [X325,X324] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X324))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X325)))) ) | $spl43),
  inference(resolution,[],[f2473,f1210])).
fof(f1210,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f897])).
fof(f3132,plain,(
  ( ! [X323,X322] : (sP11('F_0',sK25) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X322))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X323)))) ) | $spl43),
  inference(resolution,[],[f2473,f1208])).
fof(f1208,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f896])).
fof(f3269,plain,(
  ( ! [X321,X320] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X320))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X321)))))))))) ) | $spl106),
  inference(cnf_transformation,[],[f3269_D])).
fof(f3269_D,plain,(
  ( ! [X321,X320] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X320))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X321)))))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3131,plain,(
  ( ! [X321,X320] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X320))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X321)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1203])).
fof(f1203,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f893])).
fof(f3268,plain,(
  ( ! [X318,X319] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X318))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X319)))))))) ) | $spl104),
  inference(cnf_transformation,[],[f3268_D])).
fof(f3268_D,plain,(
  ( ! [X318,X319] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X318))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X319)))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3130,plain,(
  ( ! [X318,X319] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X318))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X319)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1201])).
fof(f1201,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f892])).
fof(f3267,plain,(
  ( ! [X316,X317] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X317)))))))))) ) | $spl102),
  inference(cnf_transformation,[],[f3267_D])).
fof(f3267_D,plain,(
  ( ! [X316,X317] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X317)))))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f3129,plain,(
  ( ! [X316,X317] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X316))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X317)))))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1199])).
fof(f1199,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f891])).
fof(f3266,plain,(
  ( ! [X315,X314] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X314))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X315)))))))) ) | $spl100),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  ( ! [X315,X314] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X314))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X315)))))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f3128,plain,(
  ( ! [X315,X314] : (sP11('F_0',sK25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X314))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X315)))))))) ) | $spl43),
  inference(resolution,[],[f2473,f1197])).
fof(f1197,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f890])).
fof(f3263,plain,(
  ( ! [X308,X309] : (p(s(bool,X308)) | ~sP12(X308,X309) | p(s(bool,X309))) ) | $spl96),
  inference(cnf_transformation,[],[f3263_D])).
fof(f3263_D,plain,(
  ( ! [X308,X309] : (p(s(bool,X308)) | ~sP12(X308,X309) | p(s(bool,X309))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3126,plain,(
  ( ! [X310,X311] : (sP11('F_0',sK25) | p(s(bool,X310)) | p(s(bool,X311)) | ~sP12(X310,X311)) ) | $spl43),
  inference(resolution,[],[f2473,f1188])).
fof(f3125,plain,(
  ( ! [X308,X309] : (sP11('F_0',sK25) | p(s(bool,X308)) | p(s(bool,X309)) | ~sP12(X308,X309)) ) | $spl43),
  inference(resolution,[],[f2473,f1184])).
fof(f3118,plain,(
  ( ! [X307] : (sP11('T_0',sK25) | ~p(s(bool,X307))) ) | $spl43),
  inference(resolution,[],[f2473,f1290])).
fof(f1290,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f957])).
fof(f957,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f956])).
fof(f956,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f633])).
fof(f633,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X39] : ((p(s(bool,X39)) & p(s(bool,'T_0'))) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.AND_CLAUSES|split|1')).
fof(f3117,plain,(
  ( ! [X306] : (sP11('T_0',sK25) | ~p(s(bool,X306))) ) | $spl43),
  inference(resolution,[],[f2473,f1286])).
fof(f1286,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f955])).
fof(f955,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f954])).
fof(f954,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f632])).
fof(f632,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X39] : ((p(s(bool,'T_0')) & p(s(bool,X39))) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.AND_CLAUSES|split|0')).
fof(f3116,plain,(
  ( ! [X305] : (sP11('T_0',sK25) | ~p(s(bool,X305))) ) | $spl43),
  inference(resolution,[],[f2473,f1277])).
fof(f1277,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f949])).
fof(f949,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f948])).
fof(f948,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X39] : ((p(s(bool,'T_0')) | p(s(bool,X39))) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.OR_CLAUSES|split|0')).
fof(f3115,plain,(
  ( ! [X304] : (sP11('T_0',sK25) | ~p(s(bool,X304))) ) | $spl43),
  inference(resolution,[],[f2473,f1273])).
fof(f1273,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f947])).
fof(f947,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f946])).
fof(f946,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f628])).
fof(f628,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X39] : ((p(s(bool,X39)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.OR_CLAUSES|split|1')).
fof(f3114,plain,(
  ( ! [X303] : (sP11('T_0',sK25) | p(s(bool,X303))) ) | $spl43),
  inference(resolution,[],[f2473,f1264])).
fof(f1264,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f941])).
fof(f941,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f940])).
fof(f940,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f779])).
fof(f779,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X39] : ((p(s(bool,X39)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CLAUSES|split|1')).
fof(f3113,plain,(
  ( ! [X302] : (sP11('T_0',sK25) | ~p(s(bool,X302))) ) | $spl43),
  inference(resolution,[],[f2473,f1262])).
fof(f1262,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f939])).
fof(f939,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f938])).
fof(f938,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f778])).
fof(f778,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X39] : ((p(s(bool,X39)) => p(s(bool,X39))) <=> p(s(bool,'T_0')))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CLAUSES|split|3')).
fof(f3112,plain,(
  ( ! [X301] : (sP11('T_0',sK25) | p(s(bool,X301))) ) | $spl43),
  inference(resolution,[],[f2473,f1261])).
fof(f1261,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f939])).
fof(f3111,plain,(
  ( ! [X300] : (sP11('T_0',sK25) | p(s(bool,X300))) ) | $spl43),
  inference(resolution,[],[f2473,f1258])).
fof(f1258,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f937])).
fof(f937,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f936])).
fof(f936,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f777])).
fof(f777,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X39] : ((p(s(bool,'T_0')) => p(s(bool,X39))) <=> p(s(bool,X39)))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.IMP_CLAUSES|split|0')).
fof(f3110,plain,(
  ( ! [X299] : (sP11('T_0',sK25) | ~p(s(bool,X299))) ) | $spl43),
  inference(resolution,[],[f2473,f1256])).
fof(f1256,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f935])).
fof(f3253,plain,(
  ( ! [X24] : (sP11(X24,sK25)) ) | $spl92),
  inference(cnf_transformation,[],[f3253_D])).
fof(f3253_D,plain,(
  ( ! [X24] : (sP11(X24,sK25)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3106,plain,(
  ( ! [X292,X290,X291] : (sP11(X290,sK25) | ~p(s(bool,X291)) | p(s(bool,X292)) | s(bool,X291) != s(bool,X292)) ) | $spl43),
  inference(resolution,[],[f2473,f1424])).
fof(f3250,plain,(
  ( ! [X2,X0] : (sP11(X0,sK25) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl88),
  inference(cnf_transformation,[],[f3250_D])).
fof(f3250_D,plain,(
  ( ! [X2,X0] : (sP11(X0,sK25) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3105,plain,(
  ( ! [X288,X287,X289] : (sP11(X287,sK25) | p(s(bool,X288)) | ~p(s(bool,X289)) | s(bool,X287) != s(bool,X289)) ) | $spl43),
  inference(resolution,[],[f2473,f1424])).
fof(f3005,plain,(
  ( ! [X30,X28,X31,X29] : (sP11(X28,sK25) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl43),
  inference(resolution,[],[f2473,f1018])).
fof(f3252,plain,(
  ( ! [X24,X27,X25] : (sP11(X24,sK25) | ~p(s(bool,X25)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl43),
  inference(subsumption_resolution,[],[f3004,f1424])).
fof(f3004,plain,(
  ( ! [X26,X24,X27,X25] : (sP11(X24,sK25) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X25) != s(bool,X27)) ) | $spl43),
  inference(resolution,[],[f2473,f1017])).
fof(f3003,plain,(
  ( ! [X23,X21,X22,X20] : (sP11(X20,sK25) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl43),
  inference(resolution,[],[f2473,f1017])).
fof(f3251,plain,(
  ( ! [X4,X7] : (sP11(X4,sK25) | s(bool,X4) != s(bool,X7)) ) | $spl90),
  inference(cnf_transformation,[],[f3251_D])).
fof(f3251_D,plain,(
  ( ! [X4,X7] : (sP11(X4,sK25) | s(bool,X4) != s(bool,X7)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f3002,plain,(
  ( ! [X19,X17,X18,X16] : (sP11(X16,sK25) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X16) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl43),
  inference(resolution,[],[f2473,f1016])).
fof(f3001,plain,(
  ( ! [X14,X12,X15,X13] : (sP11(X12,sK25) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl43),
  inference(resolution,[],[f2473,f1016])).
fof(f3000,plain,(
  ( ! [X10,X8,X11,X9] : (sP11(X8,sK25) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl43),
  inference(resolution,[],[f2473,f1014])).
fof(f2999,plain,(
  ( ! [X6,X4,X7,X5] : (sP11(X4,sK25) | ~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | s(bool,X5) != s(bool,X6)) ) | $spl43),
  inference(resolution,[],[f2473,f1012])).
fof(f2998,plain,(
  ( ! [X2,X0,X3,X1] : (sP11(X0,sK25) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl43),
  inference(resolution,[],[f2473,f1012])).
fof(f2983,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))))) | $spl85),
  inference(cnf_transformation,[],[f2983_D])).
fof(f2983_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f2965,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))))) | $spl53),
  inference(superposition,[],[f1771,f1059])).
fof(f1059,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f832])).
fof(f832,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X76] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X76))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.numeral.numeral_distrib|split|17')).
fof(f2964,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))))) | $spl53),
  inference(superposition,[],[f1771,f1057])).
fof(f1057,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f831])).
fof(f2981,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(cnf_transformation,[],[f2981_D])).
fof(f2981_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f2943,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) ) | $spl53),
  inference(superposition,[],[f1771,f1207])).
fof(f2942,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) ) | $spl53),
  inference(superposition,[],[f1771,f1206])).
fof(f2675,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl68),
  inference(cnf_transformation,[],[f2675_D])).
fof(f2675_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2895,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1424])).
fof(f2795,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f1466,f1018])).
fof(f2793,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1017])).
fof(f2676,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl70),
  inference(cnf_transformation,[],[f2676_D])).
fof(f2676_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2792,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f1466,f1016])).
fof(f2791,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1016])).
fof(f2790,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f1466,f1014])).
fof(f2789,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f1466,f1012])).
fof(f2788,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1012])).
fof(f1466,plain,(
  ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | $spl5),
  inference(cnf_transformation,[],[f1466_D])).
fof(f1466_D,plain,(
  ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f2754,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) | $spl80),
  inference(cnf_transformation,[],[f2754_D])).
fof(f2754_D,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2745,plain,(
  ( ! [X61,X59,X57,X54,X52,X50,X60,X58,X56,X55,X53,X51] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X61) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X51),s(bool,X50),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X56),s('type.num.num',X57),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X58),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X55),s(fun('type.fcp.cart'(bool,sK13),bool),X52),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X53))),s('type.wordSem.word_loc'(sK13),X60))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X54))),s('type.wordSem.word_loc'(sK13),X61))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X59))))))))) ) | $spl0),
  inference(superposition,[],[f1434,f1076])).
fof(f1076,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK30(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f838,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK29(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK30(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),sK29(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)),s('type.fcp.cart'(bool,X0),sK30(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f748])).
fof(f748,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X13),s('type.fcp.cart'(bool,X0),X14)))))),
  inference(flattening,[],[f747])).
fof(f747,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X13),s('type.fcp.cart'(bool,X0),X14)))))),
  inference(ennf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X13),s('type.fcp.cart'(bool,X0),X14)))))),
  inference(flattening,[],[f550])).
fof(f550,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),X13),s('type.fcp.cart'(bool,X0),X14)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X20,X19,X8,X10,X11,X12,X13,X21,X22,X16,X17,X18] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X22))),s('type.wordSem.word_loc'(X6),X19))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X21))),s('type.wordSem.word_loc'(X6),X20))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X6),'const.bool.the_value_0'))))) => ? [X23,X24] : (s('type.wordSem.word_loc'(X6),X19) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X23))) & s('type.wordSem.word_loc'(X6),X20) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X24))) & s(bool,'const.integer.int_lt_2'(s('type.integer.int',X22),s('type.integer.int',X21))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X6),X23),s('type.fcp.cart'(bool,X6),X24)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_Number_LESS')).
fof(f2753,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) | $spl78),
  inference(cnf_transformation,[],[f2753_D])).
fof(f2753_D,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2744,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40,X49] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X48) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X39),s(bool,X38),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X44),s('type.num.num',X45),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X46),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X43),s(fun('type.fcp.cart'(bool,sK13),bool),X40),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X41))),s('type.wordSem.word_loc'(sK13),X48))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X42))),s('type.wordSem.word_loc'(sK13),X49))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X47))))))))) ) | $spl0),
  inference(superposition,[],[f1434,f1075])).
fof(f1075,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK29(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f2743,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) | $spl0),
  inference(superposition,[],[f1434,f1072])).
fof(f1072,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f837])).
fof(f837,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & (s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) != s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) | s('type.integer.int',X8) = s('type.integer.int',X9)) & (s('type.integer.int',X8) != s('type.integer.int',X9) | s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) = s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f836])).
fof(f836,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & (s('type.fcp.cart'(bool,X0),X13) != s('type.fcp.cart'(bool,X0),X14) | s('type.integer.int',X8) = s('type.integer.int',X9)) & (s('type.integer.int',X8) != s('type.integer.int',X9) | s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14))))),
  inference(flattening,[],[f835])).
fof(f835,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & ((s('type.fcp.cart'(bool,X0),X13) != s('type.fcp.cart'(bool,X0),X14) | s('type.integer.int',X8) = s('type.integer.int',X9)) & (s('type.integer.int',X8) != s('type.integer.int',X9) | s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14)))))),
  inference(nnf_transformation,[],[f746])).
fof(f746,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & (s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14) <=> s('type.integer.int',X8) = s('type.integer.int',X9))))),
  inference(flattening,[],[f745])).
fof(f745,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) | ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & (s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14) <=> s('type.integer.int',X8) = s('type.integer.int',X9))))),
  inference(ennf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & (s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14) <=> s('type.integer.int',X8) = s('type.integer.int',X9))))),
  inference(flattening,[],[f548])).
fof(f548,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) => ? [X13,X14] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X13))) & s('type.wordSem.word_loc'(X0),X1) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),X14))) & (s('type.fcp.cart'(bool,X0),X13) = s('type.fcp.cart'(bool,X0),X14) <=> s('type.integer.int',X8) = s('type.integer.int',X9))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X20,X19,X8,X10,X11,X12,X13,X21,X22,X16,X17,X18] : ((p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X22))),s('type.wordSem.word_loc'(X6),X19))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X21))),s('type.wordSem.word_loc'(X6),X20))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))))) & p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X6),'const.bool.the_value_0'))))) => ? [X23,X24] : (s('type.wordSem.word_loc'(X6),X19) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X23))) & s('type.wordSem.word_loc'(X6),X20) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),X24))) & (s('type.fcp.cart'(bool,X6),X23) = s('type.fcp.cart'(bool,X6),X24) <=> s('type.integer.int',X21) = s('type.integer.int',X22))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_Number_EQ')).
fof(f2742,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) | $spl0),
  inference(superposition,[],[f1434,f1071])).
fof(f1071,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f837])).
fof(f2752,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13))))))) ) | $spl76),
  inference(cnf_transformation,[],[f2752_D])).
fof(f2752_D,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2741,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13)))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))) ) | $spl0),
  inference(superposition,[],[f1434,f1078])).
fof(f1078,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X7))))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) )),
  inference(cnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X7))))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X7))))))),
  inference(flattening,[],[f749])).
fof(f749,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))))) | (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X7))))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X7))))))),
  inference(ennf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))))) => (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X7))))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X7))))))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))))) => (s('type.wordSem.word_loc'(X0),X2) = s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X7))))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X7))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X8,X9,X10,X11,X12,X13,X15,X16,X17,X18] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X6),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))),s('type.wordSem.word_loc'(X6),X9))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8))))))) => (s('type.wordSem.word_loc'(X6),X9) = s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X15))))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.integer.int',X15))))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.memory_rel_Number_IMP')).
fof(f2688,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl74),
  inference(cnf_transformation,[],[f2688_D])).
fof(f2688_D,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1471,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1072])).
fof(f2686,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl72),
  inference(cnf_transformation,[],[f2686_D])).
fof(f2686_D,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1475,plain,(
  s('type.wordSem.word_loc'(sK13),sK14) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1076])).
fof(f2646,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,X181) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1424])).
fof(f2546,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f1466,f1018])).
fof(f2544,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1017])).
fof(f2543,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f1466,f1016])).
fof(f2542,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1016])).
fof(f2541,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f1466,f1014])).
fof(f2540,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f1466,f1012])).
fof(f2539,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) ) | $spl5),
  inference(resolution,[],[f1466,f1012])).
fof(f2500,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,sK25) != s(bool,X181)) ) | $spl43),
  inference(resolution,[],[f1762,f1424])).
fof(f2400,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl43),
  inference(resolution,[],[f1762,f1018])).
fof(f2398,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK25) != s(bool,X16)) ) | $spl43),
  inference(resolution,[],[f1762,f1017])).
fof(f2397,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,sK25) != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f2396,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f2395,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl43),
  inference(resolution,[],[f1762,f1014])).
fof(f2394,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK25) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f2393,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f2364,plain,(
  ( ! [X191,X190] : (p(s(bool,X190)) | p(s(bool,X191)) | s(bool,sK25) != s(bool,X191)) ) | $spl42),
  inference(resolution,[],[f1763,f1424])).
fof(f2264,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X21) | s(bool,X22) != s(bool,X23)) ) | $spl42),
  inference(resolution,[],[f1763,f1018])).
fof(f2263,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,sK25) != s(bool,X20)) ) | $spl42),
  inference(resolution,[],[f1763,f1018])).
fof(f2261,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,sK25) != s(bool,X14)) ) | $spl42),
  inference(resolution,[],[f1763,f1017])).
fof(f2260,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl42),
  inference(resolution,[],[f1763,f1016])).
fof(f2259,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X6) | s(bool,X7) != s(bool,X8)) ) | $spl42),
  inference(resolution,[],[f1763,f1014])).
fof(f2258,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,sK25) != s(bool,X5)) ) | $spl42),
  inference(resolution,[],[f1763,f1014])).
fof(f2257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl42),
  inference(resolution,[],[f1763,f1012])).
fof(f2255,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2106,f1759])).
fof(f2106,plain,(
  ( ! [X732,X731] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X731),s('type.integer.int',X732))))) ) | $spl40),
  inference(resolution,[],[f1759,f1167])).
fof(f1167,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) )),
  inference(cnf_transformation,[],[f877])).
fof(f877,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(nnf_transformation,[],[f575])).
fof(f575,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(rectify,[],[f368])).
fof(f368,axiom,(
  ! [X31,X30] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X30),s('type.integer.int',X31)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X30),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.int_arith.not_less')).
fof(f2254,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2103,f1759])).
fof(f2103,plain,(
  ( ! [X722,X721] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X721),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X722),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))))) ) | $spl40),
  inference(resolution,[],[f1759,f1167])).
fof(f2253,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2102,f1759])).
fof(f2102,plain,(
  ( ! [X720,X719] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X719),s('type.integer.int',X720))))) ) | $spl40),
  inference(resolution,[],[f1759,f1214])).
fof(f1214,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f770])).
fof(f770,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f596])).
fof(f596,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f595])).
fof(f595,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f440])).
fof(f440,axiom,(
  ! [X30,X31] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X30),s('type.integer.int',X31)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X31),s('type.integer.int',X30)))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.integer.INT_LT_GT')).
fof(f2252,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2101,f1759])).
fof(f2101,plain,(
  ( ! [X718,X717] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X717),s('type.integer.int',X718))))) ) | $spl40),
  inference(resolution,[],[f1759,f1214])).
fof(f2235,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2066,f1759])).
fof(f2066,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(resolution,[],[f1759,f1340])).
fof(f1340,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f979])).
fof(f979,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f663])).
fof(f663,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f85])).
fof(f85,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.NOT_CLAUSES|split|1')).
fof(f2234,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2065,f1759])).
fof(f2065,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(resolution,[],[f1759,f1338])).
fof(f1338,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f978])).
fof(f978,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f662])).
fof(f662,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f109])).
fof(f109,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bool.NOT_CLAUSES|split|2')).
fof(f2233,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2044,f1759])).
fof(f2044,plain,(
  ( ! [X315] : (~p(s(bool,X315))) ) | $spl40),
  inference(resolution,[],[f1759,f1250])).
fof(f2064,plain,(
  $false | $spl40),
  inference(resolution,[],[f1759,f1336])).
fof(f2151,plain,(
  $false | ($spl40 | $spl43)),
  inference(resolution,[],[f1759,f1762])).
fof(f1888,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | ~p(s(bool,X181)) | s(bool,sK25) != s(bool,X181)) ) | $spl43),
  inference(resolution,[],[f1762,f1424])).
fof(f1788,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,sK25) != s(bool,X22) | s(bool,X21) != s(bool,X23)) ) | $spl43),
  inference(resolution,[],[f1762,f1018])).
fof(f1786,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,sK25) != s(bool,X16)) ) | $spl43),
  inference(resolution,[],[f1762,f1017])).
fof(f1785,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,sK25) != s(bool,X14) | s(bool,X12) != s(bool,X13)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f1784,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,sK25) != s(bool,X10)) ) | $spl43),
  inference(resolution,[],[f1762,f1016])).
fof(f1783,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,sK25) != s(bool,X7) | s(bool,X6) != s(bool,X8)) ) | $spl43),
  inference(resolution,[],[f1762,f1014])).
fof(f1782,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,sK25) != s(bool,X5) | s(bool,X3) != s(bool,X4)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f1781,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,sK25) != s(bool,X1)) ) | $spl43),
  inference(resolution,[],[f1762,f1012])).
fof(f1776,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | $spl58),
  inference(cnf_transformation,[],[f1776_D])).
fof(f1776_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1629,plain,(
  ( ! [X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X9))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))))) )),
  inference(superposition,[],[f1009,f1207])).
fof(f1628,plain,(
  ( ! [X8] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.num.|48|_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK18))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8))))))) )),
  inference(superposition,[],[f1009,f1206])).
fof(f1771,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))) | $spl53),
  inference(cnf_transformation,[],[f1771_D])).
fof(f1771_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f1773,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.arithmetic.ZERO_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | $spl54),
  inference(cnf_transformation,[],[f1773_D])).
fof(f1773_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.arithmetic.ZERO_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1627,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.arithmetic.ZERO_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18)))),
  inference(superposition,[],[f1009,f1059])).
fof(f1626,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num','const.arithmetic.ZERO_0'),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK18)))),
  inference(superposition,[],[f1009,f1057])).
fof(f1762,plain,(
  ~p(s(bool,sK25)) | $spl43),
  inference(cnf_transformation,[],[f1762_D])).
fof(f1762_D,plain,(
  ~p(s(bool,sK25)) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1769,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'T_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | $spl50),
  inference(cnf_transformation,[],[f1769_D])).
fof(f1769_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'T_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1609,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'T_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,sK25))),
  inference(superposition,[],[f1009,f1328])).
fof(f1608,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'T_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,sK25))),
  inference(superposition,[],[f1009,f1326])).
fof(f1763,plain,(
  p(s(bool,sK25)) | $spl42),
  inference(cnf_transformation,[],[f1763_D])).
fof(f1763_D,plain,(
  p(s(bool,sK25)) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1767,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'F_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | $spl48),
  inference(cnf_transformation,[],[f1767_D])).
fof(f1767_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'F_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1607,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'F_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,sK25))),
  inference(superposition,[],[f1009,f1324])).
fof(f1606,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'F_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,sK25))),
  inference(superposition,[],[f1009,f1322])).
fof(f1604,plain,(
  p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,'F_0'),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,sK25))),
  inference(superposition,[],[f1009,f1252])).
fof(f1765,plain,(
  ( ! [X10] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,X10))) ) | $spl46),
  inference(cnf_transformation,[],[f1765_D])).
fof(f1765_D,plain,(
  ( ! [X10] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,X10))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1603,plain,(
  ( ! [X11] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X11),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,X11)) | ~p(s(bool,sK25))) )),
  inference(superposition,[],[f1009,f1053])).
fof(f1602,plain,(
  ( ! [X10] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | ~p(s(bool,sK25)) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f1009,f1053])).
fof(f1764,plain,(
  ( ! [X8] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X8),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,X8))) ) | $spl44),
  inference(cnf_transformation,[],[f1764_D])).
fof(f1764_D,plain,(
  ( ! [X8] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X8),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,X8))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1601,plain,(
  ( ! [X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,sK25)) | p(s(bool,X9))) )),
  inference(superposition,[],[f1009,f1050])).
fof(f1600,plain,(
  ( ! [X8] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,X8),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) | p(s(bool,X8)) | p(s(bool,sK25))) )),
  inference(superposition,[],[f1009,f1050])).
fof(f1753,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) ) | $spl28),
  inference(cnf_transformation,[],[f1753_D])).
fof(f1753_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1589,plain,(
  ( ! [X191,X192] : (p(s(bool,X191)) | p(s(bool,X192)) | s(bool,X192) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1424])).
fof(f1489,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f1009,f1018])).
fof(f1755,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) ) | $spl32),
  inference(cnf_transformation,[],[f1755_D])).
fof(f1755_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1488,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1018])).
fof(f1486,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1017])).
fof(f1485,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1016])).
fof(f1484,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15))))))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f1009,f1014])).
fof(f1483,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1014])).
fof(f1482,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))) )),
  inference(resolution,[],[f1009,f1012])).
fof(f1752,plain,(
  ( ! [X0] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))),s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X0))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',X0))))) ) | $spl26),
  inference(cnf_transformation,[],[f1752_D])).
fof(f1752_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))),s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X0))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',X0))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1481,plain,(
  ( ! [X0] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',sK24),s(bool,sK25),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),sK19),s('type.num.num',sK18),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),sK17),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),sK20),s(fun('type.fcp.cart'(bool,sK13),bool),sK23),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))),s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X0))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK22))),s('type.wordSem.word_loc'(sK13),sK16))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK21))),s('type.wordSem.word_loc'(sK13),sK14))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),sK15)))))))))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))) )),
  inference(resolution,[],[f1009,f1081])).
fof(f1081,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X6))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X6)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : (~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X6)))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))) | p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X6))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))))),
  inference(flattening,[],[f752])).
fof(f752,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X6)))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))) | p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X6))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))))),
  inference(ennf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X6)))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X6))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))))),
  inference(flattening,[],[f556])).
fof(f556,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X6)))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X8),s(bool,X9),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X4),s('type.num.num',X3),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X2),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X5),s(fun('type.fcp.cart'(bool,X0),bool),X7),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s('type.wordSem.word_loc'(X0),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X0),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X6))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))))),
  inference(rectify,[],[f203])).
fof(f203,axiom,(
  ! [X6] : ! [X8,X10,X11,X12,X13,X15,X16,X17,X18] : ((p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X6),'const.bool.the_value_0')))) & p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.integer.int',X15)))) & p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8))))) => p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X17),s(bool,X18),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X12),s('type.num.num',X11),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X6)),X10),s(fun('type.fcp.cart'(bool,X6),'type.wordSem.word_loc'(X6)),X13),s(fun('type.fcp.cart'(bool,X6),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))),s('type.wordSem.word_loc'(X6),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,X6),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',X15))))))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X6))),X8)))))))),
  file('thm.bvp_to_wordProps.memory_rel_Number_IMP_Word_2','thm.bvp_to_wordProps.IMP_memory_rel_Number')).
fof(f1751,plain,(
  p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) | $spl24),
  inference(cnf_transformation,[],[f1751_D])).
fof(f1751_D,plain,(
  p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1479,plain,(
  p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(sK13),'const.bool.the_value_0'),s('type.integer.int',sK22)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1079])).
fof(f1079,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bvp_to_wordProps.small_int_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.integer.int',X7)))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X9),s(bool,X10),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X5),s('type.num.num',X4),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X3),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X6),s(fun('type.fcp.cart'(bool,X0),bool),X8),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X1)))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0'))))) )),
  inference(cnf_transformation,[],[f750])).
fof(f1749,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',sK22))))) | $spl22),
  inference(cnf_transformation,[],[f1749_D])).
fof(f1749_D,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',sK22))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1478,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),'const.bvp_to_wordProps.Smallnum_1'(s('type.integer.int',sK22))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1078])).
fof(f1747,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl20),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1476,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK22),s('type.integer.int',sK21))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)),s('type.fcp.cart'(bool,sK13),sK30(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1077])).
fof(f1077,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X8))) = s(bool,'const.words.word_lt_2'(s('type.fcp.cart'(bool,X0),sK29(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)),s('type.fcp.cart'(bool,X0),sK30(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f1744,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl18),
  inference(cnf_transformation,[],[f1744_D])).
fof(f1744_D,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1474,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK29(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1075])).
fof(f1741,plain,(
  s('type.integer.int',sK21) != s('type.integer.int',sK22) | $spl17),
  inference(cnf_transformation,[],[f1741_D])).
fof(f1741_D,plain,(
  s('type.integer.int',sK21) != s('type.integer.int',sK22) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1739,plain,(
  s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) = s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) | $spl14),
  inference(cnf_transformation,[],[f1739_D])).
fof(f1739_D,plain,(
  s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) = s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1473,plain,(
  s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) = s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) | s('type.integer.int',sK21) != s('type.integer.int',sK22) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1074])).
fof(f1074,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) = s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) | s('type.integer.int',X8) != s('type.integer.int',X9) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f837])).
fof(f1740,plain,(
  s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) != s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) | $spl15),
  inference(cnf_transformation,[],[f1740_D])).
fof(f1740_D,plain,(
  s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) != s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1742,plain,(
  s('type.integer.int',sK21) = s('type.integer.int',sK22) | $spl16),
  inference(cnf_transformation,[],[f1742_D])).
fof(f1742_D,plain,(
  s('type.integer.int',sK21) = s('type.integer.int',sK22) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1472,plain,(
  s('type.integer.int',sK21) = s('type.integer.int',sK22) | s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) != s('type.fcp.cart'(bool,sK13),sK28(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1073])).
fof(f1073,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.integer.int',X8) = s('type.integer.int',X9) | s('type.fcp.cart'(bool,X0),sK27(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) != s('type.fcp.cart'(bool,X0),sK28(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(X0),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X11),s(bool,X12),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X6),s('type.num.num',X5),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(X0)),X4),s(fun('type.fcp.cart'(bool,X0),'type.wordSem.word_loc'(X0)),X7),s(fun('type.fcp.cart'(bool,X0),bool),X10),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X9))),s('type.wordSem.word_loc'(X0),X2))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))),s('type.wordSem.word_loc'(X0),X1))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(X0))),X3))))))))) )),
  inference(cnf_transformation,[],[f837])).
fof(f1737,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | $spl12),
  inference(cnf_transformation,[],[f1737_D])).
fof(f1737_D,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1470,plain,(
  s('type.wordSem.word_loc'(sK13),sK16) = s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),sK27(sK25,sK24,sK23,sK22,sK21,sK20,sK19,sK18,sK17,sK15,sK16,sK14,sK13)))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))),
  inference(resolution,[],[f1009,f1071])).
fof(f1469,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) | $spl10),
  inference(cnf_transformation,[],[f1469_D])).
fof(f1469_D,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1458,plain,(
  ( ! [X61,X59,X57,X54,X52,X50,X60,X58,X56,X55,X53,X51] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X61) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X51),s(bool,X50),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X56),s('type.num.num',X57),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X58),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X55),s(fun('type.fcp.cart'(bool,sK13),bool),X52),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X53))),s('type.wordSem.word_loc'(sK13),X60))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X54))),s('type.wordSem.word_loc'(sK13),X61))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X59))))))))) ) | $spl2),
  inference(superposition,[],[f1435,f1076])).
fof(f1468,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) | $spl8),
  inference(cnf_transformation,[],[f1468_D])).
fof(f1468_D,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1457,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40,X49] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X48) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X39),s(bool,X38),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X44),s('type.num.num',X45),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X46),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X43),s(fun('type.fcp.cart'(bool,sK13),bool),X40),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X41))),s('type.wordSem.word_loc'(sK13),X48))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X42))),s('type.wordSem.word_loc'(sK13),X49))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X47))))))))) ) | $spl2),
  inference(superposition,[],[f1435,f1075])).
fof(f1456,plain,(
  ( ! [X30,X28,X26,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X37) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X27),s(bool,X26),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X32),s('type.num.num',X33),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X34),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X31),s(fun('type.fcp.cart'(bool,sK13),bool),X28),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X29))),s('type.wordSem.word_loc'(sK13),X36))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X30))),s('type.wordSem.word_loc'(sK13),X37))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X35))))))))) ) | $spl2),
  inference(superposition,[],[f1435,f1072])).
fof(f1455,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X24) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0')))) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X15),s(bool,X14),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X20),s('type.num.num',X21),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X22),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X19),s(fun('type.fcp.cart'(bool,sK13),bool),X16),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))),s('type.wordSem.word_loc'(sK13),X24))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))),s('type.wordSem.word_loc'(sK13),X25))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X23))))))))) ) | $spl2),
  inference(superposition,[],[f1435,f1071])).
fof(f1467,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13))))))) ) | $spl6),
  inference(cnf_transformation,[],[f1467_D])).
fof(f1467_D,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13))))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1454,plain,(
  ( ! [X6,X4,X12,X10,X8,X7,X5,X13,X11,X9] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),X5) | ~p(s(bool,'const.bvp_to_wordProps.memory_rel_8'(s('type.bvp_to_word.config',X6),s(bool,X7),s('type.finite_map.fmap'('type.num.num','type.closSem.ref'('type.bvlSem.v')),X8),s('type.num.num',X9),s('type.finite_map.fmap'('type.wordLang.store_name','type.wordSem.word_loc'(sK13)),X10),s(fun('type.fcp.cart'(bool,sK13),'type.wordSem.word_loc'(sK13)),X11),s(fun('type.fcp.cart'(bool,sK13),bool),X12),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),'const.list.CONS_2'(s('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13)),'const.pair.,_2'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))),s('type.wordSem.word_loc'(sK13),X5))),s('type.list.list'('type.pair.prod'('type.bvlSem.v','type.wordSem.word_loc'(sK13))),X13)))))) | ~p(s(bool,'const.labProps.good_dimindex_1'(s('type.bool.itself'(sK13),'const.bool.the_value_0'))))) ) | $spl2),
  inference(superposition,[],[f1435,f1078])).
fof(f1434,plain,(
  ( ! [X13] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X13)))) ) | $spl0),
  inference(cnf_transformation,[],[f1434_D])).
fof(f1434_D,plain,(
  ( ! [X13] : (s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X13)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1435,plain,(
  ( ! [X14] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X14)))) ) | $spl2),
  inference(cnf_transformation,[],[f1435_D])).
fof(f1435_D,plain,(
  ( ! [X14] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X14)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1010,plain,(
  ( ! [X14,X13] : (s('type.wordSem.word_loc'(sK13),sK14) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X14))) | s('type.wordSem.word_loc'(sK13),sK16) != s('type.wordSem.word_loc'(sK13),'const.wordSem.Word_1'(s('type.fcp.cart'(bool,sK13),X13)))) )),
  inference(cnf_transformation,[],[f810])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 433
% Passive clauses: 3216
% Generated clauses: 7027
% Final active clauses: 370
% Final passive clauses: 1567
% Input formulas: 516
% Initial clauses: 415
% 
% Selected by SInE selection: 156
% SInE iterations: 1
% 
% Duplicate literals: 106
% Trivial inequalities: 14
% Fw subsumption resolutions: 400
% Fw demodulations: 425
% 
% Simple tautologies: 193
% Equational tautologies: 3
% Forward subsumptions: 1839
% 
% Binary resolution: 3206
% Forward superposition: 1524
% Backward superposition: 889
% Equality resolution: 6
% 
% Split clauses: 600
% Split components: 1221
% SAT solver clauses: 620
% SAT solver unit clauses: 11
% SAT solver binary clauses: 44
% 
% Sat splits: 600
% Sat splitting refutations: 19
% 
% Memory used [KB]: 6908
% Time elapsed: 0.235 s
% ------------------------------
% ------------------------------
% Success in time 11.998 s
