Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.gramProps
% (16262)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 106
% Passive clauses: 3416
% Generated clauses: 20677
% Final active clauses: 96
% Final passive clauses: 3182
% Input formulas: 44
% Initial clauses: 94
% 
% Duplicate literals: 620
% Trivial inequalities: 233
% Fw subsumption resolutions: 2626
% Fw demodulations: 3596
% Bw demodulations: 9
% 
% Simple tautologies: 350
% Equational tautologies: 59
% Forward subsumptions: 9483
% 
% Binary resolution: 194
% Factoring: 1
% Forward superposition: 6159
% Backward superposition: 6665
% Self superposition: 251
% Equality factoring: 139
% Equality resolution: 27
% 
% Split clauses: 235
% Split components: 478
% SAT solver clauses: 415
% SAT solver unit clauses: 17
% SAT solver binary clauses: 123
% 
% Sat splits: 235
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 412
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.gramProps
% (16279)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 112
% Passive clauses: 717
% Generated clauses: 16077
% Final active clauses: 91
% Final passive clauses: 527
% Input formulas: 44
% Initial clauses: 94
% 
% Duplicate literals: 506
% Trivial inequalities: 169
% Fw subsumption resolutions: 1271
% Fw demodulations: 2039
% Bw demodulations: 90
% Condensations: 1440
% Global subsumptions: 17
% 
% Simple tautologies: 181
% Equational tautologies: 36
% Forward subsumptions: 8492
% 
% Binary resolution: 148
% Factoring: 1
% Forward superposition: 4965
% Backward superposition: 4323
% Self superposition: 222
% Equality factoring: 84
% Equality resolution: 28
% 
% Split clauses: 69
% Split components: 140
% SAT solver clauses: 1363
% SAT solver unit clauses: 126
% SAT solver binary clauses: 61
% SAT solver learnt clauses: 34
% 
% Sat splits: 69
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1295
% TWLsolver calls for satisfiability: 345
% 
% Memory used [KB]: 13816
% Time elapsed: 0.731 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.gramProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.gramProps
% SZS output start Proof for thm.gramProps
fof(f118121,plain,(
  $false),
  inference(subsumption_resolution,[],[f118120,f236])).
fof(f236,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X18,X19] : s(X6,X19) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X18),s(X6,X19)))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f118120,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,'const.bool.COND_3'(s(bool,sK8(sK5,sK4,sK3,'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2)),sK1,sK0)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3))))),s(bool,'F_0')))),
  inference(forward_demodulation,[],[f118119,f3110])).
fof(f3110,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,sK9(X0,X1,X2,'const.pred_set.INSERT_2'(s('type.sum.sum'(X3,'type.num.num'),X1),s(fun('type.sum.sum'(X3,'type.num.num'),bool),X4)),X3,X5))) )),
  inference(unit_resulting_resolution,[],[f270,f1434,f243])).
fof(f243,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X18,X19] : ((p(s(bool,X18)) => p(s(bool,X19))) => ((p(s(bool,X19)) => p(s(bool,X18))) => s(bool,X18) = s(bool,X19)))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f1434,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK9(X0,X1,X2,'const.pred_set.INSERT_2'(s('type.sum.sum'(X3,'type.num.num'),X1),s(fun('type.sum.sum'(X3,'type.num.num'),bool),X4)),X3,X5)))) )),
  inference(unit_resulting_resolution,[],[f1178,f254])).
fof(f254,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,sK9(X5,X4,X3,X2,X1,X0))) | ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK8(X5,X4,X3,X2,X1,X0))) | s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))))))))) & (s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') != s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))))))))))) | p(s(bool,sK8(X5,X4,X3,X2,X1,X0)))) & ((~p(s(bool,sK9(X5,X4,X3,X2,X1,X0))) | (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))) | p(s(bool,sK9(X5,X4,X3,X2,X1,X0)))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,sK9(X5,X4,X3,X2,X1,X0)),s(bool,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))))))))) & (s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') != s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))))))))))) | p(s(bool,X6))) & ? [X7] : ((~p(s(bool,X7)) | (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))) | p(s(bool,X7))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,X7),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))))))))) & (s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') != s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))))))))))) | p(s(bool,X6)))) & ? [X7] : (((~p(s(bool,X7)) | (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) | p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))))) | p(s(bool,X7)))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,X7),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))))))))) & ? [X7] : ((p(s(bool,X7)) <=> (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,X7),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5))),s('type.sum.sum'(X1,'type.num.num'),X4))),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X0,X1)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))))))))) & ? [X7] : ((p(s(bool,X7)) <=> (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X1,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X0,X1),X5)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2)))))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,X7),s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X12,X9,X15,X11] : ? [X16] : ((p(s(bool,X16)) <=> s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.pred_set.EMPTY_0') = s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.pred_set.INTER_2'(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.finite_map.FAPPLY_2'(s('type.finite_map.fmap'('type.sum.sum'(X6,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X7,X6),X11))),s('type.sum.sum'(X6,'type.num.num'),X15))),s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.pred_set.GSPEC_1'(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),'type.pair.prod'('type.list.list'('type.grammar.symbol'(X7,X6)),bool)),'_dst_x1x1_2'(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),fun(bool,'type.pair.prod'('type.list.list'('type.grammar.symbol'(X7,X6)),bool))),'const.pair.,_0'),s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X7,X6),X11),s(fun('type.sum.sum'(X6,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X6,'type.num.num'),X15),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X12)))))))))))) & ? [X17] : ((p(s(bool,X17)) <=> (p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X6,'type.num.num'),X15),s(fun('type.sum.sum'(X6,'type.num.num'),bool),'const.finite_map.FDOM_1'(s('type.finite_map.fmap'('type.sum.sum'(X6,'type.num.num'),fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool)),'const.grammar.grammar_rules_1'(s('type.grammar.grammar'(X7,X6),X11)))))))) & ~p(s(bool,'const.bool.IN_2'(s('type.sum.sum'(X6,'type.num.num'),X15),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X12)))))) & s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X7,X6),X11),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X12))),s('type.list.list'('type.grammar.symbol'(X7,X6)),'const.list.CONS_2'(s('type.grammar.symbol'(X7,X6),'const.grammar.NT_1'(s('type.sum.sum'(X6,'type.num.num'),X15))),s('type.list.list'('type.grammar.symbol'(X7,X6)),X9))))) = s(bool,'const.bool.COND_3'(s(bool,X17),s(bool,'const.bool.COND_3'(s(bool,X16),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X7,X6),X11),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X12))),s('type.list.list'('type.grammar.symbol'(X7,X6)),X9))))),s(bool,'F_0')))))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.NTproperties.nullableML_def|split|2')).
fof(f1178,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(forward_demodulation,[],[f1137,f236])).
fof(f1137,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(unit_resulting_resolution,[],[f236,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X20,X21,X23] : (p(s(bool,'const.bool.IN_2'(s(X6,X20),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X21),s(fun(X6,bool),X23)))))) <=> (s(X6,X20) = s(X6,X21) | p(s(bool,'const.bool.IN_2'(s(X6,X20),s(fun(X6,bool),X23))))))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.pred_set.IN_INSERT')).
fof(f270,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f170,f165])).
fof(f165,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.gramProps.gramProps_7_prove_ltb',aHL_TRUTH)).
fof(f170,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f118119,plain,(
  s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK9(sK5,sK4,sK3,'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2)),sK1,sK0)),s(bool,'const.bool.COND_3'(s(bool,sK8(sK5,sK4,sK3,'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2)),sK1,sK0)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3))))),s(bool,'F_0')))),
  inference(subsumption_resolution,[],[f117856,f165])).
fof(f117856,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK9(sK5,sK4,sK3,'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2)),sK1,sK0)),s(bool,'const.bool.COND_3'(s(bool,sK8(sK5,sK4,sK3,'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2)),sK1,sK0)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3))))),s(bool,'F_0')))),
  inference(superposition,[],[f7091,f616])).
fof(f616,plain,(
  s(bool,'T_0') = s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),'const.list.CONS_2'(s('type.grammar.symbol'(sK0,sK1),'const.grammar.NT_1'(s('type.sum.sum'(sK1,'type.num.num'),sK4))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3)))))),
  inference(unit_resulting_resolution,[],[f270,f164,f343])).
fof(f343,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f335,f270])).
fof(f335,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f243,f176])).
fof(f176,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.gramProps.gramProps_7_prove_ltb',aHL_BOOL_CASES)).
fof(f164,plain,(
  s(bool,'F_0') != s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),'const.list.CONS_2'(s('type.grammar.symbol'(sK0,sK1),'const.grammar.NT_1'(s('type.sum.sum'(sK1,'type.num.num'),sK4))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3)))))),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  s(bool,'F_0') != s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(sK0,sK1),sK5),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(sK1,'type.num.num'),sK4),s(fun('type.sum.sum'(sK1,'type.num.num'),bool),sK2))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),'const.list.CONS_2'(s('type.grammar.symbol'(sK0,sK1),'const.grammar.NT_1'(s('type.sum.sum'(sK1,'type.num.num'),sK4))),s('type.list.list'('type.grammar.symbol'(sK0,sK1)),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(bool,'F_0') != s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3)))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(bool,'F_0') = s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'F_0') = s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X1,'type.num.num'),X4),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : s(bool,'F_0') = s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X6,X7),X11),s(fun('type.sum.sum'(X7,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X7,'type.num.num'),X10),s(fun('type.sum.sum'(X7,'type.num.num'),bool),X8))),s('type.list.list'('type.grammar.symbol'(X6,X7)),'const.list.CONS_2'(s('type.grammar.symbol'(X6,X7),'const.grammar.NT_1'(s('type.sum.sum'(X7,'type.num.num'),X10))),s('type.list.list'('type.grammar.symbol'(X6,X7)),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : s(bool,'F_0') = s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X6,X7),X11),s(fun('type.sum.sum'(X7,'type.num.num'),bool),'const.pred_set.INSERT_2'(s('type.sum.sum'(X7,'type.num.num'),X10),s(fun('type.sum.sum'(X7,'type.num.num'),bool),X8))),s('type.list.list'('type.grammar.symbol'(X6,X7)),'const.list.CONS_2'(s('type.grammar.symbol'(X6,X7),'const.grammar.NT_1'(s('type.sum.sum'(X7,'type.num.num'),X10))),s('type.list.list'('type.grammar.symbol'(X6,X7)),X9)))))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.gramProps.gramProps_7_prove_ltb')).
fof(f7091,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X30,X31),X32),s(fun('type.sum.sum'(X31,'type.num.num'),bool),X33),s('type.list.list'('type.grammar.symbol'(X30,X31)),'const.list.CONS_2'(s('type.grammar.symbol'(X30,X31),'const.grammar.NT_1'(s('type.sum.sum'(X31,'type.num.num'),X34))),s('type.list.list'('type.grammar.symbol'(X30,X31)),X35)))))) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK9(X32,X34,X35,X33,X31,X30)),s(bool,'const.bool.COND_3'(s(bool,sK8(X32,X34,X35,X33,X31,X30)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X30,X31),X32),s(fun('type.sum.sum'(X31,'type.num.num'),bool),X33),s('type.list.list'('type.grammar.symbol'(X30,X31)),X35))))),s(bool,'F_0')))) )),
  inference(forward_demodulation,[],[f6891,f249])).
fof(f249,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4))) = s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4))) = s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4))) = s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X2),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X3))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X7,X6] : ! [X24,X25,X26] : s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X7,X6),X24),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X25),s('type.list.list'('type.grammar.symbol'(X7,X6)),X26))) = s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X7,X6)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X7,X6),X24),s(fun('type.sum.sum'(X6,'type.num.num'),bool),X25))),s('type.list.list'('type.grammar.symbol'(X7,X6)),X26)))),
  file('thm.gramProps.gramProps_7_prove_ltb','ARITY_2_3_const.NTproperties.nullableML')).
fof(f6891,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,sK9(X32,X34,X35,X33,X31,X30)),s(bool,'const.bool.COND_3'(s(bool,sK8(X32,X34,X35,X33,X31,X30)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X30,X31),X32),s(fun('type.sum.sum'(X31,'type.num.num'),bool),X33),s('type.list.list'('type.grammar.symbol'(X30,X31)),X35))))),s(bool,'F_0'))) | ~p(s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X30,X31)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X30,X31),X32),s(fun('type.sum.sum'(X31,'type.num.num'),bool),X33))),s('type.list.list'('type.grammar.symbol'(X30,X31)),'const.list.CONS_2'(s('type.grammar.symbol'(X30,X31),'const.grammar.NT_1'(s('type.sum.sum'(X31,'type.num.num'),X34))),s('type.list.list'('type.grammar.symbol'(X30,X31)),X35))))))) )),
  inference(superposition,[],[f182,f271])).
fof(f271,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,sK9(X5,X4,X3,X2,X1,X0)),s(bool,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(bool,'F_0'),s(bool,'const.NTproperties.nullableML_3'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))) )),
  inference(forward_demodulation,[],[f256,f249])).
fof(f256,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),'const.list.CONS_2'(s('type.grammar.symbol'(X0,X1),'const.grammar.NT_1'(s('type.sum.sum'(X1,'type.num.num'),X4))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))) = s(bool,'const.bool.COND_3'(s(bool,sK9(X5,X4,X3,X2,X1,X0)),s(bool,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(bool,'F_0'),s(bool,happ(s(fun('type.list.list'('type.grammar.symbol'(X0,X1)),bool),'const.NTproperties.nullableML_2'(s('type.grammar.grammar'(X0,X1),X5),s(fun('type.sum.sum'(X1,'type.num.num'),bool),X2))),s('type.list.list'('type.grammar.symbol'(X0,X1)),X3))))),s(bool,'F_0')))) )),
  inference(cnf_transformation,[],[f163])).
fof(f182,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.gramProps.gramProps_7_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
% SZS output end Proof for thm.gramProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 344
% Passive clauses: 4647
% Generated clauses: 123140
% Final active clauses: 176
% Final passive clauses: 1012
% Input formulas: 44
% Initial clauses: 94
% Discarded non-redundant clauses: 4190
% 
% Duplicate literals: 3081
% Trivial inequalities: 1095
% Fw subsumption resolutions: 9694
% Fw demodulations: 24215
% Bw demodulations: 2666
% Condensations: 8042
% Global subsumptions: 1417
% 
% Simple tautologies: 1004
% Equational tautologies: 433
% Forward subsumptions: 67392
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 87
% 
% Unit resulting resolution: 3222
% Factoring: 2
% Forward superposition: 32677
% Backward superposition: 27833
% Self superposition: 549
% Equality factoring: 202
% Equality resolution: 75
% 
% Split clauses: 170
% Split components: 364
% SAT solver clauses: 7942
% SAT solver unit clauses: 862
% SAT solver binary clauses: 1890
% 
% Sat splits: 170
% Sat splitting refutations: 1536
% 
% Memory used [KB]: 76501
% Time elapsed: 8.009 s
% ------------------------------
% ------------------------------
% Success in time 9.39 s
