Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.modProps
% (76002)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 653
% Passive clauses: 5468
% Generated clauses: 17118
% Final active clauses: 649
% Final passive clauses: 4546
% Input formulas: 100
% Initial clauses: 273
% 
% Duplicate literals: 195
% Trivial inequalities: 69
% Fw subsumption resolutions: 929
% Fw demodulations: 2128
% Bw demodulations: 3
% 
% Simple tautologies: 226
% Equational tautologies: 31
% Forward subsumptions: 8038
% 
% Binary resolution: 5258
% Factoring: 870
% Forward superposition: 2892
% Backward superposition: 4438
% Self superposition: 42
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 165
% Split components: 332
% SAT solver clauses: 312
% SAT solver unit clauses: 7
% SAT solver binary clauses: 46
% 
% Sat splits: 165
% 
% TWLsolver clauses: 312
% TWLsolver calls for satisfiability: 266
% 
% Memory used [KB]: 12665
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.modProps
% (76031)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 380
% Passive clauses: 2307
% Generated clauses: 13217
% Final active clauses: 375
% Final passive clauses: 1823
% Input formulas: 100
% Initial clauses: 273
% 
% Duplicate literals: 102
% Trivial inequalities: 4
% Fw subsumption resolutions: 514
% Fw demodulations: 1996
% Bw demodulations: 108
% Condensations: 129
% Global subsumptions: 18
% 
% Simple tautologies: 202
% Equational tautologies: 9
% Forward subsumptions: 4980
% 
% Binary resolution: 2312
% Factoring: 307
% Forward superposition: 2467
% Backward superposition: 2374
% Self superposition: 63
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 3028
% SAT solver unit clauses: 302
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 2
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2956
% TWLsolver calls for satisfiability: 2064
% 
% Memory used [KB]: 20596
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.modProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.modProps
% SZS output start Proof for thm.modProps
fof(f17625,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3096,f3099,f3365,f3367,f3097,f3370,f3372,f3374,f3378,f3380,f3382,f3384,f3388,f3390,f3392,f3394,f3402,f3404,f3412,f3414,f3420,f3426,f3450,f3521,f3523,f3525,f3528,f3530,f3535,f3537,f3551,f3553,f3555,f3558,f3560,f3563,f3565,f3663,f3665,f3667,f3670,f3672,f3677,f3679,f3693,f3695,f3697,f3700,f3702,f3705,f3707,f3753,f3755,f3757,f3765,f3767,f3769,f3778,f3780,f3785,f3787,f3796,f3798,f3803,f3805,f3814,f3816,f3821,f3823,f3845,f3847,f3849,f3851,f4330,f4331,f3875,f3877,f3879,f3882,f3884,f3889,f3891,f3893,f3896,f3898,f3908,f3915,f3923,f3925,f3927,f3932,f3934,f3939,f3941,f3946,f3948,f3953,f3955,f3963,f3971,f3979,f4001,f4008,f4010,f4020,f4022,f4028,f4030,f4032,f4040,f4046,f4053,f4059,f4061,f4070,f4072,f4077,f4079,f4088,f4090,f4095,f4097,f4106,f4108,f4113,f4115,f4123,f4131,f4145,f4147,f4158,f4160,f4171,f4173,f4184,f4200,f4209,f4211,f4213,f4222,f4224,f4230,f4235,f4240,f4249,f4252,f4257,f4259,f4261,f4263,f4265,f4281,f4283,f4295,f4297,f4308,f4310,f4332,f4333,f4334,f4335,f4336,f4337,f4338,f4339,f4340,f4341,f4342,f4343,f4344,f4345,f4346,f4347,f4348,f4349,f4350,f4351,f4352,f4353,f4354,f4355,f4356,f4357,f4358,f4359,f4360,f4361,f4362,f4363,f4364,f4365,f4366,f4367,f4368,f4369,f4370,f4371,f4372,f4373,f4374,f4375,f4376,f4377,f4378,f4379,f4380,f4381,f4382,f4383,f4384,f4385,f4386,f4387,f4388,f4389,f4390,f4391,f4392,f4393,f4394,f4395,f4396,f4397,f4398,f4399,f4400,f4401,f4402,f4403,f4404,f4405,f4406,f4407,f4408,f4409,f4410,f4411,f4412,f4413,f4414,f4416,f4417,f4419,f4420,f4422,f4424,f4425,f4426,f4427,f4428,f4429,f4434,f4435,f4436,f4437,f4438,f4439,f4440,f4441,f4442,f4443,f4444,f4445,f4446,f4447,f4448,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f4463,f4464,f4465,f4466,f4467,f4468,f4469,f4470,f4471,f4472,f4473,f4474,f4475,f4476,f4477,f4478,f4479,f4480,f4481,f4482,f4483,f4484,f4485,f4486,f4487,f4488,f4489,f4490,f4491,f4492,f4493,f4494,f4495,f4496,f4497,f4498,f4499,f4500,f4501,f4502,f4503,f4504,f17623])).
fof(f17623,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f17620,f17232])).
fof(f17232,plain,(
  ( ! [X0] : (p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(subsumption_resolution,[],[f17231,f446])).
fof(f446,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.modProps.no_dup_top_types_eqn|split|0',aHL_TRUTH)).
fof(f17231,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(forward_demodulation,[],[f17230,f1100])).
fof(f1100,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f512,f469])).
fof(f469,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f512,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X7] : ! [X36] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X36))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.pred_set.DISJOINT_EMPTY|split|0')).
fof(f17230,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.pred_set.EMPTY_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(forward_demodulation,[],[f17229,f455])).
fof(f455,plain,(
  ( ! [X0] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X7] : s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.list.LIST_TO_SET|split|0')).
fof(f17229,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.NIL_0'))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(forward_demodulation,[],[f17228,f605])).
fof(f605,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ! [X2] : s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X2),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7,X8] : ! [X22] : s('type.list.list'(X8),'const.list.NIL_0') = s('type.list.list'(X8),'const.list.MAP_2'(s(fun(X7,X8),X22),s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.list.MAP|split|0')).
fof(f17228,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(forward_demodulation,[],[f17227,f12425])).
fof(f12425,plain,(
  s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0') = s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f12325,f454])).
fof(f454,plain,(
  ( ! [X0] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FLAT_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FLAT_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7] : s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),'const.list.FLAT_1'(s('type.list.list'('type.list.list'(X7)),'const.list.NIL_0')))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.list.FLAT|split|0')).
fof(f12325,plain,(
  s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'))) = s('type.list.list'('type.list.list'('type.string.char')),'const.list.FLAT_1'(s('type.list.list'('type.list.list'('type.list.list'('type.string.char'))),'const.list.NIL_0')))),
  inference(superposition,[],[f458,f605])).
fof(f458,plain,(
  ( ! [X0] : (s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))) = s('type.list.list'('type.list.list'('type.string.char')),'const.list.FLAT_1'(s('type.list.list'('type.list.list'('type.list.list'('type.string.char'))),'const.list.MAP_2'(s(fun('type.modLang.prompt','type.list.list'('type.list.list'('type.string.char'))),'_dst_x1x0_2'(s(fun('type.modLang.prompt',fun(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'type.list.list'('type.list.list'('type.string.char')))),'const.modLang.prompt_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'_dst_x00x10x01x00_4'(s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.list.list'('type.string.char')),fun(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'type.list.list'('type.list.list'('type.string.char'))))),'const.option.option_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),'type.option.option'('type.list.list'('type.string.char'))),'_i_0_0'),s(fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char'))),'const.modSem.decs_to_types_0'),s(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'_dst_x0_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))))))),s('type.list.list'('type.modLang.prompt'),X0)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))) = s('type.list.list'('type.list.list'('type.string.char')),'const.list.FLAT_1'(s('type.list.list'('type.list.list'('type.list.list'('type.string.char'))),'const.list.MAP_2'(s(fun('type.modLang.prompt','type.list.list'('type.list.list'('type.string.char'))),'_dst_x1x0_2'(s(fun('type.modLang.prompt',fun(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'type.list.list'('type.list.list'('type.string.char')))),'const.modLang.prompt_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'_dst_x00x10x01x00_4'(s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.list.list'('type.string.char')),fun(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'type.list.list'('type.list.list'('type.string.char'))))),'const.option.option_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),'type.option.option'('type.list.list'('type.string.char'))),'_i_0_0'),s(fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char'))),'const.modSem.decs_to_types_0'),s(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'_dst_x0_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))))))),s('type.list.list'('type.modLang.prompt'),X0)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X33] : s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X33))) = s('type.list.list'('type.list.list'('type.string.char')),'const.list.FLAT_1'(s('type.list.list'('type.list.list'('type.list.list'('type.string.char'))),'const.list.MAP_2'(s(fun('type.modLang.prompt','type.list.list'('type.list.list'('type.string.char'))),'_dst_x1x0_2'(s(fun('type.modLang.prompt',fun(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'type.list.list'('type.list.list'('type.string.char')))),'const.modLang.prompt_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char')))),'_dst_x00x10x01x00_4'(s(fun('type.option.option'('type.list.list'('type.string.char')),fun('type.list.list'('type.list.list'('type.string.char')),fun(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'type.list.list'('type.list.list'('type.string.char'))))),'const.option.option_CASE_0'),s(fun('type.option.option'('type.list.list'('type.string.char')),'type.option.option'('type.list.list'('type.string.char'))),'_i_0_0'),s(fun('type.list.list'('type.modLang.dec'),'type.list.list'('type.list.list'('type.string.char'))),'const.modSem.decs_to_types_0'),s(fun('type.list.list'('type.string.char'),'type.list.list'('type.list.list'('type.string.char'))),'_dst_x0_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0'))))))),s('type.list.list'('type.modLang.prompt'),X33)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.modSem.prog_to_top_types_def')).
fof(f17227,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(subsumption_resolution,[],[f17226,f446])).
fof(f17226,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(forward_demodulation,[],[f17118,f453])).
fof(f453,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X7] : s(bool,'T_0') = s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X7),'const.list.NIL_0')))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.list.ALL_DISTINCT|split|0')).
fof(f17118,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.list.NIL_0')))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0))))) )),
  inference(superposition,[],[f571,f12425])).
fof(f571,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1))))) )),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))) | (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))))) & (~p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1))))))),
  inference(flattening,[],[f383])).
fof(f383,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))) | (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))))) & ((~p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1))))) | p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1))))))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X0),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1)))) <=> (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0)))))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X0))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X1))))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X33,X6] : (p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),X33),s(fun('type.semanticPrimitives.tid_or_exn',bool),X6)))) <=> (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X33)))))) & p(s(bool,'const.pred_set.DISJOINT_2'(s(fun('type.semanticPrimitives.tid_or_exn',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.semanticPrimitives.tid_or_exn'),'const.list.MAP_2'(s(fun('type.list.list'('type.string.char'),'type.semanticPrimitives.tid_or_exn'),'_dst_x0x1_2'(s(fun('type.ast.id'('type.list.list'('type.string.char')),'type.semanticPrimitives.tid_or_exn'),'const.semanticPrimitives.TypeId_0'),s(fun('type.list.list'('type.string.char'),'type.ast.id'('type.list.list'('type.string.char'))),'const.ast.Short_0'))),s('type.list.list'('type.list.list'('type.string.char')),'const.modSem.prog_to_top_types_1'(s('type.list.list'('type.modLang.prompt'),X33))))))),s(fun('type.semanticPrimitives.tid_or_exn',bool),X6))))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.modSem.no_dup_top_types_def')).
fof(f17620,plain,(
  ~p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl2),
  inference(unit_resulting_resolution,[],[f759,f792,f790,f3171,f632])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f409])).
fof(f409,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f3171,plain,(
  ~sP5(sK14('_i_0_0',bool),'T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f3099,f649])).
fof(f649,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f243,f312,f311,f310,f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f310,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f311,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X38,X39,X40] : ((p(s(bool,X40)) <=> s(bool,X38) = s(bool,X39)) <=> ((p(s(bool,X40)) | p(s(bool,X39)) | p(s(bool,X38))) & (p(s(bool,X40)) | ~p(s(bool,X38)) | ~p(s(bool,X39))) & (p(s(bool,X39)) | ~p(s(bool,X38)) | ~p(s(bool,X40))) & (p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.sat.sat_2|49|_prove_ltb')).
fof(f445,plain,(
  s(bool,'T_0') != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  s(bool,'T_0') != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f267])).
fof(f267,plain,(
  ? [X0] : s(bool,'T_0') != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ~! [X0] : s(bool,'T_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'T_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),X6)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.modProps.no_dup_top_types_eqn|split|0')).
fof(f790,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f782,f637])).
fof(f637,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f412])).
fof(f412,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f411])).
fof(f411,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f311])).
fof(f782,plain,(
  ~p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))))),
  inference(unit_resulting_resolution,[],[f445,f469])).
fof(f792,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f782,f641])).
fof(f641,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f415])).
fof(f415,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f310])).
fof(f759,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f446,f645])).
fof(f645,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f309])).
fof(f4504,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3381,f3097])).
fof(f3381,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f445,f3097,f561])).
fof(f561,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f4503,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3383,f740])).
fof(f740,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f449,f446])).
fof(f449,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f3383,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f774,f3097,f561])).
fof(f774,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f740,f468])).
fof(f468,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f326])).
fof(f4502,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3391,f3097])).
fof(f3391,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f445,f3097,f561])).
fof(f4501,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3393,f740])).
fof(f3393,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f774,f3097,f561])).
fof(f4500,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3395,f740])).
fof(f3395,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f808,f3097,f587])).
fof(f587,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f392])).
fof(f392,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f808,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f740,f446,f589])).
fof(f589,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f393])).
fof(f4499,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3397,f3097])).
fof(f3397,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f858,f3097,f587])).
fof(f858,plain,(
  ( ! [X0] : (sP1('F_0','const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f740,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (sP1(X1,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f816,f446])).
fof(f816,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1(X1,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f592,f453])).
fof(f592,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f393])).
fof(f4498,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3398,f3097])).
fof(f3398,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1042,f3097,f587])).
fof(f1042,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f740,f511,f592])).
fof(f511,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X7] : ! [X36] : p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X7,bool),X36),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.pred_set.DISJOINT_EMPTY|split|1')).
fof(f4497,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3399,f740])).
fof(f3399,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f848,f3097,f587])).
fof(f848,plain,(
  ( ! [X0] : (sP1('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f740,f810])).
fof(f810,plain,(
  ( ! [X0,X1] : (sP1('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f809,f446])).
fof(f809,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1)) )),
  inference(superposition,[],[f589,f453])).
fof(f4496,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3400,f740])).
fof(f3400,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1041,f3097,f587])).
fof(f1041,plain,(
  ( ! [X0,X1] : (sP1('const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f740,f511,f589])).
fof(f4495,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3406,f740])).
fof(f3406,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f815,f3097,f587])).
fof(f815,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f740,f446,f592])).
fof(f4494,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3407,f740])).
fof(f3407,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f858,f3097,f587])).
fof(f4493,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3408,f740])).
fof(f3408,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1042,f3097,f587])).
fof(f4492,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3409,f3097])).
fof(f3409,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f848,f3097,f587])).
fof(f4491,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3410,f3097])).
fof(f3410,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1041,f3097,f587])).
fof(f4490,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3449,f3692])).
fof(f3692,plain,(
  ( ! [X0] : (sP5('T_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f755,f752,f759,f3097,f634])).
fof(f634,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f410])).
fof(f752,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f636])).
fof(f636,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f413])).
fof(f755,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f640])).
fof(f640,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f416])).
fof(f3449,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f446,f3097,f631])).
fof(f631,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f410])).
fof(f4489,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3493,f3756])).
fof(f3756,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f413])).
fof(f3493,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f758,f3097,f633])).
fof(f633,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f410])).
fof(f758,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f644])).
fof(f644,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f418])).
fof(f1350,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f774,f649])).
fof(f4488,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3494,f3756])).
fof(f3494,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f758,f3097,f633])).
fof(f1347,plain,(
  ~sP5('T_0','T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(unit_resulting_resolution,[],[f446,f445,f649])).
fof(f4487,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3496,f3756])).
fof(f3496,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f758,f3097,f633])).
fof(f756,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f740,f641])).
fof(f4486,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3497,f3756])).
fof(f3497,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f758,f3097,f633])).
fof(f4485,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3518,f3768])).
fof(f3768,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f642])).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f416])).
fof(f3518,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f752,f758,f3097,f633])).
fof(f4484,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3519,f3768])).
fof(f3519,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f752,f758,f3097,f633])).
fof(f4483,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3526,f3768])).
fof(f3526,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f753,f758,f3097,f633])).
fof(f753,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f740,f637])).
fof(f4482,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3533,f3768])).
fof(f3533,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f790,f758,f3097,f633])).
fof(f4481,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3540,f3756])).
fof(f3540,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f759,f3097,f633])).
fof(f4480,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3541,f3756])).
fof(f3541,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f759,f3097,f633])).
fof(f4479,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3542,f3756])).
fof(f3542,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f759,f3097,f633])).
fof(f4478,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3543,f3756])).
fof(f3543,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f759,f3097,f633])).
fof(f4477,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3544,f3768])).
fof(f3544,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f752,f759,f3097,f633])).
fof(f4476,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3545,f3768])).
fof(f3545,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f752,f759,f3097,f633])).
fof(f4475,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3556,f3768])).
fof(f3556,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f753,f759,f3097,f633])).
fof(f4474,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3561,f3768])).
fof(f3561,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f790,f759,f3097,f633])).
fof(f4473,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3635,f3756])).
fof(f3635,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f758,f3097,f634])).
fof(f4472,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3636,f3756])).
fof(f3636,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f758,f3097,f634])).
fof(f4471,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3638,f3756])).
fof(f3638,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f758,f3097,f634])).
fof(f4470,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3639,f3756])).
fof(f3639,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f758,f3097,f634])).
fof(f4469,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3660,f3768])).
fof(f3660,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f752,f758,f3097,f634])).
fof(f4468,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3661,f3768])).
fof(f3661,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f752,f758,f3097,f634])).
fof(f4467,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3668,f3768])).
fof(f3668,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f753,f758,f3097,f634])).
fof(f4466,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3675,f3768])).
fof(f3675,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f790,f758,f3097,f634])).
fof(f4465,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3682,f3756])).
fof(f3682,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f759,f3097,f634])).
fof(f4464,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3683,f3756])).
fof(f3683,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f759,f3097,f634])).
fof(f4463,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3684,f3756])).
fof(f3684,plain,(
  ~sP4('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f759,f3097,f634])).
fof(f4462,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3685,f3756])).
fof(f3685,plain,(
  ~sP4('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f759,f3097,f634])).
fof(f4461,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3686,f3768])).
fof(f3686,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f752,f759,f3097,f634])).
fof(f4460,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3687,f3768])).
fof(f3687,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f752,f759,f3097,f634])).
fof(f4459,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3698,f3768])).
fof(f3698,plain,(
  ~sP3('T_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f753,f759,f3097,f634])).
fof(f4458,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3703,f3768])).
fof(f3703,plain,(
  ~sP3('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f790,f759,f3097,f634])).
fof(f4457,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3752,f3756])).
fof(f3752,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f446,f3097,f635])).
fof(f635,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f413])).
fof(f4456,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3764,f3768])).
fof(f3764,plain,(
  ( ! [X0] : (~sP3('T_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f446,f3097,f639])).
fof(f639,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f416])).
fof(f4455,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3771,f3097])).
fof(f3771,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f740,f3097,f643])).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f418])).
fof(f760,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f446,f646])).
fof(f646,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f418])).
fof(f4454,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3772,f3097])).
fof(f3772,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f643])).
fof(f4453,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3774,f3097])).
fof(f3774,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f3097,f3097,f643])).
fof(f4452,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3776,f3097])).
fof(f3776,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f740,f3097,f643])).
fof(f4451,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3782,f3097])).
fof(f3782,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f3097,f3097,f643])).
fof(f4450,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3783,f3097])).
fof(f3783,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f3097,f3097,f643])).
fof(f4449,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3789,f3097])).
fof(f3789,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f740,f3097,f643])).
fof(f4448,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3791,f3097])).
fof(f3791,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f3097,f3097,f643])).
fof(f4447,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3792,f3097])).
fof(f3792,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f3097,f3097,f643])).
fof(f4446,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3793,f3097])).
fof(f3793,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f740,f3097,f643])).
fof(f4445,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3799,f3097])).
fof(f3799,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f643])).
fof(f4444,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3800,f3097])).
fof(f3800,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f3097,f3097,f643])).
fof(f4443,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3806,f3097])).
fof(f3806,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f740,f3097,f643])).
fof(f4442,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3808,f3097])).
fof(f3808,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f643])).
fof(f4441,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3810,f3097])).
fof(f3810,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f760,f3097,f3097,f643])).
fof(f4440,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3811,f3097])).
fof(f3811,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f740,f3097,f643])).
fof(f4439,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3817,f3097])).
fof(f3817,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f643])).
fof(f4438,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3818,f3097])).
fof(f3818,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f759,f3097,f3097,f643])).
fof(f4437,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3824,f4250])).
fof(f4250,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f729])).
fof(f729,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f728])).
fof(f728,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f715])).
fof(f715,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X12,X17,X13,X18] : ((s(bool,X12) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X13) = s(bool,X18))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.IMP_CONG')).
fof(f3824,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1346,f3097,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f419])).
fof(f1346,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f446,f774,f649])).
fof(f4436,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3825,f4250])).
fof(f3825,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f3097,f647])).
fof(f4435,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3826,f445])).
fof(f3826,plain,(
  s(bool,'T_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f3097,f647])).
fof(f4434,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3827,f445])).
fof(f3827,plain,(
  s(bool,'T_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1343,f3097,f647])).
fof(f1343,plain,(
  ~sP5('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f446,f445,f649])).
fof(f4429,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3856,f606])).
fof(f606,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X8,X7] : ! [X5,X19] : s(X7,X5) = s(X7,'const.combin.K_2'(s(X7,X5),s(X8,X19)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.combin.K_DEF')).
fof(f3856,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1335,f3097,f650])).
fof(f650,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f419])).
fof(f1335,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f446,f606,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f419])).
fof(f4428,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3857,f513])).
fof(f513,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X44] : ! [X45] : s(X44,X45) = s(X44,happ(s(fun(X44,X44),'_i_0_0'),s(X44,X45)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','_i0_ax')).
fof(f3857,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1336,f3097,f650])).
fof(f1336,plain,(
  ( ! [X0] : (sP5('T_0',X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f446,f513,f648])).
fof(f4427,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3860,f800])).
fof(f800,plain,(
  s(bool,'F_0') = s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(unit_resulting_resolution,[],[f782,f473])).
fof(f473,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f3860,plain,(
  s(bool,'F_0') != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1334,f3097,f650])).
fof(f1334,plain,(
  sP5('T_0','F_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(unit_resulting_resolution,[],[f446,f800,f648])).
fof(f4426,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3861,f741])).
fof(f741,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(forward_demodulation,[],[f514,f455])).
fof(f514,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X7] : ! [X12] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.NIL_0')))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.list.MEM|split|0')).
fof(f3861,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1410,f3097,f650])).
fof(f1410,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f446,f741,f648])).
fof(f4425,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3863,f800])).
fof(f3863,plain,(
  s(bool,'F_0') != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1330,f3097,f650])).
fof(f1330,plain,(
  sP5('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f446,f800,f648])).
fof(f4424,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4423,f606])).
fof(f4423,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1)))) ) | $spl0),
  inference(forward_demodulation,[],[f3864,f800])).
fof(f3864,plain,(
  ( ! [X0,X1] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1))) != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1592,f3097,f650])).
fof(f1592,plain,(
  ( ! [X6,X7] : (sP5('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'const.combin.K_2'(s(bool,'F_0'),s(X6,X7)))) )),
  inference(superposition,[],[f1335,f800])).
fof(f4422,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4421,f513])).
fof(f4421,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl0),
  inference(forward_demodulation,[],[f3865,f800])).
fof(f3865,plain,(
  s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1812,f3097,f650])).
fof(f1812,plain,(
  sP5('T_0','const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))),
  inference(superposition,[],[f1336,f800])).
fof(f4420,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3866,f513])).
fof(f3866,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1332,f3097,f650])).
fof(f1332,plain,(
  ( ! [X0] : (sP5('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f446,f513,f648])).
fof(f4419,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4418,f513])).
fof(f4418,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl0),
  inference(forward_demodulation,[],[f3867,f800])).
fof(f3867,plain,(
  s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12))) | $spl0),
  inference(unit_resulting_resolution,[],[f1788,f3097,f650])).
fof(f1788,plain,(
  sP5('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')),'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))),
  inference(superposition,[],[f1332,f800])).
fof(f4417,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3868,f606])).
fof(f3868,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1331,f3097,f650])).
fof(f1331,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.combin.K_2'(s(bool,X0),s(X1,X2)),X0)) )),
  inference(unit_resulting_resolution,[],[f446,f606,f648])).
fof(f4416,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4415,f606])).
fof(f4415,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1)))) ) | $spl0),
  inference(forward_demodulation,[],[f3869,f800])).
fof(f3869,plain,(
  ( ! [X0,X1] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1))) != s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1545,f3097,f650])).
fof(f1545,plain,(
  ( ! [X6,X7] : (sP5('T_0','const.combin.K_2'(s(bool,'F_0'),s(X6,X7)),'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) )),
  inference(superposition,[],[f1331,f800])).
fof(f4414,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3870,f741])).
fof(f3870,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1407,f3097,f650])).
fof(f1407,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f446,f741,f648])).
fof(f4413,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3907,f4275])).
fof(f4275,plain,(
  ( ! [X0,X1] : (sP7('T_0',X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f761,f3097,f732])).
fof(f732,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f657])).
fof(f657,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f421])).
fof(f421,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f420])).
fof(f420,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f761,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f659])).
fof(f659,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f423])).
fof(f423,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f3907,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f652])).
fof(f652,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f422])).
fof(f4412,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3914,f4275])).
fof(f3914,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f653])).
fof(f653,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f422])).
fof(f4411,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3922,f3926])).
fof(f3926,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f424])).
fof(f3922,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f446,f3097,f658])).
fof(f658,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f424])).
fof(f4410,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3930,f740])).
fof(f3930,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1526,f3097,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f245,f315,f314])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X38,X39,X40] : ((p(s(bool,X40)) <=> (p(s(bool,X39)) & p(s(bool,X38)))) <=> ((p(s(bool,X40)) | ~p(s(bool,X39)) | ~p(s(bool,X38))) & (p(s(bool,X39)) | ~p(s(bool,X40))) & (p(s(bool,X38)) | ~p(s(bool,X40)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.sat.sat_23_prove_ltb')).
fof(f1526,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1518,f651])).
fof(f651,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f422])).
fof(f1518,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f740,f446,f658])).
fof(f4409,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3935,f740])).
fof(f3935,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f662])).
fof(f829,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f446,f740,f652])).
fof(f4408,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3936,f3097])).
fof(f3936,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f830,f3097,f662])).
fof(f830,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f740,f653])).
fof(f4407,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3944,f740])).
fof(f3944,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1526,f3097,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f426])).
fof(f4406,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3949,f3097])).
fof(f3949,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f663])).
fof(f4405,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3950,f740])).
fof(f3950,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f830,f3097,f663])).
fof(f4404,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3970,f4114])).
fof(f4114,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f3097,f3097,f679])).
fof(f679,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f433])).
fof(f433,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f247,f318,f317])).
fof(f317,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f318,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X38,X39,X40] : ((p(s(bool,X40)) <=> (p(s(bool,X39)) | p(s(bool,X38)))) <=> ((p(s(bool,X40)) | ~p(s(bool,X39))) & (p(s(bool,X40)) | ~p(s(bool,X38))) & (p(s(bool,X39)) | p(s(bool,X38)) | ~p(s(bool,X40)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.sat.sat_25_prove_ltb')).
fof(f3970,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f668])).
fof(f668,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f428])).
fof(f428,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f427])).
fof(f427,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f318])).
fof(f4403,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3978,f4114])).
fof(f3978,plain,(
  ( ! [X0,X1] : (~sP9(X0,'T_0',X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f669])).
fof(f669,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f429])).
fof(f4402,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3994,f3097])).
fof(f3994,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f766,f859,f3097,f674])).
fof(f674,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f429])).
fof(f859,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f740,f669])).
fof(f766,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f740,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f431])).
fof(f431,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f430])).
fof(f430,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f317])).
fof(f4401,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3995,f3097])).
fof(f3995,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f765,f859,f3097,f674])).
fof(f765,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f446,f677])).
fof(f677,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f432])).
fof(f4400,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4000,f4060])).
fof(f4060,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f678])).
fof(f4000,plain,(
  ~sP8('F_0','T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f859,f740,f3097,f674])).
fof(f4399,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4007,f4060])).
fof(f4007,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f674])).
fof(f857,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f740,f668])).
fof(f4398,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4009,f4060])).
fof(f4009,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f674])).
fof(f4397,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4012,f3097])).
fof(f4012,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f764,f3097,f674])).
fof(f764,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f676])).
fof(f676,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f432])).
fof(f4396,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4015,f3097])).
fof(f4015,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f766,f3097,f674])).
fof(f4395,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4017,f4060])).
fof(f4017,plain,(
  ~sP8('F_0','F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f3097,f674])).
fof(f4394,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4024,f4060])).
fof(f4024,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f674])).
fof(f4393,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4025,f4060])).
fof(f4025,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f674])).
fof(f4392,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4035,f3097])).
fof(f4035,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f765,f446,f3097,f675])).
fof(f675,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f432])).
fof(f4391,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4039,f4060])).
fof(f4039,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f740,f3097,f675])).
fof(f4390,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4045,f4060])).
fof(f4045,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f3097,f675])).
fof(f4389,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4047,f3097])).
fof(f4047,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f764,f446,f3097,f675])).
fof(f4388,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4052,f4060])).
fof(f4052,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f740,f3097,f675])).
fof(f4387,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4058,f4060])).
fof(f4058,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097,f3097,f675])).
fof(f4386,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4063,f3097])).
fof(f4063,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f740,f3097,f679])).
fof(f4385,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4065,f3097])).
fof(f4065,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f679])).
fof(f4384,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4066,f740])).
fof(f4066,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f3097,f3097,f679])).
fof(f1686,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1679,f670])).
fof(f670,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f429])).
fof(f1679,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f740,f740,f675])).
fof(f4383,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4067,f740])).
fof(f4067,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f740,f3097,f679])).
fof(f4382,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4073,f740])).
fof(f4073,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f679])).
fof(f4381,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4074,f740])).
fof(f4074,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f679])).
fof(f4380,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4080,f3097])).
fof(f4080,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f3097,f679])).
fof(f4379,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4082,f3097])).
fof(f4082,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f679])).
fof(f4378,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4084,f740])).
fof(f4084,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f3097,f3097,f679])).
fof(f4377,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4085,f740])).
fof(f4085,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f3097,f679])).
fof(f4376,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4091,f740])).
fof(f4091,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f679])).
fof(f4375,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4092,f740])).
fof(f4092,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f679])).
fof(f4374,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4099,f740])).
fof(f4099,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f740,f3097,f679])).
fof(f4373,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4100,f3097])).
fof(f4100,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f679])).
fof(f4372,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4102,f740])).
fof(f4102,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f3097,f3097,f679])).
fof(f4371,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4104,f740])).
fof(f4104,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f740,f3097,f679])).
fof(f4370,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4110,f3097])).
fof(f4110,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f3097,f679])).
fof(f4369,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4111,f740])).
fof(f4111,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f3097,f3097,f679])).
fof(f4368,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4117,f3097])).
fof(f4117,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f873,f446,f3097,f682])).
fof(f682,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f434])).
fof(f873,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f446,f681])).
fof(f681,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f434])).
fof(f4367,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4142,f3097])).
fof(f4142,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f864,f446,f3097,f682])).
fof(f864,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f446,f680])).
fof(f680,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f434])).
fof(f4366,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4148,f740])).
fof(f4148,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f884,f3097,f685])).
fof(f685,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f435])).
fof(f435,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f884,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f740,f446,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f302,f321,f320])).
fof(f320,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X38,X39,X40] : ((p(s(bool,X40)) <=> (p(s(bool,X39)) => p(s(bool,X38)))) <=> ((p(s(bool,X40)) | p(s(bool,X39))) & (p(s(bool,X40)) | ~p(s(bool,X38))) & (~p(s(bool,X39)) | p(s(bool,X38)) | ~p(s(bool,X40)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.sat.sat_27_prove_ltb')).
fof(f4365,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4149,f3097])).
fof(f4149,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f907,f3097,f685])).
fof(f907,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f797,f446,f730])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f688])).
fof(f688,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f437])).
fof(f797,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f782,f693])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f439])).
fof(f439,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f438])).
fof(f438,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f320])).
fof(f4364,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4150,f3097])).
fof(f4150,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f890,f3097,f685])).
fof(f890,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f446,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f442])).
fof(f4363,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4153,f740])).
fof(f4153,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1027,f3097,f685])).
fof(f1027,plain,(
  ( ! [X0,X1] : (sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f740,f887])).
fof(f887,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f885,f446])).
fof(f885,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X1)) )),
  inference(superposition,[],[f697,f453])).
fof(f4362,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4154,f3097])).
fof(f4154,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1099,f3097,f685])).
fof(f1099,plain,(
  ( ! [X0,X1] : (sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f446,f893])).
fof(f893,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1,X2) | ~p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f891,f446])).
fof(f891,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP11('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1,X2)) )),
  inference(superposition,[],[f698,f453])).
fof(f4361,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4155,f3097])).
fof(f4155,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f1081,f3097,f685])).
fof(f1081,plain,(
  ( ! [X2,X0,X1] : (sP11('const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f797,f511,f730])).
fof(f4360,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4156,f3097])).
fof(f4156,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f3097,f685])).
fof(f1149,plain,(
  ( ! [X2,X0,X1] : (sP11('const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1)),X2,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f797,f512,f730])).
fof(f4359,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4164,f740])).
fof(f4164,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1827,f3097,f685])).
fof(f1827,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f769,f446,f740,f691])).
fof(f691,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f437])).
fof(f769,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f740,f695])).
fof(f695,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f440])).
fof(f4358,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4165,f3097])).
fof(f4165,plain,(
  p(s(bool,'const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)))) | $spl0),
  inference(unit_resulting_resolution,[],[f1828,f3097,f685])).
fof(f1828,plain,(
  sP11('const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f798,f446,f740,f691])).
fof(f798,plain,(
  ( ! [X0,X1] : (sP10('const.modSem.no_dup_top_types_2'(s('type.list.list'('type.modLang.prompt'),'const.list.NIL_0'),s(fun('type.semanticPrimitives.tid_or_exn',bool),sK12)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f782,f695])).
fof(f4357,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4166,f3097])).
fof(f4166,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1027,f3097,f685])).
fof(f4356,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4167,f3097])).
fof(f4167,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1099,f3097,f685])).
fof(f4355,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4168,f3097])).
fof(f4168,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1081,f3097,f685])).
fof(f4354,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4169,f3097])).
fof(f4169,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X1))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f3097,f685])).
fof(f4353,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4199,f4212])).
fof(f4212,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f695])).
fof(f4199,plain,(
  ~sP10('F_0','T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f881,f446,f3097,f691])).
fof(f881,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f740,f686])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f437])).
fof(f4352,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4208,f4212])).
fof(f4208,plain,(
  ( ! [X0] : (~sP10('T_0',X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f446,f446,f3097,f692])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f440])).
fof(f4351,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4219,f740])).
fof(f4219,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f881,f446,f3097,f696])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f442])).
fof(f4350,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4226,f740])).
fof(f4226,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f1844,f446,f3097,f696])).
fof(f1844,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1836,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f437])).
fof(f1836,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f446,f446,f740,f692])).
fof(f4349,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4280,f761])).
fof(f4280,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f732])).
fof(f4348,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4282,f761])).
fof(f4282,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f830,f3097,f732])).
fof(f4347,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4286,f3097])).
fof(f4286,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X0,X1))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1178,f3097,f743])).
fof(f743,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f572,f740])).
fof(f572,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1178,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f446,f946])).
fof(f946,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,X12)) | sP0(X15,'const.combin.K_2'(s(bool,X12),s(X13,X14)))) )),
  inference(superposition,[],[f573,f606])).
fof(f573,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f386])).
fof(f4346,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4287,f3097])).
fof(f4287,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f773,f3097,f743])).
fof(f773,plain,(
  ( ! [X2,X3] : (sP0(X3,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X2),'const.list.NIL_0')))) )),
  inference(subsumption_resolution,[],[f771,f446])).
fof(f771,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP0(X3,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X2),'const.list.NIL_0')))) )),
  inference(superposition,[],[f573,f453])).
fof(f4345,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4288,f3097])).
fof(f4288,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1035,f3097,f743])).
fof(f1035,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.pred_set.DISJOINT_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f511,f573])).
fof(f4344,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f4289,f3097])).
fof(f4289,plain,(
  p(s(bool,sK14('_i_0_0',bool))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f3108,f3097,f743])).
fof(f3108,plain,(
  ( ! [X0] : (sP0(X0,sK14('_i_0_0',bool))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f3099,f573])).
fof(f4343,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4290,f3097])).
fof(f4290,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f751,f3097,f743])).
fof(f751,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f446,f574])).
fof(f574,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f386])).
fof(f4342,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4291,f3097])).
fof(f4291,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1235,f3097,f743])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (sP0('const.combin.K_2'(s(bool,'T_0'),s(X0,X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f446,f947])).
fof(f947,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,X16)) | sP0('const.combin.K_2'(s(bool,X16),s(X17,X18)),X19)) )),
  inference(superposition,[],[f574,f606])).
fof(f4341,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4292,f3097])).
fof(f4292,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f772,f3097,f743])).
fof(f772,plain,(
  ( ! [X0,X1] : (sP0('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1)) )),
  inference(subsumption_resolution,[],[f770,f446])).
fof(f770,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP0('const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),X1)) )),
  inference(superposition,[],[f574,f453])).
fof(f4340,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4293,f3097])).
fof(f4293,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1036,f3097,f743])).
fof(f1036,plain,(
  ( ! [X2,X0,X1] : (sP0('const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f511,f574])).
fof(f4339,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4298,f3097])).
fof(f4298,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f750,f3097,f743])).
fof(f750,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f573])).
fof(f4338,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4299,f3097])).
fof(f4299,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1178,f3097,f743])).
fof(f4337,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4300,f3097])).
fof(f4300,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f773,f3097,f743])).
fof(f4336,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4301,f3097])).
fof(f4301,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1035,f3097,f743])).
fof(f4335,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f4302,f3097])).
fof(f4302,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f3108,f3097,f743])).
fof(f4334,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4304,f3097])).
fof(f4304,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X0,X1))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1235,f3097,f743])).
fof(f4333,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4305,f3097])).
fof(f4305,plain,(
  ( ! [X0] : (p(s(bool,'const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f772,f3097,f743])).
fof(f4332,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f4306,f3097])).
fof(f4306,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1036,f3097,f743])).
fof(f4310,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f750,f3097,f3097,f743])).
fof(f4308,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f750,f740,f3097,f743])).
fof(f4297,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f750,f3097,f3097,f743])).
fof(f4295,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f751,f740,f3097,f743])).
fof(f4283,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f763,f830,f3097,f732])).
fof(f763,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f740,f661])).
fof(f4281,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f762,f829,f3097,f732])).
fof(f762,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f740,f660])).
fof(f660,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f424])).
fof(f4265,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f608,f3097,f729])).
fof(f608,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X7,X8] : ! [X35,X22] : s(X8,X35) = s(X8,'const.option.option_CASE_3'(s('type.option.option'(X7),'const.option.NONE_0'),s(X8,X35),s(fun(X7,X8),X22)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.option.option_case_def|split|0')).
fof(f4263,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f744,f3097,f729])).
fof(f744,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f610,f606])).
fof(f610,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X7,X8] : ! [X41,X42] : s(X7,'const.combin.K_2'(s(X7,X41),s(X8,X42))) = s(X7,happ(s(fun(X8,X7),'const.combin.K_1'(s(X7,X41))),s(X8,X42)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','ARITY_1_2_const.combin.K')).
fof(f4261,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f607,f3097,f729])).
fof(f607,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X44,X46] : ! [X48,X49] : s(X46,X48) = s(X46,happ(s(fun(X44,X46),'_dst_x0_1'(s(X46,X48))),s(X44,X49)))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','_dst_x0_ax')).
fof(f4259,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f513,f3097,f729])).
fof(f4257,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f606,f3097,f729])).
fof(f4252,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f453,f3097,f729])).
fof(f4249,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f890,f446,f446,f3097,f699])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f442])).
fof(f4240,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1844,f446,f3097,f697])).
fof(f4235,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f881,f446,f3097,f3097,f696])).
fof(f4230,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1844,f446,f740,f3097,f696])).
fof(f4224,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f881,f3097,f446,f3097,f696])).
fof(f4222,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f881,f740,f446,f3097,f696])).
fof(f4213,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1836,f3097,f695])).
fof(f4211,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1836,f3097,f693])).
fof(f4209,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f768,f446,f446,f3097,f692])).
fof(f768,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f446,f694])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f440])).
fof(f4200,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f769,f881,f446,f3097,f691])).
fof(f4184,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f890,f446,f3097,f686])).
fof(f4173,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f884,f3097,f3097,f685])).
fof(f4171,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1827,f740,f3097,f685])).
fof(f4160,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f884,f3097,f3097,f685])).
fof(f4158,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f884,f740,f3097,f685])).
fof(f4147,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f864,f3097,f446,f3097,f682])).
fof(f4145,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f864,f740,f446,f3097,f682])).
fof(f4131,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f864,f446,f3097,f3097,f682])).
fof(f4123,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f873,f446,f740,f3097,f682])).
fof(f4115,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f3097,f679])).
fof(f4113,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f859,f740,f3097,f3097,f679])).
fof(f4108,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f740,f3097,f679])).
fof(f4106,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1686,f740,f740,f3097,f679])).
fof(f4097,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f3097,f679])).
fof(f4095,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f3097,f3097,f679])).
fof(f4090,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f740,f3097,f679])).
fof(f4088,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f740,f3097,f679])).
fof(f4079,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f3097,f3097,f3097,f679])).
fof(f4077,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f740,f3097,f3097,f679])).
fof(f4072,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f859,f3097,f740,f3097,f679])).
fof(f4070,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f859,f740,f740,f3097,f679])).
fof(f4061,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1679,f3097,f678])).
fof(f4059,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f764,f446,f3097,f3097,f675])).
fof(f4053,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f764,f446,f740,f3097,f675])).
fof(f4046,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f764,f446,f3097,f3097,f675])).
fof(f4040,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f765,f446,f740,f3097,f675])).
fof(f4032,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f766,f3097,f3097,f674])).
fof(f4030,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f859,f765,f3097,f3097,f674])).
fof(f4028,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f764,f3097,f3097,f674])).
fof(f4022,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f766,f740,f3097,f674])).
fof(f4020,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f857,f764,f740,f3097,f674])).
fof(f4010,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f766,f859,f3097,f3097,f674])).
fof(f4008,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f766,f857,f3097,f3097,f674])).
fof(f4001,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f766,f859,f740,f3097,f674])).
fof(f3979,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f873,f446,f3097,f669])).
fof(f3971,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f864,f446,f3097,f668])).
fof(f3963,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1509,f446,f446,f3097,f667])).
fof(f667,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f426])).
fof(f1509,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f761,f446,f446,f654])).
fof(f654,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f422])).
fof(f3955,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f3097,f663])).
fof(f3953,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f830,f740,f3097,f663])).
fof(f3948,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f3097,f663])).
fof(f3946,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1526,f740,f3097,f663])).
fof(f3941,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f3097,f662])).
fof(f3939,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f829,f740,f3097,f662])).
fof(f3934,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f829,f3097,f3097,f662])).
fof(f3932,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1526,f740,f3097,f662])).
fof(f3927,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1518,f3097,f661])).
fof(f3925,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1518,f3097,f660])).
fof(f3923,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f761,f446,f446,f3097,f658])).
fof(f3915,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1509,f446,f3097,f653])).
fof(f3908,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1509,f446,f3097,f652])).
fof(f3898,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1336,f513,f3097,f650])).
fof(f3896,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1335,f606,f3097,f650])).
fof(f3893,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1410,f741,f3097,f650])).
fof(f3891,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1334,f800,f3097,f650])).
fof(f3889,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1333,f453,f3097,f650])).
fof(f1333,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f446,f453,f648])).
fof(f3884,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1332,f513,f3097,f650])).
fof(f3882,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1331,f606,f3097,f650])).
fof(f3879,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1407,f741,f3097,f650])).
fof(f3877,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1330,f800,f3097,f650])).
fof(f3875,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1329,f453,f3097,f650])).
fof(f1329,plain,(
  ( ! [X0] : (sP5('T_0','const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f453,f648])).
fof(f4331,plain,(
  $false | $spl0),
  inference(trivial_inequality_removal,[],[f3862])).
fof(f3862,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2474,f3097,f650])).
fof(f2474,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f757,f753,f758,f740,f633])).
fof(f757,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f740,f642])).
fof(f4330,plain,(
  $false | $spl0),
  inference(trivial_inequality_removal,[],[f3859])).
fof(f3859,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2422,f3097,f650])).
fof(f2422,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f755,f752,f759,f446,f632])).
fof(f3851,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f774,f3097,f647])).
fof(f3849,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f445,f3097,f647])).
fof(f3847,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1346,f774,f3097,f647])).
fof(f3845,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1343,f445,f3097,f647])).
fof(f3823,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f3097,f643])).
fof(f3821,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f740,f3097,f3097,f643])).
fof(f3816,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f740,f3097,f643])).
fof(f3814,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f740,f740,f3097,f643])).
fof(f3805,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f3097,f643])).
fof(f3803,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f740,f3097,f3097,f643])).
fof(f3798,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f759,f3097,f740,f3097,f643])).
fof(f3796,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f759,f740,f740,f3097,f643])).
fof(f3787,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f3097,f3097,f643])).
fof(f3785,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f759,f740,f3097,f3097,f643])).
fof(f3780,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f758,f3097,f740,f3097,f643])).
fof(f3778,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f760,f740,f740,f3097,f643])).
fof(f3769,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1488,f3097,f642])).
fof(f1488,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f740,f446,f639])).
fof(f3767,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1488,f3097,f641])).
fof(f3765,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f755,f446,f446,f3097,f639])).
fof(f3757,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1474,f3097,f638])).
fof(f1474,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f740,f446,f635])).
fof(f3755,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1474,f3097,f637])).
fof(f3753,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1461,f446,f446,f3097,f635])).
fof(f1461,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.list.ALL_DISTINCT_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f1329,f630])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f410])).
fof(f3707,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f790,f759,f3097,f634])).
fof(f3705,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f790,f759,f3097,f634])).
fof(f3702,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f753,f759,f3097,f634])).
fof(f3700,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f753,f759,f3097,f634])).
fof(f3697,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f752,f759,f3097,f634])).
fof(f3695,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f752,f759,f3097,f634])).
fof(f3693,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f752,f759,f3097,f634])).
fof(f3679,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f790,f758,f3097,f634])).
fof(f3677,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f790,f758,f3097,f634])).
fof(f3672,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f753,f758,f3097,f634])).
fof(f3670,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f753,f758,f3097,f634])).
fof(f3667,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f752,f758,f3097,f634])).
fof(f3665,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f752,f758,f3097,f634])).
fof(f3663,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f752,f758,f3097,f634])).
fof(f3565,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f790,f759,f3097,f633])).
fof(f3563,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f790,f759,f3097,f633])).
fof(f3560,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f753,f759,f3097,f633])).
fof(f3558,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f753,f759,f3097,f633])).
fof(f3555,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f752,f759,f3097,f633])).
fof(f3553,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f752,f759,f3097,f633])).
fof(f3551,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f752,f759,f3097,f633])).
fof(f3537,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f755,f790,f758,f3097,f633])).
fof(f3535,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f790,f758,f3097,f633])).
fof(f3530,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f753,f758,f3097,f633])).
fof(f3528,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f753,f758,f3097,f633])).
fof(f3525,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1347,f792,f752,f758,f3097,f633])).
fof(f3523,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f756,f752,f758,f3097,f633])).
fof(f3521,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1350,f755,f752,f758,f3097,f633])).
fof(f3450,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f1333,f446,f446,f3097,f631])).
fof(f3426,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f805,f446,f3097,f592])).
fof(f805,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f446,f446,f588])).
fof(f588,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f393])).
fof(f3420,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f805,f446,f3097,f589])).
fof(f3414,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f808,f3097,f3097,f587])).
fof(f3412,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f815,f740,f3097,f587])).
fof(f3404,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f808,f3097,f3097,f587])).
fof(f3402,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f808,f740,f3097,f587])).
fof(f3394,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f740,f774,f3097,f561])).
fof(f3392,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f445,f3097,f561])).
fof(f3390,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f774,f3097,f561])).
fof(f3388,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f445,f3097,f561])).
fof(f3384,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f740,f774,f3097,f561])).
fof(f3382,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f445,f3097,f561])).
fof(f3380,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f774,f3097,f561])).
fof(f3378,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f3097,f445,f3097,f561])).
fof(f3374,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f774,f3097,f473])).
fof(f3372,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f453,f3097,f468])).
fof(f3370,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f446,f3097])).
fof(f3097,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) | $spl0),
  inference(cnf_transformation,[],[f3097_D])).
fof(f3097_D,plain,(
  ( ! [X5] : (~p(s(bool,X5))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f3367,plain,(
  p(s(bool,sK15('_i_0_0',bool))) | $spl4),
  inference(cnf_transformation,[],[f3367_D])).
fof(f3367_D,plain,(
  p(s(bool,sK15('_i_0_0',bool))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3365,plain,(
  ( ! [X5] : (p(s(bool,sK15('_i_0_0',bool))) | ~p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f3356,f513])).
fof(f3356,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK15('_i_0_0',bool)))))) )),
  inference(superposition,[],[f521,f513])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))) )),
  inference(cnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f355])).
fof(f355,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7] : ! [X14] : (~? [X12] : p(s(bool,happ(s(fun(X7,bool),X14),s(X7,X12)))) <=> ! [X12] : ~p(s(bool,happ(s(fun(X7,bool),X14),s(X7,X12)))))),
  file('thm.modProps.no_dup_top_types_eqn|split|0','thm.bool.NOT_EXISTS_THM')).
fof(f3099,plain,(
  p(s(bool,sK14('_i_0_0',bool))) | $spl2),
  inference(cnf_transformation,[],[f3099_D])).
fof(f3099_D,plain,(
  p(s(bool,sK14('_i_0_0',bool))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3096,plain,(
  ( ! [X5] : (p(s(bool,sK14('_i_0_0',bool))) | ~p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f3087,f513])).
fof(f3087,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK14('_i_0_0',bool)))))) )),
  inference(superposition,[],[f520,f513])).
fof(f520,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) )),
  inference(cnf_transformation,[],[f356])).
% SZS output end Proof for thm.modProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 616
% Passive clauses: 3116
% Generated clauses: 17177
% Final active clauses: 607
% Final passive clauses: 2416
% Input formulas: 100
% Initial clauses: 273
% 
% Duplicate literals: 7
% Trivial inequalities: 5
% Fw subsumption resolutions: 227
% Fw demodulations: 3179
% Bw demodulations: 37
% Condensations: 21
% 
% Simple tautologies: 77
% Equational tautologies: 27
% Forward subsumptions: 10162
% Fw demodulations to eq. taut.: 36
% 
% Unit resulting resolution: 5007
% Factoring: 1
% Forward superposition: 4659
% Backward superposition: 3677
% Self superposition: 76
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 3440
% SAT solver unit clauses: 1798
% SAT solver binary clauses: 1138
% 
% Sat splits: 2
% Sat splitting refutations: 318
% 
% Memory used [KB]: 24818
% Time elapsed: 0.943 s
% ------------------------------
% ------------------------------
% Success in time 2.28 s
