Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f91624,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f91622,f1767,f2676,f91522,f91492,f1965,f2673,f2675])).
fof(f2675,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f2675_D])).
fof(f2675_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2673,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f2673_D])).
fof(f2673_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1965,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1514])).
fof(f1514,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1513])).
fof(f1513,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1320])).
fof(f1320,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X20] : ((p(s(bool,'F_0')) => p(s(bool,X20))) <=> p(s(bool,'T_0')))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals','thm.bool.IMP_CLAUSES|split|2')).
fof(f91492,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X0))))))) ) | $spl578),
  inference(cnf_transformation,[],[f91492_D])).
fof(f91492_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X0))))))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f91522,plain,(
  ( ! [X1] : (p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X1))))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f91508,f2676])).
fof(f91508,plain,(
  ( ! [X1] : (p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X1)))))) | p(s(bool,'F_0'))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f91336,f2353])).
fof(f2353,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1652])).
fof(f1652,plain,(
  ! [X0,X1] : ((~sP9(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1)))),
  inference(rectify,[],[f1651])).
fof(f1651,plain,(
  ! [X1,X0] : ((~sP9(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP9(X1,X0)))),
  inference(flattening,[],[f1650])).
fof(f1650,plain,(
  ! [X1,X0] : ((~sP9(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP9(X1,X0)))),
  inference(nnf_transformation,[],[f1425])).
fof(f1425,plain,(
  ! [X1,X0] : (sP9(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f91336,plain,(
  ( ! [X9] : (sP9('const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X9)))),'F_0')) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f3455,f2226])).
fof(f2226,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X1))) = s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp','const.closLang.Call_2'(s('type.num.num',X0),s('type.list.list'('type.closLang.exp'),X1)))))) )),
  inference(cnf_transformation,[],[f966])).
fof(f966,plain,(
  ! [X0,X1] : s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X1))) = s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp','const.closLang.Call_2'(s('type.num.num',X0),s('type.list.list'('type.closLang.exp'),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14,X13] : s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X13))) = s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp','const.closLang.Call_2'(s('type.num.num',X14),s('type.list.list'('type.closLang.exp'),X13)))))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals','thm.clos_knownProof.set_globals_def|split|6')).
fof(f3455,plain,(
  ( ! [X9] : (sP9('const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp',X9)))),'F_0')) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f1777,f2996])).
fof(f2996,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP9(X0,'F_0')) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2992,f2673])).
fof(f2992,plain,(
  ( ! [X0] : (sP9(X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f2976,f2354])).
fof(f2354,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1652])).
fof(f2976,plain,(
  ( ! [X0] : (sP9('T_0',X0) | sP9(X0,'F_0')) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f2892,f2676])).
fof(f2892,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0) | p(s(bool,X1)) | sP9(X0,X1)) ) | $spl0),
  inference(resolution,[],[f2876,f2355])).
fof(f2355,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f1652])).
fof(f2876,plain,(
  ( ! [X4] : (p(s(bool,X4)) | sP9('T_0',X4)) ) | $spl0),
  inference(resolution,[],[f2355,f2673])).
fof(f1777,plain,(
  ( ! [X0] : (p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp',X0))))))) )),
  inference(cnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0] : p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp',X0))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X15] : p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.set_globals_1'(s('type.closLang.exp',X15))))))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals','thm.clos_knownProof.FINITE_set_globals')).
fof(f2676,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f2676_D])).
fof(f2676_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1767,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals',aHL_FALSITY)).
fof(f91622,plain,(
  $false | ($spl3 | $spl578)),
  inference(subsumption_resolution,[],[f91619,f2676])).
fof(f91619,plain,(
  p(s(bool,'F_0')) | $spl578),
  inference(superposition,[],[f91492,f4113])).
fof(f4113,plain,(
  s(bool,'F_0') = s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),sK23)))))),
  inference(resolution,[],[f1748,f1912])).
fof(f1912,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1481])).
fof(f1481,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f626])).
fof(f626,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X20] : (s(bool,'F_0') = s(bool,X20) <=> ~p(s(bool,X20)))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals','thm.bool.EQ_CLAUSES|split|2')).
fof(f1748,plain,(
  ~p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),sK23))))))),
  inference(cnf_transformation,[],[f1447])).
fof(f1447,plain,(
  ~p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),sK23))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f1295])).
fof(f1295,plain,(
  ? [X0] : ~p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X0))))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.bag.FINITE_BAG_1'(s(fun('type.num.num','type.num.num'),'const.clos_knownProof.elist_globals_1'(s('type.list.list'('type.closLang.exp'),X6))))))),
  file('thm.clos_knownProof.FINITE_BAG_elist_globals','thm.clos_knownProof.FINITE_BAG_elist_globals')).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1583
% Passive clauses: 12771
% Generated clauses: 89826
% Final active clauses: 1475
% Final passive clauses: 10596
% Input formulas: 517
% Initial clauses: 914
% 
% Duplicate literals: 278
% Trivial inequalities: 240
% Fw subsumption resolutions: 3016
% Fw demodulations: 9904
% Bw demodulations: 92
% 
% Simple tautologies: 388
% Equational tautologies: 30
% Forward subsumptions: 62210
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 8527
% Factoring: 1458
% Forward superposition: 30735
% Backward superposition: 34278
% Self superposition: 85
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 1220
% Split components: 2442
% SAT solver clauses: 1557
% SAT solver unit clauses: 43
% SAT solver binary clauses: 109
% SAT solver learnt clauses: 1
% 
% Sat splits: 1220
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1555
% TWLsolver calls for satisfiability: 661
% 
% Memory used [KB]: 55649
% Time elapsed: 2.067 s
% ------------------------------
% ------------------------------
% Success in time 2.12 s
