Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.inferProps
% (100728)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 2284
% Generated clauses: 22746
% Final active clauses: 142
% Final passive clauses: 1887
% Input formulas: 45
% Initial clauses: 80
% 
% Duplicate literals: 1303
% Trivial inequalities: 40
% Fw subsumption resolutions: 2082
% Fw demodulations: 4242
% Bw demodulations: 7
% 
% Simple tautologies: 173
% Equational tautologies: 25
% Forward subsumptions: 11900
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 23
% Forward superposition: 8706
% Backward superposition: 5876
% Self superposition: 187
% Equality factoring: 58
% Equality resolution: 7
% 
% Split clauses: 429
% Split components: 868
% SAT solver clauses: 793
% SAT solver unit clauses: 1
% SAT solver binary clauses: 26
% 
% Sat splits: 429
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 790
% TWLsolver calls for satisfiability: 81
% 
% Memory used [KB]: 15863
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.inferProps
% (100760)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 1251
% Generated clauses: 9059
% Final active clauses: 114
% Final passive clauses: 1100
% Input formulas: 45
% Initial clauses: 80
% 
% Duplicate literals: 112
% Trivial inequalities: 12
% Fw subsumption resolutions: 94
% Fw demodulations: 815
% Bw demodulations: 37
% Condensations: 242
% Global subsumptions: 7
% 
% Simple tautologies: 35
% Equational tautologies: 16
% Forward subsumptions: 4707
% 
% Binary resolution: 5
% Forward superposition: 4497
% Backward superposition: 1750
% Self superposition: 102
% Equality factoring: 30
% Equality resolution: 5
% 
% Split clauses: 28
% Split components: 61
% SAT solver clauses: 1883
% SAT solver unit clauses: 125
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 10
% 
% Sat splits: 28
% 
% TWLsolver clauses: 1855
% TWLsolver calls for satisfiability: 642
% 
% Memory used [KB]: 24818
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.inferProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.inferProps
% SZS output start Proof for thm.inferProps
fof(f110851,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f162,f244,f242,f161,f243,f241,f272,f275,f273,f881,f882,f898,f883,f884,f899,f885,f900,f897,f3335,f3336,f3337,f3338,f3339,f3357,f3358,f3359,f3360,f3361,f3365,f3367,f3369,f3370,f3371,f3372,f3373,f3374,f3414,f3463,f3416,f3464,f3418,f3419,f3420,f3421,f3422,f3466,f3462,f3423,f3465,f3461,f3480,f3481,f3482,f3495,f3483,f3498,f8379,f8380,f8381,f8382,f8383,f8384,f8385,f8386,f8387,f8388,f8389,f8390,f8391,f8392,f8393,f8394,f8395,f8396,f8397,f8398,f8399,f8400,f8401,f8402,f8403,f8404,f8405,f8406,f8407,f8408,f8409,f8410,f8411,f8412,f8413,f8414,f8415,f8416,f8417,f8418,f8419,f8420,f8421,f8422,f8423,f8424,f8428,f8433,f9022,f9027,f9144,f24391,f24407,f24392,f24393,f24408,f24394,f24411,f24409,f26858,f3494,f26859,f26860,f26880,f26861,f26881,f3496,f60623,f60608,f60660,f60565,f60560,f60551,f60631,f60550,f60699,f60700,f60735,f60661,f60496,f60697,f60777,f60491,f60778,f60488,f60779,f8430,f61068,f61069,f30303,f61078,f61077,f61141,f61145,f105622,f105623,f8429,f108597,f105632,f105640,f105642,f105645,f105647,f105653,f105655,f105657,f105663,f105665,f105667,f105669,f105673,f105675,f105677,f105679,f105681,f105683,f105685,f105687,f105689,f105691,f105693,f105695,f105700,f105702,f105707,f105709,f105711,f105713,f105716,f105719,f105721,f105723,f105728,f105730,f105732,f105734,f105736,f105738,f105743,f105745,f105747,f105749,f105751,f105753,f105755,f105757,f105759,f105761,f105763,f105765,f105767,f105769,f105771,f105773,f105802,f105952,f105954,f105956,f105958,f105960,f105962,f105964,f105966,f105968,f105970,f105972,f105974,f105976,f105978,f105980,f105982,f105984,f105986,f105988,f105990,f105992,f105994,f105996,f105998,f106000,f106002,f106004,f106006,f106008,f106010,f106012,f106014,f106016,f106018,f106020,f106022,f106024,f106026,f106028,f106030,f106032,f106034,f106036,f106038,f106040,f106042,f106044,f106046,f106048,f106050,f106052,f106054,f106056,f106058,f106060,f106062,f106064,f106066,f106068,f106070,f106072,f106074,f106076,f106078,f106080,f106082,f106084,f106086,f106088,f106090,f106092,f106094,f106096,f106098,f106100,f106102,f106104,f106106,f106108,f106110,f106112,f106114,f106116,f106118,f106120,f106122,f106124,f106126,f106128,f106130,f106132,f106134,f106136,f106138,f106140,f106142,f106144,f106146,f106148,f106150,f106152,f106154,f106156,f106158,f106160,f106162,f106164,f106166,f106168,f106170,f106172,f106174,f106176,f106178,f106180,f106182,f106184,f106186,f106188,f106190,f106192,f106194,f106196,f106198,f106200,f106202,f106204,f106206,f106208,f106210,f106212,f106214,f106216,f106218,f106220,f106222,f106224,f106226,f106228,f106230,f106232,f106234,f106236,f106238,f106240,f106242,f106245,f106247,f106249,f106251,f106253,f106255,f106257,f106259,f106261,f106263,f106265,f106267,f106269,f106271,f106273,f106275,f106277,f106279,f106281,f106283,f106285,f106287,f106289,f106291,f106293,f106295,f106297,f106299,f106301,f106303,f106305,f106307,f106309,f106311,f106313,f106315,f106317,f106319,f106321,f106323,f106325,f106327,f106329,f106331,f106333,f106335,f106337,f106339,f106341,f106343,f106345,f106347,f106349,f106351,f106353,f106355,f106357,f106359,f106361,f106363,f106365,f106367,f106369,f106371,f106373,f106375,f106377,f106379,f106381,f106383,f106385,f106387,f106389,f106391,f106393,f106395,f106397,f106399,f106401,f106403,f106405,f106407,f106409,f106411,f106413,f106415,f106417,f106419,f106421,f106423,f106425,f106427,f106429,f106431,f106433,f106435,f106437,f106439,f106441,f106443,f106445,f106447,f106449,f106451,f106453,f106455,f106457,f106459,f106461,f106463,f106465,f106467,f106469,f106471,f106473,f106475,f106477,f106479,f106481,f106483,f106485,f106487,f106489,f106491,f106493,f106495,f106497,f106499,f106501,f106503,f106505,f106507,f106509,f106511,f106513,f106515,f106517,f106519,f106521,f106523,f106525,f106527,f106529,f106531,f106533,f106535,f106537,f106543,f107001,f107003,f107006,f107008,f107014,f107016,f107018,f107024,f107026,f107028,f107030,f107034,f107036,f107038,f107040,f107042,f107044,f107046,f107048,f107050,f107052,f107054,f107056,f107061,f107063,f107068,f107070,f107072,f107074,f107077,f107080,f107082,f107084,f107089,f107091,f107093,f107095,f107097,f107099,f107104,f107106,f107108,f107110,f107112,f107114,f107116,f107118,f107120,f107122,f107124,f107126,f107128,f107130,f107132,f107134,f107163,f107313,f107315,f107317,f107319,f107321,f107323,f107325,f107327,f107329,f107331,f107333,f107335,f107337,f107339,f107341,f107343,f107345,f107347,f107349,f107351,f107353,f107355,f107357,f107359,f107361,f107363,f107365,f107367,f107369,f107371,f107373,f107375,f107377,f107379,f107381,f107383,f107385,f107387,f107389,f107391,f107393,f107395,f107397,f107399,f107401,f107403,f107405,f107407,f107409,f107411,f107413,f107415,f107417,f107419,f107421,f107423,f107425,f107427,f107429,f107431,f107433,f107435,f107437,f107439,f107441,f107443,f107445,f107447,f107449,f107451,f107453,f107455,f107457,f107459,f107461,f107463,f107465,f107467,f107469,f107471,f107473,f107475,f107477,f107479,f107481,f107483,f107485,f107487,f107489,f107491,f107493,f107495,f107497,f107499,f107501,f107503,f107505,f107507,f107509,f107511,f107513,f107515,f107517,f107519,f107521,f107523,f107525,f107527,f107529,f107531,f107533,f107535,f107537,f107539,f107541,f107543,f107545,f107547,f107549,f107551,f107553,f107555,f107557,f107559,f107561,f107563,f107565,f107567,f107569,f107571,f107573,f107575,f107577,f107579,f107581,f107583,f107585,f107587,f107589,f107591,f107593,f107595,f107597,f107599,f107601,f107603,f107606,f107608,f107610,f107612,f107614,f107616,f107618,f107620,f107622,f107624,f107626,f107628,f107630,f107632,f107634,f107636,f107638,f107640,f107642,f107644,f107646,f107648,f107650,f107652,f107654,f107656,f107658,f107660,f107662,f107664,f107666,f107668,f107670,f107672,f107674,f107676,f107678,f107680,f107682,f107684,f107686,f107688,f107690,f107692,f107694,f107696,f107698,f107700,f107702,f107704,f107706,f107708,f107710,f107712,f107714,f107716,f107718,f107720,f107722,f107724,f107726,f107728,f107730,f107732,f107734,f107736,f107738,f107740,f107742,f107744,f107746,f107748,f107750,f107752,f107754,f107756,f107758,f107760,f107762,f107764,f107766,f107768,f107770,f107772,f107774,f107776,f107778,f107780,f107782,f107784,f107786,f107788,f107790,f107792,f107794,f107796,f107798,f107800,f107802,f107804,f107806,f107808,f107810,f107812,f107814,f107816,f107818,f107820,f107822,f107824,f107826,f107828,f107830,f107832,f107834,f107836,f107838,f107840,f107842,f107844,f107846,f107848,f107850,f107852,f107854,f107856,f107858,f107860,f107862,f107864,f107866,f107868,f107870,f107872,f107874,f107876,f107878,f107880,f107882,f107884,f107886,f107888,f107890,f107892,f107894,f107896,f107898,f107904,f108610,f108611,f108612,f108613,f108614,f108615,f108616,f108617,f108618,f108619,f108620,f108621,f108622,f108623,f108624,f108625,f108626,f108627,f108628,f108629,f108630,f108631,f108632,f108633,f108634,f108635,f108636,f108637,f108638,f108639,f108640,f108641,f108642,f108643,f108644,f108645,f108646,f108647,f108648,f108649,f108650,f108651,f108652,f108653,f108654,f108655,f108656,f108657,f108658,f108659,f108660,f108661,f108662,f108663,f108664,f108665,f108666,f108667,f108668,f108669,f108670,f108671,f108672,f108673,f108674,f108675,f108676,f108677,f108678,f108679,f108680,f108681,f108682,f108683,f108684,f108685,f108686,f108687,f108688,f108689,f108690,f108691,f108692,f108693,f108694,f108695,f108696,f108697,f108698,f108699,f108700,f108701,f108702,f108703,f108704,f108705,f108706,f108707,f108708,f108709,f108710,f108711,f108712,f108713,f108714,f108715,f108716,f108717,f108718,f108719,f108720,f108721,f108722,f108723,f108724,f108725,f108726,f108727,f108728,f108729,f108730,f108731,f108732,f108733,f108734,f108735,f108736,f108737,f108738,f108739,f108740,f108741,f108742,f108743,f108744,f108745,f108746,f108747,f108748,f108749,f110821,f110850])).
fof(f110850,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f241,f110849])).
fof(f110849,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | $spl3),
  inference(forward_demodulation,[],[f244,f110819])).
fof(f110819,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(forward_demodulation,[],[f110146,f245])).
fof(f245,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(forward_demodulation,[],[f172,f169])).
fof(f169,plain,(
  ( ! [X0,X1] : (s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X35] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X35))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X35)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_1_const.infer.apply_subst_list')).
fof(f172,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X35,X36] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X35),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.apply_subst_list_0'),s('type.list.list'('type.infer_t.infer_t'),X35))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.infer.apply_subst_list')).
fof(f110146,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(superposition,[],[f250,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.st_ex_bind_2'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X1)))))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.st_ex_bind_2'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X1)))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.st_ex_bind_2'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X8] : s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_list_1'(s('type.list.list'('type.infer_t.infer_t'),X8))) = s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.st_ex_bind_2'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X8)))))))),
  file('thm.inferProps.success_eqns|split|12','thm.infer.apply_subst_list_def')).
fof(f250,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.st_ex_bind_2'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))),s(X3,X6)))) )),
  inference(forward_demodulation,[],[f221,f220])).
fof(f220,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.st_ex_bind_2'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))) = s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.st_ex_bind_2'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))) = s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5] : s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.st_ex_bind_2'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))) = s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11,X14,X12,X6] : ! [X35,X36] : s(fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)),'const.infer.st_ex_bind_2'(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X35),s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X36))) = s(fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)),happ(s(fun(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),happ(s(fun(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),fun(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)))),'const.infer.st_ex_bind_0'),s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X35))),s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.infer.st_ex_bind')).
fof(f221,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))),s(X3,X6)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),happ(s(fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),happ(s(fun(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),fun(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)))),'const.infer.st_ex_bind_0'),s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X2),X3)),X4))),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3))),X5))),s(X3,X6)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11,X14,X12,X6] : ! [X35,X36,X45] : s('type.pair.prod'('type.infer.exc'(X14,X12),X6),'const.infer.st_ex_bind_3'(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X35),s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X36),s(X6,X45))) = s('type.pair.prod'('type.infer.exc'(X14,X12),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)),happ(s(fun(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),happ(s(fun(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),fun(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)))),'const.infer.st_ex_bind_0'),s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X35))),s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X36))),s(X6,X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_3_const.infer.st_ex_bind')).
fof(f110821,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f110820,f242])).
fof(f110820,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | $spl2),
  inference(backward_demodulation,[],[f110819,f243])).
fof(f108749,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105644,f8429])).
fof(f105644,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f167])).
fof(f167,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.inferProps.success_eqns|split|12',aHL_BOOL_CASES)).
fof(f108748,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105646,f8429])).
fof(f105646,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f167])).
fof(f108747,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105649,f8429])).
fof(f105649,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f167,f167])).
fof(f108746,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105650,f275])).
fof(f105650,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108745,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105651,f8429])).
fof(f105651,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108744,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105652,f8429])).
fof(f105652,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108743,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105654,f8429])).
fof(f105654,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108742,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105656,f8429])).
fof(f105656,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108741,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105659,f8429])).
fof(f105659,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108740,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105660,f275])).
fof(f105660,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108739,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105661,f8429])).
fof(f105661,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108738,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105662,f8429])).
fof(f105662,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108737,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105664,f8429])).
fof(f105664,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108736,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105666,f8429])).
fof(f105666,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108735,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105670,f8429])).
fof(f105670,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f354])).
fof(f354,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f163,f267])).
fof(f163,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.inferProps.success_eqns|split|12',aHL_TRUTH)).
fof(f164,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.inferProps.success_eqns|split|12',aHL_FALSITY)).
fof(f108734,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105672,f107895])).
fof(f107895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f253,f8429,f248])).
fof(f248,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(subsumption_resolution,[],[f199,f165])).
fof(f165,plain,(
  ( ! [X0] : (s('type.one.one','const.one.one_0') = s('type.one.one',X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.one.one','const.one.one_0') = s('type.one.one',X0)),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7] : s('type.one.one','const.one.one_0') = s('type.one.one',X7)),
  file('thm.inferProps.success_eqns|split|12','thm.one.one')).
fof(f199,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | s(X1,X4) != s(X1,X6) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & (~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) | (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6))) & ((~p(s(bool,X2)) | s('type.one.one','const.one.one_0') != s('type.one.one',X5) | s(X1,X4) != s(X1,X6)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6))) <=> (p(s(bool,X2)) & s('type.one.one','const.one.one_0') = s('type.one.one',X5) & s(X1,X4) = s(X1,X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X11] : ! [X32,X28,X10,X7,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X11),'const.infer.guard_3'(s(bool,X32),s(X6,X28),s(X11,X10))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X6),X11),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X6),'const.infer.Success_1'(s('type.one.one',X7))),s(X11,X9))) <=> (p(s(bool,X32)) & s('type.one.one','const.one.one_0') = s('type.one.one',X7) & s(X11,X9) = s(X11,X10)))),
  file('thm.inferProps.success_eqns|split|12','thm.inferProps.inferProps_38_prove_ltb')).
fof(f253,plain,(
  ( ! [X0,X1] : (s('type.one.one',X0) = s('type.one.one',X1)) )),
  inference(superposition,[],[f165,f165])).
fof(f105672,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108733,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105674,f164])).
fof(f105674,plain,(
  p(s(bool,'F_0')) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f354])).
fof(f108732,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105676,f107895])).
fof(f105676,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108731,plain,(
  $false | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f105678,f8429])).
fof(f105678,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2)))) ) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f3461])).
fof(f108730,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105680,f8429])).
fof(f105680,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f62861])).
fof(f62861,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,X28))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl16 | $spl20)),
  inference(superposition,[],[f61142,f3461])).
fof(f61142,plain,(
  ( ! [X14,X12,X10,X11] : (s(bool,'F_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,happ(s(fun('type.one.one',X10),X12),s('type.one.one',X14)))))) ) | ($spl8 | $spl20)),
  inference(backward_demodulation,[],[f3463,f3566])).
fof(f3566,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14))) = s(bool,happ(s(fun(X10,bool),X11),s(X10,happ(s(fun('type.one.one',X10),X12),s('type.one.one',X14)))))) ) | $spl8),
  inference(superposition,[],[f207,f897])).
fof(f207,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X37,X38,X39] : ! [X41,X42,X44] : s(X37,happ(s(fun(X38,X37),X41),s(X38,happ(s(fun(X39,X38),X42),s(X39,X44))))) = s(X37,happ(s(fun(X39,X37),'_dst_x0x1_2'(s(fun(X38,X37),X41),s(fun(X39,X38),X42))),s(X39,X44)))),
  file('thm.inferProps.success_eqns|split|12','_dst_x0x1_ax')).
fof(f108729,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105682,f8429])).
fof(f105682,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f63836])).
fof(f63836,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62862,f167])).
fof(f62862,plain,(
  ( ! [X32] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'F_0')))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f3463])).
fof(f108728,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105684,f8429])).
fof(f105684,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f64447])).
fof(f64447,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f63842,f275])).
fof(f63842,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62862,f267])).
fof(f108727,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105686,f8429])).
fof(f105686,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun(X1,'type.one.one'),bool),X2),s(fun(X1,'type.one.one'),X3)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f65627])).
fof(f65627,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) = s(bool,happ(s(fun(fun(X10,'type.one.one'),bool),X11),s(fun(X10,'type.one.one'),X12))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f64992,f275])).
fof(f64992,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,happ(s(fun(fun(X10,'type.one.one'),bool),X11),s(fun(X10,'type.one.one'),X12))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62870,f267])).
fof(f62870,plain,(
  ( ! [X61,X64,X65] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X61,'type.one.one'),bool),X65),s(fun(X61,'type.one.one'),X64)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,'type.one.one'),X1) = s(fun(X0,'type.one.one'),X2)) )),
  inference(unit_resulting_resolution,[],[f253,f183])).
fof(f183,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.inferProps.success_eqns|split|12',aHL_EXT)).
fof(f108726,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105688,f8429])).
fof(f105688,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun('type.one.one',bool),bool),X1),s(fun('type.one.one',bool),X2)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f66649])).
fof(f66649,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) = s(bool,happ(s(fun(fun('type.one.one',bool),bool),X7),s(fun('type.one.one',bool),X8))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f65992,f275])).
fof(f65992,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X9) = s(bool,happ(s(fun(fun('type.one.one',bool),bool),X7),s(fun('type.one.one',bool),X8))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62888,f267])).
fof(f62888,plain,(
  ( ! [X185,X186] : (s(bool,'F_0') = s(bool,happ(s(fun(fun('type.one.one',bool),bool),X186),s(fun('type.one.one',bool),X185)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f897])).
fof(f108725,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105690,f8429])).
fof(f105690,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X1),s(bool,X0),s(fun('type.one.one',bool),X2)))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f67994])).
fof(f67994,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X33),s(bool,X32),s(fun('type.one.one',bool),X34))) | s(bool,'T_0') = s(bool,X32)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f65997,f167])).
fof(f65997,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X23),s(bool,X24),s(fun('type.one.one',bool),X25)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62888,f185])).
fof(f185,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X2),s(X1,X3),s(fun(X0,X1),X4))) = s(X1,happ(s(fun(fun(X0,X1),X1),happ(s(fun(X1,fun(fun(X0,X1),X1)),happ(s(fun('type.option.option'(X0),fun(X1,fun(fun(X0,X1),X1))),'const.option.option_CASE_0'),s('type.option.option'(X0),X2))),s(X1,X3))),s(fun(X0,X1),X4)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6,X11] : ! [X35,X36,X45] : s(X11,'const.option.option_CASE_3'(s('type.option.option'(X6),X35),s(X11,X36),s(fun(X6,X11),X45))) = s(X11,happ(s(fun(fun(X6,X11),X11),happ(s(fun(X11,fun(fun(X6,X11),X11)),happ(s(fun('type.option.option'(X6),fun(X11,fun(fun(X6,X11),X11))),'const.option.option_CASE_0'),s('type.option.option'(X6),X35))),s(X11,X36))),s(fun(X6,X11),X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_3_const.option.option_CASE')).
fof(f108724,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105692,f8429])).
fof(f105692,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X1),s(bool,X2),s(fun('type.one.one',bool),X3)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f68711])).
fof(f68711,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,X9) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X6),s(bool,X7),s(fun('type.one.one',bool),X8))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f68017,f275])).
fof(f68017,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X9) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X6),s(bool,X7),s(fun('type.one.one',bool),X8))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f65997,f267])).
fof(f108723,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105694,f8429])).
fof(f105694,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun('type.one.one','type.infer.infer_st'(bool)),X1),s('type.one.one',X2)))))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f69819])).
fof(f69819,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun('type.one.one','type.infer.infer_st'(bool)),X4),s('type.one.one',X5))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f69104,f275])).
fof(f69104,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X6) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun('type.one.one','type.infer.infer_st'(bool)),X4),s('type.one.one',X5))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62913,f267])).
fof(f62913,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun('type.one.one','type.infer.infer_st'(bool)),X41),s('type.one.one',X42)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f168])).
fof(f168,plain,(
  ( ! [X0,X1] : (s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,happ(s(fun('type.infer.infer_st'(X0),X0),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X0),X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X35] : s(X6,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X6),X35))) = s(X6,happ(s(fun('type.infer.infer_st'(X6),X6),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X6),X35)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_1_const.infer.infer_st_subst')).
fof(f108722,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105696,f8429])).
fof(f105696,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f69996,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,happ(s(fun(bool,bool),X8),s(bool,X9))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) | ($spl7 | $spl8 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f69903,f275])).
fof(f69903,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X8),s(bool,X9))) | s(bool,'T_0') = s(bool,X10)) ) | ($spl8 | $spl16 | $spl20)),
  inference(superposition,[],[f62861,f267])).
fof(f108721,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105697,f8429])).
fof(f105697,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108720,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105698,f8429])).
fof(f105698,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108719,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105699,f8429])).
fof(f105699,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108718,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105701,f8429])).
fof(f105701,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108717,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105703,f8429])).
fof(f105703,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108716,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105704,f8429])).
fof(f105704,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108715,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105705,f8429])).
fof(f105705,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108714,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105706,f8429])).
fof(f105706,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108713,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105708,f8429])).
fof(f105708,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108712,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105710,f8429])).
fof(f105710,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X1),s(bool,X0)))))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f74666])).
fof(f74666,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X1),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f73313,f167])).
fof(f73313,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X52),s(bool,'F_0')))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f63847,f168])).
fof(f63847,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,happ(s(fun(bool,X13),X15),s(bool,'F_0')))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62862,f207])).
fof(f108711,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105712,f8429])).
fof(f105712,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X1),s(bool,'F_0')))))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f75410])).
fof(f75410,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X2),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl7 | $spl8 | $spl20)),
  inference(subsumption_resolution,[],[f74674,f275])).
fof(f74674,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X2),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f73313,f267])).
fof(f108710,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105715,f8429])).
fof(f105715,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X0),s(bool,X1)))))) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f91543])).
fof(f91543,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X31),s(bool,X30))))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl8 | $spl16 | $spl20)),
  inference(superposition,[],[f69102,f3461])).
fof(f69102,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(X4,'type.infer.infer_st'(bool)),X5),s(X4,happ(s(fun('type.one.one',X4),X6),s('type.one.one',X7)))))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62913,f207])).
fof(f108709,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105717,f8429])).
fof(f105717,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f354])).
fof(f108708,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105718,f107895])).
fof(f105718,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f354])).
fof(f108707,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105720,f107895])).
fof(f105720,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108706,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105722,f107895])).
fof(f105722,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108705,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105724,f8429])).
fof(f105724,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108704,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105725,f8429])).
fof(f105725,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108703,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105726,f8429])).
fof(f105726,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108702,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f105727,f8429])).
fof(f105727,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108701,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105729,f8429])).
fof(f105729,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108700,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105731,f8429])).
fof(f105731,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108699,plain,(
  $false | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f105733,f8429])).
fof(f105733,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f3461])).
fof(f108698,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105735,f8429])).
fof(f105735,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f63836])).
fof(f108697,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105737,f8429])).
fof(f105737,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f64447])).
fof(f108696,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105739,f8429])).
fof(f105739,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108695,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105740,f8429])).
fof(f105740,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108694,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105741,f8429])).
fof(f105741,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108693,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105742,f8429])).
fof(f105742,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108692,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105744,f8429])).
fof(f105744,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108691,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105746,f8429])).
fof(f105746,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108690,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105748,f8429])).
fof(f105748,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f66649])).
fof(f108689,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105750,f8429])).
fof(f105750,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f65627])).
fof(f108688,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105752,f8429])).
fof(f105752,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f62861])).
fof(f108687,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105754,f8429])).
fof(f105754,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f286])).
fof(f286,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X2))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X2)))) )),
  inference(superposition,[],[f168,f167])).
fof(f108686,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105756,f8429])).
fof(f105756,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f69819])).
fof(f108685,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105758,f8429])).
fof(f105758,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f74666])).
fof(f108684,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105760,f8429])).
fof(f105760,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f75410])).
fof(f108683,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105762,f8429])).
fof(f105762,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f91543])).
fof(f108682,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f105764,f8429])).
fof(f105764,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f286])).
fof(f108681,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105766,f8429])).
fof(f105766,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f67994])).
fof(f108680,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f105768,f8429])).
fof(f105768,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f68711])).
fof(f108679,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107005,f8429])).
fof(f107005,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f167])).
fof(f108678,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107007,f8429])).
fof(f107007,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f167])).
fof(f108677,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107010,f8429])).
fof(f107010,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108676,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107011,f275])).
fof(f107011,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108675,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107012,f8429])).
fof(f107012,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108674,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107013,f8429])).
fof(f107013,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108673,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107015,f8429])).
fof(f107015,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108672,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107017,f8429])).
fof(f107017,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108671,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107020,f8429])).
fof(f107020,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108670,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107021,f275])).
fof(f107021,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108669,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107022,f8429])).
fof(f107022,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108668,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107023,f8429])).
fof(f107023,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108667,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107025,f8429])).
fof(f107025,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108666,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107027,f8429])).
fof(f107027,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108665,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107031,f8429])).
fof(f107031,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f354])).
fof(f108664,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107033,f107895])).
fof(f107033,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108663,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107035,f164])).
fof(f107035,plain,(
  p(s(bool,'F_0')) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f354])).
fof(f108662,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107037,f107895])).
fof(f107037,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108661,plain,(
  $false | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f107039,f8429])).
fof(f107039,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2)))) ) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f3461])).
fof(f108660,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107041,f8429])).
fof(f107041,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f62861])).
fof(f108659,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107043,f8429])).
fof(f107043,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f63836])).
fof(f108658,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107045,f8429])).
fof(f107045,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f64447])).
fof(f108657,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107047,f8429])).
fof(f107047,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun(X1,'type.one.one'),bool),X2),s(fun(X1,'type.one.one'),X3)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f65627])).
fof(f108656,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107049,f8429])).
fof(f107049,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(fun('type.one.one',bool),bool),X1),s(fun('type.one.one',bool),X2)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f66649])).
fof(f108655,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107051,f8429])).
fof(f107051,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X1),s(bool,X0),s(fun('type.one.one',bool),X2)))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f67994])).
fof(f108654,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107053,f8429])).
fof(f107053,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.option.option_CASE_3'(s('type.option.option'('type.one.one'),X1),s(bool,X2),s(fun('type.one.one',bool),X3)))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f68711])).
fof(f108653,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107055,f8429])).
fof(f107055,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun('type.one.one','type.infer.infer_st'(bool)),X1),s('type.one.one',X2)))))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f69819])).
fof(f108652,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107057,f8429])).
fof(f107057,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108651,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107058,f8429])).
fof(f107058,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108650,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107059,f8429])).
fof(f107059,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108649,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107060,f8429])).
fof(f107060,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108648,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107062,f8429])).
fof(f107062,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108647,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107064,f8429])).
fof(f107064,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108646,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107065,f8429])).
fof(f107065,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108645,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107066,f8429])).
fof(f107066,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108644,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107067,f8429])).
fof(f107067,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108643,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107069,f8429])).
fof(f107069,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108642,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107071,f8429])).
fof(f107071,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X1),s(bool,X0)))))) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f74666])).
fof(f108641,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107073,f8429])).
fof(f107073,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X1),s(bool,'F_0')))))) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f75410])).
fof(f108640,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107076,f8429])).
fof(f107076,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(bool,'type.infer.infer_st'(bool)),X0),s(bool,X1)))))) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f91543])).
fof(f108639,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107078,f8429])).
fof(f107078,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f354])).
fof(f108638,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107079,f107895])).
fof(f107079,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f354])).
fof(f108637,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107081,f107895])).
fof(f107081,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108636,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107083,f107895])).
fof(f107083,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f354])).
fof(f108635,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107085,f8429])).
fof(f107085,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108634,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107086,f8429])).
fof(f107086,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108633,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107087,f8429])).
fof(f107087,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108632,plain,(
  $false | ($spl7 | $spl36)),
  inference(subsumption_resolution,[],[f107088,f8429])).
fof(f107088,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f267])).
fof(f108631,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107090,f8429])).
fof(f107090,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108630,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107092,f8429])).
fof(f107092,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f267])).
fof(f108629,plain,(
  $false | ($spl16 | $spl36)),
  inference(subsumption_resolution,[],[f107094,f8429])).
fof(f107094,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f3461])).
fof(f108628,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107096,f8429])).
fof(f107096,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f63836])).
fof(f108627,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107098,f8429])).
fof(f107098,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f64447])).
fof(f108626,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107100,f8429])).
fof(f107100,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108625,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107101,f8429])).
fof(f107101,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108624,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107102,f8429])).
fof(f107102,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108623,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107103,f8429])).
fof(f107103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69996])).
fof(f108622,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107105,f8429])).
fof(f107105,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108621,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107107,f8429])).
fof(f107107,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69996])).
fof(f108620,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107109,f8429])).
fof(f107109,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f66649])).
fof(f108619,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107111,f8429])).
fof(f107111,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f65627])).
fof(f108618,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107113,f8429])).
fof(f107113,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f62861])).
fof(f108617,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107115,f8429])).
fof(f107115,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f286])).
fof(f108616,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107117,f8429])).
fof(f107117,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f69819])).
fof(f108615,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107119,f8429])).
fof(f107119,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f74666])).
fof(f108614,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107121,f8429])).
fof(f107121,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f75410])).
fof(f108613,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107123,f8429])).
fof(f107123,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f91543])).
fof(f108612,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f107125,f8429])).
fof(f107125,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X0)))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f286])).
fof(f108611,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107127,f8429])).
fof(f107127,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f67994])).
fof(f108610,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(subsumption_resolution,[],[f107129,f8429])).
fof(f107129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f68711])).
fof(f107904,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f243,f8429,f224])).
fof(f224,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,sK6(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK7(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,sK6(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK7(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,sK6(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK7(X8,X7,X6,X5,X4,X3,X2,X1,X0)))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X11,X12] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X11))),s(X3,X12))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X11))),s(X3,X12))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(rectify,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10))))) & (! [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))) | s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6)))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7,X8] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) <=> ? [X9,X10] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,X9))),s(X3,X10))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) & s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) = s('type.pair.prod'('type.infer.exc'(X2,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3)),happ(s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X0,X9))),s(X3,X10)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11,X12,X14,X6] : ! [X15,X16,X10,X9,X7] : (s('type.pair.prod'('type.infer.exc'(X14,X12),X6),'const.pair.,_2'(s('type.infer.exc'(X14,X12),'const.infer.Success_1'(s(X14,X7))),s(X6,X9))) = s('type.pair.prod'('type.infer.exc'(X14,X12),X6),'const.infer.st_ex_bind_3'(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X15),s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X16),s(X6,X10))) <=> ? [X13,X17] : (s('type.pair.prod'('type.infer.exc'(X11,X12),X6),'const.pair.,_2'(s('type.infer.exc'(X11,X12),'const.infer.Success_1'(s(X11,X13))),s(X6,X17))) = s('type.pair.prod'('type.infer.exc'(X11,X12),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),X15),s(X6,X10))) & s('type.pair.prod'('type.infer.exc'(X14,X12),X6),'const.pair.,_2'(s('type.infer.exc'(X14,X12),'const.infer.Success_1'(s(X14,X7))),s(X6,X9))) = s('type.pair.prod'('type.infer.exc'(X14,X12),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6)),happ(s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X14,X12),X6))),X16),s(X11,X13))),s(X6,X17)))))),
  file('thm.inferProps.success_eqns|split|12','thm.inferProps.inferProps_25_prove_ltb')).
fof(f107898,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f243,f8429,f223])).
fof(f223,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X2,X1),'const.infer.Success_1'(s(X2,X8))),s(X3,X7))) != s('type.pair.prod'('type.infer.exc'(X2,X1),X3),'const.infer.st_ex_bind_3'(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(fun(X0,fun(X3,'type.pair.prod'('type.infer.exc'(X2,X1),X3))),X5),s(X3,X6))) | s('type.pair.prod'('type.infer.exc'(X0,X1),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X0,X1),X3)),X4),s(X3,X6))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X3),'const.pair.,_2'(s('type.infer.exc'(X0,X1),'const.infer.Success_1'(s(X0,sK6(X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X3,sK7(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f107896,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f253,f8429,f248])).
fof(f107894,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f165,f8429,f248])).
fof(f107892,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f184,f8429,f248])).
fof(f184,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'('type.one.one',X0),X1)),'const.infer.guard_2'(s(bool,X2),s(X0,X3))),s(X1,X4)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X11,X6] : ! [X35,X36,X45] : s('type.pair.prod'('type.infer.exc'('type.one.one',X11),X6),'const.infer.guard_3'(s(bool,X35),s(X11,X36),s(X6,X45))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X11),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'('type.one.one',X11),X6)),'const.infer.guard_2'(s(bool,X35),s(X11,X36))),s(X6,X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_2_3_const.infer.guard')).
fof(f107890,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f186,f8429,f248])).
fof(f186,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),'const.infer.lookup_tenvC_st_ex_3'(s('type.ast.id'('type.list.list'('type.string.char')),X2),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1))),happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)))),'const.infer.lookup_tenvC_st_ex_0'),s('type.ast.id'('type.list.list'('type.string.char')),X2))),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),'const.infer.lookup_tenvC_st_ex_3'(s('type.ast.id'('type.list.list'('type.string.char')),X2),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1))),happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)))),'const.infer.lookup_tenvC_st_ex_0'),s('type.ast.id'('type.list.list'('type.string.char')),X2))),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3))),s(X1,X4)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),'const.infer.lookup_tenvC_st_ex_3'(s('type.ast.id'('type.list.list'('type.string.char')),X2),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3),s(X1,X4))) = s('type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1))),happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),fun(X1,'type.pair.prod'('type.infer.exc'(X0,'type.list.list'('type.string.char')),X1)))),'const.infer.lookup_tenvC_st_ex_0'),s('type.ast.id'('type.list.list'('type.string.char')),X2))),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X0))),X3))),s(X1,X4)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X11,X6] : ! [X35,X36,X45] : s('type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6),'const.infer.lookup_tenvC_st_ex_3'(s('type.ast.id'('type.list.list'('type.string.char')),X35),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11))),X36),s(X6,X45))) = s('type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6)),happ(s(fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11))),fun(X6,'type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6))),happ(s(fun('type.ast.id'('type.list.list'('type.string.char')),fun('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11))),fun(X6,'type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6)))),'const.infer.lookup_tenvC_st_ex_0'),s('type.ast.id'('type.list.list'('type.string.char')),X35))),s('type.pair.prod'('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11)))),'type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),X11))),X36))),s(X6,X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_3_const.infer.lookup_tenvC_st_ex')).
fof(f107888,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f208,f8429,f248])).
fof(f208,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),'const.infer.lookup_st_ex_4'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5),s(X2,X6))) = s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2)),'const.infer.lookup_st_ex_3'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5))),s(X2,X6)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),'const.infer.lookup_st_ex_4'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5),s(X2,X6))) = s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2)),'const.infer.lookup_st_ex_3'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5))),s(X2,X6)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),'const.infer.lookup_st_ex_4'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5),s(X2,X6))) = s('type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,'type.list.list'('type.string.char')),X2)),'const.infer.lookup_st_ex_3'(s(fun(X0,'type.list.list'('type.string.char')),X3),s(X0,X4),s('type.list.list'('type.pair.prod'(X0,X1)),X5))),s(X2,X6)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X12,X11,X6] : ! [X35,X36,X45,X46] : s('type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6),'const.infer.lookup_st_ex_4'(s(fun(X12,'type.list.list'('type.string.char')),X35),s(X12,X36),s('type.list.list'('type.pair.prod'(X12,X11)),X45),s(X6,X46))) = s('type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,'type.list.list'('type.string.char')),X6)),'const.infer.lookup_st_ex_3'(s(fun(X12,'type.list.list'('type.string.char')),X35),s(X12,X36),s('type.list.list'('type.pair.prod'(X12,X11)),X45))),s(X6,X46)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_3_4_const.infer.lookup_st_ex')).
fof(f107886,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f205,f8429,f248])).
fof(f205,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s('type.pair.prod'('type.infer.exc'(X0,X1),X2),'const.infer.failwith_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X0,X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X0,X1),X2)),'const.infer.failwith_1'(s(X1,X3))),s(X2,X4)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X12,X6,X11] : ! [X35,X36] : s('type.pair.prod'('type.infer.exc'(X12,X6),X11),'const.infer.failwith_2'(s(X6,X35),s(X11,X36))) = s('type.pair.prod'('type.infer.exc'(X12,X6),X11),happ(s(fun(X11,'type.pair.prod'('type.infer.exc'(X12,X6),X11)),'const.infer.failwith_1'(s(X6,X35))),s(X11,X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_1_2_const.infer.failwith')).
fof(f107884,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f204,f8429,f248])).
fof(f204,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraints_0'),s('type.list.list'('type.infer_t.infer_t'),X0))),s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraints_0'),s('type.list.list'('type.infer_t.infer_t'),X0))),s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X35,X36,X45] : s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X35),s('type.list.list'('type.infer_t.infer_t'),X36),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X45))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraints_0'),s('type.list.list'('type.infer_t.infer_t'),X35))),s('type.list.list'('type.infer_t.infer_t'),X36))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_3_const.infer.add_constraints')).
fof(f107882,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f203,f8429,f248])).
fof(f203,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraint_0'),s('type.infer_t.infer_t',X0))),s('type.infer_t.infer_t',X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraint_0'),s('type.infer_t.infer_t',X0))),s('type.infer_t.infer_t',X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X35,X36,X45] : s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X35),s('type.infer_t.infer_t',X36),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X45))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),happ(s(fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),happ(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))))),'const.infer.add_constraint_0'),s('type.infer_t.infer_t',X35))),s('type.infer_t.infer_t',X36))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X45)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_3_const.infer.add_constraint')).
fof(f107880,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f171,f8429,f248])).
fof(f171,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_1'(s('type.infer_t.infer_t',X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_1'(s('type.infer_t.infer_t',X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_1'(s('type.infer_t.infer_t',X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X35,X36] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X35),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),happ(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.apply_subst_1'(s('type.infer_t.infer_t',X35))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_1_2_const.infer.apply_subst')).
fof(f107878,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f182,f8429,f248])).
fof(f182,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),'const.infer.n_fresh_uvar_2'(s('type.num.num',X2),s('type.infer.infer_st'(X1),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1))),happ(s(fun('type.num.num',fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)))),'const.infer.n_fresh_uvar_0'),s('type.num.num',X2))),s('type.infer.infer_st'(X1),X3)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),'const.infer.n_fresh_uvar_2'(s('type.num.num',X2),s('type.infer.infer_st'(X1),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1))),happ(s(fun('type.num.num',fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)))),'const.infer.n_fresh_uvar_0'),s('type.num.num',X2))),s('type.infer.infer_st'(X1),X3)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),'const.infer.n_fresh_uvar_2'(s('type.num.num',X2),s('type.infer.infer_st'(X1),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1))),happ(s(fun('type.num.num',fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'(X1)))),'const.infer.n_fresh_uvar_0'),s('type.num.num',X2))),s('type.infer.infer_st'(X1),X3)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11,X6] : ! [X35,X36] : s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X11),'type.infer.infer_st'(X6)),'const.infer.n_fresh_uvar_2'(s('type.num.num',X35),s('type.infer.infer_st'(X6),X36))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X11),'type.infer.infer_st'(X6)),happ(s(fun('type.infer.infer_st'(X6),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X11),'type.infer.infer_st'(X6))),happ(s(fun('type.num.num',fun('type.infer.infer_st'(X6),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X11),'type.infer.infer_st'(X6)))),'const.infer.n_fresh_uvar_0'),s('type.num.num',X35))),s('type.infer.infer_st'(X6),X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.infer.n_fresh_uvar')).
fof(f107876,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f180,f8429,f248])).
fof(f180,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1)),happ(s(fun('type.infer.infer_st'(X1),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'(X1))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X1),X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X11] : ! [X35] : s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X11)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X11),X35))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X11)),happ(s(fun('type.infer.infer_st'(X11),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'(X11))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X11),X35)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_1_const.infer.fresh_uvar')).
fof(f107874,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f206,f8429,f248])).
fof(f206,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2)),happ(s(fun(X1,fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2))),'const.infer.st_ex_return_0'),s(X1,X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2)),happ(s(fun(X1,fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2))),'const.infer.st_ex_return_0'),s(X1,X3))),s(X2,X4)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2)),happ(s(fun(X1,fun(X2,'type.pair.prod'('type.infer.exc'(X1,X0),X2))),'const.infer.st_ex_return_0'),s(X1,X3))),s(X2,X4)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X12,X11,X6] : ! [X35,X36] : s('type.pair.prod'('type.infer.exc'(X11,X12),X6),'const.infer.st_ex_return_2'(s(X11,X35),s(X6,X36))) = s('type.pair.prod'('type.infer.exc'(X11,X12),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6)),happ(s(fun(X11,fun(X6,'type.pair.prod'('type.infer.exc'(X11,X12),X6))),'const.infer.st_ex_return_0'),s(X11,X35))),s(X6,X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.infer.st_ex_return')).
fof(f107872,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61610,f8429,f248])).
fof(f61610,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'F_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'F_0'),s(bool,'F_0')))) ) | $spl20),
  inference(forward_demodulation,[],[f61609,f178])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X2))),s(X1,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X11,X6] : ! [X5] : s('type.pair.prod'('type.infer.exc'(X6,X11),X6),'const.infer.read_1'(s(X6,X5))) = s('type.pair.prod'('type.infer.exc'(X6,X11),X6),'const.pair.,_2'(s('type.infer.exc'(X6,X11),'const.infer.Success_1'(s(X6,X5))),s(X6,X5)))),
  file('thm.inferProps.success_eqns|split|12','thm.infer.read_def')).
fof(f61609,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'F_0'),s(bool,'F_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'F_0'))),s(bool,'F_0')))) ) | $spl20),
  inference(forward_demodulation,[],[f61608,f3463])).
fof(f61608,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'F_0'))),s(bool,'F_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,'F_0')))) ) | $spl20),
  inference(forward_demodulation,[],[f61430,f3463])).
fof(f61430,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'F_0'))),s(bool,'F_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,happ(s(fun('type.one.one',bool),X3),s('type.one.one',X4)))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f3463,f3463,f213])).
fof(f213,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & (s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6))) & ((s(X1,X3) != s(X1,X5) | s(X2,X4) != s(X2,X6)) | s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6)))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) <=> (s(X1,X3) = s(X1,X5) & s(X2,X4) = s(X2,X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X12,X6,X11] : ! [X7,X10,X13,X9] : (s('type.pair.prod'('type.infer.exc'(X6,X12),X11),'const.infer.st_ex_return_2'(s(X6,X7),s(X11,X10))) = s('type.pair.prod'('type.infer.exc'(X6,X12),X11),'const.pair.,_2'(s('type.infer.exc'(X6,X12),'const.infer.Success_1'(s(X6,X13))),s(X11,X9))) <=> (s(X6,X7) = s(X6,X13) & s(X11,X9) = s(X11,X10)))),
  file('thm.inferProps.success_eqns|split|12','thm.inferProps.inferProps_24_prove_ltb')).
fof(f107870,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f29104,f8429,f248])).
fof(f29104,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))) ) | $spl32),
  inference(forward_demodulation,[],[f29103,f178])).
fof(f29103,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl32),
  inference(forward_demodulation,[],[f29102,f3496])).
fof(f29102,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,g(s('type.one.one',X1))))),s(bool,'T_0')))) ) | $spl32),
  inference(forward_demodulation,[],[f27940,f3496])).
fof(f27940,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,g(s('type.one.one',X1))))),s(bool,g(s('type.one.one',X2)))))) ) | $spl32),
  inference(unit_resulting_resolution,[],[f3496,f3496,f213])).
fof(f107868,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f179,f8429,f248])).
fof(f179,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'('type.infer.exc'(X1,X0),X1),'const.infer.read_1'(s(X1,X2))) = s('type.pair.prod'('type.infer.exc'(X1,X0),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X0),X1)),'const.infer.read_0'),s(X1,X2)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11,X6] : ! [X35] : s('type.pair.prod'('type.infer.exc'(X6,X11),X6),'const.infer.read_1'(s(X6,X35))) = s('type.pair.prod'('type.infer.exc'(X6,X11),X6),happ(s(fun(X6,'type.pair.prod'('type.infer.exc'(X6,X11),X6)),'const.infer.read_0'),s(X6,X35)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_1_const.infer.read')).
fof(f107866,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f178,f8429,f248])).
fof(f107864,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f250,f8429,f248])).
fof(f107862,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f243,f8429,f248])).
fof(f107860,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f245,f8429,f248])).
fof(f107858,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f176,f8429,f248])).
fof(f176,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X35,X36] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num',X36))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X35))),s('type.num.num',X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_1_2_const.arithmetic.+')).
fof(f107856,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f181,f8429,f248])).
fof(f181,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X2),s('type.list.list'(X1),X3))) = s('type.list.list'(X0),happ(s(fun('type.list.list'(X1),'type.list.list'(X0)),happ(s(fun(fun(X1,X0),fun('type.list.list'(X1),'type.list.list'(X0))),'const.list.MAP_0'),s(fun(X1,X0),X2))),s('type.list.list'(X1),X3)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11,X6] : ! [X35,X36] : s('type.list.list'(X11),'const.list.MAP_2'(s(fun(X6,X11),X35),s('type.list.list'(X6),X36))) = s('type.list.list'(X11),happ(s(fun('type.list.list'(X6),'type.list.list'(X11)),happ(s(fun(fun(X6,X11),fun('type.list.list'(X6),'type.list.list'(X11))),'const.list.MAP_0'),s(fun(X6,X11),X35))),s('type.list.list'(X6),X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.list.MAP')).
fof(f107854,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f177,f8429,f248])).
fof(f177,plain,(
  ( ! [X0,X1] : (s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.infer_t.infer_t',X1))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))),s('type.infer_t.infer_t',X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X35,X36] : s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35),s('type.infer_t.infer_t',X36))) = s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35))),s('type.infer_t.infer_t',X36)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_2_const.unify.t_walkstar')).
fof(f107852,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f166,f8429,f248])).
fof(f166,plain,(
  ( ! [X0] : (s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X0))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X35] : s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X35))) = s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X35)))),
  file('thm.inferProps.success_eqns|split|12','ARITY_0_1_const.infer_t.Infer_Tuvar')).
fof(f107850,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f169,f8429,f248])).
fof(f107848,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15105,f8429,f248])).
fof(f15105,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun('type.one.one',fun(X0,fun(X1,fun(X2,fun('type.one.one',bool))))),X3) = s(fun('type.one.one',fun(X0,fun(X1,fun(X2,fun('type.one.one',bool))))),X4)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f8129,f558])).
fof(f558,plain,(
  ( ! [X21,X19,X22,X20] : (s(X21,f(s('type.one.one',X22))) != s(X21,g(s('type.one.one',X22))) | s(fun('type.one.one',X21),X19) = s(fun('type.one.one',X21),X20)) )),
  inference(superposition,[],[f183,f253])).
fof(f8129,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun('type.one.one',bool)))),X3) = s(fun(X0,fun(X1,fun(X2,fun('type.one.one',bool)))),X4)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f3661,f183])).
fof(f3661,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun('type.one.one',bool))),X2) = s(fun(X0,fun(X1,fun('type.one.one',bool))),X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f3512,f183])).
fof(f3512,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun('type.one.one',bool)),X1) = s(fun(X0,fun('type.one.one',bool)),X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f897,f183])).
fof(f107846,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8130,f8429,f248])).
fof(f8130,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.one.one',fun(X0,fun(X1,fun('type.one.one',bool)))),X2) = s(fun('type.one.one',fun(X0,fun(X1,fun('type.one.one',bool)))),X3)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f3661,f558])).
fof(f107844,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3662,f8429,f248])).
fof(f3662,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.one.one',fun(X0,fun('type.one.one',bool))),X1) = s(fun('type.one.one',fun(X0,fun('type.one.one',bool))),X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f3512,f558])).
fof(f107842,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3513,f8429,f248])).
fof(f3513,plain,(
  ( ! [X0,X1] : (s(fun('type.one.one',fun('type.one.one',bool)),X0) = s(fun('type.one.one',fun('type.one.one',bool)),X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f897,f558])).
fof(f107840,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f897,f8429,f248])).
fof(f107838,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f220,f8429,f248])).
fof(f107836,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f743,f8429,f248])).
fof(f743,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X5) = s(fun(X0,fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),X6)) )),
  inference(unit_resulting_resolution,[],[f715,f183])).
fof(f715,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X4) = s(fun(X0,fun(X1,fun(X2,fun(X3,'type.one.one')))),X5)) )),
  inference(unit_resulting_resolution,[],[f590,f183])).
fof(f590,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X3) = s(fun(X0,fun(X1,fun(X2,'type.one.one'))),X4)) )),
  inference(unit_resulting_resolution,[],[f570,f183])).
fof(f570,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X1,'type.one.one')),X2) = s(fun(X0,fun(X1,'type.one.one')),X3)) )),
  inference(unit_resulting_resolution,[],[f552,f183])).
fof(f107834,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15104,f8429,f248])).
fof(f15104,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X1,fun(X2,fun(X3,fun('type.one.one',bool))))),X4) = s(fun(X0,fun(X1,fun(X2,fun(X3,fun('type.one.one',bool))))),X5)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f8129,f183])).
fof(f107832,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f715,f8429,f248])).
fof(f107830,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8129,f8429,f248])).
fof(f107828,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f590,f8429,f248])).
fof(f107826,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3661,f8429,f248])).
fof(f107824,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f570,f8429,f248])).
fof(f107822,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3512,f8429,f248])).
fof(f107820,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f552,f8429,f248])).
fof(f107818,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f97221,f8429,f248])).
fof(f97221,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.infer.infer_st'(bool)),happ(s(fun(bool,'type.infer.infer_st'('type.infer.infer_st'(bool))),X0),s(bool,'F_0')))))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f74672,f168])).
fof(f74672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(X0,'type.infer.infer_st'(bool)),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0')))))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f73313,f207])).
fof(f107816,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91598,f8429,f248])).
fof(f91598,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.infer.infer_st'(bool)),happ(s(fun('type.one.one','type.infer.infer_st'('type.infer.infer_st'(bool))),X0),s('type.one.one',X1)))))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f69102,f168])).
fof(f107814,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91565,f8429,f248])).
fof(f91565,plain,(
  ( ! [X125,X130,X126,X129] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(fun(X125,fun(X126,'type.one.one')),'type.infer.infer_st'(bool)),X130),s(fun(X125,fun(X126,'type.one.one')),X129)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f69102,f570])).
fof(f107812,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f74672,f8429,f248])).
fof(f107810,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f69102,f8429,f248])).
fof(f107808,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91563,f8429,f248])).
fof(f91563,plain,(
  ( ! [X118,X119,X115] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(fun(X115,fun('type.one.one',bool)),'type.infer.infer_st'(bool)),X119),s(fun(X115,fun('type.one.one',bool)),X118)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f69102,f3512])).
fof(f107806,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f85050,f8429,f248])).
fof(f85050,plain,(
  ( ! [X132,X133,X131] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(fun(X131,'type.one.one'),'type.infer.infer_st'(bool)),X132),s(fun(X131,'type.one.one'),X133)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f64998,f168])).
fof(f64998,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(fun(X36,'type.one.one'),X37),X39),s(fun(X36,'type.one.one'),X40)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62870,f207])).
fof(f107804,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f88012,f8429,f248])).
fof(f88012,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),happ(s(fun(fun('type.one.one',bool),'type.infer.infer_st'(bool)),X108),s(fun('type.one.one',bool),X109)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f65998,f168])).
fof(f65998,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(fun('type.one.one',bool),X30),X32),s(fun('type.one.one',bool),X33)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62888,f207])).
fof(f107802,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73313,f8429,f248])).
fof(f107800,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62913,f8429,f248])).
fof(f107798,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62918,f8429,f248])).
fof(f62918,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,'const.option.option_CASE_3'(s('type.option.option'(X69),X70),s(bool,X71),s(fun(X69,bool),happ(s(fun('type.one.one',fun(X69,bool)),X72),s('type.one.one',X73)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f185])).
fof(f107796,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65997,f8429,f248])).
fof(f107794,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3463,f8429,f248])).
fof(f107792,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62900,f8429,f248])).
fof(f62900,plain,(
  ( ! [X243,X246,X244,X247,X245] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.infer.exc'(X243,X244),'type.one.one'),bool),X247),s('type.pair.prod'('type.infer.exc'(X243,X244),'type.one.one'),'const.infer.st_ex_return_2'(s(X243,X245),s('type.one.one',X246)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f206])).
fof(f107790,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62899,f8429,f248])).
fof(f62899,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.infer.exc'(X238,X239),'type.one.one'),bool),X242),s('type.pair.prod'('type.infer.exc'(X238,X239),'type.one.one'),'const.infer.failwith_2'(s(X239,X240),s('type.one.one',X241)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f205])).
fof(f107788,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73296,f8429,f248])).
fof(f73296,plain,(
  ( ! [X196,X195] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.infer.exc'(bool,X195),bool),bool),X196),s('type.pair.prod'('type.infer.exc'(bool,X195),bool),'const.infer.read_1'(s(bool,'F_0')))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f63847,f179])).
fof(f107786,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62898,f8429,f248])).
fof(f62898,plain,(
  ( ! [X237,X235,X236] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.infer.exc'('type.one.one',X235),'type.one.one'),bool),X237),s('type.pair.prod'('type.infer.exc'('type.one.one',X235),'type.one.one'),'const.infer.read_1'(s('type.one.one',X236)))))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f179])).
fof(f107784,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62878,f8429,f248])).
fof(f62878,plain,(
  ( ! [X111,X107,X105,X110,X106] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X105,fun(X106,fun(X107,'type.one.one'))),bool),X111),s(fun(X105,fun(X106,fun(X107,'type.one.one'))),X110)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f590])).
fof(f107782,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62876,f8429,f248])).
fof(f62876,plain,(
  ( ! [X94,X97,X93,X98] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X93,fun(X94,fun('type.one.one',bool))),bool),X98),s(fun(X93,fun(X94,fun('type.one.one',bool))),X97)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f3661])).
fof(f107780,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62874,f8429,f248])).
fof(f62874,plain,(
  ( ! [X85,X81,X86,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X81,fun(X82,'type.one.one')),bool),X86),s(fun(X81,fun(X82,'type.one.one')),X85)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f570])).
fof(f107778,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62872,f8429,f248])).
fof(f62872,plain,(
  ( ! [X74,X71,X75] : (s(bool,'F_0') = s(bool,happ(s(fun(fun(X71,fun('type.one.one',bool)),bool),X75),s(fun(X71,fun('type.one.one',bool)),X74)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f61142,f3512])).
fof(f107776,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62870,f8429,f248])).
fof(f107774,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62888,f8429,f248])).
fof(f107772,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62862,f8429,f248])).
fof(f107770,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64997,f8429,f248])).
fof(f64997,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(fun(X31,'type.one.one'),bool),X32),s(fun(X33,fun(X31,'type.one.one')),X34))),s(X33,X35)))) ) | ($spl8 | $spl20)),
  inference(superposition,[],[f62870,f207])).
fof(f107768,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64998,f8429,f248])).
fof(f107766,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65998,f8429,f248])).
fof(f107764,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f63847,f8429,f248])).
fof(f107762,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61142,f8429,f248])).
fof(f107760,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3496,f8429,f248])).
fof(f107758,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f185,f8429,f248])).
fof(f107756,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f222,f8429,f248])).
fof(f222,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X37,X38,X39,X40] : ! [X41,X42,X43,X44] : s(X37,happ(s(fun(X39,X37),'_dst_x0x1x0_3'(s(fun(X38,fun(X40,X37)),X41),s(fun(X39,X38),X42),s(X40,X43))),s(X39,X44))) = s(X37,happ(s(fun(X40,X37),happ(s(fun(X38,fun(X40,X37)),X41),s(X38,happ(s(fun(X39,X38),X42),s(X39,X44))))),s(X40,X43)))),
  file('thm.inferProps.success_eqns|split|12','_dst_x0x1x0_ax')).
fof(f107754,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f207,f8429,f248])).
fof(f107752,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f168,f8429,f248])).
fof(f107750,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f253,f8429,f248])).
fof(f107748,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f165,f8429,f248])).
fof(f107746,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f184,f8429,f248])).
fof(f107744,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f186,f8429,f248])).
fof(f107742,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f208,f8429,f248])).
fof(f107740,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f205,f8429,f248])).
fof(f107738,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f204,f8429,f248])).
fof(f107736,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f203,f8429,f248])).
fof(f107734,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f171,f8429,f248])).
fof(f107732,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f182,f8429,f248])).
fof(f107730,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f180,f8429,f248])).
fof(f107728,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f206,f8429,f248])).
fof(f107726,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61610,f8429,f248])).
fof(f107724,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f29104,f8429,f248])).
fof(f107722,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f179,f8429,f248])).
fof(f107720,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f178,f8429,f248])).
fof(f107718,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f250,f8429,f248])).
fof(f107716,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f243,f8429,f248])).
fof(f107714,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f245,f8429,f248])).
fof(f107712,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f176,f8429,f248])).
fof(f107710,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f181,f8429,f248])).
fof(f107708,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f177,f8429,f248])).
fof(f107706,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f166,f8429,f248])).
fof(f107704,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f169,f8429,f248])).
fof(f107702,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15105,f8429,f248])).
fof(f107700,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8130,f8429,f248])).
fof(f107698,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3662,f8429,f248])).
fof(f107696,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3513,f8429,f248])).
fof(f107694,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f897,f8429,f248])).
fof(f107692,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f220,f8429,f248])).
fof(f107690,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f743,f8429,f248])).
fof(f107688,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15104,f8429,f248])).
fof(f107686,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f715,f8429,f248])).
fof(f107684,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8129,f8429,f248])).
fof(f107682,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f590,f8429,f248])).
fof(f107680,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3661,f8429,f248])).
fof(f107678,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f570,f8429,f248])).
fof(f107676,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3512,f8429,f248])).
fof(f107674,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f552,f8429,f248])).
fof(f107672,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f97221,f8429,f248])).
fof(f107670,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91598,f8429,f248])).
fof(f107668,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91565,f8429,f248])).
fof(f107666,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f74672,f8429,f248])).
fof(f107664,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f69102,f8429,f248])).
fof(f107662,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91563,f8429,f248])).
fof(f107660,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f85050,f8429,f248])).
fof(f107658,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f88012,f8429,f248])).
fof(f107656,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73313,f8429,f248])).
fof(f107654,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62913,f8429,f248])).
fof(f107652,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62918,f8429,f248])).
fof(f107650,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65997,f8429,f248])).
fof(f107648,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3463,f8429,f248])).
fof(f107646,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62900,f8429,f248])).
fof(f107644,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62899,f8429,f248])).
fof(f107642,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73296,f8429,f248])).
fof(f107640,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62898,f8429,f248])).
fof(f107638,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62878,f8429,f248])).
fof(f107636,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62876,f8429,f248])).
fof(f107634,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62874,f8429,f248])).
fof(f107632,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62872,f8429,f248])).
fof(f107630,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62870,f8429,f248])).
fof(f107628,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62888,f8429,f248])).
fof(f107626,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62862,f8429,f248])).
fof(f107624,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64997,f8429,f248])).
fof(f107622,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64998,f8429,f248])).
fof(f107620,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65998,f8429,f248])).
fof(f107618,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f63847,f8429,f248])).
fof(f107616,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61142,f8429,f248])).
fof(f107614,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3496,f8429,f248])).
fof(f107612,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f185,f8429,f248])).
fof(f107610,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f222,f8429,f248])).
fof(f107608,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f207,f8429,f248])).
fof(f107606,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f168,f8429,f248])).
fof(f107603,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f253,f8429,f213])).
fof(f107601,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f165,f8429,f213])).
fof(f107599,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f184,f8429,f213])).
fof(f107597,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f186,f8429,f213])).
fof(f107595,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f208,f8429,f213])).
fof(f107593,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f205,f8429,f213])).
fof(f107591,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f204,f8429,f213])).
fof(f107589,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f203,f8429,f213])).
fof(f107587,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f171,f8429,f213])).
fof(f107585,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f182,f8429,f213])).
fof(f107583,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f180,f8429,f213])).
fof(f107581,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f206,f8429,f213])).
fof(f107579,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61610,f8429,f213])).
fof(f107577,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f29104,f8429,f213])).
fof(f107575,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f179,f8429,f213])).
fof(f107573,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f178,f8429,f213])).
fof(f107571,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f250,f8429,f213])).
fof(f107569,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f243,f8429,f213])).
fof(f107567,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f245,f8429,f213])).
fof(f107565,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f176,f8429,f213])).
fof(f107563,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f181,f8429,f213])).
fof(f107561,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f213])).
fof(f107559,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f166,f8429,f213])).
fof(f107557,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f169,f8429,f213])).
fof(f107555,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15105,f8429,f213])).
fof(f107553,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8130,f8429,f213])).
fof(f107551,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3662,f8429,f213])).
fof(f107549,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3513,f8429,f213])).
fof(f107547,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f897,f8429,f213])).
fof(f107545,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f220,f8429,f213])).
fof(f107543,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f743,f8429,f213])).
fof(f107541,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15104,f8429,f213])).
fof(f107539,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f715,f8429,f213])).
fof(f107537,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8129,f8429,f213])).
fof(f107535,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f590,f8429,f213])).
fof(f107533,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3661,f8429,f213])).
fof(f107531,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f570,f8429,f213])).
fof(f107529,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3512,f8429,f213])).
fof(f107527,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f552,f8429,f213])).
fof(f107525,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f97221,f8429,f213])).
fof(f107523,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91598,f8429,f213])).
fof(f107521,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91565,f8429,f213])).
fof(f107519,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f74672,f8429,f213])).
fof(f107517,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f69102,f8429,f213])).
fof(f107515,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91563,f8429,f213])).
fof(f107513,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f85050,f8429,f213])).
fof(f107511,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f88012,f8429,f213])).
fof(f107509,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73313,f8429,f213])).
fof(f107507,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62913,f8429,f213])).
fof(f107505,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62918,f8429,f213])).
fof(f107503,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65997,f8429,f213])).
fof(f107501,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3463,f8429,f213])).
fof(f107499,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62900,f8429,f213])).
fof(f107497,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62899,f8429,f213])).
fof(f107495,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73296,f8429,f213])).
fof(f107493,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62898,f8429,f213])).
fof(f107491,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62878,f8429,f213])).
fof(f107489,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62876,f8429,f213])).
fof(f107487,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62874,f8429,f213])).
fof(f107485,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62872,f8429,f213])).
fof(f107483,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62870,f8429,f213])).
fof(f107481,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62888,f8429,f213])).
fof(f107479,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62862,f8429,f213])).
fof(f107477,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64997,f8429,f213])).
fof(f107475,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64998,f8429,f213])).
fof(f107473,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65998,f8429,f213])).
fof(f107471,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f63847,f8429,f213])).
fof(f107469,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61142,f8429,f213])).
fof(f107467,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3496,f8429,f213])).
fof(f107465,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f185,f8429,f213])).
fof(f107463,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f222,f8429,f213])).
fof(f107461,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f207,f8429,f213])).
fof(f107459,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f168,f8429,f213])).
fof(f107457,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f253,f8429,f213])).
fof(f107455,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f165,f8429,f213])).
fof(f107453,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f184,f8429,f213])).
fof(f107451,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f186,f8429,f213])).
fof(f107449,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f208,f8429,f213])).
fof(f107447,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f205,f8429,f213])).
fof(f107445,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f204,f8429,f213])).
fof(f107443,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f203,f8429,f213])).
fof(f107441,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f171,f8429,f213])).
fof(f107439,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f182,f8429,f213])).
fof(f107437,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f180,f8429,f213])).
fof(f107435,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f206,f8429,f213])).
fof(f107433,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61610,f8429,f213])).
fof(f107431,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f29104,f8429,f213])).
fof(f107429,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f179,f8429,f213])).
fof(f107427,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f178,f8429,f213])).
fof(f107425,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f250,f8429,f213])).
fof(f107423,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f243,f8429,f213])).
fof(f107421,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f245,f8429,f213])).
fof(f107419,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f176,f8429,f213])).
fof(f107417,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f181,f8429,f213])).
fof(f107415,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f213])).
fof(f107413,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f166,f8429,f213])).
fof(f107411,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f169,f8429,f213])).
fof(f107409,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15105,f8429,f213])).
fof(f107407,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8130,f8429,f213])).
fof(f107405,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3662,f8429,f213])).
fof(f107403,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3513,f8429,f213])).
fof(f107401,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f897,f8429,f213])).
fof(f107399,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f220,f8429,f213])).
fof(f107397,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f743,f8429,f213])).
fof(f107395,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15104,f8429,f213])).
fof(f107393,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f715,f8429,f213])).
fof(f107391,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8129,f8429,f213])).
fof(f107389,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f590,f8429,f213])).
fof(f107387,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3661,f8429,f213])).
fof(f107385,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f570,f8429,f213])).
fof(f107383,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3512,f8429,f213])).
fof(f107381,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f552,f8429,f213])).
fof(f107379,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f97221,f8429,f213])).
fof(f107377,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91598,f8429,f213])).
fof(f107375,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91565,f8429,f213])).
fof(f107373,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f74672,f8429,f213])).
fof(f107371,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f69102,f8429,f213])).
fof(f107369,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91563,f8429,f213])).
fof(f107367,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f85050,f8429,f213])).
fof(f107365,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f88012,f8429,f213])).
fof(f107363,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73313,f8429,f213])).
fof(f107361,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62913,f8429,f213])).
fof(f107359,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62918,f8429,f213])).
fof(f107357,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65997,f8429,f213])).
fof(f107355,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3463,f8429,f213])).
fof(f107353,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62900,f8429,f213])).
fof(f107351,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62899,f8429,f213])).
fof(f107349,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73296,f8429,f213])).
fof(f107347,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62898,f8429,f213])).
fof(f107345,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62878,f8429,f213])).
fof(f107343,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62876,f8429,f213])).
fof(f107341,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62874,f8429,f213])).
fof(f107339,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62872,f8429,f213])).
fof(f107337,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62870,f8429,f213])).
fof(f107335,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62888,f8429,f213])).
fof(f107333,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62862,f8429,f213])).
fof(f107331,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64997,f8429,f213])).
fof(f107329,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64998,f8429,f213])).
fof(f107327,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65998,f8429,f213])).
fof(f107325,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f63847,f8429,f213])).
fof(f107323,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61142,f8429,f213])).
fof(f107321,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3496,f8429,f213])).
fof(f107319,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f185,f8429,f213])).
fof(f107317,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f222,f8429,f213])).
fof(f107315,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f207,f8429,f213])).
fof(f107313,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f168,f8429,f213])).
fof(f107163,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f175])).
fof(f175,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) | s('type.infer_t.infer_t',X3) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1))) | s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) | (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) & s('type.infer_t.infer_t',X3) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1))))) & (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) | s('type.infer_t.infer_t',X3) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1))) | s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4)))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) | (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) & s('type.infer_t.infer_t',X3) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1))))) & ((s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) | s('type.infer_t.infer_t',X3) != s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1)))) | s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4)))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) <=> (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) & s('type.infer_t.infer_t',X3) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X0),'const.infer.Success_1'(s('type.infer_t.infer_t',X3))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) <=> (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4) & s('type.infer_t.infer_t',X3) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X20,X21,X22,X23] : (s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_2'(s('type.infer_t.infer_t',X20),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X21))) = s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.infer_t.infer_t',X6),'const.infer.Success_1'(s('type.infer_t.infer_t',X22))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))) <=> (s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X21) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23) & s('type.infer_t.infer_t',X22) = s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X21))),s('type.infer_t.infer_t',X20)))))),
  file('thm.inferProps.success_eqns|split|12','thm.inferProps.inferProps_28_prove_ltb')).
fof(f107134,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f552,f8429,f558])).
fof(f107132,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f552,f8429,f183])).
fof(f107130,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f68711])).
fof(f107128,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f67994])).
fof(f107126,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f286])).
fof(f107124,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f91543])).
fof(f107122,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f75410])).
fof(f107120,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f74666])).
fof(f107118,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69819])).
fof(f107116,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f286])).
fof(f107114,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f62861])).
fof(f107112,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f65627])).
fof(f107110,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f66649])).
fof(f107108,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107106,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107104,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f275,f8429,f69996])).
fof(f107099,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f64447])).
fof(f107097,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f63836])).
fof(f107095,plain,(
  $false | ($spl7 | $spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f3461])).
fof(f107093,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107091,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107089,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f275,f8429,f267])).
fof(f107084,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f107082,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f107080,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f164,f275,f8429,f354])).
fof(f107077,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f91543])).
fof(f107074,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f75410])).
fof(f107072,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f74666])).
fof(f107070,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107068,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107063,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107061,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f107056,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69819])).
fof(f107054,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f68711])).
fof(f107052,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f67994])).
fof(f107050,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f66649])).
fof(f107048,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f65627])).
fof(f107046,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f64447])).
fof(f107044,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f63836])).
fof(f107042,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f62861])).
fof(f107040,plain,(
  $false | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f3461])).
fof(f107038,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f107036,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f164,f275,f8429,f354])).
fof(f107034,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f107030,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f8429,f269])).
fof(f269,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f164,f167])).
fof(f107028,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107026,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107024,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f275,f8429,f267])).
fof(f107018,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107016,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f107014,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f275,f8429,f267])).
fof(f107008,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f167])).
fof(f107006,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f167])).
fof(f107003,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f29104,f8429,f1749])).
fof(f1749,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.read_1'(s(X0,X2))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.st_ex_return_2'(s(X0,X3),s(X0,X4))) | s(X0,X2) = s(X0,X4)) )),
  inference(superposition,[],[f212,f178])).
fof(f212,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X2,X4) = s(X2,X6)) )),
  inference(cnf_transformation,[],[f143])).
fof(f107001,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f29104,f8429,f1437])).
fof(f1437,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.read_1'(s(X0,X2))) != s('type.pair.prod'('type.infer.exc'(X0,X1),X0),'const.infer.st_ex_return_2'(s(X0,X3),s(X0,X4))) | s(X0,X2) = s(X0,X3)) )),
  inference(superposition,[],[f211,f178])).
fof(f211,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.infer.st_ex_return_2'(s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'(X1,X0),X2),'const.pair.,_2'(s('type.infer.exc'(X1,X0),'const.infer.Success_1'(s(X1,X5))),s(X2,X6))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f143])).
fof(f106543,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f243,f8429,f224])).
fof(f106537,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f243,f8429,f223])).
fof(f106535,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f253,f8429,f248])).
fof(f106533,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f165,f8429,f248])).
fof(f106531,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f184,f8429,f248])).
fof(f106529,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f186,f8429,f248])).
fof(f106527,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f208,f8429,f248])).
fof(f106525,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f205,f8429,f248])).
fof(f106523,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f204,f8429,f248])).
fof(f106521,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f203,f8429,f248])).
fof(f106519,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f171,f8429,f248])).
fof(f106517,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f182,f8429,f248])).
fof(f106515,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f180,f8429,f248])).
fof(f106513,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f206,f8429,f248])).
fof(f106511,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61610,f8429,f248])).
fof(f106509,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f29104,f8429,f248])).
fof(f106507,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f179,f8429,f248])).
fof(f106505,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f178,f8429,f248])).
fof(f106503,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f250,f8429,f248])).
fof(f106501,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f243,f8429,f248])).
fof(f106499,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f245,f8429,f248])).
fof(f106497,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f176,f8429,f248])).
fof(f106495,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f181,f8429,f248])).
fof(f106493,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f177,f8429,f248])).
fof(f106491,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f166,f8429,f248])).
fof(f106489,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f169,f8429,f248])).
fof(f106487,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15105,f8429,f248])).
fof(f106485,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8130,f8429,f248])).
fof(f106483,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3662,f8429,f248])).
fof(f106481,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3513,f8429,f248])).
fof(f106479,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f897,f8429,f248])).
fof(f106477,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f220,f8429,f248])).
fof(f106475,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f743,f8429,f248])).
fof(f106473,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15104,f8429,f248])).
fof(f106471,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f715,f8429,f248])).
fof(f106469,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8129,f8429,f248])).
fof(f106467,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f590,f8429,f248])).
fof(f106465,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3661,f8429,f248])).
fof(f106463,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f570,f8429,f248])).
fof(f106461,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3512,f8429,f248])).
fof(f106459,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f552,f8429,f248])).
fof(f106457,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f97221,f8429,f248])).
fof(f106455,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91598,f8429,f248])).
fof(f106453,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91565,f8429,f248])).
fof(f106451,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f74672,f8429,f248])).
fof(f106449,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f69102,f8429,f248])).
fof(f106447,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91563,f8429,f248])).
fof(f106445,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f85050,f8429,f248])).
fof(f106443,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f88012,f8429,f248])).
fof(f106441,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73313,f8429,f248])).
fof(f106439,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62913,f8429,f248])).
fof(f106437,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62918,f8429,f248])).
fof(f106435,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65997,f8429,f248])).
fof(f106433,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3463,f8429,f248])).
fof(f106431,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62900,f8429,f248])).
fof(f106429,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62899,f8429,f248])).
fof(f106427,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73296,f8429,f248])).
fof(f106425,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62898,f8429,f248])).
fof(f106423,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62878,f8429,f248])).
fof(f106421,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62876,f8429,f248])).
fof(f106419,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62874,f8429,f248])).
fof(f106417,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62872,f8429,f248])).
fof(f106415,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62870,f8429,f248])).
fof(f106413,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62888,f8429,f248])).
fof(f106411,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62862,f8429,f248])).
fof(f106409,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64997,f8429,f248])).
fof(f106407,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64998,f8429,f248])).
fof(f106405,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65998,f8429,f248])).
fof(f106403,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f63847,f8429,f248])).
fof(f106401,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61142,f8429,f248])).
fof(f106399,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3496,f8429,f248])).
fof(f106397,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f185,f8429,f248])).
fof(f106395,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f222,f8429,f248])).
fof(f106393,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f207,f8429,f248])).
fof(f106391,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f168,f8429,f248])).
fof(f106389,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f253,f8429,f248])).
fof(f106387,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f165,f8429,f248])).
fof(f106385,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f184,f8429,f248])).
fof(f106383,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f186,f8429,f248])).
fof(f106381,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f208,f8429,f248])).
fof(f106379,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f205,f8429,f248])).
fof(f106377,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f204,f8429,f248])).
fof(f106375,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f203,f8429,f248])).
fof(f106373,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f171,f8429,f248])).
fof(f106371,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f182,f8429,f248])).
fof(f106369,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f180,f8429,f248])).
fof(f106367,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f206,f8429,f248])).
fof(f106365,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61610,f8429,f248])).
fof(f106363,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f29104,f8429,f248])).
fof(f106361,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f179,f8429,f248])).
fof(f106359,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f178,f8429,f248])).
fof(f106357,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f250,f8429,f248])).
fof(f106355,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f243,f8429,f248])).
fof(f106353,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f245,f8429,f248])).
fof(f106351,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f176,f8429,f248])).
fof(f106349,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f181,f8429,f248])).
fof(f106347,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f177,f8429,f248])).
fof(f106345,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f166,f8429,f248])).
fof(f106343,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f169,f8429,f248])).
fof(f106341,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15105,f8429,f248])).
fof(f106339,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8130,f8429,f248])).
fof(f106337,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3662,f8429,f248])).
fof(f106335,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3513,f8429,f248])).
fof(f106333,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f897,f8429,f248])).
fof(f106331,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f220,f8429,f248])).
fof(f106329,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f743,f8429,f248])).
fof(f106327,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f15104,f8429,f248])).
fof(f106325,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f715,f8429,f248])).
fof(f106323,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f8129,f8429,f248])).
fof(f106321,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f590,f8429,f248])).
fof(f106319,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3661,f8429,f248])).
fof(f106317,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f570,f8429,f248])).
fof(f106315,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3512,f8429,f248])).
fof(f106313,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f552,f8429,f248])).
fof(f106311,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f97221,f8429,f248])).
fof(f106309,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91598,f8429,f248])).
fof(f106307,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91565,f8429,f248])).
fof(f106305,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f74672,f8429,f248])).
fof(f106303,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f69102,f8429,f248])).
fof(f106301,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f91563,f8429,f248])).
fof(f106299,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f85050,f8429,f248])).
fof(f106297,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f88012,f8429,f248])).
fof(f106295,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73313,f8429,f248])).
fof(f106293,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62913,f8429,f248])).
fof(f106291,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62918,f8429,f248])).
fof(f106289,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65997,f8429,f248])).
fof(f106287,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3463,f8429,f248])).
fof(f106285,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62900,f8429,f248])).
fof(f106283,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62899,f8429,f248])).
fof(f106281,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f73296,f8429,f248])).
fof(f106279,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62898,f8429,f248])).
fof(f106277,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62878,f8429,f248])).
fof(f106275,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62876,f8429,f248])).
fof(f106273,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62874,f8429,f248])).
fof(f106271,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62872,f8429,f248])).
fof(f106269,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62870,f8429,f248])).
fof(f106267,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62888,f8429,f248])).
fof(f106265,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f62862,f8429,f248])).
fof(f106263,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64997,f8429,f248])).
fof(f106261,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f64998,f8429,f248])).
fof(f106259,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f65998,f8429,f248])).
fof(f106257,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f63847,f8429,f248])).
fof(f106255,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f61142,f8429,f248])).
fof(f106253,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f163,f3496,f8429,f248])).
fof(f106251,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f185,f8429,f248])).
fof(f106249,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f222,f8429,f248])).
fof(f106247,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f207,f8429,f248])).
fof(f106245,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f168,f8429,f248])).
fof(f106242,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f253,f8429,f213])).
fof(f106240,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f165,f8429,f213])).
fof(f106238,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f184,f8429,f213])).
fof(f106236,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f186,f8429,f213])).
fof(f106234,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f208,f8429,f213])).
fof(f106232,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f205,f8429,f213])).
fof(f106230,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f204,f8429,f213])).
fof(f106228,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f203,f8429,f213])).
fof(f106226,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f171,f8429,f213])).
fof(f106224,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f182,f8429,f213])).
fof(f106222,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f180,f8429,f213])).
fof(f106220,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f206,f8429,f213])).
fof(f106218,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61610,f8429,f213])).
fof(f106216,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f29104,f8429,f213])).
fof(f106214,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f179,f8429,f213])).
fof(f106212,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f178,f8429,f213])).
fof(f106210,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f250,f8429,f213])).
fof(f106208,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f243,f8429,f213])).
fof(f106206,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f245,f8429,f213])).
fof(f106204,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f176,f8429,f213])).
fof(f106202,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f181,f8429,f213])).
fof(f106200,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f213])).
fof(f106198,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f166,f8429,f213])).
fof(f106196,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f169,f8429,f213])).
fof(f106194,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15105,f8429,f213])).
fof(f106192,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8130,f8429,f213])).
fof(f106190,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3662,f8429,f213])).
fof(f106188,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3513,f8429,f213])).
fof(f106186,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f897,f8429,f213])).
fof(f106184,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f220,f8429,f213])).
fof(f106182,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f743,f8429,f213])).
fof(f106180,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15104,f8429,f213])).
fof(f106178,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f715,f8429,f213])).
fof(f106176,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8129,f8429,f213])).
fof(f106174,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f590,f8429,f213])).
fof(f106172,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3661,f8429,f213])).
fof(f106170,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f570,f8429,f213])).
fof(f106168,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3512,f8429,f213])).
fof(f106166,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f552,f8429,f213])).
fof(f106164,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f97221,f8429,f213])).
fof(f106162,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91598,f8429,f213])).
fof(f106160,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91565,f8429,f213])).
fof(f106158,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f74672,f8429,f213])).
fof(f106156,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f69102,f8429,f213])).
fof(f106154,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91563,f8429,f213])).
fof(f106152,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f85050,f8429,f213])).
fof(f106150,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f88012,f8429,f213])).
fof(f106148,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73313,f8429,f213])).
fof(f106146,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62913,f8429,f213])).
fof(f106144,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62918,f8429,f213])).
fof(f106142,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65997,f8429,f213])).
fof(f106140,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3463,f8429,f213])).
fof(f106138,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62900,f8429,f213])).
fof(f106136,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62899,f8429,f213])).
fof(f106134,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73296,f8429,f213])).
fof(f106132,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62898,f8429,f213])).
fof(f106130,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62878,f8429,f213])).
fof(f106128,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62876,f8429,f213])).
fof(f106126,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62874,f8429,f213])).
fof(f106124,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62872,f8429,f213])).
fof(f106122,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62870,f8429,f213])).
fof(f106120,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62888,f8429,f213])).
fof(f106118,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62862,f8429,f213])).
fof(f106116,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64997,f8429,f213])).
fof(f106114,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64998,f8429,f213])).
fof(f106112,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65998,f8429,f213])).
fof(f106110,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f63847,f8429,f213])).
fof(f106108,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61142,f8429,f213])).
fof(f106106,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3496,f8429,f213])).
fof(f106104,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f185,f8429,f213])).
fof(f106102,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f222,f8429,f213])).
fof(f106100,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f207,f8429,f213])).
fof(f106098,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f168,f8429,f213])).
fof(f106096,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f253,f8429,f213])).
fof(f106094,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f165,f8429,f213])).
fof(f106092,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f184,f8429,f213])).
fof(f106090,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f186,f8429,f213])).
fof(f106088,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f208,f8429,f213])).
fof(f106086,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f205,f8429,f213])).
fof(f106084,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f204,f8429,f213])).
fof(f106082,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f203,f8429,f213])).
fof(f106080,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f171,f8429,f213])).
fof(f106078,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f182,f8429,f213])).
fof(f106076,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f180,f8429,f213])).
fof(f106074,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f206,f8429,f213])).
fof(f106072,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61610,f8429,f213])).
fof(f106070,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f29104,f8429,f213])).
fof(f106068,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f179,f8429,f213])).
fof(f106066,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f178,f8429,f213])).
fof(f106064,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f250,f8429,f213])).
fof(f106062,plain,(
  $false | ($spl2 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f243,f8429,f213])).
fof(f106060,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f245,f8429,f213])).
fof(f106058,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f176,f8429,f213])).
fof(f106056,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f181,f8429,f213])).
fof(f106054,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f213])).
fof(f106052,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f166,f8429,f213])).
fof(f106050,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f169,f8429,f213])).
fof(f106048,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15105,f8429,f213])).
fof(f106046,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8130,f8429,f213])).
fof(f106044,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3662,f8429,f213])).
fof(f106042,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3513,f8429,f213])).
fof(f106040,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f897,f8429,f213])).
fof(f106038,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f220,f8429,f213])).
fof(f106036,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f743,f8429,f213])).
fof(f106034,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f15104,f8429,f213])).
fof(f106032,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f715,f8429,f213])).
fof(f106030,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f8129,f8429,f213])).
fof(f106028,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f590,f8429,f213])).
fof(f106026,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3661,f8429,f213])).
fof(f106024,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f570,f8429,f213])).
fof(f106022,plain,(
  $false | ($spl8 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3512,f8429,f213])).
fof(f106020,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f552,f8429,f213])).
fof(f106018,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f97221,f8429,f213])).
fof(f106016,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91598,f8429,f213])).
fof(f106014,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91565,f8429,f213])).
fof(f106012,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f74672,f8429,f213])).
fof(f106010,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f69102,f8429,f213])).
fof(f106008,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f91563,f8429,f213])).
fof(f106006,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f85050,f8429,f213])).
fof(f106004,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f88012,f8429,f213])).
fof(f106002,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73313,f8429,f213])).
fof(f106000,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62913,f8429,f213])).
fof(f105998,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62918,f8429,f213])).
fof(f105996,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65997,f8429,f213])).
fof(f105994,plain,(
  $false | ($spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3463,f8429,f213])).
fof(f105992,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62900,f8429,f213])).
fof(f105990,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62899,f8429,f213])).
fof(f105988,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f73296,f8429,f213])).
fof(f105986,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62898,f8429,f213])).
fof(f105984,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62878,f8429,f213])).
fof(f105982,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62876,f8429,f213])).
fof(f105980,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62874,f8429,f213])).
fof(f105978,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62872,f8429,f213])).
fof(f105976,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62870,f8429,f213])).
fof(f105974,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62888,f8429,f213])).
fof(f105972,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f62862,f8429,f213])).
fof(f105970,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64997,f8429,f213])).
fof(f105968,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f64998,f8429,f213])).
fof(f105966,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f65998,f8429,f213])).
fof(f105964,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f63847,f8429,f213])).
fof(f105962,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f61142,f8429,f213])).
fof(f105960,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f168,f3496,f8429,f213])).
fof(f105958,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f185,f8429,f213])).
fof(f105956,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f222,f8429,f213])).
fof(f105954,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f207,f8429,f213])).
fof(f105952,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f168,f8429,f213])).
fof(f105802,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f177,f8429,f175])).
fof(f105773,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f552,f8429,f558])).
fof(f105771,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f552,f8429,f183])).
fof(f105769,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f68711])).
fof(f105767,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f67994])).
fof(f105765,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f286])).
fof(f105763,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f91543])).
fof(f105761,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f75410])).
fof(f105759,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f74666])).
fof(f105757,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f69819])).
fof(f105755,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f286])).
fof(f105753,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f62861])).
fof(f105751,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f65627])).
fof(f105749,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f66649])).
fof(f105747,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105745,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105743,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f275,f8429,f69996])).
fof(f105738,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f64447])).
fof(f105736,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f63836])).
fof(f105734,plain,(
  $false | ($spl7 | $spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f3461])).
fof(f105732,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105730,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105728,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f275,f8429,f267])).
fof(f105723,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f105721,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f105719,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f164,f275,f8429,f354])).
fof(f105716,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f91543])).
fof(f105713,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f75410])).
fof(f105711,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f74666])).
fof(f105709,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105707,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105702,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105700,plain,(
  $false | ($spl7 | $spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f69996])).
fof(f105695,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f69819])).
fof(f105693,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f68711])).
fof(f105691,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f67994])).
fof(f105689,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f66649])).
fof(f105687,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f65627])).
fof(f105685,plain,(
  $false | ($spl7 | $spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f64447])).
fof(f105683,plain,(
  $false | ($spl8 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f63836])).
fof(f105681,plain,(
  $false | ($spl8 | $spl16 | $spl20 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f62861])).
fof(f105679,plain,(
  $false | ($spl16 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f8429,f3461])).
fof(f105677,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f105675,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f164,f275,f8429,f354])).
fof(f105673,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f164,f8429,f8429,f354])).
fof(f105669,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f163,f8429,f269])).
fof(f105667,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105665,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105663,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f275,f8429,f267])).
fof(f105657,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105655,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f8429,f267])).
fof(f105653,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f8429,f275,f8429,f267])).
fof(f105647,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f8429,f8429,f167])).
fof(f105645,plain,(
  $false | ($spl7 | $spl36)),
  inference(unit_resulting_resolution,[],[f275,f8429,f167])).
fof(f105642,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f29104,f8429,f1749])).
fof(f105640,plain,(
  $false | ($spl32 | $spl36)),
  inference(unit_resulting_resolution,[],[f29104,f8429,f1437])).
fof(f105632,plain,(
  $false | $spl36),
  inference(unit_resulting_resolution,[],[f168,f8429])).
fof(f108597,plain,(
  $false | $spl36),
  inference(equality_resolution,[],[f8429])).
fof(f8429,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) | $spl36),
  inference(cnf_transformation,[],[f8429_D])).
fof(f8429_D,plain,(
  ( ! [X2,X3,X1] : (s(X1,X2) != s(X1,X3)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f105623,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) ) | $spl68),
  inference(cnf_transformation,[],[f105623_D])).
fof(f105623_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f105622,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(X2,X3) != s(X2,X4)) ) | ($spl7 | $spl8 | $spl16 | $spl20)),
  inference(condensation,[],[f105621])).
fof(f105621,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(X31,X32) != s(X31,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28)) ) | ($spl7 | $spl8 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f105221,f1173])).
fof(f1173,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X0),s(X1,X3),s(X2,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X1),'const.infer.Success_1'(s('type.one.one',X5))),s(X2,X6)))) )),
  inference(condensation,[],[f1172])).
fof(f1172,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X18),X19),'const.infer.guard_3'(s(bool,X16),s(X18,X20),s(X19,X21))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X18),X19),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X18),'const.infer.Success_1'(s('type.one.one',X22))),s(X19,X23))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1164,f269])).
fof(f1164,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X17)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X18),X19),'const.infer.guard_3'(s(bool,X16),s(X18,X20),s(X19,X21))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X18),X19),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X18),'const.infer.Success_1'(s('type.one.one',X22))),s(X19,X23))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f196,f267])).
fof(f196,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,X2),s(X0,X3),s(X1,X4))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f105221,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X29),s(bool,X30)))) | s(X31,X32) != s(X31,X33) | s('type.pair.prod'('type.infer.exc'('type.one.one',X34),X31),'const.infer.guard_3'(s(bool,X28),s(X34,X35),s(X31,X32))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X34),X31),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X34),'const.infer.Success_1'(s('type.one.one',X36))),s(X31,X33))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28)) ) | ($spl7 | $spl8 | $spl16 | $spl20)),
  inference(superposition,[],[f248,f69996])).
fof(f61145,plain,(
  $false | ($spl20 | $spl24)),
  inference(subsumption_resolution,[],[f61143,f164])).
fof(f61143,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl24)),
  inference(backward_demodulation,[],[f3463,f3465])).
fof(f61141,plain,(
  $false | ($spl24 | $spl26)),
  inference(subsumption_resolution,[],[f3466,f3465])).
fof(f61077,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,f(s('type.one.one',X21)))) ) | $spl64),
  inference(cnf_transformation,[],[f61077_D])).
fof(f61077_D,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,f(s('type.one.one',X21)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f61078,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.infer.guard_3'(s(bool,'T_0'),s(X25,X26),s(X22,X23))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X25),'const.infer.Success_1'(s('type.one.one',X27))),s(X22,X24))) | s(X22,X23) != s(X22,X24)) ) | $spl66),
  inference(cnf_transformation,[],[f61078_D])).
fof(f61078_D,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.infer.guard_3'(s(bool,'T_0'),s(X25,X26),s(X22,X23))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X25),'const.infer.Success_1'(s('type.one.one',X27))),s(X22,X24))) | s(X22,X23) != s(X22,X24)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f30303,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.infer.guard_3'(s(bool,'T_0'),s(X25,X26),s(X22,X23))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X25),'const.infer.Success_1'(s('type.one.one',X27))),s(X22,X24))) | s(X22,X23) != s(X22,X24) | s(bool,'T_0') = s(bool,f(s('type.one.one',X21)))) ) | ($spl12 | $spl32)),
  inference(forward_demodulation,[],[f30302,f3496])).
fof(f30302,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(X22,X23) != s(X22,X24) | s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X25),'const.infer.Success_1'(s('type.one.one',X27))),s(X22,X24))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.infer.guard_3'(s(bool,g(s('type.one.one',X21))),s(X25,X26),s(X22,X23))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X21)))) ) | $spl12),
  inference(subsumption_resolution,[],[f30235,f163])).
fof(f30235,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (~p(s(bool,'T_0')) | s(X22,X23) != s(X22,X24) | s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X25),'const.infer.Success_1'(s('type.one.one',X27))),s(X22,X24))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X25),X22),'const.infer.guard_3'(s(bool,g(s('type.one.one',X21))),s(X25,X26),s(X22,X23))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X21)))) ) | $spl12),
  inference(superposition,[],[f248,f26896])).
fof(f26896,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X0)))) ) | $spl12),
  inference(condensation,[],[f26895])).
fof(f26895,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f26894])).
fof(f26894,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,g(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(subsumption_resolution,[],[f26870,f267])).
fof(f26870,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,g(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(superposition,[],[f899,f267])).
fof(f61069,plain,(
  $false | $spl38),
  inference(subsumption_resolution,[],[f60822,f163])).
fof(f60822,plain,(
  ~p(s(bool,'T_0')) | $spl38),
  inference(backward_demodulation,[],[f8430,f164])).
fof(f61068,plain,(
  $false | ($spl7 | $spl18 | $spl38)),
  inference(trivial_inequality_removal,[],[f60947])).
fof(f60947,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0')))) ) | ($spl7 | $spl18 | $spl38)),
  inference(backward_demodulation,[],[f8430,f20767])).
fof(f20767,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'F_0')))) ) | ($spl7 | $spl18)),
  inference(superposition,[],[f2661,f7699])).
fof(f7699,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f7698,f178])).
fof(f7698,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,'T_0'),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f7697,f3462])).
fof(f7697,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,'T_0')))) ) | $spl18),
  inference(forward_demodulation,[],[f7579,f3462])).
fof(f7579,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.pair.,_2'(s('type.infer.exc'(bool,X0),'const.infer.Success_1'(s(bool,'T_0'))),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,happ(s(fun('type.one.one',bool),X1),s('type.one.one',X2))),s(bool,happ(s(fun('type.one.one',bool),X3),s('type.one.one',X4)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f3462,f3462,f213])).
fof(f2661,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.read_1'(s(bool,'F_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X0),bool),'const.infer.st_ex_return_2'(s(bool,X1),s(bool,'T_0')))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f275,f1749])).
fof(f8430,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f8430_D])).
fof(f8430_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f60779,plain,(
  ( ! [X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) | $spl62),
  inference(cnf_transformation,[],[f60779_D])).
fof(f60779_D,plain,(
  ( ! [X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f60488,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X80) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f60057])).
fof(f60057,plain,(
  ( ! [X80,X83,X81,X82] : (s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X82),bool),'const.infer.read_1'(s(bool,X81))) | s(bool,'T_0') = s(bool,X80) | s(bool,X81) = s(bool,happ(s(fun(bool,bool),X83),s(bool,X81))) | s(bool,'T_0') = s(bool,X80)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f20769,f10019])).
fof(f10019,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f9396,f267])).
fof(f9396,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0')))) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f3604,f3462])).
fof(f3604,plain,(
  ( ! [X14,X12,X10,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X10,bool),X11),s(X10,happ(s(fun('type.one.one',X10),X12),s('type.one.one',X14)))))) ) | ($spl8 | $spl18)),
  inference(forward_demodulation,[],[f3526,f3462])).
fof(f3526,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14))) = s(bool,happ(s(fun(X10,bool),X11),s(X10,happ(s(fun('type.one.one',X10),X12),s('type.one.one',X14)))))) ) | $spl8),
  inference(superposition,[],[f207,f897])).
fof(f20769,plain,(
  ( ! [X2,X3] : (s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,'T_0'))) != s('type.pair.prod'('type.infer.exc'(bool,X2),bool),'const.infer.read_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl18),
  inference(superposition,[],[f1749,f7699])).
fof(f60778,plain,(
  ( ! [X29,X36] : (~p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X29)))) ) | $spl60),
  inference(cnf_transformation,[],[f60778_D])).
fof(f60778_D,plain,(
  ( ! [X29,X36] : (~p(s(bool,X29)) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X29)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f60491,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f60050])).
fof(f60050,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X54))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f269,f10019])).
fof(f60777,plain,(
  ( ! [X30,X28,X31,X29,X36,X32] : (~p(s(bool,X29)) | s(X30,X31) != s(X30,X32) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X29))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f60045,f1173])).
fof(f60045,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (~p(s(bool,X29)) | s(X30,X31) != s(X30,X32) | s('type.pair.prod'('type.infer.exc'('type.one.one',X33),X30),'const.infer.guard_3'(s(bool,X28),s(X33,X34),s(X30,X31))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X33),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X33),'const.infer.Success_1'(s('type.one.one',X35))),s(X30,X32))) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X36),s(bool,X29))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f248,f10019])).
fof(f60697,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) | $spl56),
  inference(cnf_transformation,[],[f60697_D])).
fof(f60697_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f60496,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f60041])).
fof(f60041,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f167,f10019])).
fof(f60661,plain,(
  ( ! [X130,X128] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))) = s(bool,happ(s(fun(bool,bool),X130),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128)))) ) | $spl54),
  inference(cnf_transformation,[],[f60661_D])).
fof(f60661_D,plain,(
  ( ! [X130,X128] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))) = s(bool,happ(s(fun(bool,bool),X130),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f60735,plain,(
  ( ! [X261,X259,X260] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X259))) = s(bool,happ(s(fun(bool,bool),X261),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X259))))) | s(bool,'T_0') = s(bool,X260) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X259)))) ) | ($spl8 | $spl18)),
  inference(forward_demodulation,[],[f60514,f168])).
fof(f60514,plain,(
  ( ! [X261,X259,X260] : (s(bool,'T_0') = s(bool,X260) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X259))) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X259))) = s(bool,happ(s(fun(bool,bool),X261),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X259)))))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59944])).
fof(f59944,plain,(
  ( ! [X261,X259,X260] : (s(bool,'T_0') = s(bool,X260) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X259))) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X259))) = s(bool,happ(s(fun(bool,bool),X261),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X259))))) | s(bool,'T_0') = s(bool,X260)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f286,f10019])).
fof(f60700,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10)))) ) | $spl58),
  inference(cnf_transformation,[],[f60700_D])).
fof(f60700_D,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f60699,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f59892,f269])).
fof(f59892,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (p(s(bool,X11)) | s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.infer.guard_3'(s(bool,X10),s(X12,X14),s(X13,X15))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X12),X13),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X12),'const.infer.Success_1'(s('type.one.one',X16))),s(X13,X17))) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f196,f10019])).
fof(f60550,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59790])).
fof(f59790,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X29),s(bool,X28))) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f167])).
fof(f60631,plain,(
  ( ! [X37,X35,X38,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37)))) ) | $spl52),
  inference(cnf_transformation,[],[f60631_D])).
fof(f60631_D,plain,(
  ( ! [X37,X35,X38,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,X34) = s(bool,X37) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f60551,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X14)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59788])).
fof(f59788,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),X18),s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f10019])).
fof(f60560,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X32),s(bool,X31))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X30)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59731])).
fof(f59731,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,happ(s(fun(bool,bool),X32),s(bool,X31))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X33),s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f10019])).
fof(f60565,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59724])).
fof(f59724,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f167])).
fof(f60660,plain,(
  ( ! [X130,X128,X129] : (s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))) = s(bool,happ(s(fun(bool,bool),X130),s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128)))) ) | ($spl8 | $spl18)),
  inference(forward_demodulation,[],[f60583,f168])).
fof(f60583,plain,(
  ( ! [X130,X128,X129] : (s(bool,'T_0') = s(bool,X129) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X128))) = s(bool,happ(s(fun(bool,bool),X130),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X128))))) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128)))) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59700])).
fof(f59700,plain,(
  ( ! [X130,X128,X129] : (s(bool,'T_0') = s(bool,X129) | s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X128))) = s(bool,happ(s(fun(bool,bool),X130),s(bool,happ(s(fun('type.infer.infer_st'(bool),bool),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(bool),X128))))) | s(bool,'T_0') = s(bool,X129) | s(bool,'F_0') = s(bool,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(bool),X128)))) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f286])).
fof(f60608,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X19))) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X22))) | s(bool,X19) = s(bool,X22)) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59678])).
fof(f59678,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X21),s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),X23),s(bool,X22))) | s(bool,X19) = s(bool,X22)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f10019])).
fof(f60623,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))) | s(bool,X34) = s(bool,X37)) ) | ($spl8 | $spl18)),
  inference(duplicate_literal_removal,[],[f59615])).
fof(f59615,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X34) = s(bool,happ(s(fun(bool,bool),X35),s(bool,X34))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,happ(s(fun(bool,bool),X38),s(bool,X37))) | s(bool,X34) = s(bool,X37)) ) | ($spl8 | $spl18)),
  inference(superposition,[],[f10019,f10019])).
fof(f3496,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl32),
  inference(cnf_transformation,[],[f3496_D])).
fof(f3496_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f26881,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl50),
  inference(cnf_transformation,[],[f26881_D])).
fof(f26881_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f26861,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl12),
  inference(superposition,[],[f899,f165])).
fof(f26880,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl48),
  inference(cnf_transformation,[],[f26880_D])).
fof(f26880_D,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f26860,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) != s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl12),
  inference(superposition,[],[f899,f253])).
fof(f26859,plain,(
  ( ! [X4,X2,X3] : (s(bool,X4) != s(bool,f(s('type.one.one',X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,g(s('type.one.one',X2)))) ) | $spl12),
  inference(superposition,[],[f899,f253])).
fof(f3494,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) | $spl28),
  inference(cnf_transformation,[],[f3494_D])).
fof(f3494_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f26858,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl12),
  inference(superposition,[],[f899,f165])).
fof(f24409,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) | $spl44),
  inference(cnf_transformation,[],[f24409_D])).
fof(f24409_D,plain,(
  ( ! [X4,X2] : (s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f24411,plain,(
  s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl47),
  inference(cnf_transformation,[],[f24411_D])).
fof(f24411_D,plain,(
  s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f24394,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,X9) = s(bool,g(s('type.one.one',X8)))) ) | $spl10),
  inference(superposition,[],[f898,f165])).
fof(f24408,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl42),
  inference(cnf_transformation,[],[f24408_D])).
fof(f24408_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f24393,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,X7) = s(bool,g(s('type.one.one',X5)))) ) | $spl10),
  inference(superposition,[],[f898,f253])).
fof(f24392,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,g(s('type.one.one',X2)))) ) | $spl10),
  inference(superposition,[],[f898,f253])).
fof(f24407,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl40),
  inference(cnf_transformation,[],[f24407_D])).
fof(f24407_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f24391,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl10),
  inference(superposition,[],[f898,f165])).
fof(f9144,plain,(
  $false | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f9031,f163])).
fof(f9031,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl18)),
  inference(backward_demodulation,[],[f9013,f164])).
fof(f9013,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | ($spl16 | $spl18)),
  inference(duplicate_literal_removal,[],[f8881])).
fof(f8881,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | ($spl16 | $spl18)),
  inference(superposition,[],[f3461,f3462])).
fof(f9027,plain,(
  $false | ($spl7 | $spl16 | $spl18)),
  inference(subsumption_resolution,[],[f9026,f275])).
fof(f9026,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl16 | $spl18)),
  inference(condensation,[],[f9025])).
fof(f9025,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X21)) ) | ($spl16 | $spl18)),
  inference(forward_demodulation,[],[f9015,f3462])).
fof(f9015,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X22),s('type.one.one',X23))) | s(bool,'T_0') = s(bool,X21)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f8861])).
fof(f8861,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X22),s('type.one.one',X23))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl16),
  inference(superposition,[],[f3461,f167])).
fof(f9022,plain,(
  $false | ($spl7 | $spl16 | $spl18)),
  inference(subsumption_resolution,[],[f9021,f275])).
fof(f9021,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl16 | $spl18)),
  inference(forward_demodulation,[],[f8855,f3462])).
fof(f8855,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X0),s('type.one.one',X1)))) ) | ($spl7 | $spl16)),
  inference(unit_resulting_resolution,[],[f275,f3461])).
fof(f8433,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8432,f3509])).
fof(f3509,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,'T_0'),s(X29,X31),s(X30,X32))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34)))) ) | ($spl18 | $spl22)),
  inference(backward_demodulation,[],[f3462,f3464])).
fof(f8432,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8431])).
fof(f8431,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8372,f354])).
fof(f8372,plain,(
  ( ! [X47,X54,X52,X50,X48,X53,X51,X49] : (~p(s(bool,X47)) | s(X48,X49) != s(X48,X50) | s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.infer.guard_3'(s(bool,'T_0'),s(X51,X52),s(X48,X49))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X51),X48),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X51),'const.infer.Success_1'(s('type.one.one',X53))),s(X48,X50))) | s(bool,X47) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f248,f267])).
fof(f8428,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8427,f3509])).
fof(f8427,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X4,X5),s(X1,X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X4),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X4),'const.infer.Success_1'(s('type.one.one',X6))),s(X1,X3)))) )),
  inference(condensation,[],[f8426])).
fof(f8426,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f8371,f354])).
fof(f8371,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (~p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.infer.guard_3'(s(bool,'T_0'),s(X43,X44),s(X40,X41))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X43),X40),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X43),'const.infer.Success_1'(s('type.one.one',X45))),s(X40,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f248,f267])).
fof(f8424,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8320,f3509])).
fof(f8320,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f168,f248])).
fof(f8423,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8321,f3509])).
fof(f8321,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f163,f207,f248])).
fof(f8422,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8322,f3509])).
fof(f8322,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(bool,happ(s(fun('type.one.one',bool),X3),s('type.one.one',X4)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f163,f3462,f248])).
fof(f8421,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8323,f3509])).
fof(f8323,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f163,f552,f248])).
fof(f8420,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8324,f3509])).
fof(f8324,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3512,f248])).
fof(f8419,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8325,f3509])).
fof(f8325,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f163,f570,f248])).
fof(f8418,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8326,f3509])).
fof(f8326,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun('type.one.one',bool)))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,fun('type.one.one',bool))),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun('type.one.one',bool)))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,fun('type.one.one',bool))),X6)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3661,f248])).
fof(f8417,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8327,f3509])).
fof(f8327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f163,f590,f248])).
fof(f8416,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8328,f3509])).
fof(f8328,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f163,f715,f248])).
fof(f8415,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8329,f3509])).
fof(f8329,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f163,f743,f248])).
fof(f8414,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8330,f3509])).
fof(f8330,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f897,f248])).
fof(f8413,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8331,f3509])).
fof(f8331,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',fun('type.one.one',bool)),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',fun('type.one.one',bool)),X4)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3513,f248])).
fof(f8412,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8332,f3509])).
fof(f8332,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f163,f166,f248])).
fof(f8411,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8333,f3509])).
fof(f8333,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f177,f248])).
fof(f8410,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8334,f3509])).
fof(f8334,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f163,f181,f248])).
fof(f8409,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8335,f3509])).
fof(f8335,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f176,f248])).
fof(f8408,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8336,f3509])).
fof(f8336,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f178,f248])).
fof(f8407,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8337,f3509])).
fof(f8337,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f179,f248])).
fof(f8406,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8338,f3509])).
fof(f8338,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f180,f248])).
fof(f8405,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8339,f3509])).
fof(f8339,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f205,f248])).
fof(f8404,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8340,f3509])).
fof(f8340,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f184,f248])).
fof(f8403,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8341,f3509])).
fof(f8341,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one','const.one.one_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.one.one',X3)))) )),
  inference(unit_resulting_resolution,[],[f163,f165,f248])).
fof(f8402,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8342,f3509])).
fof(f8342,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f163,f253,f248])).
fof(f8401,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8343,f3509])).
fof(f8343,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(X1,happ(s(fun('type.infer.infer_st'(X1),X1),'const.infer.infer_st_subst_0'),s('type.infer.infer_st'(X1),X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f168,f248])).
fof(f8400,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8344,f3509])).
fof(f8344,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s(X1,happ(s(fun(X3,X1),X4),s(X3,happ(s(fun(X5,X3),X6),s(X5,X7))))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),X1),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X8),s(X1,happ(s(fun(X5,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X5,X3),X6))),s(X5,X7)))))) )),
  inference(unit_resulting_resolution,[],[f163,f207,f248])).
fof(f8399,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8345,f3509])).
fof(f8345,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s(bool,'T_0'))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),bool),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(bool,happ(s(fun('type.one.one',bool),X3),s('type.one.one',X4)))))) ) | $spl18),
  inference(unit_resulting_resolution,[],[f163,f3462,f248])).
fof(f8398,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8346,f3509])).
fof(f8346,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,'type.one.one'),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,'type.one.one')),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,'type.one.one'),X5)))) )),
  inference(unit_resulting_resolution,[],[f163,f552,f248])).
fof(f8397,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8347,f3509])).
fof(f8347,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X2),s(fun(X1,fun('type.one.one',bool)),X3))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s(fun(X1,fun('type.one.one',bool)),X5)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3512,f248])).
fof(f8396,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8348,f3509])).
fof(f8348,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,'type.one.one')),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,'type.one.one'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,'type.one.one')),X6)))) )),
  inference(unit_resulting_resolution,[],[f163,f570,f248])).
fof(f8395,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8349,f3509])).
fof(f8349,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun('type.one.one',bool)))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s(fun(X1,fun(X2,fun('type.one.one',bool))),X4))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun('type.one.one',bool)))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X5))),s(fun(X1,fun(X2,fun('type.one.one',bool))),X6)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3661,f248])).
fof(f8394,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8350,f3509])).
fof(f8350,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X5))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,'type.one.one')))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X6))),s(fun(X1,fun(X2,fun(X3,'type.one.one'))),X7)))) )),
  inference(unit_resulting_resolution,[],[f163,f590,f248])).
fof(f8393,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8351,f3509])).
fof(f8351,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X6))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one'))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X7))),s(fun(X1,fun(X2,fun(X3,fun(X4,'type.one.one')))),X8)))) )),
  inference(unit_resulting_resolution,[],[f163,f715,f248])).
fof(f8392,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8352,f3509])).
fof(f8352,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X7))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one')))))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X8))),s(fun(X1,fun(X2,fun(X3,fun(X4,fun(X5,'type.one.one'))))),X9)))) )),
  inference(unit_resulting_resolution,[],[f163,f743,f248])).
fof(f8391,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8353,f3509])).
fof(f8353,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',bool),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',bool)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',bool),X4)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f897,f248])).
fof(f8390,plain,(
  $false | ($spl8 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8354,f3509])).
fof(f8354,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',fun('type.one.one',bool))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s(fun('type.one.one',fun('type.one.one',bool)),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),fun('type.one.one',fun('type.one.one',bool))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s(fun('type.one.one',fun('type.one.one',bool)),X4)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f163,f3513,f248])).
fof(f8389,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8355,f3509])).
fof(f8355,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.infer_t.Infer_Tuvar_1'(s('type.num.num',X2))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X3),s('type.infer_t.infer_t',happ(s(fun('type.num.num','type.infer_t.infer_t'),'const.infer_t.Infer_Tuvar_0'),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f163,f166,f248])).
fof(f8388,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8356,f3509])).
fof(f8356,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.infer_t.infer_t','const.unify.t_walkstar_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2),s('type.infer_t.infer_t',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.infer_t.infer_t'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.infer_t.infer_t',happ(s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),happ(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s('type.infer_t.infer_t',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f177,f248])).
fof(f8387,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8357,f3509])).
fof(f8357,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X2))),s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X3,X1),X4),s('type.list.list'(X3),X5))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.list.list'(X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X6),s('type.list.list'(X1),happ(s(fun('type.list.list'(X3),'type.list.list'(X1)),happ(s(fun(fun(X3,X1),fun('type.list.list'(X3),'type.list.list'(X1))),'const.list.MAP_0'),s(fun(X3,X1),X4))),s('type.list.list'(X3),X5)))))) )),
  inference(unit_resulting_resolution,[],[f163,f181,f248])).
fof(f8386,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8358,f3509])).
fof(f8358,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.num.num'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X4),s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.arithmetic.+_1'(s('type.num.num',X2))),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f176,f248])).
fof(f8385,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8359,f3509])).
fof(f8359,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.pair.,_2'(s('type.infer.exc'(X1,X2),'const.infer.Success_1'(s(X1,X4))),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f178,f248])).
fof(f8384,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8360,f3509])).
fof(f8360,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),'const.infer.read_1'(s(X1,X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'(X1,X2),X1),happ(s(fun(X1,'type.pair.prod'('type.infer.exc'(X1,X2),X1)),'const.infer.read_0'),s(X1,X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f179,f248])).
fof(f8383,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8361,f3509])).
fof(f8361,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),'const.infer.fresh_uvar_1'(s('type.infer.infer_st'(X2),X4))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X5),s('type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2)),happ(s(fun('type.infer.infer_st'(X2),'type.pair.prod'('type.infer.exc'('type.infer_t.infer_t',X1),'type.infer.infer_st'(X2))),'const.infer.fresh_uvar_0'),s('type.infer.infer_st'(X2),X4)))))) )),
  inference(unit_resulting_resolution,[],[f163,f180,f248])).
fof(f8382,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8362,f3509])).
fof(f8362,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X4))),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),'const.infer.failwith_2'(s(X2,X5),s(X3,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'(X1,X2),X3),happ(s(fun(X3,'type.pair.prod'('type.infer.exc'(X1,X2),X3)),'const.infer.failwith_1'(s(X2,X5))),s(X3,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f205,f248])).
fof(f8381,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8363,f3509])).
fof(f8363,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),'const.infer.guard_3'(s(bool,X4),s(X1,X5),s(X2,X6))))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X7),s('type.pair.prod'('type.infer.exc'('type.one.one',X1),X2),happ(s(fun(X2,'type.pair.prod'('type.infer.exc'('type.one.one',X1),X2)),'const.infer.guard_2'(s(bool,X4),s(X1,X5))),s(X2,X6)))))) )),
  inference(unit_resulting_resolution,[],[f163,f184,f248])).
fof(f8380,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8364,f3509])).
fof(f8364,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one','const.one.one_0')))) )),
  inference(unit_resulting_resolution,[],[f163,f165,f248])).
fof(f8379,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f8365,f3509])).
fof(f8365,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.infer.guard_3'(s(bool,'T_0'),s(X0,X1),s('type.one.one',X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one',X0),'type.one.one'),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X0),'const.infer.Success_1'(s('type.one.one',X3))),s('type.one.one',X4)))) )),
  inference(unit_resulting_resolution,[],[f163,f253,f248])).
fof(f3498,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | $spl35),
  inference(cnf_transformation,[],[f3498_D])).
fof(f3498_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f3483,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,f(s('type.one.one','const.one.one_0'))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X5)))) ) | $spl14),
  inference(superposition,[],[f900,f165])).
fof(f3495,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) | $spl30),
  inference(cnf_transformation,[],[f3495_D])).
fof(f3495_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3482,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X3)))) ) | $spl14),
  inference(superposition,[],[f900,f253])).
fof(f3481,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X2))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X1)))) ) | $spl14),
  inference(superposition,[],[f900,f253])).
fof(f3480,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,g(s('type.one.one','const.one.one_0')))) ) | $spl14),
  inference(superposition,[],[f900,f165])).
fof(f3461,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl16),
  inference(cnf_transformation,[],[f3461_D])).
fof(f3461_D,plain,(
  ( ! [X26,X29,X25] : (s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3465,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) | $spl24),
  inference(cnf_transformation,[],[f3465_D])).
fof(f3465_D,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3423,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl8),
  inference(superposition,[],[f354,f3209])).
fof(f3209,plain,(
  ( ! [X269,X271,X268,X270] : (s(bool,happ(s(fun('type.one.one',bool),X270),s('type.one.one',X268))) = s(bool,happ(s(fun('type.one.one',bool),X271),s('type.one.one',X269)))) ) | $spl8),
  inference(superposition,[],[f2033,f253])).
fof(f2033,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,happ(s(fun('type.one.one',bool),X143),s('type.one.one',X144))) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun('type.one.one',X140),X142),s('type.one.one',X144)))))) ) | $spl8),
  inference(superposition,[],[f207,f897])).
fof(f3462,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) | $spl18),
  inference(cnf_transformation,[],[f3462_D])).
fof(f3462_D,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3466,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) | $spl26),
  inference(cnf_transformation,[],[f3466_D])).
fof(f3466_D,plain,(
  ( ! [X68,X67] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3422,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl8),
  inference(superposition,[],[f269,f3209])).
fof(f3421,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3420,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3419,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3418,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3464,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl22),
  inference(cnf_transformation,[],[f3464_D])).
fof(f3464_D,plain,(
  ( ! [X30,X26,X33,X31,X29,X25,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3416,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl8),
  inference(superposition,[],[f196,f3209])).
fof(f3463,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) | $spl20),
  inference(cnf_transformation,[],[f3463_D])).
fof(f3463_D,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3414,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl8),
  inference(superposition,[],[f167,f3209])).
fof(f3374,plain,(
  ( ! [X70,X72,X71,X69,X73] : (p(s(bool,happ(s(fun('type.one.one',bool),X71),s('type.one.one',X72)))) | s(bool,X73) = s(bool,happ(s(fun('type.one.one',bool),X69),s('type.one.one',X70))) | s(bool,'T_0') = s(bool,X73)) ) | $spl8),
  inference(superposition,[],[f354,f3209])).
fof(f3373,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,happ(s(fun('type.one.one',bool),X67),s('type.one.one',X68)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X65),s('type.one.one',X66)))) ) | $spl8),
  inference(superposition,[],[f269,f3209])).
fof(f3372,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X62),s('type.one.one',X63))) | s(bool,X64) = s(bool,happ(s(fun('type.one.one',bool),X60),s('type.one.one',X61))) | s(bool,'T_0') = s(bool,X64)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3371,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X59) = s(bool,happ(s(fun('type.one.one',bool),X57),s('type.one.one',X58))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X55),s('type.one.one',X56))) | s(bool,'T_0') = s(bool,X59)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3370,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,X54) = s(bool,happ(s(fun('type.one.one',bool),X52),s('type.one.one',X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X50),s('type.one.one',X51)))) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3369,plain,(
  ( ! [X47,X45,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,X49) = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46))) | s(bool,'T_0') = s(bool,X49)) ) | $spl8),
  inference(superposition,[],[f267,f3209])).
fof(f3367,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X34,X32] : (p(s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28)))) | s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.pair.,_2'(s('type.infer.exc'('type.one.one',X29),'const.infer.Success_1'(s('type.one.one',X33))),s(X30,X34))) != s('type.pair.prod'('type.infer.exc'('type.one.one',X29),X30),'const.infer.guard_3'(s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))),s(X29,X31),s(X30,X32)))) ) | $spl8),
  inference(superposition,[],[f196,f3209])).
fof(f3365,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X15),s('type.one.one',X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X13),s('type.one.one',X14)))) ) | $spl8),
  inference(superposition,[],[f167,f3209])).
fof(f3361,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl8),
  inference(superposition,[],[f3209,f167])).
fof(f3360,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3359,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3358,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3357,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3339,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,happ(s(fun('type.one.one',bool),X47),s('type.one.one',X48))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X45),s('type.one.one',X46)))) ) | $spl8),
  inference(superposition,[],[f3209,f167])).
fof(f3338,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.one.one',bool),X43),s('type.one.one',X44))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X40),s('type.one.one',X41)))) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3337,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X37) = s(bool,happ(s(fun('type.one.one',bool),X38),s('type.one.one',X39))) | s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X35),s('type.one.one',X36))) | s(bool,'T_0') = s(bool,X37)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3336,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X32),s('type.one.one',X33))) | s(bool,X34) = s(bool,happ(s(fun('type.one.one',bool),X30),s('type.one.one',X31))) | s(bool,'T_0') = s(bool,X34)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f3335,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.one.one',bool),X27),s('type.one.one',X28))) | s(bool,X29) = s(bool,happ(s(fun('type.one.one',bool),X25),s('type.one.one',X26))) | s(bool,'T_0') = s(bool,X29)) ) | $spl8),
  inference(superposition,[],[f3209,f267])).
fof(f897,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) | $spl8),
  inference(cnf_transformation,[],[f897_D])).
fof(f897_D,plain,(
  ( ! [X2,X1] : (s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f900,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) | $spl14),
  inference(cnf_transformation,[],[f900_D])).
fof(f900_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f885,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s('type.one.one',X16))) | s(fun('type.one.one',bool),X17) = s(fun('type.one.one',bool),X18) | s(bool,'T_0') = s(bool,g(s('type.one.one',X16)))) )),
  inference(superposition,[],[f558,f167])).
fof(f899,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) | $spl12),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  ( ! [X8,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f884,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X13) != s(bool,f(s('type.one.one',X12))) | s(fun('type.one.one',bool),X14) = s(fun('type.one.one',bool),X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,g(s('type.one.one',X12)))) )),
  inference(superposition,[],[f558,f267])).
fof(f883,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,f(s('type.one.one',X8))) | s(fun('type.one.one',bool),X10) = s(fun('type.one.one',bool),X11) | s(bool,'T_0') = s(bool,g(s('type.one.one',X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f558,f267])).
fof(f898,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) | $spl10),
  inference(cnf_transformation,[],[f898_D])).
fof(f898_D,plain,(
  ( ! [X0,X3] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,g(s('type.one.one',X0)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f882,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X4))) | s(fun('type.one.one',bool),X5) = s(fun('type.one.one',bool),X6) | s(bool,X7) = s(bool,g(s('type.one.one',X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f558,f267])).
fof(f881,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s('type.one.one',X0))) | s(fun('type.one.one',bool),X1) = s(fun('type.one.one',bool),X2) | s(bool,X3) = s(bool,g(s('type.one.one',X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f558,f267])).
fof(f273,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f275,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f272,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f167])).
fof(f241,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | $spl0),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f243,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) | $spl2),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f161,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3)))),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4)))) & (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f126])).
fof(f126,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4)))) & (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4)))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <~> s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X4))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X8),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) <=> s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X8))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),X8),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) <=> s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),X6),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),X8))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10))))),
  file('thm.inferProps.success_eqns|split|12','thm.inferProps.success_eqns|split|12')).
fof(f242,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) | $spl1),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f244,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) | $spl3),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f162,plain,(
  s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.st_ex_bind_3'(s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')))),'const.infer.read_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'_dst_x0x1_2'(s(fun('type.list.list'('type.infer_t.infer_t'),fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))))),'const.infer.st_ex_return_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.list.list'('type.infer_t.infer_t')),'_dst_x0x1x0_3'(s(fun(fun('type.infer_t.infer_t','type.infer_t.infer_t'),fun('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t'))),'const.list.MAP_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),fun('type.infer_t.infer_t','type.infer_t.infer_t')),'const.unify.t_walkstar_0'),s(fun('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_0'))),s('type.list.list'('type.infer_t.infer_t'),sK2))))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) | s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.apply_subst_list_2'(s('type.list.list'('type.infer_t.infer_t'),sK2),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK4))) != s('type.pair.prod'('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.list.list'('type.infer_t.infer_t'),sK0),'const.infer.Success_1'(s('type.list.list'('type.infer_t.infer_t'),sK1))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK3)))),
  inference(cnf_transformation,[],[f127])).
% SZS output end Proof for thm.inferProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 215
% Passive clauses: 4947
% Generated clauses: 110749
% Final active clauses: 106
% Final passive clauses: 750
% Input formulas: 45
% Initial clauses: 80
% Discarded non-redundant clauses: 13766
% 
% Duplicate literals: 613
% Trivial inequalities: 98
% Fw subsumption resolutions: 379
% Fw demodulations: 49252
% Bw demodulations: 348
% Condensations: 872
% Global subsumptions: 1
% 
% Simple tautologies: 31
% Equational tautologies: 49
% Forward subsumptions: 44047
% Fw demodulations to eq. taut.: 1329
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 38180
% Forward superposition: 9775
% Backward superposition: 10140
% Self superposition: 254
% Equality factoring: 58
% Equality resolution: 14
% 
% Split clauses: 63
% Split components: 127
% SAT solver clauses: 5973
% SAT solver unit clauses: 1229
% SAT solver binary clauses: 1699
% 
% Sat splits: 63
% Sat splitting refutations: 897
% 
% Memory used [KB]: 93900
% Time elapsed: 4.295 s
% ------------------------------
% ------------------------------
% Success in time 5.63 s
