Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exhSem
% (103458)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 177
% Passive clauses: 3473
% Generated clauses: 24426
% Final active clauses: 137
% Final passive clauses: 3037
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 942
% Trivial inequalities: 172
% Fw subsumption resolutions: 1831
% Fw demodulations: 3769
% Bw demodulations: 33
% 
% Simple tautologies: 217
% Equational tautologies: 51
% Forward subsumptions: 13280
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 427
% Factoring: 3
% Forward superposition: 9751
% Backward superposition: 6975
% Self superposition: 218
% Equality factoring: 104
% Equality resolution: 20
% 
% Split clauses: 233
% Split components: 466
% SAT solver clauses: 666
% SAT solver unit clauses: 41
% SAT solver binary clauses: 111
% 
% Sat splits: 233
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 666
% TWLsolver calls for satisfiability: 143
% 
% Memory used [KB]: 20468
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exhSem
% (103478)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 129
% Passive clauses: 1172
% Generated clauses: 15457
% Final active clauses: 103
% Final passive clauses: 880
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 439
% Trivial inequalities: 25
% Fw subsumption resolutions: 530
% Fw demodulations: 1544
% Bw demodulations: 183
% Condensations: 922
% 
% Simple tautologies: 55
% Equational tautologies: 40
% Forward subsumptions: 8584
% 
% Binary resolution: 262
% Factoring: 1
% Forward superposition: 5374
% Backward superposition: 5346
% Self superposition: 187
% Equality factoring: 78
% Equality resolution: 1
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 1863
% SAT solver unit clauses: 130
% SAT solver binary clauses: 37
% 
% Sat splits: 33
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1863
% TWLsolver calls for satisfiability: 247
% 
% Memory used [KB]: 23539
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exhSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exhSem
% SZS output start Proof for thm.exhSem
fof(f27885,plain,(
  $false),
  inference(subsumption_resolution,[],[f27884,f214])).
fof(f214,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exhSem.state_fupdfupds_comp|split|3',aHL_TRUTH)).
fof(f27884,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f27882,f2819])).
fof(f2819,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2489,f296])).
fof(f296,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f216,f217])).
fof(f217,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.exhSem.state_fupdfupds_comp|split|3',aHL_BOOL_CASES)).
fof(f216,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.exhSem.state_fupdfupds_comp|split|3',aHL_FALSITY)).
fof(f2489,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f2404,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X14,X6] : ! [X17,X20] : s(X6,X17) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X14),'const.pair.,_2'(s(X6,X17),s(X14,X20)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|3','thm.pair.FST')).
fof(f2404,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f260,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.exhSem.state_fupdfupds_comp|split|3',arity_0_HL_EQ)).
fof(f27882,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.exhSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.exhSem.state'(sK1)),fun(fun(sK0,'type.exhSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))))),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f27817,f7577])).
fof(f7577,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.exhSem.state'(sK1)),bool),happ(s(fun(fun(sK0,'type.exhSem.state'(sK1)),fun(fun(sK0,'type.exhSem.state'(sK1)),bool)),'=_0'),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4))),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))))))),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f213,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f213,plain,(
  s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))) != s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4))),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2))) != s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK4))),s(fun(sK0,'type.exhSem.state'(sK1)),'const.combin.o_2'(s(fun('type.exhSem.state'(sK1),'type.exhSem.state'(sK1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),sK3))),s(fun(sK0,'type.exhSem.state'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))))),s(fun(X0,'type.exhSem.state'(X1)),X2))) != s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4))),s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))),s(fun(X0,'type.exhSem.state'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))))),s(fun(X0,'type.exhSem.state'(X1)),X2))) = s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4))),s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))),s(fun(X0,'type.exhSem.state'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))))),s(fun(X0,'type.exhSem.state'(X1)),X2))) = s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X4))),s(fun(X0,'type.exhSem.state'(X1)),'const.combin.o_2'(s(fun('type.exhSem.state'(X1),'type.exhSem.state'(X1)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X3))),s(fun(X0,'type.exhSem.state'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9))))),s(fun(X6,'type.exhSem.state'(X7)),X8))) = s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10))),s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9))),s(fun(X6,'type.exhSem.state'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9))))),s(fun(X6,'type.exhSem.state'(X7)),X8))) = s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10))),s(fun(X6,'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9))),s(fun(X6,'type.exhSem.state'(X7)),X8)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|3','thm.exhSem.state_fupdfupds_comp|split|3')).
fof(f27817,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2))))),s(fun(X3,'type.exhSem.state'(X0)),X4))) = s(fun(X3,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1))),s(fun(X3,'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2))),s(fun(X3,'type.exhSem.state'(X0)),X4)))))) )),
  inference(superposition,[],[f273,f229])).
fof(f229,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1))))) = s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2))),s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1))))) = s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2))),s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1))))) = s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.combin.o_2'(s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X2))),s(fun('type.exhSem.state'(X0),'type.exhSem.state'(X0)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),'const.combin.o_2'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10),s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9))))) = s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.combin.o_2'(s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X10))),s(fun('type.exhSem.state'(X7),'type.exhSem.state'(X7)),'const.exhSem.state_refs_fupd_1'(s(fun('type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v')),'type.list.list'('type.semanticPrimitives.store_v'('type.exhSem.v'))),X9)))))),
  file('thm.exhSem.state_fupdfupds_comp|split|3','thm.exhSem.state_fupdfupds_comp|split|2')).
fof(f273,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X14,X6,X22,X15] : ! [X10,X9,X8] : s(fun(X22,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X22,X6),'const.combin.o_2'(s(fun(X15,X6),X9),s(fun(X22,X15),X8))))) = s(fun(X22,X14),'const.combin.o_2'(s(fun(X15,X14),'const.combin.o_2'(s(fun(X6,X14),X10),s(fun(X15,X6),X9))),s(fun(X22,X15),X8)))),
  file('thm.exhSem.state_fupdfupds_comp|split|3','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.exhSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 185
% Passive clauses: 2381
% Generated clauses: 28404
% Final active clauses: 113
% Final passive clauses: 1215
% Input formulas: 68
% Initial clauses: 77
% 
% Duplicate literals: 856
% Trivial inequalities: 52
% Fw subsumption resolutions: 668
% Fw demodulations: 4228
% Bw demodulations: 881
% Condensations: 1104
% Global subsumptions: 2
% 
% Simple tautologies: 67
% Equational tautologies: 130
% Forward subsumptions: 18790
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 44
% 
% Unit resulting resolution: 776
% Factoring: 1
% Forward superposition: 9310
% Backward superposition: 9152
% Self superposition: 288
% Equality factoring: 116
% Equality resolution: 2
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 2473
% SAT solver unit clauses: 538
% SAT solver binary clauses: 1170
% 
% Sat splits: 26
% Sat splitting refutations: 36
% 
% Memory used [KB]: 37611
% Time elapsed: 1.523 s
% ------------------------------
% ------------------------------
% Success in time 2.9 s
