Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.patSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.patSem
% SZS output start Proof for thm.patSem
fof(f5404,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5395,f3750,f125,f188,f126,f190,f127,f192,f128,f194,f186,f3853,f129,f195,f193,f191,f189,f187,f3816,f3814])).
fof(f3814,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f3813,f195])).
fof(f3813,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) = s('type.list.list'('type.option.option'('type.patSem.v')),sK8) | ($spl0 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f3789,f168])).
fof(f168,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),'const.patSem.state_globals_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),'const.patSem.state_globals_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),'const.patSem.state_globals_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X32,X39,X31,X40] : s('type.list.list'('type.option.option'('type.patSem.v')),X40) = s('type.list.list'('type.option.option'('type.patSem.v')),'const.patSem.state_globals_1'(s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X39),s('type.ffi.ffi_state'(X6),X31),s('type.list.list'('type.option.option'('type.patSem.v')),X40)))))),
  file('thm.patSem.state_|49||49|','thm.patSem.state_globals')).
fof(f3789,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK8) = s('type.list.list'('type.option.option'('type.patSem.v')),'const.patSem.state_globals_1'(s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))))) | ($spl0 | $spl2 | $spl6)),
  inference(superposition,[],[f168,f3788])).
fof(f3788,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | ($spl0 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f3751,f188])).
fof(f3751,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f186,f192])).
fof(f3816,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3815,f191])).
fof(f3815,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | ($spl0 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f3790,f167])).
fof(f167,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'const.patSem.state_refs_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'const.patSem.state_refs_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'const.patSem.state_refs_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X32,X39,X31,X40] : s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X39) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'const.patSem.state_refs_1'(s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X39),s('type.ffi.ffi_state'(X6),X31),s('type.list.list'('type.option.option'('type.patSem.v')),X40)))))),
  file('thm.patSem.state_|49||49|','thm.patSem.state_refs')).
fof(f3790,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),'const.patSem.state_refs_1'(s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))))) | ($spl0 | $spl2 | $spl6)),
  inference(superposition,[],[f167,f3788])).
fof(f187,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | $spl1),
  inference(cnf_transformation,[],[f187_D])).
fof(f187_D,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f189,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) | $spl3),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f191,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | $spl5),
  inference(cnf_transformation,[],[f191_D])).
fof(f191_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f193,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | $spl7),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f195,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) != s('type.list.list'('type.option.option'('type.patSem.v')),sK8) | $spl9),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) != s('type.list.list'('type.option.option'('type.patSem.v')),sK8) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f129,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) != s('type.list.list'('type.option.option'('type.patSem.v')),sK8) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8)))),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  (s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | (s('type.num.num',sK1) = s('type.num.num',sK5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) & s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) & s('type.list.list'('type.option.option'('type.patSem.v')),sK4) = s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) & (s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | s('type.num.num',sK1) != s('type.num.num',sK5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | s('type.ffi.ffi_state'(sK0),sK3) != s('type.ffi.ffi_state'(sK0),sK7) | s('type.list.list'('type.option.option'('type.patSem.v')),sK4) != s('type.list.list'('type.option.option'('type.patSem.v')),sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) = s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),X8))) & (s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) != s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) | s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X7) | s('type.list.list'('type.option.option'('type.patSem.v')),X4) != s('type.list.list'('type.option.option'('type.patSem.v')),X8)))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) = s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) | (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),X8))) & (s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) != s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) | (s('type.num.num',X1) != s('type.num.num',X5) | s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) != s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) | s('type.ffi.ffi_state'(X0),X3) != s('type.ffi.ffi_state'(X0),X7) | s('type.list.list'('type.option.option'('type.patSem.v')),X4) != s('type.list.list'('type.option.option'('type.patSem.v')),X8))))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) = s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) <~> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),X8)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) = s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),X8)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4))) = s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6),s('type.ffi.ffi_state'(X0),X7),s('type.list.list'('type.option.option'('type.patSem.v')),X8))) <=> (s('type.num.num',X1) = s('type.num.num',X5) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X6) & s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),X7) & s('type.list.list'('type.option.option'('type.patSem.v')),X4) = s('type.list.list'('type.option.option'('type.patSem.v')),X8)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.list.list'('type.option.option'('type.patSem.v')),X10))) = s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X12),s('type.ffi.ffi_state'(X6),X13),s('type.list.list'('type.option.option'('type.patSem.v')),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X12) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X13) & s('type.list.list'('type.option.option'('type.patSem.v')),X10) = s('type.list.list'('type.option.option'('type.patSem.v')),X14)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14] : (s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X7),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X8),s('type.ffi.ffi_state'(X6),X9),s('type.list.list'('type.option.option'('type.patSem.v')),X10))) = s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X11),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X12),s('type.ffi.ffi_state'(X6),X13),s('type.list.list'('type.option.option'('type.patSem.v')),X14))) <=> (s('type.num.num',X7) = s('type.num.num',X11) & s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X8) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X12) & s('type.ffi.ffi_state'(X6),X9) = s('type.ffi.ffi_state'(X6),X13) & s('type.list.list'('type.option.option'('type.patSem.v')),X10) = s('type.list.list'('type.option.option'('type.patSem.v')),X14)))),
  file('thm.patSem.state_|49||49|','thm.patSem.state_|49||49|')).
fof(f3853,plain,(
  $false | ($spl0 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f3852,f193])).
fof(f3852,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3828,f166])).
fof(f166,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.patSem.state_ffi_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.patSem.state_ffi_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.ffi.ffi_state'(X0),X3) = s('type.ffi.ffi_state'(X0),'const.patSem.state_ffi_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X32,X39,X31,X40] : s('type.ffi.ffi_state'(X6),X31) = s('type.ffi.ffi_state'(X6),'const.patSem.state_ffi_1'(s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X39),s('type.ffi.ffi_state'(X6),X31),s('type.list.list'('type.option.option'('type.patSem.v')),X40)))))),
  file('thm.patSem.state_|49||49|','thm.patSem.state_ffi')).
fof(f3828,plain,(
  s('type.ffi.ffi_state'(sK0),sK7) = s('type.ffi.ffi_state'(sK0),'const.patSem.state_ffi_1'(s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))))) | ($spl0 | $spl2)),
  inference(superposition,[],[f166,f3825])).
fof(f3825,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f186,f188])).
fof(f186,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) | $spl0),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f194,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) = s('type.list.list'('type.option.option'('type.patSem.v')),sK8) | $spl8),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) = s('type.list.list'('type.option.option'('type.patSem.v')),sK8) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f128,plain,(
  s('type.list.list'('type.option.option'('type.patSem.v')),sK4) = s('type.list.list'('type.option.option'('type.patSem.v')),sK8) | s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8)))),
  inference(cnf_transformation,[],[f108])).
fof(f192,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | $spl6),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f127,plain,(
  s('type.ffi.ffi_state'(sK0),sK3) = s('type.ffi.ffi_state'(sK0),sK7) | s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8)))),
  inference(cnf_transformation,[],[f108])).
fof(f190,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | $spl4),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f126,plain,(
  s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2) = s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6) | s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8)))),
  inference(cnf_transformation,[],[f108])).
fof(f188,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | $spl2),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f125,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) = s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK8)))),
  inference(cnf_transformation,[],[f108])).
fof(f3750,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(trivial_inequality_removal,[],[f3749])).
fof(f3749,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3748,f188])).
fof(f3748,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3747,f190])).
fof(f3747,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) | ($spl1 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f3746,f192])).
fof(f3746,plain,(
  s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) != s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK5),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK6),s('type.ffi.ffi_state'(sK0),sK7),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))) | ($spl1 | $spl8)),
  inference(forward_demodulation,[],[f187,f194])).
fof(f5395,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f5394,f189])).
fof(f5394,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK5) | $spl0),
  inference(forward_demodulation,[],[f5369,f165])).
fof(f165,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num',X1) = s('type.num.num','const.patSem.state_clock_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.patSem.state_clock_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.num.num',X1) = s('type.num.num','const.patSem.state_clock_1'(s('type.patSem.state'(X0),'const.patSem.recordtype|46|state_4'(s('type.num.num',X1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X2),s('type.ffi.ffi_state'(X0),X3),s('type.list.list'('type.option.option'('type.patSem.v')),X4)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X32,X39,X31,X40] : s('type.num.num',X32) = s('type.num.num','const.patSem.state_clock_1'(s('type.patSem.state'(X6),'const.patSem.recordtype|46|state_4'(s('type.num.num',X32),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),X39),s('type.ffi.ffi_state'(X6),X31),s('type.list.list'('type.option.option'('type.patSem.v')),X40)))))),
  file('thm.patSem.state_|49||49|','thm.patSem.state_clock')).
fof(f5369,plain,(
  s('type.num.num',sK5) = s('type.num.num','const.patSem.state_clock_1'(s('type.patSem.state'(sK0),'const.patSem.recordtype|46|state_4'(s('type.num.num',sK1),s('type.list.list'('type.semanticPrimitives.store_v'('type.patSem.v')),sK2),s('type.ffi.ffi_state'(sK0),sK3),s('type.list.list'('type.option.option'('type.patSem.v')),sK4))))) | $spl0),
  inference(superposition,[],[f165,f186])).
% SZS output end Proof for thm.patSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 1155
% Generated clauses: 5171
% Final active clauses: 70
% Final passive clauses: 953
% Input formulas: 37
% Initial clauses: 60
% 
% Duplicate literals: 253
% Trivial inequalities: 73
% Fw subsumption resolutions: 620
% Fw demodulations: 287
% 
% Simple tautologies: 134
% Equational tautologies: 44
% Forward subsumptions: 2429
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 171
% Factoring: 1
% Forward superposition: 1606
% Backward superposition: 1798
% Self superposition: 129
% Equality factoring: 96
% Equality resolution: 14
% 
% Split clauses: 137
% Split components: 281
% SAT solver clauses: 278
% SAT solver unit clauses: 22
% SAT solver binary clauses: 92
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 137
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 274
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 4093
% Time elapsed: 0.184 s
% ------------------------------
% ------------------------------
% Success in time 0.185 s
