Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 148 next slice time: 44
dis+11_7_41 on thm.clos_annotateProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_annotateProof
% SZS output start Proof for thm.clos_annotateProof
fof(f165791,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f165763,f1797,f2983,f2981,f6123,f1798,f2984,f2982,f1815,f2986,f1951,f2987,f2985])).
fof(f2985,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(cnf_transformation,[],[f2985_D])).
fof(f2985_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2987,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f2987_D])).
fof(f2987_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1951,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1502])).
fof(f1502,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1501])).
fof(f1501,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1344])).
fof(f1344,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.bool.IMP_CLAUSES|split|2')).
fof(f2986,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f2986_D])).
fof(f2986_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1815,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1',aHL_FALSITY)).
fof(f2982,plain,(
  ~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) | $spl1),
  inference(cnf_transformation,[],[f2982_D])).
fof(f2982_D,plain,(
  ~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2984,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK56) | $spl3),
  inference(cnf_transformation,[],[f2984_D])).
fof(f2984_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK56) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1798,plain,(
  s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK56) | ~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56))))),
  inference(cnf_transformation,[],[f1455])).
fof(f1455,plain,(
  (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56)) & (~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK56))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK56])],[f1454])).
fof(f1454,plain,(
  ? [X0] : ((p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0)) & (~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0)))),
  inference(nnf_transformation,[],[f1337])).
fof(f1337,plain,(
  ? [X0] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) <~> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X0)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X6)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X6)))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X6))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.clos_annotateProof.v_rel_Unit|split|1')).
fof(f6123,plain,(
  $false | ($spl1 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6122,f2986])).
fof(f6122,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f6121,f3037])).
fof(f3037,plain,(
  s(bool,'F_0') = s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0'))) | ($spl1 | $spl2)),
  inference(resolution,[],[f1895,f2996])).
fof(f2996,plain,(
  ~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2982,f2983])).
fof(f1895,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1466])).
fof(f1466,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f593])).
fof(f593,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f185])).
fof(f185,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f6121,plain,(
  p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | $spl2),
  inference(subsumption_resolution,[],[f6110,f2983])).
fof(f6110,plain,(
  p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',sK56) | $spl2),
  inference(superposition,[],[f1903,f2983])).
fof(f1903,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0)) )),
  inference(cnf_transformation,[],[f1470])).
fof(f1470,plain,(
  ! [X0] : ((~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0)) & (s('type.closSem.v','const.closSem.Unit_0') != s('type.closSem.v',X0) | p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0'))))))),
  inference(nnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Unit_0')))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v',X6),s('type.closSem.v','const.closSem.Unit_0')))) <=> s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X6))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.clos_annotateProof.v_rel_Unit|split|0')).
fof(f2981,plain,(
  p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) | $spl0),
  inference(cnf_transformation,[],[f2981_D])).
fof(f2981_D,plain,(
  p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2983,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56) | $spl2),
  inference(cnf_transformation,[],[f2983_D])).
fof(f2983_D,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1797,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56) | p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56))))),
  inference(cnf_transformation,[],[f1455])).
fof(f165763,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f165762,f2984])).
fof(f165762,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f165710,f2987])).
fof(f165710,plain,(
  ~p(s(bool,'T_0')) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',sK56) | $spl0),
  inference(superposition,[],[f94148,f6138])).
fof(f6138,plain,(
  s(bool,'T_0') = s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',sK56))) | $spl0),
  inference(resolution,[],[f2981,f1891])).
fof(f1891,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1464])).
fof(f1464,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f94148,plain,(
  ( ! [X3] : (~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X3)))) | s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v',X3)) )),
  inference(forward_demodulation,[],[f94075,f94035])).
fof(f94035,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0')))),
  inference(forward_demodulation,[],[f93965,f3291])).
fof(f3291,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(forward_demodulation,[],[f1808,f1802])).
fof(f1802,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.dec_to_exh.tuple_tag_0')),
  inference(cnf_transformation,[],[f111])).
fof(f111,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.dec_to_exh.tuple_tag_0')),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.dec_to_exh.tuple_tag_def')).
fof(f1808,plain,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  s('type.closSem.v','const.closSem.Unit_0') = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.dec_to_exh.tuple_tag_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.closSem.Unit_def')).
fof(f93965,plain,(
  s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(superposition,[],[f89791,f1987])).
fof(f1987,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f665])).
fof(f665,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f664])).
fof(f664,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f159])).
fof(f159,axiom,(
  ! [X49] : ! [X127,X128] : s(X49,X127) = s(X49,'const.bool.COND_3'(s(bool,'T_0'),s(X49,X127),s(X49,X128)))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f89791,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f1869,f1801])).
fof(f1801,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.conLang.true_tag_0')),
  inference(cnf_transformation,[],[f157])).
fof(f157,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.conLang.true_tag_0')),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.conLang.true_tag_def')).
fof(f1869,plain,(
  ( ! [X0] : (s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0] : s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  inference(rectify,[],[f161])).
fof(f161,axiom,(
  ! [X41] : s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X41))) = s('type.closSem.v','const.closSem.Block_2'(s('type.num.num','const.bool.COND_3'(s(bool,X41),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.closSem.v'),'const.list.NIL_0')))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.closSem.Boolv_def')).
fof(f94075,plain,(
  ( ! [X3] : (~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Unit_0'),s('type.closSem.v',X3)))) | s('type.closSem.v',X3) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0')))) )),
  inference(superposition,[],[f2547,f94035])).
fof(f2547,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1)))) )),
  inference(cnf_transformation,[],[f1696])).
fof(f1696,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v',X0)))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v',X0))))))),
  inference(nnf_transformation,[],[f881])).
fof(f881,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v',X0)))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X41] : (p(s(bool,'const.clos_annotateProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X41))),s('type.closSem.v',X6)))) <=> s('type.closSem.v',X6) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X41))))),
  file('thm.clos_annotateProof.v_rel_Unit|split|1','thm.clos_annotateProof.v_rel_Boolv|split|1')).
% SZS output end Proof for thm.clos_annotateProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1183
% Passive clauses: 18232
% Generated clauses: 164111
% Final active clauses: 869
% Final passive clauses: 15103
% Input formulas: 517
% Initial clauses: 1176
% 
% Duplicate literals: 492
% Trivial inequalities: 2478
% Fw subsumption resolutions: 9981
% Fw demodulations: 34759
% Bw demodulations: 243
% 
% Simple tautologies: 1016
% Equational tautologies: 46
% Forward subsumptions: 96879
% Fw demodulations to eq. taut.: 38
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 2996
% Factoring: 1
% Forward superposition: 48070
% Backward superposition: 63186
% Self superposition: 201
% Equality factoring: 40
% Equality resolution: 69
% 
% Split clauses: 311
% Split components: 624
% SAT solver clauses: 1510
% SAT solver unit clauses: 17
% SAT solver binary clauses: 496
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 5
% 
% Sat splits: 311
% Sat splitting refutations: 70
% 
% TWLsolver clauses: 1505
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 91085
% Time elapsed: 2.913 s
% ------------------------------
% ------------------------------
% Success in time 3.06 s
