Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.infer_eComplete
% (111725)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 634
% Passive clauses: 8255
% Generated clauses: 17441
% Final active clauses: 632
% Final passive clauses: 7273
% Input formulas: 133
% Initial clauses: 359
% 
% Duplicate literals: 221
% Trivial inequalities: 39
% Fw subsumption resolutions: 755
% Fw demodulations: 517
% 
% Simple tautologies: 202
% Equational tautologies: 50
% Forward subsumptions: 6617
% 
% Binary resolution: 9007
% Factoring: 1132
% Forward superposition: 2887
% Backward superposition: 2354
% Self superposition: 64
% Equality factoring: 2
% Equality resolution: 29
% 
% Split clauses: 176
% Split components: 355
% SAT solver clauses: 372
% SAT solver unit clauses: 12
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 176
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 371
% TWLsolver calls for satisfiability: 299
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.infer_eComplete
% (111747)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 410
% Passive clauses: 2951
% Generated clauses: 13968
% Final active clauses: 369
% Final passive clauses: 2075
% Input formulas: 133
% Initial clauses: 359
% 
% Duplicate literals: 103
% Trivial inequalities: 52
% Fw subsumption resolutions: 718
% Fw demodulations: 1063
% Bw demodulations: 503
% Condensations: 124
% Global subsumptions: 29
% 
% Simple tautologies: 203
% Equational tautologies: 14
% Forward subsumptions: 5514
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2023
% Factoring: 98
% Forward superposition: 3290
% Backward superposition: 2531
% Self superposition: 68
% Equality factoring: 4
% Equality resolution: 22
% 
% Split clauses: 25
% Split components: 51
% SAT solver clauses: 3399
% SAT solver unit clauses: 433
% SAT solver binary clauses: 159
% SAT solver learnt clauses: 59
% SAT solver learnt literals: 31
% 
% Sat splits: 25
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3282
% TWLsolver calls for satisfiability: 3021
% 
% Memory used [KB]: 20980
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.infer_eComplete
% (111769)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1212
% Passive clauses: 11713
% Generated clauses: 163965
% Final active clauses: 1030
% Final passive clauses: 2444
% Input formulas: 133
% Initial clauses: 359
% Discarded non-redundant clauses: 50331
% 
% Duplicate literals: 134
% Trivial inequalities: 54
% Fw subsumption resolutions: 994
% Fw demodulations: 53550
% Bw demodulations: 1436
% Condensations: 819
% Global subsumptions: 281
% 
% Simple tautologies: 99
% Equational tautologies: 110
% Forward subsumptions: 49318
% Fw demodulations to eq. taut.: 14889
% Bw demodulations to eq. taut.: 70
% 
% Unit resulting resolution: 82709
% Factoring: 1
% Forward superposition: 12741
% Backward superposition: 10254
% Self superposition: 149
% Equality factoring: 28
% Equality resolution: 31
% 
% Split clauses: 69
% Split components: 139
% SAT solver clauses: 12437
% SAT solver unit clauses: 2690
% SAT solver binary clauses: 5868
% 
% Sat splits: 69
% Sat splitting refutations: 303
% 
% Memory used [KB]: 172321
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.infer_eComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.infer_eComplete
% SZS output start Proof for thm.infer_eComplete
fof(f44746,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f698,f701,f699,f433,f432,f431,f430,f700,f429,f703,f428,f705,f427,f707,f426,f711,f709,f819,f890,f892,f893,f895,f897,f988,f1066,f1070,f1072,f1073,f1074,f1077,f1168,f1233,f1240,f1242,f1243,f1245,f1247,f1338,f1381,f1482,f1486,f1488,f1489,f1490,f1493,f1584,f1649,f1656,f1658,f1659,f1661,f1663,f1754,f2049,f2055,f2057,f2056,f2051,f2059,f2058,f2094,f2095,f2096,f2178,f2180,f2181,f2183,f2185,f2276,f2337,f2467,f2469,f2470,f2471,f2474,f2565,f2634,f3281,f3810,f3812,f3358,f3359,f3360,f3821,f3822,f3824,f3825,f3856,f3858,f3859,f3860,f3863,f3954,f4067,f4069,f4070,f4072,f4074,f4165,f4210,f4341,f4343,f4344,f4346,f4348,f4439,f4484,f4773,f4774,f4910,f4907,f4906,f4889,f4888,f4887,f5007,f5008,f5009,f5010,f5011,f5012,f5050,f5187,f5183,f5164,f5276,f5277,f5278,f5660,f5665,f5667,f5715,f5720,f5724,f2463,f5726,f2460,f5728,f2459,f5730,f5734,f5736,f5737,f5738,f5741,f5832,f6228,f6234,f6229,f6230,f6402,f6440,f6403,f6443,f6404,f6442,f7778,f8404,f8406,f8407,f8409,f8411,f8502,f8547,f9807,f9809,f9810,f9811,f9808,f10478,f10461,f10460,f10459,f10580,f10581,f10582,f10583,f10584,f10585,f10668,f6235,f10669,f10691,f10670,f10690,f10835,f10853,f1230,f10838,f6236,f10995,f10996,f10997,f10998,f1229,f11002,f11004,f11005,f11006,f11009,f11100,f3997,f13179,f13162,f13161,f13160,f13282,f13283,f13284,f13285,f13286,f13287,f13582,f13583,f13621,f13622,f13810,f13813,f13815,f13879,f13880,f13881,f13882,f13901,f13907,f13908,f14175,f14177,f14178,f14179,f14176,f14937,f15392,f3815,f15563,f15574,f15567,f15576,f17894,f17896,f17897,f17898,f17901,f17992,f9591,f1648,f18052,f18054,f18055,f18057,f18059,f18150,f18204,f18233,f18209,f18210,f5723,f18211,f5674,f18212,f18273,f18297,f18319,f18298,f1797,f18304,f18322,f18324,f18326,f18327,f18329,f18331,f18422,f18477,f18724,f18721,f18699,f18815,f18816,f18817,f20601,f20577,f20693,f20694,f20695,f25233,f25237,f25240,f25244,f25246,f25361,f1034,f30298,f30300,f30301,f30302,f30305,f30396,f30457,f30540,f30546,f30561,f1232,f30723,f30565,f30567,f30568,f30569,f30572,f30663,f30724,f30729,f30731,f30732,f30734,f30736,f30827,f30886,f30891,f30892,f30893,f30896,f30897,f30898,f30936,f30899,f30900,f30901,f30997,f30998,f31031,f30874,f31032,f31062,f31034,f31035,f31039,f32692,f32693,f32695,f32696,f32739,f32740,f32754,f32755,f32756,f32757,f32769,f32772,f32790,f32791,f32792,f32793,f32812,f32770,f30294,f32839,f1646,f20459,f32841,f1645,f20433,f3817,f32893,f32896,f32897,f32898,f32929,f32899,f32900,f32901,f32930,f32904,f32905,f32906,f32967,f32969,f32970,f32971,f32974,f33065,f18047,f33159,f18294,f33160,f33181,f3816,f33186,f33187,f32931,f33190,f33313,f33315,f33316,f33317,f33320,f33411,f5875,f33575,f33578,f33579,f33580,f33636,f33581,f33582,f33583,f33639,f33586,f33587,f33640,f33588,f33641,f33638,f33676,f33678,f33679,f33680,f33677,f33683,f33685,f33686,f33687,f33684,f27482,f27478,f27477,f27475,f27047,f33689,f33691,f33692,f27032,f27031,f33693,f33694,f33695,f33696,f25411,f25166,f33697,f25165,f33698,f33699,f33700,f33708,f32867,f14374,f33709,f33710,f33711,f33712,f33713,f33715,f33716,f33717,f33718,f33719,f14147,f33720,f33721,f33722,f33723,f32962,f13983,f13700,f3820,f33875,f14167,f30934,f14166,f14165,f3282,f33952,f33954,f34812,f34859,f30935,f34868,f30933,f34892,f33635,f34894,f5722,f37779,f13811,f37780,f33642,f37781,f13816,f37782,f13812,f13814,f37816,f13909,f37811,f31064,f31025,f37818,f30931,f37847,f31026,f38022,f38154,f38155,f38156,f38157,f38159,f38158,f38142,f38165,f38143,f38166,f38144,f38145,f38167,f38153,f38200,f38626,f31063,f40908,f30930,f42955,f42957,f42958,f42960,f42962,f43053,f43114,f43278,f43280,f43281,f43283,f43285,f43376,f43437,f43589,f43591,f43592,f43594,f43596,f43687,f43732,f43744,f43746,f43747,f43749,f43751,f43842,f43903,f43933,f3823,f43934,f43964,f43936,f43963,f44102,f44104,f44105,f1033,f44107,f1036,f44109,f1035,f44200,f44261,f1037,f44661,f44705,f44707,f44743,f44744,f44745])).
fof(f44745,plain,(
  $false | ($spl2 | $spl44)),
  inference(trivial_inequality_removal,[],[f44729])).
fof(f44729,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ($spl2 | $spl44)),
  inference(superposition,[],[f44700,f700])).
fof(f44700,plain,(
  ( ! [X48] : (s('type.num.num',sK12) != s('type.num.num',X48)) ) | $spl44),
  inference(superposition,[],[f2058,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X2) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : s('type.num.num',X2) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num',X2) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X51] : ! [X9,X50] : s('type.num.num',X50) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X51),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X50))),s('type.infer.infer_st'(X51),X9)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.infer_st_rewrs|split|4')).
fof(f44744,plain,(
  $false | ($spl44 | $spl208)),
  inference(trivial_inequality_removal,[],[f44730])).
fof(f44730,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ($spl44 | $spl208)),
  inference(superposition,[],[f44700,f32867])).
fof(f44743,plain,(
  $false | $spl44),
  inference(equality_resolution,[],[f44700])).
fof(f44707,plain,(
  $false | ($spl2 | $spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44706,f2058])).
fof(f44706,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f707,f700])).
fof(f44705,plain,(
  $false | ($spl4 | $spl44)),
  inference(trivial_inequality_removal,[],[f44694])).
fof(f44694,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ($spl4 | $spl44)),
  inference(superposition,[],[f2058,f703])).
fof(f44661,plain,(
  $false | ($spl12 | $spl46)),
  inference(subsumption_resolution,[],[f44498,f2461])).
fof(f2461,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) | $spl12),
  inference(resolution,[],[f711,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(X3,X2,X1,X0))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(X3,X2,X1,X0))))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(X3,X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & (! [X4] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ? [X5] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & (! [X4] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),
  inference(rectify,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ? [X4] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))))) & (! [X4] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) <=> ? [X4] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X34,X13,X42,X14] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X13),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X14))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X42)))) <=> ? [X36] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X13),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X36)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X36),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X14),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X42))))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.pure_add_constraints_append')).
fof(f44498,plain,(
  ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) | ($spl12 | $spl46)),
  inference(resolution,[],[f2059,f2462])).
fof(f2462,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl12),
  inference(resolution,[],[f711,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(X3,X2,X1,X0)),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f1037,plain,(
  ( ! [X23] : (p(s(bool,X23))) ) | $spl22),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  ( ! [X23] : (p(s(bool,X23))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f44261,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl308),
  inference(cnf_transformation,[],[f44261_D])).
fof(f44261_D,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f44200,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X270)))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X272) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X270),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f1612,f691])).
fof(f691,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X32,X6,X33] : ((s(bool,X7) = s(bool,X32) & (p(s(bool,X32)) => s(bool,X6) = s(bool,X33))) => ((p(s(bool,X7)) => p(s(bool,X6))) <=> (p(s(bool,X32)) => p(s(bool,X33)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CONG')).
fof(f1612,plain,(
  ( ! [X1] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) ) | $spl10),
  inference(superposition,[],[f709,f679])).
fof(f679,plain,(
  ( ! [X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0) = s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))) )),
  inference(cnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0) = s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)) & (s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0) != s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) <=> s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0) = s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X21,X20] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X21)))) <=> s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20) = s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X21))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.infer.pure_add_constraints_def|split|0')).
fof(f1035,plain,(
  ( ! [X13,X11] : (~p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) | $spl18),
  inference(cnf_transformation,[],[f1035_D])).
fof(f1035_D,plain,(
  ( ! [X13,X11] : (~p(s(bool,X11)) | s(bool,X11) != s(bool,X13)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f44109,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34)))) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X35) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f1612,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f1036,plain,(
  ( ! [X25] : (~p(s(bool,X25))) ) | $spl20),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  ( ! [X25] : (~p(s(bool,X25))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f44107,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26)))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f1612,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f1033,plain,(
  ( ! [X5,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X5)) ) | $spl14),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  ( ! [X5,X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,X5)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f44105,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18)))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f1612,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f44104,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X14)))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X14),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X15) != s(bool,X17)) ) | $spl10),
  inference(resolution,[],[f1612,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f44102,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6)))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f1612,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f308])).
fof(f43963,plain,(
  ( ! [X45,X44] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X44,'type.num.num'),X45,'type.num.num')) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X45),s('type.num.num',sK12)))) ) | $spl304),
  inference(cnf_transformation,[],[f43963_D])).
fof(f43963_D,plain,(
  ( ! [X45,X44] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X44,'type.num.num'),X45,'type.num.num')) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X45),s('type.num.num',sK12)))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f43936,plain,(
  ( ! [X54,X56,X55] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X54,'type.num.num'),X55,'type.num.num')) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X55),s('type.num.num',sK12))) | ~p(s(bool,X56))) ) | ($spl10 | $spl208)),
  inference(resolution,[],[f43601,f34808])).
fof(f34808,plain,(
  ( ! [X12,X11] : (~sP2(X11,X12,'F_0') | s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X11))) ) | $spl10),
  inference(subsumption_resolution,[],[f34807,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f34807,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X12) | ~sP2(X11,X12,'F_0') | ~sP0(X12,'F_0',X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f34806,f623])).
fof(f623,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X23] : ((p(s(bool,X23)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CLAUSES|split|4')).
fof(f34806,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X12) | p(s(bool,'F_0')) | ~sP2(X11,X12,'F_0') | ~sP0(X12,'F_0',X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f34791,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f316])).
fof(f316,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f34791,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X12) | p(s(bool,'F_0')) | ~sP2(X11,X12,'F_0') | ~sP1('F_0',X12,X11) | ~sP0(X12,'F_0',X11)) ) | $spl10),
  inference(resolution,[],[f33757,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f33757,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f9219,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f143,f280,f279,f278,f277])).
fof(f279,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X58,X56,X57] : ((p(s(bool,X57)) <=> s(bool,X56) = s(bool,X58)) <=> ((p(s(bool,X57)) | p(s(bool,X56)) | p(s(bool,X58))) & (p(s(bool,X57)) | ~p(s(bool,X58)) | ~p(s(bool,X56))) & (p(s(bool,X56)) | ~p(s(bool,X58)) | ~p(s(bool,X57))) & (p(s(bool,X58)) | ~p(s(bool,X56)) | ~p(s(bool,X57)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f9219,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl10),
  inference(backward_demodulation,[],[f1582,f1580])).
fof(f1580,plain,(
  s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl10),
  inference(resolution,[],[f709,f652])).
fof(f652,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X23] : (s(bool,'F_0') = s(bool,X23) <=> ~p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f1582,plain,(
  s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl10),
  inference(resolution,[],[f709,f657])).
fof(f657,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X23] : (s(bool,'T_0') = s(bool,X23) <=> p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f43601,plain,(
  ( ! [X61,X59,X62,X60] : (sP2(X61,happ(s(fun('type.num.num',bool),X60),s('type.num.num',sK12)),X62) | s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X59,'type.num.num'),X60,'type.num.num'))) ) | $spl208),
  inference(resolution,[],[f34240,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f279])).
fof(f34240,plain,(
  ( ! [X121,X120] : (~p(s(bool,happ(s(fun('type.num.num',bool),X121),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X120,'type.num.num'),X121,'type.num.num'))) ) | $spl208),
  inference(forward_demodulation,[],[f34209,f32867])).
fof(f34209,plain,(
  ( ! [X121,X120] : (~p(s(bool,happ(s(fun('type.num.num',bool),X121),s('type.num.num',sK12)))) | s('type.num.num',sK22(sK13,X120,'type.num.num')) = s('type.num.num',sK22(sK22(sK13,X120,'type.num.num'),X121,'type.num.num'))) ) | $spl208),
  inference(superposition,[],[f490,f32867])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK22(X2,X1,X0)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK22(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X19] : ! [X37,X29] : (? [X7] : (s(X19,X7) = s(X19,X29) & p(s(bool,happ(s(fun(X19,bool),X37),s(X19,X7))))) <=> p(s(bool,happ(s(fun(X19,bool),X37),s(X19,X29)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.UNWIND_THM2')).
fof(f43964,plain,(
  ( ! [X48,X49] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X48,'type.num.num'),X49,'type.num.num')) | s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',sK12)))) ) | $spl306),
  inference(cnf_transformation,[],[f43964_D])).
fof(f43964_D,plain,(
  ( ! [X48,X49] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X48,'type.num.num'),X49,'type.num.num')) | s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',sK12)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f43934,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X48,'type.num.num'),X49,'type.num.num')) | s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',sK12))) | ~p(s(bool,X50))) ) | ($spl52 | $spl208)),
  inference(resolution,[],[f43601,f34542])).
fof(f34542,plain,(
  ( ! [X12,X11] : (~sP2(X11,X12,'F_0') | s(bool,'F_0') = s(bool,X12) | ~p(s(bool,X11))) ) | $spl52),
  inference(subsumption_resolution,[],[f34529,f469])).
fof(f34529,plain,(
  ( ! [X12,X11] : (~sP2(X11,X12,'F_0') | ~sP0(X12,'F_0',X11) | s(bool,'F_0') = s(bool,X12) | ~p(s(bool,X11))) ) | $spl52),
  inference(resolution,[],[f4698,f474])).
fof(f4698,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f4696,f666])).
fof(f666,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb',aHL_FALSITY)).
fof(f4696,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl52),
  inference(resolution,[],[f4693,f458])).
fof(f4693,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) ) | $spl52),
  inference(subsumption_resolution,[],[f4688,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f317])).
fof(f4688,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) ) | $spl52),
  inference(resolution,[],[f4647,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f4647,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f4643,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f317])).
fof(f4643,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl52),
  inference(resolution,[],[f4637,f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f4637,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl52),
  inference(subsumption_resolution,[],[f4636,f657])).
fof(f4636,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl52),
  inference(subsumption_resolution,[],[f4633,f4623])).
fof(f4623,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f4618,f469])).
fof(f4618,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl52),
  inference(resolution,[],[f3830,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f3830,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl52),
  inference(superposition,[],[f3355,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f320])).
fof(f3355,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl52),
  inference(equality_resolution,[],[f3282])).
fof(f4633,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl52),
  inference(resolution,[],[f4624,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f311])).
fof(f4624,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl52),
  inference(subsumption_resolution,[],[f4619,f465])).
fof(f4619,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl52),
  inference(resolution,[],[f3830,f454])).
fof(f3823,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33)) ) | $spl62),
  inference(cnf_transformation,[],[f3823_D])).
fof(f3823_D,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f43933,plain,(
  ( ! [X47,X45,X46,X44] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X44,'type.num.num'),X45,'type.num.num')) | ~p(s(bool,X46)) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X45),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X47)) ) | ($spl52 | $spl208)),
  inference(resolution,[],[f43601,f40612])).
fof(f40612,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X5,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X3)) ) | $spl52),
  inference(subsumption_resolution,[],[f40611,f469])).
fof(f40611,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP2(X4,X5,X3) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f40610,f465])).
fof(f40610,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP2(X4,X5,X3) | ~sP1(X3,X5,X4) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f40560,f652])).
fof(f40560,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | ~sP2(X4,X5,X3) | ~sP1(X3,X5,X4) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(resolution,[],[f34652,f458])).
fof(f34652,plain,(
  ( ! [X10,X11,X9] : (~sP3(X11,X9,X10) | s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl52),
  inference(superposition,[],[f3282,f474])).
fof(f43903,plain,(
  ( ! [X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) | $spl302),
  inference(cnf_transformation,[],[f43903_D])).
fof(f43903_D,plain,(
  ( ! [X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f43842,plain,(
  ( ! [X271,X273,X272] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X271)))) | p(s(bool,X272)) | ~p(s(bool,X273)) | s(bool,X273) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X271)))) ) | $spl10),
  inference(resolution,[],[f1611,f691])).
fof(f1611,plain,(
  ( ! [X0] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))) ) | $spl10),
  inference(superposition,[],[f709,f679])).
fof(f43751,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35)))) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35))) | s(bool,X36) != s(bool,X38)) ) | $spl10),
  inference(resolution,[],[f1611,f452])).
fof(f43749,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27)))) | p(s(bool,X28)) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27)))) ) | $spl10),
  inference(resolution,[],[f1611,f451])).
fof(f43747,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19)))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19)))) ) | $spl10),
  inference(resolution,[],[f1611,f450])).
fof(f43746,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X15)))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X15))) | s(bool,X16) != s(bool,X18)) ) | $spl10),
  inference(resolution,[],[f1611,f448])).
fof(f43744,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) | $spl10),
  inference(resolution,[],[f1611,f446])).
fof(f43732,plain,(
  ( ! [X6,X4,X3] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X3,'type.num.num'),X4,'type.num.num')) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',sK12)))) ) | $spl300),
  inference(cnf_transformation,[],[f43732_D])).
fof(f43732_D,plain,(
  ( ! [X6,X4,X3] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X3,'type.num.num'),X4,'type.num.num')) | ~p(s(bool,X6)) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',sK12)))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f43687,plain,(
  ( ! [X366,X368,X365,X367] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X365,'type.num.num'),X366,'type.num.num')) | p(s(bool,X367)) | ~p(s(bool,X368)) | s(bool,X368) != s(bool,happ(s(fun('type.num.num',bool),X366),s('type.num.num',sK12)))) ) | $spl208),
  inference(resolution,[],[f34240,f691])).
fof(f43596,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X38,'type.num.num'),X39,'type.num.num')) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun('type.num.num',bool),X39),s('type.num.num',sK12))) | s(bool,X40) != s(bool,X42)) ) | $spl208),
  inference(resolution,[],[f34240,f452])).
fof(f43594,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X28,'type.num.num'),X29,'type.num.num')) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun('type.num.num',bool),X29),s('type.num.num',sK12)))) ) | $spl208),
  inference(resolution,[],[f34240,f451])).
fof(f43592,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X18,'type.num.num'),X19,'type.num.num')) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun('type.num.num',bool),X19),s('type.num.num',sK12)))) ) | $spl208),
  inference(resolution,[],[f34240,f450])).
fof(f43591,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X13,'type.num.num'),X14,'type.num.num')) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK12))) | s(bool,X15) != s(bool,X17)) ) | $spl208),
  inference(resolution,[],[f34240,f448])).
fof(f43589,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK22(sK13,X3,'type.num.num'),X4,'type.num.num')) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun('type.num.num',bool),X4),s('type.num.num',sK12)))) ) | $spl208),
  inference(resolution,[],[f34240,f446])).
fof(f43437,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl298),
  inference(cnf_transformation,[],[f43437_D])).
fof(f43437_D,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f43376,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X270),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X271)) | ~p(s(bool,X272)) | s(bool,X272) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X270),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl10),
  inference(resolution,[],[f1594,f691])).
fof(f1594,plain,(
  ( ! [X1] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl10),
  inference(superposition,[],[f709,f679])).
fof(f43285,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X34),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X35) != s(bool,X37)) ) | $spl10),
  inference(resolution,[],[f1594,f452])).
fof(f43283,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X26),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl10),
  inference(resolution,[],[f1594,f451])).
fof(f43281,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl10),
  inference(resolution,[],[f1594,f450])).
fof(f43280,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X14),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X14),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X15) != s(bool,X17)) ) | $spl10),
  inference(resolution,[],[f1594,f448])).
fof(f43278,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl10),
  inference(resolution,[],[f1594,f446])).
fof(f43114,plain,(
  ( ! [X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) | $spl296),
  inference(cnf_transformation,[],[f43114_D])).
fof(f43114_D,plain,(
  ( ! [X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f43053,plain,(
  ( ! [X271,X273,X272] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X271),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X272)) | ~p(s(bool,X273)) | s(bool,X273) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X271)))) ) | $spl10),
  inference(resolution,[],[f1593,f691])).
fof(f1593,plain,(
  ( ! [X0] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl10),
  inference(superposition,[],[f709,f679])).
fof(f42962,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X37) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X35))) | s(bool,X36) != s(bool,X38)) ) | $spl10),
  inference(resolution,[],[f1593,f452])).
fof(f42960,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X28)) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27)))) ) | $spl10),
  inference(resolution,[],[f1593,f451])).
fof(f42958,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X19)))) ) | $spl10),
  inference(resolution,[],[f1593,f450])).
fof(f42957,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X15),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X15))) | s(bool,X16) != s(bool,X18)) ) | $spl10),
  inference(resolution,[],[f1593,f448])).
fof(f42955,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X7)))) ) | $spl10),
  inference(resolution,[],[f1593,f446])).
fof(f30930,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl167),
  inference(cnf_transformation,[],[f30930_D])).
fof(f30930_D,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f40908,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'F_0') != s(bool,X3)) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f40612,f30741])).
fof(f30741,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X33)) ) | $spl161),
  inference(resolution,[],[f30723,f462])).
fof(f31063,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X15) | p(s(bool,X14))) ) | $spl184),
  inference(cnf_transformation,[],[f31063_D])).
fof(f31063_D,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X15) | p(s(bool,X14))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f38626,plain,(
  ( ! [X39,X40] : (sP3(X39,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X40) | p(s(bool,X39)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl87 | $spl122 | $spl161)),
  inference(resolution,[],[f27637,f30745])).
fof(f30745,plain,(
  ( ! [X41,X40] : (sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X40,X41)) ) | $spl161),
  inference(resolution,[],[f30723,f467])).
fof(f27637,plain,(
  ( ! [X6,X8,X7] : (~sP1(X6,X8,X7) | sP3(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27624,f463])).
fof(f27624,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP3(X7,X6,X8) | p(s(bool,X7)) | ~sP2(X7,X8,X6) | ~sP1(X6,X8,X7)) ) | ($spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f27608,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f311])).
fof(f27608,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | s(bool,'T_0') != s(bool,X1)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27607,f469])).
fof(f27607,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | sP0(X0,X1,X2)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27560,f14662])).
fof(f14662,plain,(
  ( ! [X6,X8,X7] : (sP0(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f14661,f469])).
fof(f14661,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP0(X7,X6,X8)) ) | $spl52),
  inference(subsumption_resolution,[],[f14645,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f319])).
fof(f14645,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP0(X7,X6,X8)) ) | $spl52),
  inference(resolution,[],[f14566,f453])).
fof(f14566,plain,(
  ( ! [X10,X8,X9] : (sP3(X8,X9,X10) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,X10) | p(s(bool,X8))) ) | $spl52),
  inference(resolution,[],[f14550,f463])).
fof(f14550,plain,(
  ( ! [X6,X8,X7] : (~sP2(X8,X6,X7) | sP3(X8,X7,X6) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl52),
  inference(subsumption_resolution,[],[f14549,f467])).
fof(f14549,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP3(X8,X7,X6) | p(s(bool,X7)) | ~sP2(X8,X6,X7) | ~sP1(X7,X6,X8)) ) | $spl52),
  inference(subsumption_resolution,[],[f14544,f653])).
fof(f653,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f411])).
fof(f14544,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | sP3(X8,X7,X6) | p(s(bool,X7)) | ~sP2(X8,X6,X7) | ~sP1(X7,X6,X8)) ) | $spl52),
  inference(resolution,[],[f14536,f458])).
fof(f14536,plain,(
  ( ! [X6,X8,X7] : (sP0(X7,X6,X8) | s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl52),
  inference(subsumption_resolution,[],[f14520,f469])).
fof(f14520,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP0(X7,X6,X8)) ) | $spl52),
  inference(resolution,[],[f5532,f453])).
fof(f5532,plain,(
  ( ! [X4,X5,X3] : (sP3(X5,X3,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl52),
  inference(superposition,[],[f3282,f472])).
fof(f27560,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | sP0(X0,X1,X2)) ) | ($spl87 | $spl122)),
  inference(resolution,[],[f27462,f453])).
fof(f27462,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2))) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27388,f10990])).
fof(f10990,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl87),
  inference(trivial_inequality_removal,[],[f10985])).
fof(f10985,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl87),
  inference(superposition,[],[f6236,f690])).
fof(f690,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb',aHL_BOOL_CASES)).
fof(f27388,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X0) | sP3(X2,X1,X0) | p(s(bool,X2))) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f472])).
fof(f14390,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl87 | $spl122)),
  inference(trivial_inequality_removal,[],[f14376])).
fof(f14376,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X67) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14168,f10990])).
fof(f14168,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl122),
  inference(trivial_inequality_removal,[],[f14143])).
fof(f14143,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl122),
  inference(superposition,[],[f13583,f690])).
fof(f38200,plain,(
  $false | $spl282),
  inference(equality_resolution,[],[f38153])).
fof(f38153,plain,(
  ( ! [X12] : (s('type.one.one','const.one.one_0') != s('type.one.one',X12)) ) | $spl282),
  inference(cnf_transformation,[],[f38153_D])).
fof(f38153_D,plain,(
  ( ! [X12] : (s('type.one.one','const.one.one_0') != s('type.one.one',X12)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f38167,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X30),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X28))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26)))) ) | $spl294),
  inference(cnf_transformation,[],[f38167_D])).
fof(f38167_D,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X30),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X28))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26)))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f38145,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X36),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X33),s('type.list.list'('type.infer_t.infer_t'),X34),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X35))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X32))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X35))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X33),s('type.list.list'('type.infer_t.infer_t'),X34))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X32)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X32))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X35))) | s('type.one.one','const.one.one_0') != s('type.one.one',X31) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X33))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X34)))) ) | $spl4),
  inference(superposition,[],[f720,f499])).
fof(f499,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1)))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14,X15,X16,X7] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X13),s('type.list.list'('type.infer_t.infer_t'),X14),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X14))) & s('type.one.one','const.one.one_0') = s('type.one.one',X7) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X13),s('type.list.list'('type.infer_t.infer_t'),X14))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))))))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.success_eqns|split|8')).
fof(f720,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.pair.prod'(X11,'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s(X11,X13),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'(X11,'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s(X11,X12),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X10)))) ) | $spl4),
  inference(superposition,[],[f703,f477])).
fof(f477,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X3) = s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X19,X52] : ! [X7,X6,X29,X28] : (s('type.pair.prod'(X19,X52),'const.pair.,_2'(s(X19,X7),s(X52,X6))) = s('type.pair.prod'(X19,X52),'const.pair.,_2'(s(X19,X29),s(X52,X28))) <=> (s(X19,X7) = s(X19,X29) & s(X52,X6) = s(X52,X28)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f38144,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X30),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X27),s('type.list.list'('type.infer_t.infer_t'),X28))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X26))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X29))) | s('type.one.one','const.one.one_0') != s('type.one.one',X25) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X28)))) ) | $spl4),
  inference(superposition,[],[f720,f504])).
fof(f504,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1)))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X1))) & s('type.one.one','const.one.one_0') = s('type.one.one',X4) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X0),s('type.list.list'('type.infer_t.infer_t'),X1))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X13,X14,X9,X12,X7] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X13),s('type.list.list'('type.infer_t.infer_t'),X14),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X13))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X14))) & s('type.one.one','const.one.one_0') = s('type.one.one',X7) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X13),s('type.list.list'('type.infer_t.infer_t'),X14))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))))))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.inferProps_3|48|_prove_ltb')).
fof(f38166,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X24),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23)))) ) | $spl292),
  inference(cnf_transformation,[],[f38166_D])).
fof(f38166_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X24),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f38143,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X24),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X23))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X21),s('type.infer_t.infer_t',X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X20)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X19)) ) | $spl4),
  inference(superposition,[],[f720,f508])).
fof(f508,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4)) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))) & (s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))))) & ((s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X17,X18,X9,X12,X7] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X17),s('type.infer_t.infer_t',X18),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) <=> (s('type.one.one','const.one.one_0') = s('type.one.one',X7) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X17),s('type.infer_t.infer_t',X18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))))) & s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.infer_eComplete.infer_eComplete_25_prove_ltb')).
fof(f38165,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X17),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))),s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13)))) ) | $spl290),
  inference(cnf_transformation,[],[f38165_D])).
fof(f38165_D,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X17),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))),s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13)))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f38142,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X17),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X13) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X16))),s('type.infer_t.infer_t',X14),s('type.infer_t.infer_t',X15))) | s('type.one.one','const.one.one_0') != s('type.one.one',X12)) ) | $spl4),
  inference(superposition,[],[f720,f684])).
fof(f684,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) | s('type.one.one','const.one.one_0') != s('type.one.one',X4)) )),
  inference(cnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK35(X4,X3,X2,X1,X0)))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK35(X4,X3,X2,X1,X0)))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & (s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ! [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & ? [X6] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X6))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & (s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ! [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(rectify,[],[f423])).
fof(f423,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & ? [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & (s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ! [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(flattening,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) | (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & ? [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & ((s('type.one.one','const.one.one_0') != s('type.one.one',X4) | ! [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3)))))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X4))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))) <=> (s('type.one.one','const.one.one_0') = s('type.one.one',X4) & ? [X5] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))),s('type.infer_t.infer_t',X0),s('type.infer_t.infer_t',X1))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X5))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X17,X18,X9,X12,X7] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X17),s('type.infer_t.infer_t',X18),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))) = s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'const.infer.Success_1'(s('type.one.one',X7))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) <=> (s('type.one.one','const.one.one_0') = s('type.one.one',X7) & ? [X20] : (s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20))) = s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.infer_t.infer_t',X17),s('type.infer_t.infer_t',X18))) & s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12) = s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.add_constraint_success')).
fof(f38158,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X26))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X29),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28)))) ) | $spl288),
  inference(cnf_transformation,[],[f38158_D])).
fof(f38158_D,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X26))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X29),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28)))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f38159,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X35),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X32),s('type.list.list'('type.infer_t.infer_t'),X33),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X32),s('type.list.list'('type.infer_t.infer_t'),X33))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X31) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X32))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X33)))) ) | $spl4),
  inference(forward_demodulation,[],[f38025,f703])).
fof(f38025,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X35),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X32),s('type.list.list'('type.infer_t.infer_t'),X33),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X36))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X32),s('type.list.list'('type.infer_t.infer_t'),X33))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X34))) | s('type.one.one','const.one.one_0') != s('type.one.one',X31) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X32))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X33)))) ) | $spl4),
  inference(superposition,[],[f720,f499])).
fof(f38157,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X29),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X25) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X26))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27)))) ) | $spl4),
  inference(forward_demodulation,[],[f38024,f703])).
fof(f38024,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X29),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraints_3'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X30))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.ZIP_1'(s('type.pair.prod'('type.list.list'('type.infer_t.infer_t'),'type.list.list'('type.infer_t.infer_t')),'const.pair.,_2'(s('type.list.list'('type.infer_t.infer_t'),X26),s('type.list.list'('type.infer_t.infer_t'),X27))))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X28))) | s('type.one.one','const.one.one_0') != s('type.one.one',X25) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X26))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.infer_t.infer_t'),X27)))) ) | $spl4),
  inference(superposition,[],[f720,f504])).
fof(f38156,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X23),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22)))) ) | $spl286),
  inference(cnf_transformation,[],[f38156_D])).
fof(f38156_D,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X23),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22)))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f38155,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X23),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X19)) ) | $spl4),
  inference(forward_demodulation,[],[f38023,f703])).
fof(f38023,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X23),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X24))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X22))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.CONS_2'(s('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),'const.pair.,_2'(s('type.infer_t.infer_t',X20),s('type.infer_t.infer_t',X21))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.NIL_0'))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s('type.one.one','const.one.one_0') != s('type.one.one',X19)) ) | $spl4),
  inference(superposition,[],[f720,f508])).
fof(f38154,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X16),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))),s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17)))) ) | $spl284),
  inference(cnf_transformation,[],[f38154_D])).
fof(f38154_D,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X16),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))),s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f38022,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.pair.,_2'(s('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),X16),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17))) != s('type.pair.prod'('type.infer.exc'('type.one.one','type.list.list'('type.string.char')),'type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'))),'const.infer.add_constraint_3'(s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X17))) | s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18) != s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))) | s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.option.SOME_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X18))) != s('type.option.option'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.unify.t_unify_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X15))),s('type.infer_t.infer_t',X13),s('type.infer_t.infer_t',X14))) | s('type.one.one','const.one.one_0') != s('type.one.one',X12)) ) | $spl4),
  inference(superposition,[],[f720,f684])).
fof(f31026,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X12))) ) | $spl180),
  inference(cnf_transformation,[],[f31026_D])).
fof(f31026_D,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X12))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f37847,plain,(
  ( ! [X33,X34] : (sP3(X33,X34,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X33)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f34340,f30744])).
fof(f30744,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X39)) ) | $spl161),
  inference(resolution,[],[f30723,f466])).
fof(f34340,plain,(
  ( ! [X6,X8,X7] : (~sP1(X8,X6,X7) | sP3(X7,X8,X6) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl52),
  inference(subsumption_resolution,[],[f34333,f463])).
fof(f34333,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP3(X7,X8,X6) | p(s(bool,X7)) | ~sP2(X7,X6,X8) | ~sP1(X8,X6,X7)) ) | $spl52),
  inference(resolution,[],[f14662,f459])).
fof(f30931,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X9))) ) | $spl168),
  inference(cnf_transformation,[],[f30931_D])).
fof(f30931_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X9))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f37818,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X3)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f34339,f30741])).
fof(f34339,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X3,X5) | sP3(X4,X5,X3) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X3)) ) | $spl52),
  inference(subsumption_resolution,[],[f34332,f467])).
fof(f34332,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP3(X4,X5,X3) | p(s(bool,X5)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | $spl52),
  inference(resolution,[],[f14662,f458])).
fof(f31025,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl178),
  inference(cnf_transformation,[],[f31025_D])).
fof(f31025_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f31064,plain,(
  ( ! [X22] : (sP3(X22,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),'T_0') | p(s(bool,X22))) ) | $spl186),
  inference(cnf_transformation,[],[f31064_D])).
fof(f31064_D,plain,(
  ( ! [X22] : (sP3(X22,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),'T_0') | p(s(bool,X22))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f37811,plain,(
  ( ! [X33,X34] : (sP3(X33,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),'T_0') | p(s(bool,X33)) | s(bool,X34) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f34330,f30745])).
fof(f34330,plain,(
  ( ! [X6,X8,X7] : (~sP1(X7,'T_0',X8) | sP3(X8,X7,'T_0') | p(s(bool,X8)) | s(bool,X6) != s(bool,X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f34324,f463])).
fof(f34324,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | sP3(X8,X7,'T_0') | p(s(bool,X8)) | ~sP2(X8,'T_0',X7) | ~sP1(X7,'T_0',X8)) ) | $spl52),
  inference(resolution,[],[f13713,f459])).
fof(f13713,plain,(
  ( ! [X6,X8,X7] : (sP0('T_0',X7,X8) | s(bool,X6) != s(bool,X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f13705,f469])).
fof(f13705,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | sP0('T_0',X7,X8) | p(s(bool,X8))) ) | $spl52),
  inference(resolution,[],[f13620,f463])).
fof(f13620,plain,(
  ( ! [X6,X8,X7] : (~sP2(X8,'T_0',X7) | s(bool,X6) != s(bool,X7) | sP0('T_0',X7,X8)) ) | $spl52),
  inference(subsumption_resolution,[],[f13610,f470])).
fof(f13610,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~sP2(X8,'T_0',X7) | sP0('T_0',X7,X8)) ) | $spl52),
  inference(resolution,[],[f13598,f453])).
fof(f13598,plain,(
  ( ! [X6,X8,X7] : (sP3(X8,X7,'T_0') | s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~sP2(X8,'T_0',X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f13597,f9801])).
fof(f9801,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X43) | p(s(bool,X42))) ) | $spl52),
  inference(trivial_inequality_removal,[],[f9776])).
fof(f9776,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | s(bool,'F_0') = s(bool,X43) | p(s(bool,X42))) ) | $spl52),
  inference(superposition,[],[f5509,f655])).
fof(f655,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X23] : (s(bool,'F_0') = s(bool,X23) <=> ~p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f5509,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl52),
  inference(superposition,[],[f3282,f690])).
fof(f13597,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6) | sP3(X8,X7,'T_0') | p(s(bool,X7)) | ~sP2(X8,'T_0',X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f13592,f467])).
fof(f13592,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6) | sP3(X8,X7,'T_0') | p(s(bool,X7)) | ~sP2(X8,'T_0',X7) | ~sP1(X7,'T_0',X8)) ) | $spl52),
  inference(resolution,[],[f13584,f458])).
fof(f13584,plain,(
  ( ! [X6,X8,X7] : (sP0('T_0',X7,X8) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl52),
  inference(subsumption_resolution,[],[f13571,f469])).
fof(f13571,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | sP0('T_0',X7,X8)) ) | $spl52),
  inference(resolution,[],[f5512,f453])).
fof(f5512,plain,(
  ( ! [X6,X8,X7] : (sP3(X8,X6,'T_0') | s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | p(s(bool,X8))) ) | $spl52),
  inference(superposition,[],[f3282,f472])).
fof(f13909,plain,(
  ( ! [X24] : (sP3(X24,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X24))) ) | $spl136),
  inference(cnf_transformation,[],[f13909_D])).
fof(f13909_D,plain,(
  ( ! [X24] : (sP3(X24,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X24))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f37816,plain,(
  ( ! [X31,X32] : (s(bool,'F_0') != s(bool,X32) | sP3(X31,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X31))) ) | ($spl7 | $spl52)),
  inference(forward_demodulation,[],[f37810,f2271])).
fof(f2271,plain,(
  s(bool,'F_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl7),
  inference(resolution,[],[f704,f653])).
fof(f704,plain,(
  ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl7),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f37810,plain,(
  ( ! [X31,X32] : (sP3(X31,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X31)) | s(bool,X32) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl7 | $spl52)),
  inference(resolution,[],[f34330,f2194])).
fof(f2194,plain,(
  ( ! [X41,X40] : (sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X40,X41)) ) | $spl7),
  inference(resolution,[],[f704,f467])).
fof(f13814,plain,(
  ( ! [X14] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X14)) ) | $spl130),
  inference(cnf_transformation,[],[f13814_D])).
fof(f13814_D,plain,(
  ( ! [X14] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X14)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f13812,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,X13)) ) | $spl128),
  inference(cnf_transformation,[],[f13812_D])).
fof(f13812_D,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,X13)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f37782,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | sP3(X18,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X18)) ) | ($spl7 | $spl32 | $spl52)),
  inference(forward_demodulation,[],[f37774,f33409])).
fof(f33409,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f657])).
fof(f32849,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ($spl7 | $spl32)),
  inference(resolution,[],[f5815,f2339])).
fof(f2339,plain,(
  ( ! [X157] : (~sP11('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X157) | p(s(bool,X157))) ) | $spl7),
  inference(subsumption_resolution,[],[f2260,f642])).
fof(f642,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f403])).
fof(f403,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X23] : ((p(s(bool,'F_0')) & p(s(bool,X23))) <=> p(s(bool,'F_0')))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.AND_CLAUSES|split|2')).
fof(f2260,plain,(
  ( ! [X157] : (p(s(bool,'F_0')) | p(s(bool,X157)) | ~sP11('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X157)) ) | $spl7),
  inference(resolution,[],[f704,f592])).
fof(f592,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f371])).
fof(f371,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f5815,plain,(
  ( ! [X164] : (sP11(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f593])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f372])).
fof(f37774,plain,(
  ( ! [X19,X18] : (sP3(X18,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | s(bool,X19) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X18)) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f34328,f5747])).
fof(f5747,plain,(
  ( ! [X43,X42] : (sP2(X42,X43,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f34328,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,'T_0',X1) | sP3(X2,X1,'T_0') | s(bool,X0) != s(bool,X1) | ~sP1(X1,'T_0',X2)) ) | $spl52),
  inference(subsumption_resolution,[],[f34322,f695])).
fof(f695,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f618])).
fof(f618,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X19] : ! [X7] : (s(X19,X7) = s(X19,X7) <=> p(s(bool,'T_0')))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f34322,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP3(X2,X1,'T_0') | ~p(s(bool,'T_0')) | ~sP2(X2,'T_0',X1) | ~sP1(X1,'T_0',X2)) ) | $spl52),
  inference(resolution,[],[f13713,f457])).
fof(f13816,plain,(
  ( ! [X16] : (sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) | $spl132),
  inference(cnf_transformation,[],[f13816_D])).
fof(f13816_D,plain,(
  ( ! [X16] : (sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f37781,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) | ($spl24 | $spl52 | $spl87)),
  inference(forward_demodulation,[],[f37773,f10990])).
fof(f37773,plain,(
  ( ! [X17,X16] : (sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | s(bool,X17) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) | ($spl24 | $spl52)),
  inference(resolution,[],[f34328,f3869])).
fof(f3869,plain,(
  ( ! [X43,X42] : (sP2(X42,X43,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl24),
  inference(resolution,[],[f1229,f461])).
fof(f33642,plain,(
  ( ! [X32] : (sP3(X32,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0',X32)) ) | $spl258),
  inference(cnf_transformation,[],[f33642_D])).
fof(f33642_D,plain,(
  ( ! [X32] : (sP3(X32,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0',X32)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f37780,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | sP3(X14,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0',X14)) ) | ($spl12 | $spl52)),
  inference(forward_demodulation,[],[f37772,f2563])).
fof(f2563,plain,(
  s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl12),
  inference(resolution,[],[f711,f657])).
fof(f37772,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | s(bool,X15) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0',X14)) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f34328,f2480])).
fof(f2480,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f711,f461])).
fof(f13811,plain,(
  ( ! [X12] : (sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) | $spl126),
  inference(cnf_transformation,[],[f13811_D])).
fof(f13811_D,plain,(
  ( ! [X12] : (sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f37779,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) | ($spl10 | $spl52)),
  inference(forward_demodulation,[],[f37771,f1582])).
fof(f37771,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) | ($spl10 | $spl52)),
  inference(resolution,[],[f34328,f1499])).
fof(f1499,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f709,f461])).
fof(f5722,plain,(
  s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl73),
  inference(cnf_transformation,[],[f5722_D])).
fof(f5722_D,plain,(
  s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f34894,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f34543,f5753])).
fof(f5753,plain,(
  ( ! [X54,X55] : (sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X54,X55)) ) | $spl32),
  inference(resolution,[],[f1645,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f319])).
fof(f34543,plain,(
  ( ! [X14,X13] : (~sP0(X14,'F_0',X13) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14)) ) | $spl52),
  inference(subsumption_resolution,[],[f34530,f463])).
fof(f34530,plain,(
  ( ! [X14,X13] : (~sP2(X13,X14,'F_0') | ~sP0(X14,'F_0',X13) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,X14)) ) | $spl52),
  inference(resolution,[],[f4698,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f320])).
fof(f33635,plain,(
  s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl247),
  inference(cnf_transformation,[],[f33635_D])).
fof(f33635_D,plain,(
  s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) <=> ~$spl247),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl247])])).
fof(f34892,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f34543,f2486])).
fof(f2486,plain,(
  ( ! [X54,X53] : (sP0('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X53,X54)) ) | $spl12),
  inference(resolution,[],[f711,f471])).
fof(f30933,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl170),
  inference(cnf_transformation,[],[f30933_D])).
fof(f30933_D,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f34868,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~p(s(bool,X1))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f34542,f30741])).
fof(f30935,plain,(
  ( ! [X31] : (sP3(X31,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl174),
  inference(cnf_transformation,[],[f30935_D])).
fof(f30935_D,plain,(
  ( ! [X31] : (sP3(X31,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f34859,plain,(
  ( ! [X1] : (sP3(X1,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f34319,f30741])).
fof(f34319,plain,(
  ( ! [X2,X3] : (~sP2(X3,X2,'F_0') | sP3(X3,'F_0',X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl52),
  inference(subsumption_resolution,[],[f34318,f666])).
fof(f34318,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | p(s(bool,'F_0')) | ~sP2(X3,X2,'F_0')) ) | $spl52),
  inference(subsumption_resolution,[],[f34311,f467])).
fof(f34311,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | p(s(bool,'F_0')) | ~sP2(X3,X2,'F_0') | ~sP1('F_0',X2,X3)) ) | $spl52),
  inference(resolution,[],[f4623,f458])).
fof(f34812,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~p(s(bool,X1))) ) | ($spl10 | $spl161)),
  inference(resolution,[],[f34808,f30741])).
fof(f33954,plain,(
  $false | ($spl7 | $spl62)),
  inference(trivial_inequality_removal,[],[f33937])).
fof(f33937,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl7 | $spl62)),
  inference(superposition,[],[f3823,f2271])).
fof(f33952,plain,(
  $false | $spl62),
  inference(equality_resolution,[],[f3823])).
fof(f3282,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f3282_D])).
fof(f3282_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f14165,plain,(
  ( ! [X59,X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,'F_0') != s(bool,X60) | p(s(bool,X59))) ) | $spl122),
  inference(trivial_inequality_removal,[],[f14146])).
fof(f14146,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,X59) | s(bool,'T_0') != s(bool,X60) | s(bool,'F_0') != s(bool,X60) | p(s(bool,X59))) ) | $spl122),
  inference(superposition,[],[f13583,f653])).
fof(f14166,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,X58) | s(bool,'F_0') != s(bool,X58) | p(s(bool,X57))) ) | $spl122),
  inference(trivial_inequality_removal,[],[f14145])).
fof(f14145,plain,(
  ( ! [X57,X58] : (s(bool,X57) != s(bool,X57) | s(bool,'T_0') != s(bool,X58) | s(bool,'F_0') != s(bool,X58) | p(s(bool,X57))) ) | $spl122),
  inference(superposition,[],[f13583,f655])).
fof(f30934,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | $spl172),
  inference(cnf_transformation,[],[f30934_D])).
fof(f30934_D,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f14167,plain,(
  ( ! [X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,'T_0') = s(bool,X55)) ) | $spl122),
  inference(trivial_inequality_removal,[],[f14144])).
fof(f14144,plain,(
  ( ! [X56,X55] : (s(bool,X55) != s(bool,X55) | s(bool,'T_0') != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,'T_0') = s(bool,X55)) ) | $spl122),
  inference(superposition,[],[f13583,f690])).
fof(f33875,plain,(
  $false | ($spl10 | $spl58)),
  inference(subsumption_resolution,[],[f3819,f9219])).
fof(f3819,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl58),
  inference(cnf_transformation,[],[f3819_D])).
fof(f3819_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f3820,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl59),
  inference(cnf_transformation,[],[f3820_D])).
fof(f3820_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f13700,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl124),
  inference(duplicate_literal_removal,[],[f13695])).
fof(f13695,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X9))) ) | $spl124),
  inference(resolution,[],[f13683,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f320])).
fof(f13683,plain,(
  ( ! [X3] : (~sP3(X3,'F_0','T_0') | ~p(s(bool,X3))) ) | $spl124),
  inference(resolution,[],[f13622,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X2,X1) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f13983,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl134),
  inference(duplicate_literal_removal,[],[f13976])).
fof(f13976,plain,(
  ( ! [X6] : (p(s(bool,X6)) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl134),
  inference(resolution,[],[f13907,f475])).
fof(f32962,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) != s(bool,X11)) ) | $spl222),
  inference(cnf_transformation,[],[f32962_D])).
fof(f32962_D,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,X10) != s(bool,X11)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f33723,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X45) | s(bool,X44) != s(bool,X45)) ) | $spl122),
  inference(subsumption_resolution,[],[f14140,f14169])).
fof(f14169,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X43) | s(bool,X42) != s(bool,X43) | p(s(bool,X42))) ) | $spl122),
  inference(trivial_inequality_removal,[],[f14139])).
fof(f14139,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X43) | s(bool,X42) != s(bool,X43) | p(s(bool,X42))) ) | $spl122),
  inference(superposition,[],[f13583,f655])).
fof(f14140,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X45) | s(bool,X44) != s(bool,X45) | ~p(s(bool,X44))) ) | $spl122),
  inference(superposition,[],[f13583,f657])).
fof(f33722,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X47) | s(bool,X46) != s(bool,X47)) ) | $spl122),
  inference(subsumption_resolution,[],[f14141,f14169])).
fof(f14141,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X47) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X46))) ) | $spl122),
  inference(superposition,[],[f13583,f659])).
fof(f659,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X23] : (s(bool,'T_0') = s(bool,X23) <=> p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f33721,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X49) | s(bool,X48) != s(bool,X49)) ) | $spl122),
  inference(subsumption_resolution,[],[f14142,f13583])).
fof(f14142,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X49) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,X48)) ) | $spl122),
  inference(superposition,[],[f13583,f690])).
fof(f33720,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,X61) | s(bool,X61) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl280),
  inference(cnf_transformation,[],[f33720_D])).
fof(f33720_D,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,X61) | s(bool,X61) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f14147,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X61) | s(bool,X61) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl10 | $spl122)),
  inference(superposition,[],[f13583,f1582])).
fof(f33719,plain,(
  ( ! [X64] : (s(bool,X64) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl278),
  inference(cnf_transformation,[],[f33719_D])).
fof(f33719_D,plain,(
  ( ! [X64] : (s(bool,X64) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f33718,plain,(
  ( ! [X64] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X64) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl87 | $spl122)),
  inference(subsumption_resolution,[],[f14150,f14390])).
fof(f14150,plain,(
  ( ! [X64] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X64) | s(bool,X64) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f13583,f10990])).
fof(f33717,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | p(s(bool,X41))) ) | $spl122),
  inference(subsumption_resolution,[],[f14214,f480])).
fof(f480,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17,X18] : (s(bool,X17) = s(bool,X18) <=> ((p(s(bool,X17)) => p(s(bool,X18))) & (p(s(bool,X18)) => p(s(bool,X17)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.EQ_IMP_THM')).
fof(f14214,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | p(s(bool,X40))) ) | $spl122),
  inference(superposition,[],[f14169,f653])).
fof(f33716,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | p(s(bool,X43))) ) | $spl122),
  inference(subsumption_resolution,[],[f14215,f480])).
fof(f14215,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | p(s(bool,X43)) | p(s(bool,X42))) ) | $spl122),
  inference(superposition,[],[f14169,f655])).
fof(f33715,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | p(s(bool,X51))) ) | $spl122),
  inference(subsumption_resolution,[],[f14219,f14169])).
fof(f14219,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | p(s(bool,X51)) | s(bool,'T_0') = s(bool,X50)) ) | $spl122),
  inference(superposition,[],[f14169,f690])).
fof(f33713,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | p(s(bool,X40))) ) | $spl122),
  inference(subsumption_resolution,[],[f14361,f14169])).
fof(f14361,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | s(bool,'T_0') = s(bool,X41) | p(s(bool,X40))) ) | $spl122),
  inference(superposition,[],[f14168,f653])).
fof(f33712,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | p(s(bool,X42))) ) | $spl122),
  inference(subsumption_resolution,[],[f14362,f14169])).
fof(f14362,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X43) | p(s(bool,X42))) ) | $spl122),
  inference(superposition,[],[f14168,f655])).
fof(f33711,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl276),
  inference(cnf_transformation,[],[f33711_D])).
fof(f33711_D,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f33710,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl122),
  inference(subsumption_resolution,[],[f14366,f14168])).
fof(f14366,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X50)) ) | $spl122),
  inference(superposition,[],[f14168,f690])).
fof(f33709,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'T_0') = s(bool,X65)) ) | $spl274),
  inference(cnf_transformation,[],[f33709_D])).
fof(f33709_D,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'T_0') = s(bool,X65)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f14374,plain,(
  ( ! [X65] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X65) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl7 | $spl122)),
  inference(superposition,[],[f14168,f2271])).
fof(f32867,plain,(
  ( ! [X2] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X2,'type.num.num'))) ) | $spl208),
  inference(cnf_transformation,[],[f32867_D])).
fof(f32867_D,plain,(
  ( ! [X2] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X2,'type.num.num'))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f33708,plain,(
  ( ! [X12] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X12,'type.num.num')) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl2 | $spl134)),
  inference(subsumption_resolution,[],[f15644,f888])).
fof(f888,plain,(
  ( ! [X92] : (~p(s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X92,'type.num.num'))) ) | $spl2),
  inference(forward_demodulation,[],[f850,f700])).
fof(f850,plain,(
  ( ! [X92] : (~p(s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',sK12)))) | s('type.num.num',sK13) = s('type.num.num',sK22(sK13,X92,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f490,f700])).
fof(f15644,plain,(
  ( ! [X12] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X12,'type.num.num')) | p(s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',sK12)))) | s(bool,'T_0') != s(bool,'F_0')) ) | ($spl2 | $spl134)),
  inference(resolution,[],[f14003,f475])).
fof(f14003,plain,(
  ( ! [X4] : (sP3(happ(s(fun('type.num.num',bool),X4),s('type.num.num',sK12)),'F_0','T_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X4,'type.num.num'))) ) | ($spl2 | $spl134)),
  inference(resolution,[],[f13996,f903])).
fof(f903,plain,(
  ( ! [X50,X51,X49] : (sP2(happ(s(fun('type.num.num',bool),X49),s('type.num.num',sK12)),X50,X51) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X49,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f463])).
fof(f13996,plain,(
  ( ! [X3] : (~sP2(X3,'T_0','F_0') | sP3(X3,'F_0','T_0')) ) | $spl134),
  inference(subsumption_resolution,[],[f13995,f666])).
fof(f13995,plain,(
  ( ! [X3] : (sP3(X3,'F_0','T_0') | p(s(bool,'F_0')) | ~sP2(X3,'T_0','F_0')) ) | $spl134),
  inference(subsumption_resolution,[],[f13992,f467])).
fof(f13992,plain,(
  ( ! [X3] : (sP3(X3,'F_0','T_0') | p(s(bool,'F_0')) | ~sP2(X3,'T_0','F_0') | ~sP1('F_0','T_0',X3)) ) | $spl134),
  inference(resolution,[],[f13984,f458])).
fof(f13984,plain,(
  ( ! [X7] : (sP0('T_0','F_0',X7)) ) | $spl134),
  inference(subsumption_resolution,[],[f13977,f469])).
fof(f13977,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP0('T_0','F_0',X7)) ) | $spl134),
  inference(resolution,[],[f13907,f453])).
fof(f33700,plain,(
  ( ! [X681] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X681,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25135,f888])).
fof(f25135,plain,(
  ( ! [X681] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,happ(s(fun('type.num.num',bool),X681),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X681,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f656,f983])).
fof(f983,plain,(
  ( ! [X261] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X261),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X261,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f653])).
fof(f656,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f413])).
fof(f33699,plain,(
  ( ! [X685] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X685,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25139,f888])).
fof(f25139,plain,(
  ( ! [X685] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,happ(s(fun('type.num.num',bool),X685),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X685,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f658,f983])).
fof(f658,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f414])).
fof(f33698,plain,(
  ( ! [X741,X742] : (s(bool,X742) != s(bool,happ(s(fun('type.num.num',bool),X741),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X741,'type.num.num')) | s(bool,'T_0') = s(bool,X742)) ) | $spl272),
  inference(cnf_transformation,[],[f33698_D])).
fof(f33698_D,plain,(
  ( ! [X741,X742] : (s(bool,X742) != s(bool,happ(s(fun('type.num.num',bool),X741),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X741,'type.num.num')) | s(bool,'T_0') = s(bool,X742)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f25165,plain,(
  ( ! [X741,X742] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X742) != s(bool,happ(s(fun('type.num.num',bool),X741),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,X742) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X741,'type.num.num'))) ) | ($spl2 | $spl122)),
  inference(superposition,[],[f14168,f983])).
fof(f33697,plain,(
  ( ! [X743,X744] : (s(bool,X744) != s(bool,happ(s(fun('type.num.num',bool),X743),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X743,'type.num.num')) | p(s(bool,X744))) ) | $spl270),
  inference(cnf_transformation,[],[f33697_D])).
fof(f33697_D,plain,(
  ( ! [X743,X744] : (s(bool,X744) != s(bool,happ(s(fun('type.num.num',bool),X743),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X743,'type.num.num')) | p(s(bool,X744))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f25166,plain,(
  ( ! [X743,X744] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X744) != s(bool,happ(s(fun('type.num.num',bool),X743),s('type.num.num',sK12))) | p(s(bool,X744)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X743,'type.num.num'))) ) | ($spl2 | $spl122)),
  inference(superposition,[],[f14169,f983])).
fof(f25411,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X0,'type.num.num'))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f25385])).
fof(f25385,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X0,'type.num.num')) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X0,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f985,f983])).
fof(f985,plain,(
  ( ! [X263] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X263),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X263,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f656])).
fof(f33696,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X29,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25398,f888])).
fof(f25398,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X29,'type.num.num')) | p(s(bool,happ(s(fun('type.num.num',bool),X29),s('type.num.num',sK12))))) ) | $spl2),
  inference(superposition,[],[f985,f653])).
fof(f33695,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X30,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25399,f888])).
fof(f25399,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X30,'type.num.num')) | p(s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK12))))) ) | $spl2),
  inference(superposition,[],[f985,f655])).
fof(f33694,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X34,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25403,f985])).
fof(f25403,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X34,'type.num.num')) | s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X34),s('type.num.num',sK12)))) ) | $spl2),
  inference(superposition,[],[f985,f690])).
fof(f33693,plain,(
  ( ! [X57,X56,X55] : (s(bool,X56) != s(bool,X57) | p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X55) != s(bool,X57)) ) | $spl268),
  inference(cnf_transformation,[],[f33693_D])).
fof(f33693_D,plain,(
  ( ! [X57,X56,X55] : (s(bool,X56) != s(bool,X57) | p(s(bool,X55)) | p(s(bool,X56)) | s(bool,X55) != s(bool,X57)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f27031,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X53) != s(bool,X54) | s(bool,X52) != s(bool,X54) | p(s(bool,X53)) | p(s(bool,X52))) ) | $spl122),
  inference(superposition,[],[f14102,f653])).
fof(f14102,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X5) | s(bool,X3) != s(bool,X4) | s(bool,X4) != s(bool,X5) | p(s(bool,X3))) ) | $spl122),
  inference(superposition,[],[f13583,f655])).
fof(f27032,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X56) != s(bool,X57) | s(bool,X55) != s(bool,X57) | p(s(bool,X56)) | p(s(bool,X55))) ) | $spl122),
  inference(superposition,[],[f14102,f655])).
fof(f33692,plain,(
  ( ! [X68,X69,X67] : (s(bool,X68) != s(bool,X69) | p(s(bool,X68)) | s(bool,X67) != s(bool,X69)) ) | $spl266),
  inference(cnf_transformation,[],[f33692_D])).
fof(f33692_D,plain,(
  ( ! [X68,X69,X67] : (s(bool,X68) != s(bool,X69) | p(s(bool,X68)) | s(bool,X67) != s(bool,X69)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f33691,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X68) != s(bool,X69) | s(bool,X67) != s(bool,X69) | p(s(bool,X68))) ) | $spl122),
  inference(subsumption_resolution,[],[f27036,f14102])).
fof(f27036,plain,(
  ( ! [X68,X69,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X68) != s(bool,X69) | s(bool,X67) != s(bool,X69) | p(s(bool,X68)) | s(bool,'T_0') = s(bool,X67)) ) | $spl122),
  inference(superposition,[],[f14102,f690])).
fof(f33689,plain,(
  ( ! [X97,X98,X96] : (s(bool,X97) != s(bool,X98) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X96,'type.num.num')) | p(s(bool,X97)) | s(bool,X98) != s(bool,happ(s(fun('type.num.num',bool),X96),s('type.num.num',sK12)))) ) | $spl264),
  inference(cnf_transformation,[],[f33689_D])).
fof(f33689_D,plain,(
  ( ! [X97,X98,X96] : (s(bool,X97) != s(bool,X98) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X96,'type.num.num')) | p(s(bool,X97)) | s(bool,X98) != s(bool,happ(s(fun('type.num.num',bool),X96),s('type.num.num',sK12)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f27047,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X97) != s(bool,X98) | s(bool,X98) != s(bool,happ(s(fun('type.num.num',bool),X96),s('type.num.num',sK12))) | p(s(bool,X97)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X96,'type.num.num'))) ) | ($spl2 | $spl122)),
  inference(superposition,[],[f14102,f983])).
fof(f27475,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27474,f10990])).
fof(f27474,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | p(s(bool,X28))) ) | ($spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27400,f656])).
fof(f27400,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X28) | p(s(bool,X28))) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f653])).
fof(f27477,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X29))) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27476,f10990])).
fof(f27476,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | p(s(bool,X29))) ) | ($spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27401,f656])).
fof(f27401,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X29) | p(s(bool,X29))) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f655])).
fof(f27478,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X33)) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27461,f10990])).
fof(f27461,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl87 | $spl122)),
  inference(duplicate_literal_removal,[],[f27405])).
fof(f27405,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f690])).
fof(f27482,plain,(
  ( ! [X40] : (s(bool,'T_0') != s(bool,'F_0') | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X40,'type.num.num'))) ) | ($spl2 | $spl87 | $spl122)),
  inference(forward_demodulation,[],[f27481,f10990])).
fof(f27481,plain,(
  ( ! [X40] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X40,'type.num.num'))) ) | ($spl2 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27409,f985])).
fof(f27409,plain,(
  ( ! [X40] : (s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X40,'type.num.num'))) ) | ($spl2 | $spl87 | $spl122)),
  inference(superposition,[],[f14390,f983])).
fof(f33684,plain,(
  ( ! [X39,X37] : (s(bool,X37) != s(bool,X39) | ~p(s(bool,X37)) | s(bool,'T_0') = s(bool,X39)) ) | $spl262),
  inference(cnf_transformation,[],[f33684_D])).
fof(f33684_D,plain,(
  ( ! [X39,X37] : (s(bool,X37) != s(bool,X39) | ~p(s(bool,X37)) | s(bool,'T_0') = s(bool,X39)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f33687,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl122),
  inference(subsumption_resolution,[],[f14355,f657])).
fof(f14355,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl122),
  inference(superposition,[],[f14168,f484])).
fof(f484,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f324])).
fof(f33686,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl122),
  inference(subsumption_resolution,[],[f14356,f657])).
fof(f14356,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl122),
  inference(superposition,[],[f14168,f484])).
fof(f33685,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl122),
  inference(subsumption_resolution,[],[f14359,f657])).
fof(f14359,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | $spl122),
  inference(superposition,[],[f14168,f488])).
fof(f488,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X17,X18] : ((p(s(bool,X17)) => p(s(bool,X18))) => ((p(s(bool,X18)) => p(s(bool,X17))) => s(bool,X17) = s(bool,X18)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f33683,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) != s(bool,X39) | s(bool,'T_0') = s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl122),
  inference(subsumption_resolution,[],[f14360,f657])).
fof(f14360,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,X37) != s(bool,X39) | s(bool,'T_0') = s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl122),
  inference(superposition,[],[f14168,f488])).
fof(f33677,plain,(
  ( ! [X50,X48,X51] : (s(bool,X50) != s(bool,X51) | ~p(s(bool,X48)) | p(s(bool,X50)) | s(bool,X48) != s(bool,X51)) ) | $spl260),
  inference(cnf_transformation,[],[f33677_D])).
fof(f33677_D,plain,(
  ( ! [X50,X48,X51] : (s(bool,X50) != s(bool,X51) | ~p(s(bool,X48)) | p(s(bool,X50)) | s(bool,X48) != s(bool,X51)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f33680,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X28) != s(bool,X31) | p(s(bool,X30)) | ~p(s(bool,X28))) ) | $spl122),
  inference(subsumption_resolution,[],[f27025,f656])).
fof(f27025,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') != s(bool,X29) | s(bool,X30) != s(bool,X31) | s(bool,X28) != s(bool,X31) | p(s(bool,X30)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl122),
  inference(superposition,[],[f14102,f484])).
fof(f33679,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X34) != s(bool,X35) | s(bool,X32) != s(bool,X35) | p(s(bool,X34)) | ~p(s(bool,X32))) ) | $spl122),
  inference(subsumption_resolution,[],[f27026,f656])).
fof(f27026,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X34) != s(bool,X35) | s(bool,X32) != s(bool,X35) | p(s(bool,X34)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl122),
  inference(superposition,[],[f14102,f484])).
fof(f33678,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | s(bool,X44) != s(bool,X47) | p(s(bool,X46)) | ~p(s(bool,X44))) ) | $spl122),
  inference(subsumption_resolution,[],[f27029,f656])).
fof(f27029,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') != s(bool,X45) | s(bool,X46) != s(bool,X47) | s(bool,X44) != s(bool,X47) | p(s(bool,X46)) | ~p(s(bool,X44)) | ~p(s(bool,X45))) ) | $spl122),
  inference(superposition,[],[f14102,f488])).
fof(f33676,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,X49) | s(bool,X50) != s(bool,X51) | s(bool,X48) != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X48))) ) | $spl122),
  inference(subsumption_resolution,[],[f27030,f656])).
fof(f27030,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,X49) | s(bool,X50) != s(bool,X51) | s(bool,X48) != s(bool,X51) | p(s(bool,X50)) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl122),
  inference(superposition,[],[f14102,f488])).
fof(f33638,plain,(
  s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl250),
  inference(cnf_transformation,[],[f33638_D])).
fof(f33638_D,plain,(
  s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f33641,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl256),
  inference(cnf_transformation,[],[f33641_D])).
fof(f33641_D,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f33588,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | s(bool,X37) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f2480,f28244])).
fof(f28244,plain,(
  ( ! [X83,X81,X82] : (~sP2(X82,'T_0',X81) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | s(bool,'T_0') = s(bool,X81)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28225,f657])).
fof(f28225,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,X81) | ~p(s(bool,X82)) | s(bool,X81) != s(bool,X83) | p(s(bool,X81)) | ~sP2(X82,'T_0',X81)) ) | ($spl52 | $spl87 | $spl122)),
  inference(trivial_inequality_removal,[],[f28198])).
fof(f28198,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') = s(bool,X81) | ~p(s(bool,X82)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X81) != s(bool,X83) | p(s(bool,X81)) | ~sP2(X82,'T_0',X81)) ) | ($spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f27465,f13598])).
fof(f27465,plain,(
  ( ! [X10,X11,X9] : (~sP3(X11,X9,X10) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27391,f10990])).
fof(f27391,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X11)) | ~sP3(X11,X9,X10)) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f474])).
fof(f33640,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,X26) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl254),
  inference(cnf_transformation,[],[f33640_D])).
fof(f33640_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,X26) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f33587,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,X35) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f14459])).
fof(f14459,plain,(
  ( ! [X24,X23,X22] : (~sP2(X23,'T_0',X24) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,'F_0') != s(bool,X22)) ) | $spl52),
  inference(subsumption_resolution,[],[f14458,f695])).
fof(f14458,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP2(X23,'T_0',X24)) ) | $spl52),
  inference(subsumption_resolution,[],[f14457,f471])).
fof(f14457,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP2(X23,'T_0',X24) | ~sP0('T_0',X24,X23)) ) | $spl52),
  inference(subsumption_resolution,[],[f14440,f465])).
fof(f14440,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP2(X23,'T_0',X24) | ~sP1(X24,'T_0',X23) | ~sP0('T_0',X24,X23)) ) | $spl52),
  inference(resolution,[],[f5514,f457])).
fof(f5514,plain,(
  ( ! [X14,X12,X13] : (~sP3(X14,X12,'T_0') | s(bool,'F_0') != s(bool,X13) | ~p(s(bool,X14)) | s(bool,X12) != s(bool,X13)) ) | $spl52),
  inference(superposition,[],[f3282,f474])).
fof(f33586,plain,(
  ( ! [X33,X32] : (sP3(X32,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | s(bool,X33) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0',X32)) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f13719])).
fof(f13719,plain,(
  ( ! [X4,X5,X3] : (~sP2(X5,'T_0',X4) | sP3(X5,X4,'T_0') | s(bool,X3) != s(bool,X4) | ~sP1(X4,'T_0',X5)) ) | $spl52),
  inference(subsumption_resolution,[],[f13715,f695])).
fof(f13715,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | sP3(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP2(X5,'T_0',X4) | ~sP1(X4,'T_0',X5)) ) | $spl52),
  inference(resolution,[],[f13713,f457])).
fof(f33639,plain,(
  ( ! [X27] : (sP3(X27,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0')) ) | $spl252),
  inference(cnf_transformation,[],[f33639_D])).
fof(f33639_D,plain,(
  ( ! [X27] : (sP3(X27,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0')) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f33583,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | sP3(X27,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | s(bool,X26) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f13604])).
fof(f13604,plain,(
  ( ! [X4,X5,X3] : (~sP2(X5,'T_0',X4) | s(bool,'F_0') != s(bool,X3) | sP3(X5,X4,'T_0') | s(bool,X3) != s(bool,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f13603,f695])).
fof(f13603,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | sP3(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP2(X5,'T_0',X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f13600,f471])).
fof(f13600,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | sP3(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP2(X5,'T_0',X4) | ~sP0('T_0',X4,X5)) ) | $spl52),
  inference(resolution,[],[f13585,f457])).
fof(f13585,plain,(
  ( ! [X10,X11,X9] : (sP1(X10,'T_0',X11) | s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl52),
  inference(subsumption_resolution,[],[f13572,f465])).
fof(f13572,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | sP1(X10,'T_0',X11)) ) | $spl52),
  inference(resolution,[],[f5512,f454])).
fof(f33582,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f2480,f28239])).
fof(f28239,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X5,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28238,f469])).
fof(f28238,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP2(X4,X5,X3) | ~sP0(X5,X3,X4)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28237,f657])).
fof(f28237,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | ~sP2(X4,X5,X3) | ~sP0(X5,X3,X4)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28172,f27976])).
fof(f27976,plain,(
  ( ! [X4,X5,X3] : (sP1(X4,X3,X5) | s(bool,'T_0') = s(bool,X4)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27975,f466])).
fof(f27975,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | sP1(X4,X3,X5)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27925,f465])).
fof(f27925,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | p(s(bool,X5)) | sP1(X4,X3,X5)) ) | ($spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f27913,f454])).
fof(f27913,plain,(
  ( ! [X12,X10,X11] : (sP3(X10,X11,X12) | ~p(s(bool,X12)) | s(bool,'T_0') = s(bool,X11) | p(s(bool,X10))) ) | ($spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f27906,f463])).
fof(f27906,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X0,X1) | sP3(X2,X1,X0) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27905,f657])).
fof(f27905,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X1) | sP3(X2,X1,X0) | ~p(s(bool,X0)) | ~sP2(X2,X0,X1)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27885,f14662])).
fof(f27885,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X1) | sP3(X2,X1,X0) | ~p(s(bool,X0)) | ~sP2(X2,X0,X1) | ~sP0(X0,X1,X2)) ) | ($spl87 | $spl122)),
  inference(resolution,[],[f27873,f457])).
fof(f27873,plain,(
  ( ! [X4,X5,X3] : (sP1(X3,X4,X5) | s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27827,f465])).
fof(f27827,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP1(X3,X4,X5)) ) | ($spl87 | $spl122)),
  inference(resolution,[],[f27463,f454])).
fof(f27463,plain,(
  ( ! [X4,X5,X3] : (sP3(X5,X3,X4) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | ($spl87 | $spl122)),
  inference(forward_demodulation,[],[f27389,f10990])).
fof(f27389,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') = s(bool,X3) | sP3(X5,X3,X4) | p(s(bool,X5))) ) | ($spl87 | $spl122)),
  inference(superposition,[],[f14390,f472])).
fof(f28172,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | ~sP2(X4,X5,X3) | ~sP1(X3,X5,X4) | ~sP0(X5,X3,X4)) ) | ($spl87 | $spl122)),
  inference(resolution,[],[f27465,f458])).
fof(f33581,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f2480,f28236])).
fof(f28236,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28235,f471])).
fof(f28235,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP0(X2,X0,X1)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28234,f657])).
fof(f28234,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP0(X2,X0,X1)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f28171,f27976])).
fof(f28171,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP1(X0,X2,X1) | ~sP0(X2,X0,X1)) ) | ($spl87 | $spl122)),
  inference(resolution,[],[f27465,f457])).
fof(f33636,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X11) | ~p(s(bool,X11))) ) | $spl248),
  inference(cnf_transformation,[],[f33636_D])).
fof(f33636_D,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X11) | ~p(s(bool,X11))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f33580,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f2480,f27906])).
fof(f33579,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f15058])).
fof(f15058,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X5,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X3)) ) | $spl52),
  inference(subsumption_resolution,[],[f15057,f469])).
fof(f15057,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP2(X4,X5,X3) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f15056,f465])).
fof(f15056,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP2(X4,X5,X3) | ~sP1(X3,X5,X4) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f15015,f652])).
fof(f15015,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | ~sP2(X4,X5,X3) | ~sP1(X3,X5,X4) | ~sP0(X5,X3,X4)) ) | $spl52),
  inference(resolution,[],[f5534,f458])).
fof(f5534,plain,(
  ( ! [X10,X11,X9] : (~sP3(X11,X9,X10) | s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl52),
  inference(superposition,[],[f3282,f474])).
fof(f33578,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f15055])).
fof(f15055,plain,(
  ( ! [X2,X0,X1] : (~sP2(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X0)) ) | $spl52),
  inference(subsumption_resolution,[],[f15054,f471])).
fof(f15054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP0(X2,X0,X1)) ) | $spl52),
  inference(subsumption_resolution,[],[f15053,f465])).
fof(f15053,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP1(X0,X2,X1) | ~sP0(X2,X0,X1)) ) | $spl52),
  inference(subsumption_resolution,[],[f15014,f657])).
fof(f15014,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~sP2(X1,X2,X0) | ~sP1(X0,X2,X1) | ~sP0(X2,X0,X1)) ) | $spl52),
  inference(resolution,[],[f5534,f457])).
fof(f33575,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl12 | $spl52)),
  inference(resolution,[],[f2480,f14560])).
fof(f14560,plain,(
  ( ! [X4,X5,X3] : (~sP2(X5,X3,X4) | sP3(X5,X4,X3) | ~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f14559,f657])).
fof(f14559,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP3(X5,X4,X3) | ~p(s(bool,X3)) | ~sP2(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f14552,f14536])).
fof(f14552,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP3(X5,X4,X3) | ~p(s(bool,X3)) | ~sP2(X5,X3,X4) | ~sP0(X3,X4,X5)) ) | $spl52),
  inference(resolution,[],[f14537,f457])).
fof(f14537,plain,(
  ( ! [X10,X11,X9] : (sP1(X9,X10,X11) | s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl52),
  inference(subsumption_resolution,[],[f14521,f465])).
fof(f14521,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') != s(bool,X10) | p(s(bool,X11)) | sP1(X9,X10,X11)) ) | $spl52),
  inference(resolution,[],[f5532,f454])).
fof(f5875,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl82),
  inference(cnf_transformation,[],[f5875_D])).
fof(f5875_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f33411,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f691])).
fof(f33320,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X30) != s(bool,X31)) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f452])).
fof(f33317,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f451])).
fof(f33316,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f450])).
fof(f33315,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X15) != s(bool,X16)) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f448])).
fof(f33313,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl7 | $spl32)),
  inference(resolution,[],[f32849,f446])).
fof(f33190,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5753,f5459])).
fof(f5459,plain,(
  ( ! [X4,X5] : (~sP0(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl52),
  inference(subsumption_resolution,[],[f5453,f463])).
fof(f5453,plain,(
  ( ! [X4,X5] : (~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl52),
  inference(resolution,[],[f4698,f475])).
fof(f32931,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,X26) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl218),
  inference(cnf_transformation,[],[f32931_D])).
fof(f32931_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,X26) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f33187,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5753,f14915])).
fof(f14915,plain,(
  ( ! [X4,X5,X3] : (~sP0(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f14900])).
fof(f14900,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP0(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl52),
  inference(resolution,[],[f14810,f475])).
fof(f14810,plain,(
  ( ! [X10,X11,X9] : (sP3(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP0(X11,X9,X10)) ) | $spl52),
  inference(subsumption_resolution,[],[f14806,f463])).
fof(f14806,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP3(X10,X9,X11) | p(s(bool,X10)) | ~sP2(X10,X11,X9) | ~sP0(X11,X9,X10)) ) | $spl52),
  inference(resolution,[],[f14799,f459])).
fof(f14799,plain,(
  ( ! [X10,X11,X9] : (sP1(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl52),
  inference(subsumption_resolution,[],[f14798,f466])).
fof(f14798,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP1(X10,X9,X11)) ) | $spl52),
  inference(subsumption_resolution,[],[f14782,f465])).
fof(f14782,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP1(X10,X9,X11)) ) | $spl52),
  inference(resolution,[],[f14581,f454])).
fof(f14581,plain,(
  ( ! [X10,X8,X9] : (sP3(X8,X9,X10) | ~p(s(bool,X10)) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl52),
  inference(resolution,[],[f14560,f463])).
fof(f33186,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5753,f14870])).
fof(f14870,plain,(
  ( ! [X4,X5,X3] : (~sP0(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f14855])).
fof(f14855,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP0(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl52),
  inference(resolution,[],[f14814,f475])).
fof(f14814,plain,(
  ( ! [X10,X8,X9] : (sP3(X9,X8,X10) | p(s(bool,X8)) | ~sP0(X10,X8,X9) | p(s(bool,X9))) ) | $spl52),
  inference(resolution,[],[f14809,f463])).
fof(f14809,plain,(
  ( ! [X6,X8,X7] : (~sP2(X7,X8,X6) | p(s(bool,X6)) | sP3(X7,X6,X8) | ~sP0(X8,X6,X7)) ) | $spl52),
  inference(subsumption_resolution,[],[f14805,f653])).
fof(f14805,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP3(X7,X6,X8) | p(s(bool,X6)) | ~sP2(X7,X8,X6) | ~sP0(X8,X6,X7)) ) | $spl52),
  inference(resolution,[],[f14799,f458])).
fof(f3816,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl56),
  inference(cnf_transformation,[],[f3816_D])).
fof(f3816_D,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f33181,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X13) | p(s(bool,X12))) ) | $spl236),
  inference(cnf_transformation,[],[f33181_D])).
fof(f33181_D,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X13) | p(s(bool,X12))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f33160,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X13) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f5752,f28005])).
fof(f28005,plain,(
  ( ! [X6,X8,X7] : (~sP0(X8,X6,X7) | sP3(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) ) | ($spl52 | $spl87 | $spl122)),
  inference(subsumption_resolution,[],[f27990,f463])).
fof(f27990,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,X6) | sP3(X7,X6,X8) | p(s(bool,X7)) | ~sP2(X7,X8,X6) | ~sP0(X8,X6,X7)) ) | ($spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f27976,f459])).
fof(f5752,plain,(
  ( ! [X52,X53] : (sP0(X52,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X53)) ) | $spl32),
  inference(resolution,[],[f1645,f470])).
fof(f18294,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl148),
  inference(cnf_transformation,[],[f18294_D])).
fof(f18294_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f33159,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5752,f14915])).
fof(f18047,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl144),
  inference(cnf_transformation,[],[f18047_D])).
fof(f18047_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f33065,plain,(
  ( ! [X175,X174] : (p(s(bool,X174)) | p(s(bool,X175)) | s(bool,X175) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f691])).
fof(f32974,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X24) != s(bool,X25)) ) | $spl78),
  inference(resolution,[],[f5728,f452])).
fof(f32971,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f451])).
fof(f32970,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f450])).
fof(f32969,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X9) != s(bool,X10)) ) | $spl78),
  inference(resolution,[],[f5728,f448])).
fof(f32967,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f446])).
fof(f32906,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | s(bool,X37) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f5747,f28244])).
fof(f32905,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,X35) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f14459])).
fof(f32904,plain,(
  ( ! [X33,X32] : (sP3(X32,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | s(bool,X33) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X32)) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f13719])).
fof(f32930,plain,(
  ( ! [X27] : (sP3(X27,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0')) ) | $spl216),
  inference(cnf_transformation,[],[f32930_D])).
fof(f32930_D,plain,(
  ( ! [X27] : (sP3(X27,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0')) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f32901,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | sP3(X27,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | s(bool,X26) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f13604])).
fof(f32900,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f5747,f28239])).
fof(f32899,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f5747,f28236])).
fof(f32929,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X11) | ~p(s(bool,X11))) ) | $spl214),
  inference(cnf_transformation,[],[f32929_D])).
fof(f32929_D,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X11) | ~p(s(bool,X11))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f32898,plain,(
  ( ! [X21,X20] : (sP3(X20,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52 | $spl87 | $spl122)),
  inference(resolution,[],[f5747,f27906])).
fof(f32897,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f15058])).
fof(f32896,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f15055])).
fof(f32893,plain,(
  ( ! [X10,X11] : (sP3(X10,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5747,f14560])).
fof(f3817,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl57),
  inference(cnf_transformation,[],[f3817_D])).
fof(f3817_D,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f20433,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl152),
  inference(duplicate_literal_removal,[],[f20411])).
fof(f20411,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl152),
  inference(resolution,[],[f18322,f475])).
fof(f1645,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl32),
  inference(cnf_transformation,[],[f1645_D])).
fof(f1645_D,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f32841,plain,(
  ( ! [X4] : (sP3(X4,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP2(X4,'T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl206),
  inference(cnf_transformation,[],[f32841_D])).
fof(f32841_D,plain,(
  ( ! [X4] : (sP3(X4,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP2(X4,'T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f20459,plain,(
  ( ! [X4] : (sP3(X4,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~sP2(X4,'T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl152),
  inference(subsumption_resolution,[],[f20443,f467])).
fof(f20443,plain,(
  ( ! [X4] : (sP3(X4,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~sP2(X4,'T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X4)) ) | $spl152),
  inference(resolution,[],[f20434,f458])).
fof(f20434,plain,(
  ( ! [X4] : (sP0('T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X4)) ) | $spl152),
  inference(subsumption_resolution,[],[f20406,f469])).
fof(f20406,plain,(
  ( ! [X4] : (p(s(bool,X4)) | sP0('T_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X4)) ) | $spl152),
  inference(resolution,[],[f18322,f453])).
fof(f1646,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl33),
  inference(cnf_transformation,[],[f1646_D])).
fof(f1646_D,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f32839,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | $spl204),
  inference(cnf_transformation,[],[f32839_D])).
fof(f32839_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f30294,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl12),
  inference(resolution,[],[f2461,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0)))))),
  inference(flattening,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) | p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0)))))),
  inference(ennf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X0)))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X11,X20,X27] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27))))) => p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.infer_eComplete.t_compat_pure_add_constraints_2')).
fof(f32770,plain,(
  ( ! [X507,X506] : (sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X506,X507) | ~sP8(X506,X507,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl198),
  inference(cnf_transformation,[],[f32770_D])).
fof(f32770_D,plain,(
  ( ! [X507,X506] : (sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X506,X507) | ~sP8(X506,X507,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f32812,plain,(
  ( ! [X700,X702,X701] : (p(s(bool,X700)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X701,X702) | ~sP8(X701,X702,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32050,f656])).
fof(f32050,plain,(
  ( ! [X700,X702,X701] : (p(s(bool,X700)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X701,X702) | ~sP8(X701,X702,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X700)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f694,f29922])).
fof(f29922,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl10 | $spl33)),
  inference(trivial_inequality_removal,[],[f29895])).
fof(f29895,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X33) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X33)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f29850,f690])).
fof(f29850,plain,(
  ( ! [X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,X51) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f29849,f695])).
fof(f29849,plain,(
  ( ! [X51] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X51) | s(bool,X51) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(forward_demodulation,[],[f29793,f1582])).
fof(f29793,plain,(
  ( ! [X51] : (s(bool,'F_0') != s(bool,X51) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | s(bool,X51) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(resolution,[],[f19225,f1507])).
fof(f1507,plain,(
  ( ! [X57,X58] : (~sP3('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X58,X57) | s(bool,X57) = s(bool,X58)) ) | $spl10),
  inference(resolution,[],[f709,f474])).
fof(f19225,plain,(
  ( ! [X165,X164] : (sP3(X165,X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') != s(bool,X164) | ~p(s(bool,X165))) ) | $spl33),
  inference(superposition,[],[f473,f18145])).
fof(f18145,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl33),
  inference(resolution,[],[f1646,f653])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f557])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP8(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP8(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f354])).
fof(f354,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP8(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (sP8(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f32793,plain,(
  ( ! [X631,X632,X630] : (~p(s(bool,happ(s(fun(X630,bool),X631),s(X630,X632)))) | p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK28('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X631,X630)))))) ) | $spl202),
  inference(cnf_transformation,[],[f32793_D])).
fof(f32793_D,plain,(
  ( ! [X631,X632,X630] : (~p(s(bool,happ(s(fun(X630,bool),X631),s(X630,X632)))) | p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK28('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X631,X630)))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f32792,plain,(
  ( ! [X629,X631,X632,X630] : (p(s(bool,X629)) | ~p(s(bool,happ(s(fun(X630,bool),X631),s(X630,X632)))) | p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK28('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X631,X630)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32008,f656])).
fof(f32008,plain,(
  ( ! [X629,X631,X632,X630] : (p(s(bool,X629)) | ~p(s(bool,happ(s(fun(X630,bool),X631),s(X630,X632)))) | p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK28('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X631,X630))))) | s(bool,'T_0') = s(bool,X629)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f606,f29922])).
fof(f606,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1,X2] : (((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f377])).
fof(f377,plain,(
  ! [X0,X1,X2] : (((? [X6] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ! [X5] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f376])).
fof(f376,plain,(
  ! [X0,X1,X2] : (((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0,X1,X2] : (((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X19] : ! [X37,X81] : ((? [X7] : p(s(bool,happ(s(fun(X19,bool),X37),s(X19,X7)))) => p(s(bool,X81))) <=> ! [X7] : (p(s(bool,happ(s(fun(X19,bool),X37),s(X19,X7)))) => p(s(bool,X81))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.PULL_EXISTS|split|0')).
fof(f32791,plain,(
  ( ! [X627,X628,X626] : (~p(s(bool,happ(s(fun(X626,bool),X627),s(X626,X628)))) | p(s(bool,happ(s(fun(X626,bool),X627),s(X626,sK27('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X627,X626)))))) ) | $spl200),
  inference(cnf_transformation,[],[f32791_D])).
fof(f32791_D,plain,(
  ( ! [X627,X628,X626] : (~p(s(bool,happ(s(fun(X626,bool),X627),s(X626,X628)))) | p(s(bool,happ(s(fun(X626,bool),X627),s(X626,sK27('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X627,X626)))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f32790,plain,(
  ( ! [X625,X627,X628,X626] : (p(s(bool,X625)) | ~p(s(bool,happ(s(fun(X626,bool),X627),s(X626,X628)))) | p(s(bool,happ(s(fun(X626,bool),X627),s(X626,sK27('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X627,X626)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32007,f656])).
fof(f32007,plain,(
  ( ! [X625,X627,X628,X626] : (p(s(bool,X625)) | ~p(s(bool,happ(s(fun(X626,bool),X627),s(X626,X628)))) | p(s(bool,happ(s(fun(X626,bool),X627),s(X626,sK27('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X627,X626))))) | s(bool,'T_0') = s(bool,X625)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f604,f29922])).
fof(f604,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f32772,plain,(
  ( ! [X511,X512,X513] : (p(s(bool,X511)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X512,X513) | ~sP8(X512,X513,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32771,f656])).
fof(f32771,plain,(
  ( ! [X511,X512,X513] : (p(s(bool,X511)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X512,X513) | ~sP8(X512,X513,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X511)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31961,f18119])).
fof(f18119,plain,(
  ( ! [X136,X137] : (sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X136,X137) | p(s(bool,X137))) ) | $spl33),
  inference(resolution,[],[f1646,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f358])).
fof(f358,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f174,f289,f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : (sP8(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X58,X56,X57] : ((p(s(bool,X57)) <=> (p(s(bool,X56)) & p(s(bool,X58)))) <=> ((p(s(bool,X57)) | ~p(s(bool,X56)) | ~p(s(bool,X58))) & (p(s(bool,X56)) | ~p(s(bool,X57))) & (p(s(bool,X58)) | ~p(s(bool,X57)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.sat.sat_23_prove_ltb')).
fof(f31961,plain,(
  ( ! [X511,X512,X513] : (p(s(bool,X511)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X512,X513) | ~p(s(bool,X513)) | ~sP8(X512,X513,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X511)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f556,f29922])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f32769,plain,(
  ( ! [X505,X507,X506] : (p(s(bool,X505)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X506,X507) | ~sP8(X506,X507,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32768,f656])).
fof(f32768,plain,(
  ( ! [X505,X507,X506] : (p(s(bool,X505)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X506,X507) | ~sP8(X506,X507,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X505)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31959,f18121])).
fof(f18121,plain,(
  ( ! [X140,X141] : (sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X140,X141) | p(s(bool,X140))) ) | $spl33),
  inference(resolution,[],[f1646,f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f359])).
fof(f31959,plain,(
  ( ! [X505,X507,X506] : (p(s(bool,X505)) | sP9('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X506,X507) | ~p(s(bool,X506)) | ~sP8(X506,X507,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X505)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f555,f29922])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP8(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f355])).
fof(f32757,plain,(
  ( ! [X426,X425] : (sP7(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X426)) | ~sP6(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl196),
  inference(cnf_transformation,[],[f32757_D])).
fof(f32757_D,plain,(
  ( ! [X426,X425] : (sP7(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X426)) | ~sP6(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f32756,plain,(
  ( ! [X424,X426,X425] : (p(s(bool,X424)) | sP7(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP6(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X426))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31932,f656])).
fof(f31932,plain,(
  ( ! [X424,X426,X425] : (p(s(bool,X424)) | sP7(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP6(X425,X426,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X426)) | s(bool,'T_0') = s(bool,X424)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f540,f29922])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f346])).
fof(f346,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f32755,plain,(
  ( ! [X422,X423] : (sP7(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | p(s(bool,X423)) | ~sP6(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423)) ) | $spl194),
  inference(cnf_transformation,[],[f32755_D])).
fof(f32755_D,plain,(
  ( ! [X422,X423] : (sP7(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | p(s(bool,X423)) | ~sP6(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f32754,plain,(
  ( ! [X422,X421,X423] : (p(s(bool,X421)) | sP7(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | ~sP6(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | p(s(bool,X423))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31931,f656])).
fof(f31931,plain,(
  ( ! [X422,X421,X423] : (p(s(bool,X421)) | sP7(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | ~sP6(X422,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X423) | p(s(bool,X423)) | s(bool,'T_0') = s(bool,X421)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f540,f29922])).
fof(f32740,plain,(
  ( ! [X333,X332] : (sP5(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~p(s(bool,X333)) | ~sP4(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333)) ) | $spl192),
  inference(cnf_transformation,[],[f32740_D])).
fof(f32740_D,plain,(
  ( ! [X333,X332] : (sP5(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~p(s(bool,X333)) | ~sP4(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f32739,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,X331)) | sP5(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~sP4(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~p(s(bool,X333))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31901,f656])).
fof(f31901,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,X331)) | sP5(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~sP4(X332,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X333) | ~p(s(bool,X333)) | s(bool,'T_0') = s(bool,X331)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f523,f29922])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~sP4(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP4(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP4(X0,X1,X2) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f338])).
fof(f338,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f337])).
fof(f337,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP4(X2,X0,X1)) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP4(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f32696,plain,(
  ( ! [X167,X168] : (sP3('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X167,X168) | ~sP0(X168,X167,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1(X167,X168,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl190),
  inference(cnf_transformation,[],[f32696_D])).
fof(f32696_D,plain,(
  ( ! [X167,X168] : (sP3('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X167,X168) | ~sP0(X168,X167,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1(X167,X168,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f32695,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,X166)) | sP3('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X167,X168) | ~sP1(X167,X168,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP0(X168,X167,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32694,f1669])).
fof(f1669,plain,(
  ( ! [X35,X34] : (sP2('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X34,X35)) ) | $spl33),
  inference(resolution,[],[f1646,f463])).
fof(f32694,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,X166)) | sP3('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X167,X168) | ~sP2('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X168,X167) | ~sP1(X167,X168,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP0(X168,X167,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31838,f656])).
fof(f31838,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,X166)) | sP3('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X167,X168) | ~sP2('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X168,X167) | ~sP1(X167,X168,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP0(X168,X167,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') = s(bool,X166)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f459,f29922])).
fof(f32693,plain,(
  ( ! [X165,X164] : (sP3(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165) | ~sP0(X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X164) | ~sP2(X164,X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl188),
  inference(cnf_transformation,[],[f32693_D])).
fof(f32693_D,plain,(
  ( ! [X165,X164] : (sP3(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165) | ~sP0(X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X164) | ~sP2(X164,X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f32692,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X163)) | sP3(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165) | ~sP2(X164,X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP0(X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X164)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f32691,f1672])).
fof(f1672,plain,(
  ( ! [X41,X40] : (sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X40,X41)) ) | $spl33),
  inference(resolution,[],[f1646,f467])).
fof(f32691,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X163)) | sP3(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165) | ~sP2(X164,X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165,X164) | ~sP0(X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X164)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f31837,f656])).
fof(f31837,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X163)) | sP3(X164,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165) | ~sP2(X164,X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X165,X164) | ~sP0(X165,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X164) | s(bool,'T_0') = s(bool,X163)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f458,f29922])).
fof(f31039,plain,(
  ( ! [X23,X22] : (sP3(X22,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),'T_0') | p(s(bool,X22)) | s(bool,X23) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30745,f13720])).
fof(f13720,plain,(
  ( ! [X10,X11,X9] : (~sP1(X10,'T_0',X11) | sP3(X11,X10,'T_0') | p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl52),
  inference(subsumption_resolution,[],[f13717,f463])).
fof(f13717,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | sP3(X11,X10,'T_0') | p(s(bool,X11)) | ~sP2(X11,'T_0',X10) | ~sP1(X10,'T_0',X11)) ) | $spl52),
  inference(resolution,[],[f13713,f459])).
fof(f31035,plain,(
  ( ! [X17,X16] : (sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X17) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl87 | $spl122 | $spl161)),
  inference(resolution,[],[f30745,f27637])).
fof(f31034,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))),X15) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl33 | $spl52 | $spl161)),
  inference(resolution,[],[f30745,f26452])).
fof(f26452,plain,(
  ( ! [X17,X15,X16] : (~sP1(X15,X17,X16) | sP3(X16,X15,X17) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26443,f463])).
fof(f26443,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | sP3(X16,X15,X17) | p(s(bool,X16)) | ~sP2(X16,X17,X15) | ~sP1(X15,X17,X16)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26421,f459])).
fof(f26421,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,X0,X2) | s(bool,'F_0') = s(bool,X0)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26420,f469])).
fof(f26420,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X2)) | sP0(X1,X0,X2)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26377,f26266])).
fof(f26266,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,X0,X2) | s(bool,'F_0') = s(bool,X1)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26265,f469])).
fof(f26265,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | sP0(X1,X0,X2)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26218,f470])).
fof(f26218,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X1) | p(s(bool,X2)) | sP0(X1,X0,X2)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26207,f453])).
fof(f26207,plain,(
  ( ! [X12,X10,X11] : (sP3(X10,X11,X12) | p(s(bool,X11)) | s(bool,'F_0') = s(bool,X12) | p(s(bool,X10))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26201,f463])).
fof(f26201,plain,(
  ( ! [X14,X12,X13] : (~sP2(X14,X13,X12) | sP3(X14,X12,X13) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26200,f467])).
fof(f26200,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X13) | sP3(X14,X12,X13) | p(s(bool,X12)) | ~sP2(X14,X13,X12) | ~sP1(X12,X13,X14)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26189,f653])).
fof(f26189,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,'F_0') = s(bool,X13) | sP3(X14,X12,X13) | p(s(bool,X12)) | ~sP2(X14,X13,X12) | ~sP1(X12,X13,X14)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26113,f458])).
fof(f26113,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26068,f469])).
fof(f26068,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2)) | sP0(X0,X1,X2)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26032,f453])).
fof(f26032,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | ($spl33 | $spl52)),
  inference(forward_demodulation,[],[f25957,f18145])).
fof(f25957,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') = s(bool,X0) | sP3(X2,X1,X0) | p(s(bool,X2))) ) | ($spl33 | $spl52)),
  inference(superposition,[],[f19467,f472])).
fof(f19467,plain,(
  ( ! [X444] : (s(bool,X444) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') = s(bool,X444)) ) | ($spl33 | $spl52)),
  inference(trivial_inequality_removal,[],[f19414])).
fof(f19414,plain,(
  ( ! [X444] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X444) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') = s(bool,X444)) ) | ($spl33 | $spl52)),
  inference(superposition,[],[f5509,f18145])).
fof(f26377,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2)) | sP0(X1,X0,X2)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26033,f453])).
fof(f26033,plain,(
  ( ! [X4,X5,X3] : (sP3(X5,X3,X4) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5))) ) | ($spl33 | $spl52)),
  inference(forward_demodulation,[],[f25958,f18145])).
fof(f25958,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') = s(bool,X3) | sP3(X5,X3,X4) | p(s(bool,X5))) ) | ($spl33 | $spl52)),
  inference(superposition,[],[f19467,f472])).
fof(f31062,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl182),
  inference(cnf_transformation,[],[f31062_D])).
fof(f31062_D,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f31032,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30745,f14893])).
fof(f14893,plain,(
  ( ! [X4,X5,X3] : (~sP1(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f14878])).
fof(f14878,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP1(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl52),
  inference(resolution,[],[f14674,f475])).
fof(f14674,plain,(
  ( ! [X10,X11,X9] : (sP3(X10,X11,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10)) | ~sP1(X11,X9,X10)) ) | $spl52),
  inference(subsumption_resolution,[],[f14671,f463])).
fof(f14671,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP3(X10,X11,X9) | p(s(bool,X10)) | ~sP2(X10,X9,X11) | ~sP1(X11,X9,X10)) ) | $spl52),
  inference(resolution,[],[f14662,f459])).
fof(f30874,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl164),
  inference(cnf_transformation,[],[f30874_D])).
fof(f30874_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f31031,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30745,f14774])).
fof(f14774,plain,(
  ( ! [X4,X5,X3] : (~sP1(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f14758])).
fof(f14758,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP1(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl52),
  inference(resolution,[],[f14678,f475])).
fof(f14678,plain,(
  ( ! [X10,X8,X9] : (sP3(X9,X10,X8) | ~p(s(bool,X8)) | ~sP1(X10,X8,X9) | p(s(bool,X9))) ) | $spl52),
  inference(resolution,[],[f14673,f463])).
fof(f14673,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X3,X5) | ~p(s(bool,X3)) | sP3(X4,X5,X3) | ~sP1(X5,X3,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f14669,f657])).
fof(f14669,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP3(X4,X5,X3) | ~p(s(bool,X3)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | $spl52),
  inference(resolution,[],[f14662,f457])).
fof(f30998,plain,(
  ( ! [X12,X13] : (sP3(X12,X13,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl33 | $spl52 | $spl161)),
  inference(resolution,[],[f30744,f26291])).
fof(f26291,plain,(
  ( ! [X17,X15,X16] : (~sP1(X17,X15,X16) | sP3(X16,X17,X15) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26283,f463])).
fof(f26283,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | sP3(X16,X17,X15) | p(s(bool,X16)) | ~sP2(X16,X15,X17) | ~sP1(X17,X15,X16)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26266,f459])).
fof(f30997,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,X11) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30744,f14893])).
fof(f30901,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~p(s(bool,X36))) ) | ($spl59 | $spl161)),
  inference(resolution,[],[f30741,f5659])).
fof(f5659,plain,(
  ( ! [X14,X15] : (~sP2(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl59),
  inference(subsumption_resolution,[],[f5658,f469])).
fof(f5658,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl59),
  inference(subsumption_resolution,[],[f5657,f623])).
fof(f5657,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl59),
  inference(subsumption_resolution,[],[f5650,f467])).
fof(f5650,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP1('F_0',X15,X14) | ~sP0(X15,'F_0',X14)) ) | $spl59),
  inference(resolution,[],[f5023,f458])).
fof(f5023,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl59),
  inference(superposition,[],[f3820,f474])).
fof(f30900,plain,(
  ( ! [X35] : (sP3(X35,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30741,f5345])).
fof(f5345,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl52),
  inference(subsumption_resolution,[],[f5344,f666])).
fof(f5344,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) ) | $spl52),
  inference(subsumption_resolution,[],[f5340,f467])).
fof(f5340,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) ) | $spl52),
  inference(resolution,[],[f4623,f458])).
fof(f30899,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | ~p(s(bool,X34))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30741,f5458])).
fof(f5458,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl52),
  inference(subsumption_resolution,[],[f5452,f469])).
fof(f5452,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~sP0(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl52),
  inference(resolution,[],[f4698,f474])).
fof(f30936,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'F_0') = s(bool,X30)) ) | $spl176),
  inference(cnf_transformation,[],[f30936_D])).
fof(f30936_D,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'F_0') = s(bool,X30)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f30898,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | s(bool,'F_0') = s(bool,X33) | s(bool,X33) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl10 | $spl33 | $spl161)),
  inference(resolution,[],[f30741,f30521])).
fof(f30521,plain,(
  ( ! [X23,X21,X22] : (~sP2(X23,X21,'F_0') | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,X22) | s(bool,X21) != s(bool,X22)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f30520,f469])).
fof(f30520,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,X22) | ~sP2(X23,X21,'F_0') | ~sP0(X21,'F_0',X23)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f30519,f623])).
fof(f30519,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,X22) | p(s(bool,'F_0')) | ~sP2(X23,X21,'F_0') | ~sP0(X21,'F_0',X23)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f30469,f467])).
fof(f30469,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X22) | ~p(s(bool,X23)) | s(bool,'F_0') = s(bool,X22) | p(s(bool,'F_0')) | ~sP2(X23,X21,'F_0') | ~sP1('F_0',X21,X23) | ~sP0(X21,'F_0',X23)) ) | ($spl10 | $spl33)),
  inference(resolution,[],[f29930,f458])).
fof(f29930,plain,(
  ( ! [X17,X15,X16] : (~sP3(X17,'F_0',X15) | s(bool,X15) != s(bool,X16) | ~p(s(bool,X17)) | s(bool,'F_0') = s(bool,X16)) ) | ($spl10 | $spl33)),
  inference(forward_demodulation,[],[f29859,f18145])).
fof(f29859,plain,(
  ( ! [X17,X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,X16) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X17)) | ~sP3(X17,'F_0',X15)) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f29850,f474])).
fof(f30897,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | sP3(X31,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl10 | $spl33 | $spl52 | $spl161)),
  inference(resolution,[],[f30741,f30230])).
fof(f30230,plain,(
  ( ! [X6,X7,X5] : (~sP2(X7,X6,'F_0') | s(bool,X5) != s(bool,X6) | sP3(X7,'F_0',X6) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f30229,f4693])).
fof(f30229,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,X5) != s(bool,X6) | sP3(X7,'F_0',X6) | ~sP2(X7,X6,'F_0') | ~sP1('F_0',X6,X7)) ) | ($spl10 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f30216,f9801])).
fof(f30216,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,X5) != s(bool,X6) | sP3(X7,'F_0',X6) | ~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | ~sP1('F_0',X6,X7)) ) | ($spl10 | $spl33)),
  inference(resolution,[],[f30201,f457])).
fof(f30201,plain,(
  ( ! [X10,X8,X9] : (sP0(X8,'F_0',X10) | s(bool,'F_0') = s(bool,X9) | s(bool,X8) != s(bool,X9)) ) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f30162,f469])).
fof(f30162,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X9) | p(s(bool,X10)) | sP0(X8,'F_0',X10)) ) | ($spl10 | $spl33)),
  inference(resolution,[],[f29928,f453])).
fof(f29928,plain,(
  ( ! [X10,X11,X9] : (sP3(X11,'F_0',X9) | s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | p(s(bool,X11))) ) | ($spl10 | $spl33)),
  inference(forward_demodulation,[],[f29857,f18145])).
fof(f29857,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | sP3(X11,'F_0',X9) | p(s(bool,X11))) ) | ($spl10 | $spl33)),
  inference(superposition,[],[f29850,f472])).
fof(f30896,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl52 | $spl87 | $spl122 | $spl161)),
  inference(resolution,[],[f30741,f28239])).
fof(f30893,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl33 | $spl52 | $spl161)),
  inference(resolution,[],[f30741,f26757])).
fof(f26757,plain,(
  ( ! [X4,X5,X3] : (~sP2(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X3)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26756,f465])).
fof(f26756,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26755,f652])).
fof(f26755,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4)) ) | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f26695,f26113])).
fof(f26695,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP2(X4,X3,X5) | ~sP1(X5,X3,X4) | ~sP0(X3,X5,X4)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f26034,f458])).
fof(f26034,plain,(
  ( ! [X6,X8,X7] : (~sP3(X8,X7,X6) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | ($spl33 | $spl52)),
  inference(forward_demodulation,[],[f25959,f18145])).
fof(f25959,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X8)) | ~sP3(X8,X7,X6)) ) | ($spl33 | $spl52)),
  inference(superposition,[],[f19467,f474])).
fof(f30892,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X21)) | s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl33 | $spl52 | $spl161)),
  inference(resolution,[],[f30741,f26201])).
fof(f30891,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30741,f15058])).
fof(f30886,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | ($spl52 | $spl161)),
  inference(resolution,[],[f30741,f14550])).
fof(f30827,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl161),
  inference(resolution,[],[f30723,f691])).
fof(f30736,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,X21) != s(bool,X23)) ) | $spl161),
  inference(resolution,[],[f30723,f452])).
fof(f30734,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl161),
  inference(resolution,[],[f30723,f451])).
fof(f30732,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl161),
  inference(resolution,[],[f30723,f450])).
fof(f30731,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,X6) != s(bool,X8)) ) | $spl161),
  inference(resolution,[],[f30723,f448])).
fof(f30729,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl161),
  inference(resolution,[],[f30723,f446])).
fof(f30724,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl162),
  inference(cnf_transformation,[],[f30724_D])).
fof(f30724_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f30663,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f2462,f691])).
fof(f30572,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X29) != s(bool,X30)) ) | $spl12),
  inference(resolution,[],[f2462,f452])).
fof(f30569,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f2462,f451])).
fof(f30568,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f2462,f450])).
fof(f30567,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X14) != s(bool,X15)) ) | $spl12),
  inference(resolution,[],[f2462,f448])).
fof(f30565,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f2462,f446])).
fof(f30723,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) | $spl161),
  inference(cnf_transformation,[],[f30723_D])).
fof(f30723_D,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f1232,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))) | $spl26),
  inference(cnf_transformation,[],[f1232_D])).
fof(f1232_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f30561,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) | $spl12),
  inference(resolution,[],[f2462,f678])).
fof(f30546,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X17) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X18,'type.num.num')) | s(bool,X17) != s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK12)))) ) | $spl158),
  inference(cnf_transformation,[],[f30546_D])).
fof(f30546_D,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X17) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X18,'type.num.num')) | s(bool,X17) != s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK12)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f30540,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X18,'type.num.num'))) ) | ($spl2 | $spl10 | $spl33)),
  inference(resolution,[],[f30521,f902])).
fof(f902,plain,(
  ( ! [X47,X48,X46] : (sP2(X47,happ(s(fun('type.num.num',bool),X46),s('type.num.num',sK12)),X48) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X46,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f462])).
fof(f30457,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl156),
  inference(cnf_transformation,[],[f30457_D])).
fof(f30457_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f30396,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl12),
  inference(resolution,[],[f2461,f691])).
fof(f30305,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,X29) != s(bool,X30)) ) | $spl12),
  inference(resolution,[],[f2461,f452])).
fof(f30302,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl12),
  inference(resolution,[],[f2461,f451])).
fof(f30301,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl12),
  inference(resolution,[],[f2461,f450])).
fof(f30300,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | s(bool,X14) != s(bool,X15)) ) | $spl12),
  inference(resolution,[],[f2461,f448])).
fof(f30298,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK19(sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)),sK17,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))))) ) | $spl12),
  inference(resolution,[],[f2461,f446])).
fof(f1034,plain,(
  ( ! [X4,X2] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X2,'type.num.num')) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl16),
  inference(cnf_transformation,[],[f1034_D])).
fof(f1034_D,plain,(
  ( ! [X4,X2] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X2,'type.num.num')) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f25361,plain,(
  ( ! [X699,X700,X698] : (p(s(bool,X699)) | ~p(s(bool,X700)) | s(bool,X700) != s(bool,happ(s(fun('type.num.num',bool),X698),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X698,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f25147,f642])).
fof(f25147,plain,(
  ( ! [X699,X700,X698] : (p(s(bool,'F_0')) | p(s(bool,X699)) | ~p(s(bool,X700)) | s(bool,X700) != s(bool,happ(s(fun('type.num.num',bool),X698),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X698,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f691,f983])).
fof(f25246,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,X153)) | ~p(s(bool,X154)) | s(bool,X154) != s(bool,happ(s(fun('type.num.num',bool),X152),s('type.num.num',sK12))) | s(bool,X153) != s(bool,X155) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X152,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f24935,f623])).
fof(f24935,plain,(
  ( ! [X154,X152,X155,X153] : (p(s(bool,'F_0')) | ~p(s(bool,X153)) | ~p(s(bool,X154)) | s(bool,X154) != s(bool,happ(s(fun('type.num.num',bool),X152),s('type.num.num',sK12))) | s(bool,X153) != s(bool,X155) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X152,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f452,f983])).
fof(f25244,plain,(
  ( ! [X123,X121,X122,X120] : (p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X123)) | s(bool,X122) != s(bool,happ(s(fun('type.num.num',bool),X120),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X120,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f24927,f642])).
fof(f24927,plain,(
  ( ! [X123,X121,X122,X120] : (p(s(bool,'F_0')) | p(s(bool,X121)) | ~p(s(bool,X122)) | ~p(s(bool,X123)) | s(bool,X122) != s(bool,happ(s(fun('type.num.num',bool),X120),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X120,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f451,f983])).
fof(f25240,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X100,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f24922,f642])).
fof(f24922,plain,(
  ( ! [X103,X101,X102,X100] : (p(s(bool,'F_0')) | p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X101) != s(bool,X103) | s(bool,X102) != s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X100,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f450,f983])).
fof(f25237,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun('type.num.num',bool),X72),s('type.num.num',sK12))) | s(bool,X73) != s(bool,X75) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X72,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f24915,f623])).
fof(f24915,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,'F_0')) | ~p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X74) != s(bool,happ(s(fun('type.num.num',bool),X72),s('type.num.num',sK12))) | s(bool,X73) != s(bool,X75) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X72,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f448,f983])).
fof(f25233,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,happ(s(fun('type.num.num',bool),X44),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X44,'type.num.num'))) ) | $spl2),
  inference(subsumption_resolution,[],[f24908,f642])).
fof(f24908,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'F_0')) | p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,happ(s(fun('type.num.num',bool),X44),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X44,'type.num.num'))) ) | $spl2),
  inference(superposition,[],[f446,f983])).
fof(f20695,plain,(
  $false | ($spl33 | $spl82)),
  inference(subsumption_resolution,[],[f20579,f20464])).
fof(f20464,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | ($spl33 | $spl82)),
  inference(subsumption_resolution,[],[f20463,f655])).
fof(f20463,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9))) ) | ($spl33 | $spl82)),
  inference(forward_demodulation,[],[f5875,f18145])).
fof(f20579,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl82)),
  inference(resolution,[],[f20464,f670])).
fof(f670,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f416])).
fof(f416,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f224])).
fof(f224,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f56])).
fof(f56,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f20694,plain,(
  $false | ($spl33 | $spl82)),
  inference(subsumption_resolution,[],[f20578,f20464])).
fof(f20578,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl82)),
  inference(resolution,[],[f20464,f668])).
fof(f668,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f415])).
fof(f415,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f223])).
fof(f223,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f57])).
fof(f57,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f20693,plain,(
  $false | ($spl33 | $spl82)),
  inference(subsumption_resolution,[],[f20572,f20464])).
fof(f20572,plain,(
  ( ! [X278] : (~p(s(bool,X278))) ) | ($spl33 | $spl82)),
  inference(resolution,[],[f20464,f623])).
fof(f20577,plain,(
  $false | ($spl33 | $spl82)),
  inference(resolution,[],[f20464,f666])).
fof(f20601,plain,(
  $false | ($spl33 | $spl79 | $spl82)),
  inference(resolution,[],[f20464,f5727])).
fof(f5727,plain,(
  ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl79),
  inference(cnf_transformation,[],[f5727_D])).
fof(f5727_D,plain,(
  ~p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f18817,plain,(
  $false | ($spl7 | $spl28)),
  inference(subsumption_resolution,[],[f18701,f18587])).
fof(f18587,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | ($spl7 | $spl28)),
  inference(subsumption_resolution,[],[f18586,f655])).
fof(f18586,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8))) ) | ($spl7 | $spl28)),
  inference(forward_demodulation,[],[f1233,f2271])).
fof(f18701,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl28)),
  inference(resolution,[],[f18587,f670])).
fof(f18816,plain,(
  $false | ($spl7 | $spl28)),
  inference(subsumption_resolution,[],[f18700,f18587])).
fof(f18700,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl28)),
  inference(resolution,[],[f18587,f668])).
fof(f18815,plain,(
  $false | ($spl7 | $spl28)),
  inference(subsumption_resolution,[],[f18694,f18587])).
fof(f18694,plain,(
  ( ! [X278] : (~p(s(bool,X278))) ) | ($spl7 | $spl28)),
  inference(resolution,[],[f18587,f623])).
fof(f18699,plain,(
  $false | ($spl7 | $spl28)),
  inference(resolution,[],[f18587,f666])).
fof(f18721,plain,(
  $false | ($spl7 | $spl28 | $spl33)),
  inference(resolution,[],[f18587,f1646])).
fof(f18724,plain,(
  $false | ($spl7 | $spl28 | $spl79)),
  inference(resolution,[],[f18587,f5727])).
fof(f18477,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl154),
  inference(cnf_transformation,[],[f18477_D])).
fof(f18477_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f18422,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl79),
  inference(resolution,[],[f5727,f691])).
fof(f18331,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X21) != s(bool,X23)) ) | $spl79),
  inference(resolution,[],[f5727,f452])).
fof(f18329,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl79),
  inference(resolution,[],[f5727,f451])).
fof(f18327,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl79),
  inference(resolution,[],[f5727,f450])).
fof(f18326,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X6) != s(bool,X8)) ) | $spl79),
  inference(resolution,[],[f5727,f448])).
fof(f18324,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl79),
  inference(resolution,[],[f5727,f446])).
fof(f18322,plain,(
  ( ! [X14] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | p(s(bool,X14))) ) | $spl152),
  inference(cnf_transformation,[],[f18322_D])).
fof(f18322_D,plain,(
  ( ! [X14] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | p(s(bool,X14))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f18304,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1672,f13720])).
fof(f1797,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl38),
  inference(cnf_transformation,[],[f1797_D])).
fof(f1797_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f18298,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1672,f14774])).
fof(f18319,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl150),
  inference(cnf_transformation,[],[f18319_D])).
fof(f18319_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f18297,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1672,f14893])).
fof(f18273,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1671,f14893])).
fof(f1671,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X39)) ) | $spl33),
  inference(resolution,[],[f1646,f466])).
fof(f18212,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X22))) ) | ($spl33 | $spl59)),
  inference(resolution,[],[f1668,f5659])).
fof(f1668,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),X33)) ) | $spl33),
  inference(resolution,[],[f1646,f462])).
fof(f5674,plain,(
  ( ! [X0] : (sP3(X0,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl70),
  inference(cnf_transformation,[],[f5674_D])).
fof(f5674_D,plain,(
  ( ! [X0] : (sP3(X0,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f18211,plain,(
  ( ! [X21] : (sP3(X21,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1668,f5345])).
fof(f5723,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl72),
  inference(cnf_transformation,[],[f5723_D])).
fof(f5723_D,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f18210,plain,(
  ( ! [X20] : (s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X20))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1668,f5458])).
fof(f18209,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1668,f15058])).
fof(f18233,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X9))) ) | $spl146),
  inference(cnf_transformation,[],[f18233_D])).
fof(f18233_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X9))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f18204,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f1668,f14550])).
fof(f18150,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f691])).
fof(f18059,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1646,f452])).
fof(f18057,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f451])).
fof(f18055,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f450])).
fof(f18054,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1646,f448])).
fof(f18052,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f446])).
fof(f1648,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | $spl34),
  inference(cnf_transformation,[],[f1648_D])).
fof(f1648_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9591,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl10),
  inference(subsumption_resolution,[],[f9272,f695])).
fof(f9272,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl10),
  inference(superposition,[],[f678,f1582])).
fof(f17992,plain,(
  ( ! [X175,X174] : (p(s(bool,X174)) | p(s(bool,X175)) | s(bool,X175) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f691])).
fof(f17901,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X24) != s(bool,X25)) ) | $spl78),
  inference(resolution,[],[f5728,f452])).
fof(f17898,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f451])).
fof(f17897,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f450])).
fof(f17896,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X9) != s(bool,X10)) ) | $spl78),
  inference(resolution,[],[f5728,f448])).
fof(f17894,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl78),
  inference(resolution,[],[f5728,f446])).
fof(f15576,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X46,'type.num.num')) | s(bool,X45) != s(bool,happ(s(fun('type.num.num',bool),X46),s('type.num.num',sK12)))) ) | $spl142),
  inference(cnf_transformation,[],[f15576_D])).
fof(f15576_D,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X46,'type.num.num')) | s(bool,X45) != s(bool,happ(s(fun('type.num.num',bool),X46),s('type.num.num',sK12)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f15567,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | s(bool,'T_0') != s(bool,X45) | s(bool,X45) != s(bool,happ(s(fun('type.num.num',bool),X46),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X46,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f14893,f906])).
fof(f906,plain,(
  ( ! [X59,X60,X58] : (sP1(happ(s(fun('type.num.num',bool),X58),s('type.num.num',sK12)),X59,X60) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X58,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f467])).
fof(f15574,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X35,'type.num.num')) | s(bool,X36) != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12)))) ) | $spl140),
  inference(cnf_transformation,[],[f15574_D])).
fof(f15574_D,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X35,'type.num.num')) | s(bool,X36) != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f15563,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12))) | s(bool,X36) != s(bool,happ(s(fun('type.num.num',bool),X35),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X35,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f14893,f905])).
fof(f905,plain,(
  ( ! [X57,X56,X55] : (sP1(X56,happ(s(fun('type.num.num',bool),X55),s('type.num.num',sK12)),X57) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X55,'type.num.num'))) ) | $spl2),
  inference(resolution,[],[f888,f466])).
fof(f3815,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X8,'type.num.num'))) ) | $spl54),
  inference(cnf_transformation,[],[f3815_D])).
fof(f3815_D,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X8,'type.num.num'))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f15392,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',sK12))) | s(bool,'F_0') != s(bool,X16) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X15,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f15058,f902])).
fof(f14937,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,happ(s(fun('type.num.num',bool),X46),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X46,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f14774,f906])).
fof(f14176,plain,(
  ( ! [X18,X16] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,X16)) | s(bool,X16) != s(bool,X18)) ) | $spl138),
  inference(cnf_transformation,[],[f14176_D])).
fof(f14176_D,plain,(
  ( ! [X18,X16] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,X16)) | s(bool,X16) != s(bool,X18)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f14179,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X31) != s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl122),
  inference(subsumption_resolution,[],[f14135,f653])).
fof(f14135,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') != s(bool,X33) | s(bool,X31) != s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl122),
  inference(superposition,[],[f13583,f485])).
fof(f485,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f14178,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X30) | s(bool,X28) != s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl122),
  inference(subsumption_resolution,[],[f14134,f653])).
fof(f14134,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,X30) | s(bool,X28) != s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl122),
  inference(superposition,[],[f13583,f485])).
fof(f14177,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl122),
  inference(subsumption_resolution,[],[f14131,f653])).
fof(f14131,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl122),
  inference(superposition,[],[f13583,f481])).
fof(f481,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f14175,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl122),
  inference(subsumption_resolution,[],[f14130,f653])).
fof(f14130,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl122),
  inference(superposition,[],[f13583,f481])).
fof(f13908,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | sP3(X24,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X24))) ) | ($spl7 | $spl52)),
  inference(forward_demodulation,[],[f13902,f2271])).
fof(f13902,plain,(
  ( ! [X24,X25] : (sP3(X24,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),'T_0') | p(s(bool,X24)) | s(bool,X25) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl7 | $spl52)),
  inference(resolution,[],[f13720,f2194])).
fof(f13907,plain,(
  ( ! [X22] : (sP3(X22,'F_0','T_0') | p(s(bool,X22))) ) | $spl134),
  inference(cnf_transformation,[],[f13907_D])).
fof(f13907_D,plain,(
  ( ! [X22] : (sP3(X22,'F_0','T_0') | p(s(bool,X22))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f13901,plain,(
  ( ! [X23,X22] : (sP3(X22,'F_0','T_0') | p(s(bool,X22)) | s(bool,'F_0') != s(bool,X23)) ) | $spl52),
  inference(resolution,[],[f13720,f4693])).
fof(f13882,plain,(
  $false | ($spl10 | $spl128)),
  inference(trivial_inequality_removal,[],[f13863])).
fof(f13863,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl10 | $spl128)),
  inference(superposition,[],[f13812,f1582])).
fof(f13881,plain,(
  $false | ($spl32 | $spl128)),
  inference(trivial_inequality_removal,[],[f13865])).
fof(f13865,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl32 | $spl128)),
  inference(superposition,[],[f13812,f5830])).
fof(f5830,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | $spl32),
  inference(resolution,[],[f1645,f657])).
fof(f13880,plain,(
  $false | ($spl87 | $spl128)),
  inference(trivial_inequality_removal,[],[f13866])).
fof(f13866,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl87 | $spl128)),
  inference(superposition,[],[f13812,f10990])).
fof(f13879,plain,(
  $false | $spl128),
  inference(equality_resolution,[],[f13812])).
fof(f13815,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) | ($spl24 | $spl52 | $spl87)),
  inference(forward_demodulation,[],[f13805,f10990])).
fof(f13805,plain,(
  ( ! [X17,X16] : (sP3(X16,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | s(bool,X17) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X16)) ) | ($spl24 | $spl52)),
  inference(resolution,[],[f13719,f3869])).
fof(f13813,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X14)) ) | ($spl32 | $spl52)),
  inference(forward_demodulation,[],[f13804,f5830])).
fof(f13804,plain,(
  ( ! [X14,X15] : (sP3(X14,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0') | s(bool,X15) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~sP1('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))),'T_0',X14)) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f13719,f5747])).
fof(f13810,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) | ($spl10 | $spl52)),
  inference(forward_demodulation,[],[f13803,f1582])).
fof(f13803,plain,(
  ( ! [X12,X13] : (sP3(X12,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0') | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),'T_0',X12)) ) | ($spl10 | $spl52)),
  inference(resolution,[],[f13719,f1499])).
fof(f13622,plain,(
  ( ! [X19] : (~sP2(X19,'T_0','F_0') | ~p(s(bool,X19))) ) | $spl124),
  inference(cnf_transformation,[],[f13622_D])).
fof(f13622_D,plain,(
  ( ! [X19] : (~sP2(X19,'T_0','F_0') | ~p(s(bool,X19))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f13621,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X18) | ~sP2(X19,'T_0','F_0') | ~p(s(bool,X19))) ) | ($spl52 | $spl59)),
  inference(subsumption_resolution,[],[f13618,f623])).
fof(f13618,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'F_0')) | ~sP2(X19,'T_0','F_0') | ~p(s(bool,X19))) ) | ($spl52 | $spl59)),
  inference(trivial_inequality_removal,[],[f13614])).
fof(f13614,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'F_0')) | ~sP2(X19,'T_0','F_0') | ~p(s(bool,X19)) | s(bool,'T_0') != s(bool,'T_0')) ) | ($spl52 | $spl59)),
  inference(resolution,[],[f13598,f5023])).
fof(f13583,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) | $spl122),
  inference(cnf_transformation,[],[f13583_D])).
fof(f13583_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f13582,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f13570])).
fof(f13570,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl52),
  inference(resolution,[],[f5512,f475])).
fof(f13287,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13159,f13005])).
fof(f13005,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13004,f659])).
fof(f13004,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl32 | $spl38)),
  inference(forward_demodulation,[],[f1797,f5830])).
fof(f13159,plain,(
  p(s(bool,'F_0')) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f669])).
fof(f669,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f416])).
fof(f13286,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13158,f13005])).
fof(f13158,plain,(
  p(s(bool,'F_0')) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f667])).
fof(f667,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f415])).
fof(f13285,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13154,f13005])).
fof(f13154,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f634])).
fof(f634,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f399])).
fof(f399,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X23] : ((p(s(bool,X23)) => p(s(bool,X23))) <=> p(s(bool,'T_0')))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CLAUSES|split|3')).
fof(f13284,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13153,f13005])).
fof(f13153,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f631])).
fof(f631,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f397])).
fof(f397,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X23] : ((p(s(bool,X23)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CLAUSES|split|1')).
fof(f13283,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13152,f13005])).
fof(f13152,plain,(
  ( ! [X265] : (p(s(bool,X265))) ) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f628])).
fof(f628,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X23] : ((p(s(bool,'T_0')) => p(s(bool,X23))) <=> p(s(bool,X23)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CLAUSES|split|0')).
fof(f13282,plain,(
  $false | ($spl32 | $spl38)),
  inference(subsumption_resolution,[],[f13150,f13005])).
fof(f13150,plain,(
  p(s(bool,'F_0')) | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f625])).
fof(f625,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f393])).
fof(f393,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X23] : ((p(s(bool,'F_0')) => p(s(bool,X23))) <=> p(s(bool,'T_0')))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.IMP_CLAUSES|split|2')).
fof(f13160,plain,(
  $false | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f671])).
fof(f671,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,axiom,(
  p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.bool.TRUTH')).
fof(f13161,plain,(
  $false | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f672])).
fof(f672,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb',aHL_TRUTH)).
fof(f13162,plain,(
  $false | ($spl32 | $spl38)),
  inference(resolution,[],[f13005,f695])).
fof(f13179,plain,(
  $false | ($spl12 | $spl32 | $spl38)),
  inference(resolution,[],[f13005,f711])).
fof(f3997,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl64),
  inference(cnf_transformation,[],[f3997_D])).
fof(f3997_D,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f11100,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl87),
  inference(resolution,[],[f10991,f691])).
fof(f10991,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl87),
  inference(trivial_inequality_removal,[],[f10981])).
fof(f10981,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl87),
  inference(superposition,[],[f6236,f655])).
fof(f11009,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X30) != s(bool,X31)) ) | $spl87),
  inference(resolution,[],[f10991,f452])).
fof(f11006,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl87),
  inference(resolution,[],[f10991,f451])).
fof(f11005,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl87),
  inference(resolution,[],[f10991,f450])).
fof(f11004,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X15) != s(bool,X16)) ) | $spl87),
  inference(resolution,[],[f10991,f448])).
fof(f11002,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl87),
  inference(resolution,[],[f10991,f446])).
fof(f1229,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl24),
  inference(cnf_transformation,[],[f1229_D])).
fof(f1229_D,plain,(
  p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f10998,plain,(
  ( ! [X13] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X13))) ) | $spl87),
  inference(subsumption_resolution,[],[f10977,f653])).
fof(f10977,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X13))) ) | $spl87),
  inference(superposition,[],[f6236,f485])).
fof(f10997,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl87),
  inference(subsumption_resolution,[],[f10976,f653])).
fof(f10976,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,X12)) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl87),
  inference(superposition,[],[f6236,f485])).
fof(f10996,plain,(
  ( ! [X9] : (p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X9))) ) | $spl87),
  inference(subsumption_resolution,[],[f10973,f653])).
fof(f10973,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | p(s(bool,X9))) ) | $spl87),
  inference(superposition,[],[f6236,f481])).
fof(f10995,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl87),
  inference(subsumption_resolution,[],[f10972,f653])).
fof(f10972,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))) ) | $spl87),
  inference(superposition,[],[f6236,f481])).
fof(f6236,plain,(
  s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl87),
  inference(cnf_transformation,[],[f6236_D])).
fof(f6236_D,plain,(
  s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f10838,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl24 | $spl52)),
  inference(resolution,[],[f3875,f5459])).
fof(f3875,plain,(
  ( ! [X54,X55] : (sP0('const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X54,X55)) ) | $spl24),
  inference(resolution,[],[f1229,f471])).
fof(f1230,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl25),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f10853,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1(X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X2) | ~sP2(X2,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X3)) ) | $spl112),
  inference(cnf_transformation,[],[f10853_D])).
fof(f10853_D,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~sP1(X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X2) | ~sP2(X2,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X3)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10835,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~sP2(X2,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X3) | ~sP1(X3,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X2)) ) | $spl24),
  inference(resolution,[],[f3875,f457])).
fof(f10690,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl101),
  inference(cnf_transformation,[],[f10690_D])).
fof(f10690_D,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f10670,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~p(s(bool,X10))) ) | ($spl25 | $spl59)),
  inference(resolution,[],[f1252,f5659])).
fof(f1252,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X33)) ) | $spl25),
  inference(resolution,[],[f1230,f462])).
fof(f10691,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl102),
  inference(cnf_transformation,[],[f10691_D])).
fof(f10691_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f10669,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl25 | $spl52)),
  inference(resolution,[],[f1252,f5345])).
fof(f6235,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl86),
  inference(cnf_transformation,[],[f6235_D])).
fof(f6235_D,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10668,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | ~p(s(bool,X8))) ) | ($spl25 | $spl52)),
  inference(resolution,[],[f1252,f5458])).
fof(f10585,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10458,f10347])).
fof(f10347,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10346,f659])).
fof(f10346,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl24 | $spl30)),
  inference(forward_demodulation,[],[f1381,f3952])).
fof(f3952,plain,(
  s(bool,'T_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl24),
  inference(resolution,[],[f1229,f657])).
fof(f10458,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f669])).
fof(f10584,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10457,f10347])).
fof(f10457,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f667])).
fof(f10583,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10453,f10347])).
fof(f10453,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f634])).
fof(f10582,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10452,f10347])).
fof(f10452,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f631])).
fof(f10581,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10451,f10347])).
fof(f10451,plain,(
  ( ! [X265] : (p(s(bool,X265))) ) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f628])).
fof(f10580,plain,(
  $false | ($spl24 | $spl30)),
  inference(subsumption_resolution,[],[f10449,f10347])).
fof(f10449,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f625])).
fof(f10459,plain,(
  $false | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f671])).
fof(f10460,plain,(
  $false | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f672])).
fof(f10461,plain,(
  $false | ($spl24 | $spl30)),
  inference(resolution,[],[f10347,f695])).
fof(f10478,plain,(
  $false | ($spl12 | $spl24 | $spl30)),
  inference(resolution,[],[f10347,f711])).
fof(f9808,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X18)) ) | $spl96),
  inference(cnf_transformation,[],[f9808_D])).
fof(f9808_D,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9811,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl52),
  inference(subsumption_resolution,[],[f9772,f653])).
fof(f9772,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl52),
  inference(superposition,[],[f5509,f485])).
fof(f9810,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl52),
  inference(subsumption_resolution,[],[f9771,f653])).
fof(f9771,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,X28) != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl52),
  inference(superposition,[],[f5509,f485])).
fof(f9809,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl52),
  inference(subsumption_resolution,[],[f9768,f653])).
fof(f9768,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | s(bool,'F_0') = s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl52),
  inference(superposition,[],[f5509,f481])).
fof(f9807,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl52),
  inference(subsumption_resolution,[],[f9767,f653])).
fof(f9767,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl52),
  inference(superposition,[],[f5509,f481])).
fof(f8547,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num'))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl94),
  inference(cnf_transformation,[],[f8547_D])).
fof(f8547_D,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num'))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8502,plain,(
  ( ! [X267,X268,X266] : (p(s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X266,'type.num.num'))))) | p(s(bool,X267)) | ~p(s(bool,X268)) | s(bool,X268) != s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f778,f691])).
fof(f778,plain,(
  ( ! [X145] : (~p(s(bool,happ(s(fun('type.num.num',bool),X145),s('type.num.num',sK12)))) | p(s(bool,happ(s(fun('type.num.num',bool),X145),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X145,'type.num.num')))))) ) | $spl4),
  inference(superposition,[],[f491,f703])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f8411,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X30,'type.num.num'))))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK12))) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f778,f452])).
fof(f8409,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X22,'type.num.num'))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f778,f451])).
fof(f8407,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X14,'type.num.num'))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f778,f450])).
fof(f8406,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X10,'type.num.num'))))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK12))) | s(bool,X11) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f778,f448])).
fof(f8404,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num'))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f778,f446])).
fof(f7778,plain,(
  $false | ($spl32 | $spl57)),
  inference(subsumption_resolution,[],[f3817,f5830])).
fof(f6442,plain,(
  s(bool,'T_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl91),
  inference(cnf_transformation,[],[f6442_D])).
fof(f6442_D,plain,(
  s(bool,'T_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f6404,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | ~p(s(bool,X10))) ) | ($spl7 | $spl59)),
  inference(resolution,[],[f2190,f5659])).
fof(f2190,plain,(
  ( ! [X33,X32] : (sP2(X32,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),X33)) ) | $spl7),
  inference(resolution,[],[f704,f462])).
fof(f6443,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl92),
  inference(cnf_transformation,[],[f6443_D])).
fof(f6443_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6403,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,'T_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | ($spl7 | $spl52)),
  inference(resolution,[],[f2190,f5345])).
fof(f6440,plain,(
  s(bool,'F_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | $spl88),
  inference(cnf_transformation,[],[f6440_D])).
fof(f6440_D,plain,(
  s(bool,'F_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6402,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | ~p(s(bool,X8))) ) | ($spl7 | $spl52)),
  inference(resolution,[],[f2190,f5458])).
fof(f6230,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl32 | $spl52)),
  inference(resolution,[],[f5459,f5753])).
fof(f6229,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl24 | $spl52)),
  inference(resolution,[],[f5459,f3875])).
fof(f6234,plain,(
  s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl85),
  inference(cnf_transformation,[],[f6234_D])).
fof(f6234_D,plain,(
  s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f6228,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | ($spl10 | $spl52)),
  inference(resolution,[],[f5459,f1505])).
fof(f1505,plain,(
  ( ! [X54,X53] : (sP0('const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),X53,X54)) ) | $spl10),
  inference(resolution,[],[f709,f471])).
fof(f5832,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f691])).
fof(f5741,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X30) != s(bool,X31)) ) | $spl32),
  inference(resolution,[],[f1645,f452])).
fof(f5738,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f451])).
fof(f5737,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f450])).
fof(f5736,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X15) != s(bool,X16)) ) | $spl32),
  inference(resolution,[],[f1645,f448])).
fof(f5734,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl32),
  inference(resolution,[],[f1645,f446])).
fof(f5730,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) | $spl80),
  inference(cnf_transformation,[],[f5730_D])).
fof(f5730_D,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2459,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl12),
  inference(resolution,[],[f711,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(X3,X2,X1,X0))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(X3,X2,X1,X0))))) & p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(X3,X2,X1,X0)),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ? [X4] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) & p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))) | ? [X4] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) & p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))) => ? [X4] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4)))) & p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X27,X20,X28,X29] : ((p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20)))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X29),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X28))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27))))) => ? [X30] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X20),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X28),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X29))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X30)))) & p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X30),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X27))))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.infer_eComplete.infer_eComplete_4_prove_ltb')).
fof(f5728,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl78),
  inference(cnf_transformation,[],[f5728_D])).
fof(f5728_D,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2460,plain,(
  p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(sK17,sK16,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl12),
  inference(resolution,[],[f711,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.unify.t_compat_2'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),sK20(X3,X2,X1,X0)),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X2))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X1))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f5726,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))))) | $spl76),
  inference(cnf_transformation,[],[f5726_D])).
fof(f5726_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2463,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl12),
  inference(resolution,[],[f711,f678])).
fof(f5724,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X9,'type.num.num'))) ) | $spl74),
  inference(cnf_transformation,[],[f5724_D])).
fof(f5724_D,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',sK12))) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X9,'type.num.num'))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5720,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',sK12))) | ~p(s(bool,X10)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X9,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f5458,f902])).
fof(f5715,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X0))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f5458,f1668])).
fof(f5667,plain,(
  ( ! [X0] : (sP3(X0,'F_0','const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f5345,f1668])).
fof(f5665,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',sK12))) | ~p(s(bool,X10)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X9,'type.num.num'))) ) | ($spl2 | $spl59)),
  inference(resolution,[],[f5659,f902])).
fof(f5660,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X0))) ) | ($spl33 | $spl59)),
  inference(resolution,[],[f5659,f1668])).
fof(f5278,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f5166,f1037])).
fof(f5166,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(resolution,[],[f1037,f670])).
fof(f5277,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f5165,f1037])).
fof(f5165,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(resolution,[],[f1037,f668])).
fof(f5276,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f5159,f1037])).
fof(f5159,plain,(
  ( ! [X278] : (~p(s(bool,X278))) ) | $spl22),
  inference(resolution,[],[f1037,f623])).
fof(f5164,plain,(
  $false | $spl22),
  inference(resolution,[],[f1037,f666])).
fof(f5183,plain,(
  $false | ($spl7 | $spl22)),
  inference(resolution,[],[f1037,f704])).
fof(f5187,plain,(
  $false | ($spl22 | $spl33)),
  inference(resolution,[],[f1037,f1646])).
fof(f5050,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl59),
  inference(duplicate_literal_removal,[],[f5045])).
fof(f5045,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl59),
  inference(resolution,[],[f5020,f475])).
fof(f5020,plain,(
  ( ! [X4,X3] : (sP3(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl59),
  inference(superposition,[],[f3820,f472])).
fof(f5012,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4886,f1036])).
fof(f4886,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1036,f669])).
fof(f5011,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4885,f1036])).
fof(f4885,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1036,f667])).
fof(f5010,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4881,f1036])).
fof(f4881,plain,(
  ( ! [X267] : (p(s(bool,X267))) ) | $spl20),
  inference(resolution,[],[f1036,f634])).
fof(f5009,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4880,f1036])).
fof(f4880,plain,(
  ( ! [X266] : (p(s(bool,X266))) ) | $spl20),
  inference(resolution,[],[f1036,f631])).
fof(f5008,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4879,f1036])).
fof(f4879,plain,(
  ( ! [X265] : (p(s(bool,X265))) ) | $spl20),
  inference(resolution,[],[f1036,f628])).
fof(f5007,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f4877,f1036])).
fof(f4877,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1036,f625])).
fof(f4887,plain,(
  $false | $spl20),
  inference(resolution,[],[f1036,f671])).
fof(f4888,plain,(
  $false | $spl20),
  inference(resolution,[],[f1036,f672])).
fof(f4889,plain,(
  $false | $spl20),
  inference(resolution,[],[f1036,f695])).
fof(f4906,plain,(
  $false | ($spl10 | $spl20)),
  inference(resolution,[],[f1036,f709])).
fof(f4907,plain,(
  $false | ($spl12 | $spl20)),
  inference(resolution,[],[f1036,f711])).
fof(f4910,plain,(
  $false | ($spl20 | $spl24)),
  inference(resolution,[],[f1036,f1229])).
fof(f4774,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X10))) ) | ($spl33 | $spl52)),
  inference(resolution,[],[f4768,f1668])).
fof(f4768,plain,(
  ( ! [X14,X15] : (~sP2(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl52),
  inference(subsumption_resolution,[],[f4767,f469])).
fof(f4767,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl52),
  inference(subsumption_resolution,[],[f4766,f623])).
fof(f4766,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl52),
  inference(subsumption_resolution,[],[f4759,f467])).
fof(f4759,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP1('F_0',X15,X14) | ~sP0(X15,'F_0',X14)) ) | $spl52),
  inference(resolution,[],[f3832,f458])).
fof(f3832,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl52),
  inference(superposition,[],[f3355,f474])).
fof(f4773,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',sK12))) | ~p(s(bool,X9)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X8,'type.num.num'))) ) | ($spl2 | $spl52)),
  inference(resolution,[],[f4768,f902])).
fof(f4484,plain,(
  ( ! [X4,X2] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num')) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl68),
  inference(cnf_transformation,[],[f4484_D])).
fof(f4484_D,plain,(
  ( ! [X4,X2] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num')) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4439,plain,(
  ( ! [X267,X268,X266] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X266,'type.num.num')) | p(s(bool,X267)) | ~p(s(bool,X268)) | s(bool,X268) != s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f817,f691])).
fof(f817,plain,(
  ( ! [X144] : (~p(s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',sK12)))) | s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X144,'type.num.num'))) ) | $spl4),
  inference(forward_demodulation,[],[f777,f703])).
fof(f777,plain,(
  ( ! [X144] : (~p(s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',sK12)))) | s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X144,'type.num.num'))) ) | $spl4),
  inference(superposition,[],[f490,f703])).
fof(f4348,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X30,'type.num.num')) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK12))) | s(bool,X31) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f817,f452])).
fof(f4346,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X22,'type.num.num')) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f817,f451])).
fof(f4344,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X14,'type.num.num')) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f817,f450])).
fof(f4343,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X10,'type.num.num')) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK12))) | s(bool,X11) != s(bool,X13)) ) | $spl4),
  inference(resolution,[],[f817,f448])).
fof(f4341,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.num.num',sK12) = s('type.num.num',sK22('const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)),X2,'type.num.num')) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl4),
  inference(resolution,[],[f817,f446])).
fof(f4210,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22(sK13,X2,'type.num.num'))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl66),
  inference(cnf_transformation,[],[f4210_D])).
fof(f4210_D,plain,(
  ( ! [X4,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22(sK13,X2,'type.num.num'))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4165,plain,(
  ( ! [X267,X268,X266] : (p(s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK22(sK13,X266,'type.num.num'))))) | p(s(bool,X267)) | ~p(s(bool,X268)) | s(bool,X268) != s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f851,f691])).
fof(f851,plain,(
  ( ! [X93] : (~p(s(bool,happ(s(fun('type.num.num',bool),X93),s('type.num.num',sK12)))) | p(s(bool,happ(s(fun('type.num.num',bool),X93),s('type.num.num',sK22(sK13,X93,'type.num.num')))))) ) | $spl2),
  inference(superposition,[],[f491,f700])).
fof(f4074,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK22(sK13,X30,'type.num.num'))))) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK12))) | s(bool,X31) != s(bool,X33)) ) | $spl2),
  inference(resolution,[],[f851,f452])).
fof(f4072,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK22(sK13,X22,'type.num.num'))))) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f851,f451])).
fof(f4070,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK22(sK13,X14,'type.num.num'))))) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f851,f450])).
fof(f4069,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK22(sK13,X10,'type.num.num'))))) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK12))) | s(bool,X11) != s(bool,X13)) ) | $spl2),
  inference(resolution,[],[f851,f448])).
fof(f4067,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK22(sK13,X2,'type.num.num'))))) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f851,f446])).
fof(f3954,plain,(
  ( ! [X180,X181] : (p(s(bool,X180)) | p(s(bool,X181)) | s(bool,X181) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl24),
  inference(resolution,[],[f1229,f691])).
fof(f3863,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X30) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f1229,f452])).
fof(f3860,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl24),
  inference(resolution,[],[f1229,f451])).
fof(f3859,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl24),
  inference(resolution,[],[f1229,f450])).
fof(f3858,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X15) != s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f1229,f448])).
fof(f3856,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl24),
  inference(resolution,[],[f1229,f446])).
fof(f3825,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl52),
  inference(subsumption_resolution,[],[f3330,f653])).
fof(f3330,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X28) | p(s(bool,X28))) ) | $spl52),
  inference(superposition,[],[f3282,f653])).
fof(f3824,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X29))) ) | $spl52),
  inference(subsumption_resolution,[],[f3331,f653])).
fof(f3331,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X29) | p(s(bool,X29))) ) | $spl52),
  inference(superposition,[],[f3282,f655])).
fof(f3822,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X33)) ) | $spl52),
  inference(subsumption_resolution,[],[f3335,f3282])).
fof(f3335,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl52),
  inference(superposition,[],[f3282,f690])).
fof(f3821,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,X34)) ) | $spl60),
  inference(cnf_transformation,[],[f3821_D])).
fof(f3821_D,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,X34)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3360,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X34)) ) | $spl52),
  inference(trivial_inequality_removal,[],[f3336])).
fof(f3336,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,X34) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X34)) ) | $spl52),
  inference(superposition,[],[f3282,f690])).
fof(f3359,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) ) | $spl52),
  inference(trivial_inequality_removal,[],[f3337])).
fof(f3337,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) ) | $spl52),
  inference(superposition,[],[f3282,f659])).
fof(f3358,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X36))) ) | $spl52),
  inference(trivial_inequality_removal,[],[f3338])).
fof(f3338,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,X36) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X36))) ) | $spl52),
  inference(superposition,[],[f3282,f657])).
fof(f3812,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | ~p(s(bool,X11))) ) | ($spl25 | $spl33)),
  inference(resolution,[],[f3805,f1668])).
fof(f3805,plain,(
  ( ! [X14,X15] : (~sP2(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl25),
  inference(subsumption_resolution,[],[f3804,f469])).
fof(f3804,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl25),
  inference(subsumption_resolution,[],[f3803,f623])).
fof(f3803,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP0(X15,'F_0',X14)) ) | $spl25),
  inference(subsumption_resolution,[],[f3796,f467])).
fof(f3796,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP2(X14,X15,'F_0') | ~sP1('F_0',X15,X14) | ~sP0(X15,'F_0',X14)) ) | $spl25),
  inference(resolution,[],[f3254,f458])).
fof(f3254,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl25),
  inference(superposition,[],[f3203,f474])).
fof(f3203,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl25),
  inference(superposition,[],[f1335,f1333])).
fof(f1333,plain,(
  s(bool,'F_0') = s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl25),
  inference(resolution,[],[f1230,f653])).
fof(f1335,plain,(
  s(bool,'T_0') != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | $spl25),
  inference(resolution,[],[f1230,f656])).
fof(f3810,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',sK12))) | ~p(s(bool,X9)) | s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X8,'type.num.num'))) ) | ($spl2 | $spl25)),
  inference(resolution,[],[f3805,f902])).
fof(f3281,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl25),
  inference(duplicate_literal_removal,[],[f3276])).
fof(f3276,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl25),
  inference(resolution,[],[f3251,f475])).
fof(f3251,plain,(
  ( ! [X4,X3] : (sP3(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl25),
  inference(superposition,[],[f3203,f472])).
fof(f2634,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl50),
  inference(cnf_transformation,[],[f2634_D])).
fof(f2634_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2565,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f711,f691])).
fof(f2474,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X29) != s(bool,X30)) ) | $spl12),
  inference(resolution,[],[f711,f452])).
fof(f2471,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f711,f451])).
fof(f2470,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f711,f450])).
fof(f2469,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X14) != s(bool,X15)) ) | $spl12),
  inference(resolution,[],[f711,f448])).
fof(f2467,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl12),
  inference(resolution,[],[f711,f446])).
fof(f2337,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl48),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2276,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl7),
  inference(resolution,[],[f704,f691])).
fof(f2185,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f704,f452])).
fof(f2183,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl7),
  inference(resolution,[],[f704,f451])).
fof(f2181,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl7),
  inference(resolution,[],[f704,f450])).
fof(f2180,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f704,f448])).
fof(f2178,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl7),
  inference(resolution,[],[f704,f446])).
fof(f2096,plain,(
  $false | ($spl4 | $spl40)),
  inference(trivial_inequality_removal,[],[f2070])).
fof(f2070,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ($spl4 | $spl40)),
  inference(superposition,[],[f2056,f703])).
fof(f2095,plain,(
  $false | ($spl2 | $spl40)),
  inference(trivial_inequality_removal,[],[f2081])).
fof(f2081,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ($spl2 | $spl40)),
  inference(superposition,[],[f2056,f700])).
fof(f2094,plain,(
  $false | $spl40),
  inference(equality_resolution,[],[f2056])).
fof(f2058,plain,(
  ( ! [X30,X32] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X30),X32)))) ) | $spl44),
  inference(cnf_transformation,[],[f2058_D])).
fof(f2058_D,plain,(
  ( ! [X30,X32] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X30),X32)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2059,plain,(
  ( ! [X31] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31))))) ) | $spl46),
  inference(cnf_transformation,[],[f2059_D])).
fof(f2059_D,plain,(
  ( ! [X31] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2051,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31)))) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X30),X32)))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f2050,f686])).
fof(f686,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),'const.infer.infer_st_subst_fupd_2'(s(fun(X0,X1),'const.combin.K_1'(s(X1,X3))),s('type.infer.infer_st'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),'const.infer.infer_st_subst_fupd_2'(s(fun(X0,X1),'const.combin.K_1'(s(X1,X3))),s('type.infer.infer_st'(X0),X2)))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X1),'const.infer.infer_st_subst_fupd_2'(s(fun(X0,X1),'const.combin.K_1'(s(X1,X3))),s('type.infer.infer_st'(X0),X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X45,X46] : ! [X9,X20] : s(X46,X20) = s(X46,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X46),'const.infer.infer_st_subst_fupd_2'(s(fun(X45,X46),'const.combin.K_1'(s(X46,X20))),s('type.infer.infer_st'(X45),X9)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.infer_st_rewrs|split|2')).
fof(f2050,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31)))) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X30),X32))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun(X30,'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31))),s('type.infer.infer_st'(X30),X32))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f2025,f686])).
fof(f2025,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X30),X32))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun(X30,'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31))),s('type.infer.infer_st'(X30),X32)))))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_subst_fupd_2'(s(fun(X30,'type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.combin.K_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X31))),s('type.infer.infer_st'(X30),X32))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f1895,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_subst_fupd_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X3))),s('type.infer.infer_st'(X1),X2)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_subst_fupd_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X3))),s('type.infer.infer_st'(X1),X2)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X1),X2))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_subst_fupd_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X3))),s('type.infer.infer_st'(X1),X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X47,X48] : ! [X9,X20] : s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X48),X9))) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'(X47),'const.infer.infer_st_subst_fupd_2'(s(fun(X48,X47),'const.combin.K_1'(s(X47,X20))),s('type.infer.infer_st'(X48),X9)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.infer_st_rewrs|split|3')).
fof(f1895,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(duplicate_literal_removal,[],[f1894])).
fof(f1894,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f699,f700])).
fof(f2056,plain,(
  ( ! [X42] : (s('type.num.num',sK12) != s('type.num.num',X42)) ) | $spl40),
  inference(cnf_transformation,[],[f2056_D])).
fof(f2056_D,plain,(
  ( ! [X42] : (s('type.num.num',sK12) != s('type.num.num',X42)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2057,plain,(
  ( ! [X43] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))))))) ) | $spl42),
  inference(cnf_transformation,[],[f2057_D])).
fof(f2057_D,plain,(
  ( ! [X43] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2055,plain,(
  ( ! [X43,X42] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43)))))) | s('type.num.num',sK12) != s('type.num.num',X42)) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f2054,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),X1))) = s(X0,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X0),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X2))),s('type.infer.infer_st'(X0),X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X49] : ! [X9,X50] : s(X49,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X49),X9))) = s(X49,'const.infer.infer_st_subst_1'(s('type.infer.infer_st'(X49),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X50))),s('type.infer.infer_st'(X49),X9)))))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.inferProps.infer_st_rewrs|split|5')).
fof(f2054,plain,(
  ( ! [X43,X42] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43)))))) | s('type.num.num',sK12) != s('type.num.num',X42) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X42))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f2030,f687])).
fof(f2030,plain,(
  ( ! [X43,X42] : (s('type.num.num',sK12) != s('type.num.num',X42) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X42))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43)))))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),'const.infer.infer_st_next_uvar_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num',X42))),s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X43))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f1895,f515])).
fof(f2049,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f2048,f705])).
fof(f2048,plain,(
  ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f2035,f709])).
fof(f2035,plain,(
  ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ($spl0 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f2024])).
fof(f2024,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f1895,f703])).
fof(f1754,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f691])).
fof(f1663,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1646,f452])).
fof(f1661,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f451])).
fof(f1659,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f450])).
fof(f1658,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1646,f448])).
fof(f1656,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) ) | $spl33),
  inference(resolution,[],[f1646,f446])).
fof(f1649,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl36),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1584,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f709,f691])).
fof(f1493,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X29) != s(bool,X30)) ) | $spl10),
  inference(resolution,[],[f709,f452])).
fof(f1490,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f709,f451])).
fof(f1489,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f709,f450])).
fof(f1488,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X14) != s(bool,X15)) ) | $spl10),
  inference(resolution,[],[f709,f448])).
fof(f1486,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl10),
  inference(resolution,[],[f709,f446])).
fof(f1482,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15)))))) | $spl10),
  inference(resolution,[],[f709,f678])).
fof(f1381,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl30),
  inference(cnf_transformation,[],[f1381_D])).
fof(f1381_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1338,plain,(
  ( ! [X167,X166] : (p(s(bool,X166)) | ~p(s(bool,X167)) | s(bool,X167) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl25),
  inference(resolution,[],[f1230,f691])).
fof(f1247,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X21) != s(bool,X23)) ) | $spl25),
  inference(resolution,[],[f1230,f452])).
fof(f1245,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl25),
  inference(resolution,[],[f1230,f451])).
fof(f1243,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl25),
  inference(resolution,[],[f1230,f450])).
fof(f1242,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))) | s(bool,X6) != s(bool,X8)) ) | $spl25),
  inference(resolution,[],[f1230,f448])).
fof(f1240,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) ) | $spl25),
  inference(resolution,[],[f1230,f446])).
fof(f1233,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl28),
  inference(cnf_transformation,[],[f1233_D])).
fof(f1233_D,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1168,plain,(
  ( ! [X180,X179] : (p(s(bool,X179)) | p(s(bool,X180)) | s(bool,X180) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl6),
  inference(resolution,[],[f705,f691])).
fof(f1077,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X29) != s(bool,X30)) ) | $spl6),
  inference(resolution,[],[f705,f452])).
fof(f1074,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl6),
  inference(resolution,[],[f705,f451])).
fof(f1073,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl6),
  inference(resolution,[],[f705,f450])).
fof(f1072,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))) | s(bool,X14) != s(bool,X15)) ) | $spl6),
  inference(resolution,[],[f705,f448])).
fof(f1070,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) ) | $spl6),
  inference(resolution,[],[f705,f446])).
fof(f1066,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t'),bool),'const.pair.UNCURRY_1'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'_dst_x00x10x01_3'(s(fun('type.infer_t.infer_t',fun('type.infer_t.infer_t',bool)),'=_0'),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))),s(fun('type.infer_t.infer_t','type.infer_t.infer_t'),'const.unify.t_walkstar_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16)))) | ~p(s(bool,'const.unify.t_wfs_1'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl6),
  inference(resolution,[],[f705,f678])).
fof(f988,plain,(
  ( ! [X267,X268,X266] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X266,'type.num.num')) | p(s(bool,X267)) | ~p(s(bool,X268)) | s(bool,X268) != s(bool,happ(s(fun('type.num.num',bool),X266),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f888,f691])).
fof(f897,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X30,'type.num.num')) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,happ(s(fun('type.num.num',bool),X30),s('type.num.num',sK12))) | s(bool,X31) != s(bool,X33)) ) | $spl2),
  inference(resolution,[],[f888,f452])).
fof(f895,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X22,'type.num.num')) | p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f888,f451])).
fof(f893,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X14,'type.num.num')) | p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,happ(s(fun('type.num.num',bool),X14),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f888,f450])).
fof(f892,plain,(
  ( ! [X12,X10,X13,X11] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X10,'type.num.num')) | ~p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X12) != s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',sK12))) | s(bool,X11) != s(bool,X13)) ) | $spl2),
  inference(resolution,[],[f888,f448])).
fof(f890,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.num.num',sK12) = s('type.num.num',sK22(sK13,X2,'type.num.num')) | p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK12)))) ) | $spl2),
  inference(resolution,[],[f888,f446])).
fof(f819,plain,(
  $false | ($spl3 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f818,f701])).
fof(f818,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f707,f703])).
fof(f709,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) | $spl10),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f711,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl12),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f426,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ((p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))))) & s('type.num.num',sK13) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) & s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) & s('type.num.num',sK12) = s('type.num.num',sK13))) & (! [X6] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',sK13) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK12) != s('type.num.num',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f297])).
fof(f297,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((? [X7] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X7)))))) & s('type.num.num',X1) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X7))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X7))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X7)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num',X1))) & (! [X6] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',X1) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(rectify,[],[f296])).
fof(f296,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((? [X6] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) & s('type.num.num',X1) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num',X1))) & (! [X6] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',X1) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((? [X6] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) & s('type.num.num',X1) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num',X1))) & (! [X6] : (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',X1) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) | (~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) | s('type.num.num',X0) != s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (? [X6] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) & s('type.num.num',X1) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) <~> (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (? [X6] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) & s('type.num.num',X1) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))) <=> (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X3))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X5),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X4))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X2)))))) & s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10,X11] : (? [X12] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))))) & s('type.num.num',X7) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X10),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X8)))))) & s('type.num.num',X6) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))) <=> (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X10))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X8)))))) & s('type.num.num',X6) = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10,X11] : (? [X12] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))))) & s('type.num.num',X7) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))) & p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X10),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X8)))))) & s('type.num.num',X6) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X12)))) <=> (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X9))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X11),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X10))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X8)))))) & s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.infer_eComplete.infer_eComplete_26_prove_ltb','thm.infer_eComplete.infer_eComplete_26_prove_ltb')).
fof(f707,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) | $spl8),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  s('type.num.num',sK13) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f427,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK13) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),
  inference(cnf_transformation,[],[f298])).
fof(f705,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | $spl6),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f428,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))),
  inference(cnf_transformation,[],[f298])).
fof(f703,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) | $spl4),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f429,plain,(
  p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),
  inference(cnf_transformation,[],[f298])).
fof(f700,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | $spl2),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f430,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))))))),
  inference(cnf_transformation,[],[f298])).
fof(f431,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | s('type.num.num',sK13) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),
  inference(cnf_transformation,[],[f298])).
fof(f432,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))),
  inference(cnf_transformation,[],[f298])).
fof(f433,plain,(
  s('type.num.num',sK12) = s('type.num.num',sK13) | s('type.num.num',sK12) = s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK18)))),
  inference(cnf_transformation,[],[f298])).
fof(f699,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',sK13) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) | $spl0),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6)))))) | s('type.num.num',sK13) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14))))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f701,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) | $spl3),
  inference(cnf_transformation,[],[f701_D])).
fof(f701_D,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f698,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num',sK13) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK13) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))))))) )),
  inference(subsumption_resolution,[],[f434,f437])).
fof(f437,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X3),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X2)))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X0),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),X1),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),X4))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f434,plain,(
  ( ! [X6] : (s('type.num.num',sK12) != s('type.num.num',sK13) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),'const.list.APPEND_2'(s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16))),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK12) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK16),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK14)))))) | s('type.num.num',sK13) != s('type.num.num','const.infer.infer_st_next_uvar_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))) | ~p(s(bool,'const.infer.pure_add_constraints_3'(s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),sK15))),s('type.list.list'('type.pair.prod'('type.infer_t.infer_t','type.infer_t.infer_t')),sK17),s('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t'),'const.infer.infer_st_subst_1'(s('type.infer.infer_st'('type.finite_map.fmap'('type.num.num','type.infer_t.infer_t')),X6))))))) )),
  inference(cnf_transformation,[],[f298])).
% SZS output end Proof for thm.infer_eComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1016
% Passive clauses: 16598
% Generated clauses: 44905
% Final active clauses: 584
% Final passive clauses: 9851
% Input formulas: 133
% Initial clauses: 265
% 
% Selected by SInE selection: 72
% SInE iterations: 1
% 
% Duplicate literals: 2070
% Trivial inequalities: 313
% Fw subsumption resolutions: 4702
% Fw demodulations: 4620
% Bw demodulations: 16
% 
% Simple tautologies: 3792
% Equational tautologies: 82
% Forward subsumptions: 12503
% Bw demodulations to eq. taut.: 30
% 
% Binary resolution: 13341
% Forward superposition: 14801
% Backward superposition: 4027
% Self superposition: 2
% Equality factoring: 2
% Equality resolution: 41
% 
% Split clauses: 459
% Split components: 986
% SAT solver clauses: 534
% SAT solver unit clauses: 18
% SAT solver binary clauses: 53
% 
% Sat splits: 459
% Sat splitting refutations: 69
% 
% Memory used [KB]: 16119
% Time elapsed: 1.630 s
% ------------------------------
% ------------------------------
% Success in time 13.37 s
