Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.PairingHeap
% (18610)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 2566
% Generated clauses: 18790
% Final active clauses: 95
% Final passive clauses: 2160
% Input formulas: 13
% Initial clauses: 37
% 
% Duplicate literals: 1136
% Trivial inequalities: 229
% Fw subsumption resolutions: 2642
% Fw demodulations: 1257
% Bw demodulations: 6
% 
% Simple tautologies: 381
% Equational tautologies: 101
% Forward subsumptions: 7956
% 
% Binary resolution: 144
% Forward superposition: 7803
% Backward superposition: 4663
% Self superposition: 378
% Equality factoring: 276
% Equality resolution: 17
% 
% Split clauses: 1123
% Split components: 2264
% SAT solver clauses: 2467
% SAT solver unit clauses: 254
% SAT solver binary clauses: 574
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 4
% 
% Sat splits: 1123
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2456
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.PairingHeap
% (18635)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 97
% Passive clauses: 351
% Generated clauses: 11169
% Final active clauses: 58
% Final passive clauses: 153
% Input formulas: 13
% Initial clauses: 37
% 
% Duplicate literals: 1207
% Trivial inequalities: 121
% Fw subsumption resolutions: 900
% Fw demodulations: 435
% Bw demodulations: 127
% Condensations: 2111
% Global subsumptions: 4
% 
% Simple tautologies: 183
% Equational tautologies: 69
% Forward subsumptions: 4865
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 134
% Factoring: 7
% Forward superposition: 3481
% Backward superposition: 1717
% Self superposition: 213
% Equality factoring: 179
% Equality resolution: 20
% 
% Split clauses: 96
% Split components: 199
% SAT solver clauses: 776
% SAT solver unit clauses: 32
% SAT solver binary clauses: 82
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 12
% 
% Sat splits: 96
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 756
% TWLsolver calls for satisfiability: 380
% 
% Memory used [KB]: 6268
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.PairingHeap
% (18657)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 584
% Passive clauses: 3906
% Generated clauses: 152190
% Final active clauses: 124
% Final passive clauses: 409
% Input formulas: 13
% Initial clauses: 37
% 
% Duplicate literals: 9894
% Trivial inequalities: 949
% Fw subsumption resolutions: 14273
% Fw demodulations: 13151
% Bw demodulations: 2313
% Condensations: 20174
% Global subsumptions: 312
% 
% Simple tautologies: 1328
% Equational tautologies: 2082
% Forward subsumptions: 85436
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 58
% 
% Unit resulting resolution: 7133
% Forward superposition: 38287
% Backward superposition: 28568
% Self superposition: 1077
% Equality factoring: 561
% Equality resolution: 102
% 
% Split clauses: 689
% Split components: 1393
% SAT solver clauses: 5904
% SAT solver unit clauses: 12
% SAT solver binary clauses: 202
% 
% Sat splits: 689
% Sat splitting refutations: 306
% 
% Memory used [KB]: 59999
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.PairingHeap
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.PairingHeap
% SZS output start Proof for thm.PairingHeap
fof(f966,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f94,f98,f96,f101,f97,f95,f102,f104,f105,f107,f112,f113,f261,f236,f234,f232,f230,f211,f209,f207,f205,f165,f164,f263,f163,f162,f264,f146,f145,f265,f144,f143,f266,f100,f269,f117,f115,f270,f116,f114,f487,f578,f581,f748,f585,f751,f428,f752,f753,f754,f384,f755,f756,f757,f99,f964])).
fof(f964,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl18)),
  inference(subsumption_resolution,[],[f963,f96])).
fof(f963,plain,(
  s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10) | ($spl2 | $spl6 | $spl18)),
  inference(subsumption_resolution,[],[f957,f92])).
fof(f92,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f68])).
fof(f68,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X14] : (s(X6,X14) = s(X6,X14) <=> p(s(bool,'T_0')))),
  file('thm.PairingHeap.PairingHeap_39_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f957,plain,(
  ~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10) | ($spl2 | $spl6 | $spl18)),
  inference(superposition,[],[f944,f811])).
fof(f811,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK10))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f806,f91])).
fof(f91,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.PairingHeap.PairingHeap_39_prove_ltb',aHL_FALSITY)).
fof(f806,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK10))) | ($spl2 | $spl6)),
  inference(superposition,[],[f790,f76])).
fof(f76,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.PairingHeap.PairingHeap_39_prove_ltb',aHL_BOOL_CASES)).
fof(f790,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK10)))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f104,f97])).
fof(f944,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',X0)))) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X0)) ) | $spl18),
  inference(equality_resolution,[],[f265])).
fof(f99,plain,(
  sP0(sK5,sK6) | $spl4),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  sP0(sK5,sK6) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f757,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK2(sK6,sK5)) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',X0))))) ) | $spl30),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK2(sK6,sK5)) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',X0))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f756,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK2(sK6,sK5)) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',X0)))) | sP0(sK5,sK6)) ) | $spl10),
  inference(forward_demodulation,[],[f383,f115])).
fof(f383,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',X0)))) | s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | sP0(sK5,sK6)) )),
  inference(resolution,[],[f187,f58])).
fof(f58,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),sK3(X1,X0)),s('type.semanticPrimitives.v',sK4(X1,X0)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X6,X7,X8,X9] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X6),s('type.semanticPrimitives.v',X7)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X8),s('type.semanticPrimitives.v',X9)))) | ((s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(X0),X6) = s('type.PairingHeap.heap'(X0),X8)) & (s('type.PairingHeap.heap'(X0),X6) != s('type.PairingHeap.heap'(X0),X8) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9))))) & ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),sK1(X1,X0)),s('type.semanticPrimitives.v',sK2(X1,X0))))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),sK3(X1,X0)),s('type.semanticPrimitives.v',sK4(X1,X0))))) & (s('type.semanticPrimitives.v',sK2(X1,X0)) = s('type.semanticPrimitives.v',sK4(X1,X0)) | s('type.PairingHeap.heap'(X0),sK1(X1,X0)) = s('type.PairingHeap.heap'(X0),sK3(X1,X0))) & (s('type.semanticPrimitives.v',sK2(X1,X0)) != s('type.semanticPrimitives.v',sK4(X1,X0)) | s('type.PairingHeap.heap'(X0),sK1(X1,X0)) != s('type.PairingHeap.heap'(X0),sK3(X1,X0)))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X6,X7,X8,X9] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X6),s('type.semanticPrimitives.v',X7)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X8),s('type.semanticPrimitives.v',X9)))) | ((s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(X0),X6) = s('type.PairingHeap.heap'(X0),X8)) & (s('type.PairingHeap.heap'(X0),X6) != s('type.PairingHeap.heap'(X0),X8) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9))))) & (? [X2,X3,X4,X5] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X2),s('type.semanticPrimitives.v',X3)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X4),s('type.semanticPrimitives.v',X5)))) & (s('type.semanticPrimitives.v',X3) = s('type.semanticPrimitives.v',X5) | s('type.PairingHeap.heap'(X0),X2) = s('type.PairingHeap.heap'(X0),X4)) & (s('type.semanticPrimitives.v',X3) != s('type.semanticPrimitives.v',X5) | s('type.PairingHeap.heap'(X0),X2) != s('type.PairingHeap.heap'(X0),X4))) | sP0(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X7,X8,X9,X10] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) | ((s('type.semanticPrimitives.v',X8) != s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9)) & (s('type.PairingHeap.heap'(X0),X7) != s('type.PairingHeap.heap'(X0),X9) | s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10))))) & (? [X7,X8,X9,X10] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) & (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9)) & (s('type.semanticPrimitives.v',X8) != s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) != s('type.PairingHeap.heap'(X0),X9))) | sP0(X0,X1)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X7,X8,X9,X10] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) | ((s('type.semanticPrimitives.v',X8) != s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9)) & (s('type.PairingHeap.heap'(X0),X7) != s('type.PairingHeap.heap'(X0),X9) | s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10))))) & (? [X7,X8,X9,X10] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) & ((s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9)) & (s('type.semanticPrimitives.v',X8) != s('type.semanticPrimitives.v',X10) | s('type.PairingHeap.heap'(X0),X7) != s('type.PairingHeap.heap'(X0),X9)))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ! [X7,X8,X9,X10] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) | (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) <=> s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f187,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X0)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X1)))) | s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',X1)) )),
  inference(equality_resolution,[],[f108])).
fof(f108,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.PairingHeap.heap'(sK5),X6) != s('type.PairingHeap.heap'(sK5),X8) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X6),s('type.semanticPrimitives.v',X7))))) )),
  inference(subsumption_resolution,[],[f62,f56])).
fof(f56,plain,(
  ( ! [X6,X0,X8,X7,X1,X9] : (s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(X0),X6) != s('type.PairingHeap.heap'(X0),X8) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X6),s('type.semanticPrimitives.v',X7)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f36])).
fof(f62,plain,(
  ( ! [X6,X8,X7,X9] : (sP0(sK5,sK6) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(sK5),X6) != s('type.PairingHeap.heap'(sK5),X8) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X6),s('type.semanticPrimitives.v',X7))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ((! [X6,X7,X8,X9] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X6),s('type.semanticPrimitives.v',X7)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X8),s('type.semanticPrimitives.v',X9)))) | ((s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(sK5),X6) = s('type.PairingHeap.heap'(sK5),X8)) & (s('type.PairingHeap.heap'(sK5),X6) != s('type.PairingHeap.heap'(sK5),X8) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9)))) & p(s(bool,'T_0'))) | sP0(sK5,sK6)) & ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK10)))) & (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9)) & (s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',sK10) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),sK9))) | ~p(s(bool,'T_0')) | ~sP0(sK5,sK6))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7,sK8,sK9,sK10])],[f39])).
fof(f39,plain,(
  ? [X0,X1] : (((! [X6,X7,X8,X9] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X6),s('type.semanticPrimitives.v',X7)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X8),s('type.semanticPrimitives.v',X9)))) | ((s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(X0),X6) = s('type.PairingHeap.heap'(X0),X8)) & (s('type.PairingHeap.heap'(X0),X6) != s('type.PairingHeap.heap'(X0),X8) | s('type.semanticPrimitives.v',X7) = s('type.semanticPrimitives.v',X9)))) & p(s(bool,'T_0'))) | sP0(X0,X1)) & (? [X2,X3,X4,X5] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X2),s('type.semanticPrimitives.v',X3)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X4),s('type.semanticPrimitives.v',X5)))) & (s('type.semanticPrimitives.v',X3) = s('type.semanticPrimitives.v',X5) | s('type.PairingHeap.heap'(X0),X2) = s('type.PairingHeap.heap'(X0),X4)) & (s('type.semanticPrimitives.v',X3) != s('type.semanticPrimitives.v',X5) | s('type.PairingHeap.heap'(X0),X2) != s('type.PairingHeap.heap'(X0),X4))) | ~p(s(bool,'T_0')) | ~sP0(X0,X1)))),
  inference(rectify,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (((! [X3,X4,X5,X6] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) | ((s('type.semanticPrimitives.v',X4) != s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5)) & (s('type.PairingHeap.heap'(X0),X3) != s('type.PairingHeap.heap'(X0),X5) | s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6)))) & p(s(bool,'T_0'))) | sP0(X0,X1)) & (? [X3,X4,X5,X6] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) & (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5)) & (s('type.semanticPrimitives.v',X4) != s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) != s('type.PairingHeap.heap'(X0),X5))) | ~p(s(bool,'T_0')) | ~sP0(X0,X1)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ? [X0,X1] : (((! [X3,X4,X5,X6] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) | ((s('type.semanticPrimitives.v',X4) != s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5)) & (s('type.PairingHeap.heap'(X0),X3) != s('type.PairingHeap.heap'(X0),X5) | s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6)))) & p(s(bool,'T_0'))) | sP0(X0,X1)) & ((? [X3,X4,X5,X6] : (p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) & ((s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5)) & (s('type.semanticPrimitives.v',X4) != s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(X0),X3) != s('type.PairingHeap.heap'(X0),X5)))) | ~p(s(bool,'T_0'))) | ~sP0(X0,X1)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ? [X0,X1] : ((! [X3,X4,X5,X6] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) | (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) <=> s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5))) & p(s(bool,'T_0'))) <~> sP0(X0,X1))),
  inference(definition_folding,[],[f29,f31])).
fof(f29,plain,(
  ? [X0,X1] : ((! [X3,X4,X5,X6] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6)))) | (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) <=> s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5))) & p(s(bool,'T_0'))) <~> ! [X7,X8,X9,X10] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10)))) | (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) <=> s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ? [X0,X1] : ((! [X3,X4,X5,X6] : ((~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6))))) | (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) <=> s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5))) & p(s(bool,'T_0'))) <~> ! [X7,X8,X9,X10] : ((~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10))))) | (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) <=> s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1] : ((! [X3,X4,X5,X6] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6))))) => (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) <=> s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5))) & p(s(bool,'T_0'))) <=> ! [X7,X8,X9,X10] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10))))) => (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) <=> s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ~! [X0] : ! [X1] : ((! [X3,X4,X5,X6] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X3),s('type.semanticPrimitives.v',X4)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X5),s('type.semanticPrimitives.v',X6))))) => (s('type.semanticPrimitives.v',X4) = s('type.semanticPrimitives.v',X6) <=> s('type.PairingHeap.heap'(X0),X3) = s('type.PairingHeap.heap'(X0),X5))) & p(s(bool,'T_0'))) <=> ! [X7,X8,X9,X10] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X7),s('type.semanticPrimitives.v',X8)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X9),s('type.semanticPrimitives.v',X10))))) => (s('type.semanticPrimitives.v',X8) = s('type.semanticPrimitives.v',X10) <=> s('type.PairingHeap.heap'(X0),X7) = s('type.PairingHeap.heap'(X0),X9))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ((! [X8,X9,X10,X11] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X8),s('type.semanticPrimitives.v',X9)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X10),s('type.semanticPrimitives.v',X11))))) => (s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X11) <=> s('type.PairingHeap.heap'(X6),X8) = s('type.PairingHeap.heap'(X6),X10))) & ! [X12] : p(s(bool,'T_0'))) <=> ! [X8,X9,X10,X11] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X8),s('type.semanticPrimitives.v',X9)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X10),s('type.semanticPrimitives.v',X11))))) => (s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X11) <=> s('type.PairingHeap.heap'(X6),X8) = s('type.PairingHeap.heap'(X6),X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ((! [X8,X9,X10,X11] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X8),s('type.semanticPrimitives.v',X9)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X10),s('type.semanticPrimitives.v',X11))))) => (s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X11) <=> s('type.PairingHeap.heap'(X6),X8) = s('type.PairingHeap.heap'(X6),X10))) & ! [X12] : p(s(bool,'T_0'))) <=> ! [X8,X9,X10,X11] : ((p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X8),s('type.semanticPrimitives.v',X9)))) & p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X6,fun('type.semanticPrimitives.v',bool)),X7),s('type.PairingHeap.heap'(X6),X10),s('type.semanticPrimitives.v',X11))))) => (s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X11) <=> s('type.PairingHeap.heap'(X6),X8) = s('type.PairingHeap.heap'(X6),X10))))),
  file('thm.PairingHeap.PairingHeap_39_prove_ltb','thm.PairingHeap.PairingHeap_39_prove_ltb')).
fof(f755,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',X1)))) | s('type.semanticPrimitives.v',X1) = s('type.semanticPrimitives.v',sK2(sK6,sK5))) ) | $spl28),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',X1)))) | s('type.semanticPrimitives.v',X1) = s('type.semanticPrimitives.v',sK2(sK6,sK5))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f384,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',X1)))) | s('type.semanticPrimitives.v',X1) = s('type.semanticPrimitives.v',sK2(sK6,sK5)) | sP0(sK5,sK6)) )),
  inference(resolution,[],[f187,f57])).
fof(f57,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),sK1(X1,X0)),s('type.semanticPrimitives.v',sK2(X1,X0)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f754,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X0) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5))) ) | $spl26),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X0) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f753,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X0) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | sP0(sK5,sK6)) ) | $spl10),
  inference(forward_demodulation,[],[f427,f115])).
fof(f427,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK4(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X0) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | sP0(sK5,sK6)) )),
  inference(resolution,[],[f279,f58])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X2)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X1),s('type.semanticPrimitives.v',X2)))) | s('type.PairingHeap.heap'(sK5),X0) = s('type.PairingHeap.heap'(sK5),X1)) )),
  inference(equality_resolution,[],[f109])).
fof(f109,plain,(
  ( ! [X6,X8,X7,X9] : (s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | s('type.PairingHeap.heap'(sK5),X6) = s('type.PairingHeap.heap'(sK5),X8) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X6),s('type.semanticPrimitives.v',X7))))) )),
  inference(subsumption_resolution,[],[f61,f55])).
fof(f55,plain,(
  ( ! [X6,X0,X8,X7,X1,X9] : (s('type.PairingHeap.heap'(X0),X6) = s('type.PairingHeap.heap'(X0),X8) | s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(X0,fun('type.semanticPrimitives.v',bool)),X1),s('type.PairingHeap.heap'(X0),X6),s('type.semanticPrimitives.v',X7)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f36])).
fof(f61,plain,(
  ( ! [X6,X8,X7,X9] : (sP0(sK5,sK6) | s('type.PairingHeap.heap'(sK5),X6) = s('type.PairingHeap.heap'(sK5),X8) | s('type.semanticPrimitives.v',X7) != s('type.semanticPrimitives.v',X9) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X8),s('type.semanticPrimitives.v',X9)))) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X6),s('type.semanticPrimitives.v',X7))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f752,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X1),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X1) = s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5))) ) | $spl24),
  inference(cnf_transformation,[],[f752_D])).
fof(f752_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X1),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X1) = s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f428,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X1),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X1) = s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) | sP0(sK5,sK6)) )),
  inference(resolution,[],[f279,f57])).
fof(f751,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | $spl22),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f585,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | sP0(sK5,sK6) | $spl10),
  inference(superposition,[],[f58,f115])).
fof(f748,plain,(
  $false | ($spl5 | $spl10 | $spl13)),
  inference(subsumption_resolution,[],[f747,f116])).
fof(f747,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f743,f92])).
fof(f743,plain,(
  ~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | ($spl5 | $spl10)),
  inference(superposition,[],[f449,f642])).
fof(f642,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5)))) | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f641,f91])).
fof(f641,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5)))) | ($spl5 | $spl10)),
  inference(superposition,[],[f624,f76])).
fof(f624,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | ($spl5 | $spl10)),
  inference(forward_demodulation,[],[f268,f115])).
fof(f268,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',sK4(sK6,sK5))))) | $spl5),
  inference(resolution,[],[f100,f58])).
fof(f449,plain,(
  ( ! [X1] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X1),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | s('type.PairingHeap.heap'(sK5),X1) = s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5))) ) | $spl5),
  inference(subsumption_resolution,[],[f428,f100])).
fof(f581,plain,(
  $false | ($spl5 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f580,f114])).
fof(f580,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | ($spl5 | $spl12)),
  inference(subsumption_resolution,[],[f573,f92])).
fof(f573,plain,(
  ~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | ($spl5 | $spl12)),
  inference(superposition,[],[f404,f329])).
fof(f329,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5)))) | $spl5),
  inference(subsumption_resolution,[],[f328,f91])).
fof(f328,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5)))) | $spl5),
  inference(superposition,[],[f267,f76])).
fof(f267,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',sK2(sK6,sK5))))) | $spl5),
  inference(resolution,[],[f100,f57])).
fof(f404,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)),s('type.semanticPrimitives.v',X0)))) | s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK4(sK6,sK5))) ) | ($spl5 | $spl12)),
  inference(forward_demodulation,[],[f403,f117])).
fof(f403,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)),s('type.semanticPrimitives.v',X0)))) | s('type.semanticPrimitives.v',X0) = s('type.semanticPrimitives.v',sK4(sK6,sK5))) ) | $spl5),
  inference(subsumption_resolution,[],[f383,f100])).
fof(f578,plain,(
  $false | ($spl5 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f577,f100])).
fof(f577,plain,(
  sP0(sK5,sK6) | ($spl5 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f562,f114])).
fof(f562,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | sP0(sK5,sK6) | ($spl5 | $spl12)),
  inference(resolution,[],[f404,f57])).
fof(f487,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f486,f98])).
fof(f486,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | ($spl0 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f482,f92])).
fof(f482,plain,(
  ~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | ($spl0 | $spl6 | $spl8)),
  inference(superposition,[],[f452,f180])).
fof(f180,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK8))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f175,f91])).
fof(f175,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK8))) | ($spl0 | $spl6)),
  inference(superposition,[],[f142,f76])).
fof(f142,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK8)))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f104,f95])).
fof(f452,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK8)))) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f441,f92])).
fof(f441,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK8)))) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0)) ) | $spl8),
  inference(superposition,[],[f279,f161])).
fof(f161,plain,(
  s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8))) | $spl8),
  inference(subsumption_resolution,[],[f156,f91])).
fof(f156,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8))) | $spl8),
  inference(superposition,[],[f107,f76])).
fof(f114,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) != s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl11),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) != s('type.semanticPrimitives.v',sK4(sK6,sK5)) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f116,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) != s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | $spl13),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) != s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f270,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) != s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | s('type.semanticPrimitives.v',sK2(sK6,sK5)) != s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl5),
  inference(resolution,[],[f100,f60])).
fof(f60,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | s('type.PairingHeap.heap'(X0),sK1(X1,X0)) != s('type.PairingHeap.heap'(X0),sK3(X1,X0)) | s('type.semanticPrimitives.v',sK2(X1,X0)) != s('type.semanticPrimitives.v',sK4(X1,X0))) )),
  inference(cnf_transformation,[],[f36])).
fof(f115,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl10),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f117,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | $spl12),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f269,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl5),
  inference(resolution,[],[f100,f59])).
fof(f59,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | s('type.PairingHeap.heap'(X0),sK1(X1,X0)) = s('type.PairingHeap.heap'(X0),sK3(X1,X0)) | s('type.semanticPrimitives.v',sK2(X1,X0)) = s('type.semanticPrimitives.v',sK4(X1,X0))) )),
  inference(cnf_transformation,[],[f36])).
fof(f100,plain,(
  ~sP0(sK5,sK6) | $spl5),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  ~sP0(sK5,sK6) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f266,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X2) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3)) ) | $spl20),
  inference(cnf_transformation,[],[f266_D])).
fof(f266_D,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X2) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f143,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(resolution,[],[f107,f55])).
fof(f144,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(resolution,[],[f107,f55])).
fof(f265,plain,(
  ( ! [X6,X7] : (s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X7) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6))))) ) | $spl18),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  ( ! [X6,X7] : (s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X7) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f145,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(resolution,[],[f107,f56])).
fof(f146,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(resolution,[],[f107,f56])).
fof(f264,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X2) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3)) ) | $spl16),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X2) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f162,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f142,f55])).
fof(f163,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f142,f55])).
fof(f263,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6))))) ) | $spl14),
  inference(cnf_transformation,[],[f263_D])).
fof(f263_D,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f164,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f142,f56])).
fof(f165,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f142,f56])).
fof(f205,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(subsumption_resolution,[],[f196,f92])).
fof(f196,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(superposition,[],[f55,f161])).
fof(f207,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(subsumption_resolution,[],[f197,f92])).
fof(f197,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(superposition,[],[f55,f161])).
fof(f209,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(subsumption_resolution,[],[f198,f92])).
fof(f198,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(superposition,[],[f56,f161])).
fof(f211,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(subsumption_resolution,[],[f199,f92])).
fof(f199,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | $spl8),
  inference(superposition,[],[f56,f161])).
fof(f230,plain,(
  ( ! [X0,X1] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f221,f92])).
fof(f221,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X0) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f55,f180])).
fof(f232,plain,(
  ( ! [X2,X3] : (s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f222,f92])).
fof(f222,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK9) = s('type.PairingHeap.heap'(sK5),X2) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X3) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X2),s('type.semanticPrimitives.v',X3)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f55,f180])).
fof(f234,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f223,f92])).
fof(f223,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X4) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X5) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X5),s('type.semanticPrimitives.v',X4)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f56,f180])).
fof(f236,plain,(
  ( ! [X6,X7] : (s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f224,f92])).
fof(f224,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',X6) | s('type.PairingHeap.heap'(sK5),sK9) != s('type.PairingHeap.heap'(sK5),X7) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X7),s('type.semanticPrimitives.v',X6)))) | ~sP0(sK5,sK6)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f56,f180])).
fof(f261,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f260,f98])).
fof(f260,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | ($spl0 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f257,f92])).
fof(f257,plain,(
  ~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | ($spl0 | $spl4 | $spl6 | $spl8)),
  inference(superposition,[],[f244,f180])).
fof(f244,plain,(
  ( ! [X0] : (~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',sK8)))) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0)) ) | ($spl4 | $spl8)),
  inference(equality_resolution,[],[f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',X1) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),X0) | ~p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),X0),s('type.semanticPrimitives.v',X1))))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f143,f99])).
fof(f113,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) != s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | s('type.semanticPrimitives.v',sK2(sK6,sK5)) != s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl5),
  inference(resolution,[],[f100,f60])).
fof(f112,plain,(
  s('type.PairingHeap.heap'(sK5),sK1(sK6,sK5)) = s('type.PairingHeap.heap'(sK5),sK3(sK6,sK5)) | s('type.semanticPrimitives.v',sK2(sK6,sK5)) = s('type.semanticPrimitives.v',sK4(sK6,sK5)) | $spl5),
  inference(resolution,[],[f100,f59])).
fof(f107,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8)))) | $spl8),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f105,plain,(
  ~sP0(sK5,sK6) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8))))),
  inference(subsumption_resolution,[],[f64,f92])).
fof(f64,plain,(
  ~sP0(sK5,sK6) | ~p(s(bool,'T_0')) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK7),s('type.semanticPrimitives.v',sK8))))),
  inference(cnf_transformation,[],[f40])).
fof(f104,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK10)))) | $spl6),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK10)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f102,plain,(
  ~sP0(sK5,sK6) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK10))))),
  inference(subsumption_resolution,[],[f65,f92])).
fof(f65,plain,(
  ~sP0(sK5,sK6) | ~p(s(bool,'T_0')) | p(s(bool,'const.PairingHeap.PAIRINGHEAP_HEAP_TYPE_3'(s(fun(sK5,fun('type.semanticPrimitives.v',bool)),sK6),s('type.PairingHeap.heap'(sK5),sK9),s('type.semanticPrimitives.v',sK10))))),
  inference(cnf_transformation,[],[f40])).
fof(f95,plain,(
  s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10) | $spl0),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f97,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | $spl2),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f101,plain,(
  ~sP0(sK5,sK6) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10)),
  inference(subsumption_resolution,[],[f66,f92])).
fof(f66,plain,(
  ~sP0(sK5,sK6) | ~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) = s('type.PairingHeap.heap'(sK5),sK9) | s('type.semanticPrimitives.v',sK8) = s('type.semanticPrimitives.v',sK10)),
  inference(cnf_transformation,[],[f40])).
fof(f96,plain,(
  s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',sK10) | $spl1),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',sK10) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f98,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),sK9) | $spl3),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),sK9) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f94,plain,(
  ~sP0(sK5,sK6) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),sK9) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',sK10)),
  inference(subsumption_resolution,[],[f67,f92])).
fof(f67,plain,(
  ~sP0(sK5,sK6) | ~p(s(bool,'T_0')) | s('type.PairingHeap.heap'(sK5),sK7) != s('type.PairingHeap.heap'(sK5),sK9) | s('type.semanticPrimitives.v',sK8) != s('type.semanticPrimitives.v',sK10)),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for thm.PairingHeap
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 69
% Passive clauses: 356
% Generated clauses: 876
% Final active clauses: 34
% Final passive clauses: 102
% Input formulas: 13
% Initial clauses: 37
% 
% Selected by SInE selection: 13
% SInE iterations: 1
% 
% Duplicate literals: 5
% Trivial inequalities: 8
% Fw subsumption resolutions: 95
% Fw demodulations: 80
% Bw demodulations: 7
% 
% Simple tautologies: 37
% Equational tautologies: 3
% Forward subsumptions: 278
% 
% Binary resolution: 90
% Forward superposition: 367
% Backward superposition: 168
% Equality resolution: 6
% 
% Split clauses: 29
% Split components: 60
% SAT solver clauses: 36
% SAT solver binary clauses: 4
% 
% Sat splits: 29
% Sat splitting refutations: 6
% 
% Memory used [KB]: 1279
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
% Success in time 11.784 s
