Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (15728)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 2688
% Generated clauses: 19685
% Final active clauses: 103
% Final passive clauses: 2316
% Input formulas: 132
% Initial clauses: 139
% 
% Duplicate literals: 708
% Trivial inequalities: 226
% Fw subsumption resolutions: 2318
% Fw demodulations: 3518
% Bw demodulations: 30
% 
% Simple tautologies: 235
% Equational tautologies: 72
% Forward subsumptions: 9415
% 
% Binary resolution: 304
% Factoring: 1
% Forward superposition: 5821
% Backward superposition: 6100
% Self superposition: 261
% Equality factoring: 160
% Equality resolution: 16
% 
% Split clauses: 235
% Split components: 470
% SAT solver clauses: 673
% SAT solver unit clauses: 19
% SAT solver binary clauses: 92
% 
% Sat splits: 235
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 671
% TWLsolver calls for satisfiability: 130
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (15755)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 902
% Generated clauses: 13423
% Final active clauses: 90
% Final passive clauses: 751
% Input formulas: 132
% Initial clauses: 139
% 
% Duplicate literals: 542
% Trivial inequalities: 82
% Fw subsumption resolutions: 661
% Fw demodulations: 1368
% Bw demodulations: 59
% Condensations: 1047
% 
% Simple tautologies: 136
% Equational tautologies: 43
% Forward subsumptions: 6452
% 
% Binary resolution: 194
% Factoring: 1
% Forward superposition: 4047
% Backward superposition: 4025
% Self superposition: 213
% Equality factoring: 98
% Equality resolution: 9
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 1642
% SAT solver unit clauses: 258
% SAT solver binary clauses: 45
% 
% Sat splits: 40
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 1642
% TWLsolver calls for satisfiability: 291
% 
% Memory used [KB]: 19573
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f18555,plain,(
  $false),
  inference(subsumption_resolution,[],[f18554,f404])).
fof(f404,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.state_fupdfupds_comp|split|21',aHL_TRUTH)).
fof(f18554,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f18552,f2537])).
fof(f2537,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f2435,f547])).
fof(f547,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f405,f406])).
fof(f406,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.stackSem.state_fupdfupds_comp|split|21',aHL_BOOL_CASES)).
fof(f405,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.stackSem.state_fupdfupds_comp|split|21',aHL_FALSITY)).
fof(f2435,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f2356,f1178])).
fof(f1178,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X0,X0),'_i_0_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1138,f463])).
fof(f463,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7,X20] : ! [X11,X10,X18] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X11),s(fun(X20,X7),X10),s(X20,X18))) = s(X6,happ(s(fun(X7,X6),X11),s(X7,happ(s(fun(X20,X7),X10),s(X20,X18)))))),
  file('thm.stackSem.state_fupdfupds_comp|split|21','thm.combin.o_THM')).
fof(f1138,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f407,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f402])).
fof(f402,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X49,X50,X51] : (p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X49,fun(X49,bool)),'=_0'),s(X49,X50))),s(X49,X51)))) <=> s(X49,X50) = s(X49,X51))),
  file('thm.stackSem.state_fupdfupds_comp|split|21',arity_0_HL_EQ)).
fof(f407,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X27] : ! [X29] : s(X27,X29) = s(X27,happ(s(fun(X27,X27),'_i_0_0'),s(X27,X29)))),
  file('thm.stackSem.state_fupdfupds_comp|split|21','_i0_ax')).
fof(f2356,plain,(
  ( ! [X70,X68,X71,X69] : (s(X70,happ(s(fun(X68,X70),X71),s(X68,X69))) = s(X70,'const.combin.o_3'(s(fun(X68,X70),X71),s(fun(X68,X68),'_i_0_0'),s(X68,X69)))) )),
  inference(superposition,[],[f463,f407])).
fof(f18552,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),sK5),s(fun(bool,bool),sK4))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),sK5),s(fun(bool,bool),sK4))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))))))),
  inference(backward_demodulation,[],[f18505,f6060])).
fof(f6060,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),bool),happ(s(fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),fun(fun(sK0,'type.stackSem.state'(sK1,sK2)),bool)),'=_0'),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK5))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK4))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),sK5),s(fun(bool,bool),sK4))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))))))),
  inference(unit_resulting_resolution,[],[f403,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f402])).
fof(f403,plain,(
  s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),sK5),s(fun(bool,bool),sK4))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK5))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK4))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3)))))),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),sK5),s(fun(bool,bool),sK4))))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3))) != s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK5))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),'const.combin.o_2'(s(fun('type.stackSem.state'(sK1,sK2),'type.stackSem.state'(sK1,sK2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),sK4))),s(fun(sK0,'type.stackSem.state'(sK1,sK2)),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f388])).
fof(f388,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4))))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3))) != s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X5))),s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X4))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4))))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3))) = s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X5))),s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X4))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4))))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3))) = s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X5))),s(fun(X0,'type.stackSem.state'(X1,X2)),'const.combin.o_2'(s(fun('type.stackSem.state'(X1,X2),'type.stackSem.state'(X1,X2)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X4))),s(fun(X0,'type.stackSem.state'(X1,X2)),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X11),s(fun(bool,bool),X10))))),s(fun(X6,'type.stackSem.state'(X7,X8)),X9))) = s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X11))),s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X10))),s(fun(X6,'type.stackSem.state'(X7,X8)),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X11),s(fun(bool,bool),X10))))),s(fun(X6,'type.stackSem.state'(X7,X8)),X9))) = s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X11))),s(fun(X6,'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X10))),s(fun(X6,'type.stackSem.state'(X7,X8)),X9)))))),
  file('thm.stackSem.state_fupdfupds_comp|split|21','thm.stackSem.state_fupdfupds_comp|split|21')).
fof(f18505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X4,'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X3))))),s(fun(X4,'type.stackSem.state'(X0,X1)),X5))) = s(fun(X4,'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X2))),s(fun(X4,'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X3))),s(fun(X4,'type.stackSem.state'(X0,X1)),X5)))))) )),
  inference(superposition,[],[f476,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X2))))) = s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X3))),s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X2)))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X2))))) = s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X3))),s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X2)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ! [X2,X3] : s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X2))))) = s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.combin.o_2'(s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X3))),s(fun('type.stackSem.state'(X0,X1),'type.stackSem.state'(X0,X1)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X2)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : ! [X10,X11] : s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),'const.combin.o_2'(s(fun(bool,bool),X11),s(fun(bool,bool),X10))))) = s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.combin.o_2'(s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X11))),s(fun('type.stackSem.state'(X7,X8),'type.stackSem.state'(X7,X8)),'const.stackSem.state_use_alloc_fupd_1'(s(fun(bool,bool),X10)))))),
  file('thm.stackSem.state_fupdfupds_comp|split|21','thm.stackSem.state_fupdfupds_comp|split|20')).
fof(f476,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7,X22,X20] : ! [X11,X10,X9] : s(fun(X22,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X22,X7),'const.combin.o_2'(s(fun(X20,X7),X10),s(fun(X22,X20),X9))))) = s(fun(X22,X6),'const.combin.o_2'(s(fun(X20,X6),'const.combin.o_2'(s(fun(X7,X6),X11),s(fun(X20,X7),X10))),s(fun(X22,X20),X9)))),
  file('thm.stackSem.state_fupdfupds_comp|split|21','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 161
% Passive clauses: 1491
% Generated clauses: 18561
% Final active clauses: 92
% Final passive clauses: 765
% Input formulas: 132
% Initial clauses: 139
% 
% Duplicate literals: 721
% Trivial inequalities: 120
% Fw subsumption resolutions: 739
% Fw demodulations: 2932
% Bw demodulations: 524
% Condensations: 1155
% Global subsumptions: 1
% 
% Simple tautologies: 149
% Equational tautologies: 165
% Forward subsumptions: 10998
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 821
% Factoring: 1
% Forward superposition: 5696
% Backward superposition: 4752
% Self superposition: 226
% Equality factoring: 110
% Equality resolution: 10
% 
% Split clauses: 34
% Split components: 68
% SAT solver clauses: 1580
% SAT solver unit clauses: 403
% SAT solver binary clauses: 513
% 
% Sat splits: 34
% Sat splitting refutations: 18
% 
% Memory used [KB]: 18677
% Time elapsed: 0.789 s
% ------------------------------
% ------------------------------
% Success in time 2.16 s
