Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.misc
% (72266)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 158
% Passive clauses: 2059
% Generated clauses: 9051
% Final active clauses: 132
% Final passive clauses: 1721
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 354
% Trivial inequalities: 281
% Fw subsumption resolutions: 1057
% Fw demodulations: 1059
% Bw demodulations: 21
% 
% Simple tautologies: 268
% Equational tautologies: 22
% Forward subsumptions: 3568
% Fw demodulations to eq. taut.: 8
% 
% Binary resolution: 321
% Forward superposition: 3326
% Backward superposition: 2656
% Self superposition: 57
% Equality factoring: 36
% Equality resolution: 49
% 
% Split clauses: 263
% Split components: 537
% SAT solver clauses: 545
% SAT solver unit clauses: 23
% SAT solver binary clauses: 144
% 
% Sat splits: 263
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 542
% TWLsolver calls for satisfiability: 157
% 
% Memory used [KB]: 5245
% Time elapsed: 0.620 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.misc
% (72291)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 121
% Passive clauses: 1260
% Generated clauses: 13592
% Final active clauses: 103
% Final passive clauses: 961
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 733
% Trivial inequalities: 314
% Fw subsumption resolutions: 1093
% Fw demodulations: 1274
% Bw demodulations: 177
% Condensations: 1016
% 
% Simple tautologies: 218
% Equational tautologies: 44
% Forward subsumptions: 5454
% Fw demodulations to eq. taut.: 33
% Bw demodulations to eq. taut.: 16
% 
% Binary resolution: 178
% Forward superposition: 3269
% Backward superposition: 3958
% Self superposition: 133
% Equality factoring: 100
% Equality resolution: 38
% 
% Split clauses: 161
% Split components: 322
% SAT solver clauses: 2134
% SAT solver unit clauses: 39
% SAT solver binary clauses: 55
% 
% Sat splits: 161
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2134
% TWLsolver calls for satisfiability: 268
% 
% Memory used [KB]: 10874
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.misc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.misc
% SZS output start Proof for thm.misc
fof(f4870,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f165,f175,f173,f481,f485,f483,f526,f527,f528,f529,f531,f797,f804,f802,f1624,f1626,f1622,f1630,f1628,f1656,f1658,f1790,f1791,f1792,f1793,f1795,f1800,f1802,f1801,f1804,f2220,f2232,f2234,f2235,f2236,f2237,f2238,f2239,f2240,f2241,f2242,f2243,f2246,f2247,f2248,f2249,f2252,f2253,f2254,f2255,f2256,f2257,f2258,f2259,f2261,f2263,f2264,f2265,f2266,f2268,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f3939,f3949,f3938,f3950,f482,f4494,f4495,f4496,f4497,f4498,f4499,f4500,f4565,f4566,f4503,f4560,f4504,f4505,f4569,f4567,f4506,f4568,f4555,f4507,f4513,f4514,f4515,f4563,f4516,f4557,f4562,f4521,f4579,f4532,f4564,f4590,f4592,f4561,f4593,f4556,f4559,f4639,f4642,f4851,f4852,f4854,f4858,f4869])).
fof(f4869,plain,(
  $false | ($spl0 | $spl13 | $spl14)),
  inference(subsumption_resolution,[],[f4868,f4643])).
fof(f4643,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4621,f241])).
fof(f241,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X4,bool),X5),s('type.option.option'(X4),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f237,f82])).
fof(f82,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.misc.OPTION_EVERY_cong','thm.bool.TRUTH')).
fof(f237,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X4,bool),X5),s('type.option.option'(X4),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f151,f94])).
fof(f94,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.NONE_0')))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.NONE_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X22] : s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X22),s('type.option.option'(X6),'const.option.NONE_0')))),
  file('thm.misc.OPTION_EVERY_cong','thm.misc.OPTION_EVERY_def|split|0')).
fof(f151,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f90,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.misc.OPTION_EVERY_cong',aHL_BOOL_CASES)).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.misc.OPTION_EVERY_cong','thm.bool.EQ_CLAUSES|split|2')).
fof(f4621,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f4611,f2302])).
fof(f2302,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f2296])).
fof(f2296,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f2277,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f85,f85])).
fof(f2277,plain,(
  s(bool,'T_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | $spl0),
  inference(forward_demodulation,[],[f139,f173])).
fof(f139,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1)))),
  inference(backward_demodulation,[],[f79,f81])).
fof(f81,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK2)))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK2) & ! [X5] : (s('type.option.option'(sK0),sK2) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X5))) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X5))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X5)))) & s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f52])).
fof(f52,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),X2) & ! [X5] : (s('type.option.option'(X0),X2) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) & s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X3),s('type.option.option'(X0),X1))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X4),s('type.option.option'(X0),X2))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.option.option'(X0),X1) = s('type.option.option'(X0),X2) & ! [X5] : (s('type.option.option'(X0),X2) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) | s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) & s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X3),s('type.option.option'(X0),X1))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X4),s('type.option.option'(X0),X2))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3,X4] : ((s('type.option.option'(X0),X1) = s('type.option.option'(X0),X2) & ! [X5] : (s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X3),s('type.option.option'(X0),X1))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X4),s('type.option.option'(X0),X2))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((s('type.option.option'(X0),X1) = s('type.option.option'(X0),X2) & ! [X5] : (s('type.option.option'(X0),X2) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X5))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X3),s('type.option.option'(X0),X1))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X4),s('type.option.option'(X0),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((s('type.option.option'(X6),X7) = s('type.option.option'(X6),X8) & ! [X11] : (s('type.option.option'(X6),X8) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X11))) => s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))) = s(bool,happ(s(fun(X6,bool),X10),s(X6,X11))))) => s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X9),s('type.option.option'(X6),X7))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X10),s('type.option.option'(X6),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((s('type.option.option'(X6),X7) = s('type.option.option'(X6),X8) & ! [X11] : (s('type.option.option'(X6),X8) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X11))) => s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))) = s(bool,happ(s(fun(X6,bool),X10),s(X6,X11))))) => s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X9),s('type.option.option'(X6),X7))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X10),s('type.option.option'(X6),X8))))),
  file('thm.misc.OPTION_EVERY_cong','thm.misc.OPTION_EVERY_cong')).
fof(f79,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK2)),
  inference(cnf_transformation,[],[f60])).
fof(f4611,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | $spl14),
  inference(global_subsumption,[],[f81,f80,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f1628,f1631,f1632,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f2192,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f2251,f1515,f2260,f2262,f658,f263,f139,f265,f400,f2269,f2267,f455,f807,f845,f937,f940,f1146,f1151,f2982,f2986,f2987,f2988,f2993,f2995,f2996,f3053,f3054,f3060,f3061,f3062,f3063,f3035,f3037,f3039,f3040,f3210,f3218,f3233,f3241,f3617,f3626,f3904,f4610])).
fof(f4610,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))),
  inference(forward_demodulation,[],[f2757,f79])).
fof(f2757,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK2)),
  inference(trivial_inequality_removal,[],[f2755])).
fof(f2755,plain,(
  s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),sK1) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK2)),
  inference(superposition,[],[f459,f79])).
fof(f3904,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X37) != s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X37))) | s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X34,bool),X35),s('type.option.option'(X34),'const.option.SOME_1'(s(X34,X36)))))) )),
  inference(superposition,[],[f305,f98])).
fof(f3626,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,sK5(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f269,f95])).
fof(f3617,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))))) | s(bool,X35) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X32,bool),X33),s('type.option.option'(X32),'const.option.SOME_1'(s(X32,X34))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f269,f98])).
fof(f3241,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,sK5(X0,bool)) | s(bool,'T_0') = s(bool,X1) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f268,f95])).
fof(f3233,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X36))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X33,bool),X34),s(X33,X35))))) | s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X33,bool),X34),s('type.option.option'(X33),'const.option.SOME_1'(s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f268,f98])).
fof(f3218,plain,(
  ( ! [X0,X1] : (s('type.option.option'(bool),X0) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK5(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) )),
  inference(superposition,[],[f268,f95])).
fof(f3210,plain,(
  ( ! [X33,X31,X34,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X34))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X31,bool),X32),s('type.option.option'(X31),'const.option.SOME_1'(s(X31,X33)))))) )),
  inference(superposition,[],[f268,f98])).
fof(f3040,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,'T_0')))) | s(bool,X106) != s(bool,happ(s(fun(X102,bool),X103),s(X102,X104))) | p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X106)))) | s(X102,X104) = s(X102,sK6(X104,X103,X102))) )),
  inference(superposition,[],[f105,f2269])).
fof(f3039,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0')))) | s(bool,happ(s(fun(X98,bool),X99),s(X98,X100))) = s(bool,sK6(happ(s(fun(X98,bool),X99),s(X98,X100)),X101,bool)) | s(X98,X100) = s(X98,sK6(X100,X99,X98))) )),
  inference(superposition,[],[f103,f2269])).
fof(f3037,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,'T_0') != s(bool,X93) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X93))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X90,bool),X91),s(X90,X92))))) | s(X90,X92) = s(X90,sK6(X92,X91,X90))) )),
  inference(superposition,[],[f102,f2269])).
fof(f3035,plain,(
  ( ! [X85,X83,X84,X82] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X85))) | s(bool,X85) = s(bool,happ(s(fun(X82,bool),X83),s(X82,X84))) | s(X82,X84) = s(X82,sK6(X84,X83,X82))) )),
  inference(superposition,[],[f101,f2269])).
fof(f3063,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X47) != s(bool,happ(s(fun(X44,bool),X45),s(X44,X46))) | s(bool,'T_0') = s(bool,X47) | s(X44,X46) = s(X44,sK6(X46,X45,X44))) )),
  inference(subsumption_resolution,[],[f3023,f82])).
fof(f3023,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | s(bool,X47) != s(bool,happ(s(fun(X44,bool),X45),s(X44,X46))) | s(bool,'T_0') = s(bool,X47) | s(X44,X46) = s(X44,sK6(X46,X45,X44))) )),
  inference(superposition,[],[f151,f2269])).
fof(f3062,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X29) != s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) | p(s(bool,X29)) | s(X26,X28) = s(X26,sK6(X28,X27,X26))) )),
  inference(subsumption_resolution,[],[f3020,f82])).
fof(f3020,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X29) != s(bool,happ(s(fun(X26,bool),X27),s(X26,X28))) | p(s(bool,X29)) | s(X26,X28) = s(X26,sK6(X28,X27,X26))) )),
  inference(superposition,[],[f136,f2269])).
fof(f3061,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK6(X9,X8,X7)) = s(X7,sK6(sK6(X9,X8,X7),X8,X7))) )),
  inference(subsumption_resolution,[],[f3013,f82])).
fof(f3013,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | s(X7,sK6(X9,X8,X7)) = s(X7,sK6(sK6(X9,X8,X7),X8,X7))) )),
  inference(superposition,[],[f104,f2269])).
fof(f3060,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(subsumption_resolution,[],[f3011,f82])).
fof(f3011,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(superposition,[],[f105,f2269])).
fof(f3054,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X41,bool),X42),s('type.option.option'(X41),'const.option.SOME_1'(s(X41,X43)))),X44,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X44),s(bool,happ(s(fun(X41,bool),X42),s(X41,X43)))))) )),
  inference(forward_demodulation,[],[f2998,f98])).
fof(f2998,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X44),s(bool,happ(s(fun(X41,bool),X42),s(X41,X43))))) | s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X41,bool),X42),s('type.option.option'(X41),'const.option.SOME_1'(s(X41,X43))))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X41,bool),X42),s('type.option.option'(X41),'const.option.SOME_1'(s(X41,X43)))),X44,bool))) )),
  inference(superposition,[],[f2269,f98])).
fof(f3053,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X38,bool),X39),s('type.option.option'(X38),'const.option.NONE_0')),X40,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0')))) )),
  inference(forward_demodulation,[],[f2997,f94])).
fof(f2997,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0'))) | s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X38,bool),X39),s('type.option.option'(X38),'const.option.NONE_0'))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X38,bool),X39),s('type.option.option'(X38),'const.option.NONE_0')),X40,bool))) )),
  inference(superposition,[],[f2269,f94])).
fof(f2996,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0'))) | s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) = s(bool,sK6(happ(s(fun(X34,bool),X35),s(X34,X36)),X37,bool)) | s(X34,X36) = s(X34,sK6(X36,X35,X34))) )),
  inference(superposition,[],[f2269,f2269])).
fof(f2995,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,X32))) | s(bool,'F_0') = s(bool,sK6('F_0',X33,bool)) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f2269,f85])).
fof(f2993,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),X27),s(bool,X26))) | s(bool,'T_0') = s(bool,sK6('T_0',X27,bool)) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f2269,f149])).
fof(f2988,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,X14))) | s(bool,X13) = s(bool,sK6(X13,X15,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2269,f149])).
fof(f2987,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,X10) = s(bool,sK6(X10,X11,bool)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2269,f149])).
fof(f2986,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0'))) | s(bool,X8) = s(bool,sK6(X8,X9,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2269,f85])).
fof(f2982,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(X0),bool),X2),s('type.option.option'(X0),X1))) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK5(X1,X0)))) = s('type.option.option'(X0),sK6('const.option.SOME_1'(s(X0,sK5(X1,X0))),X2,'type.option.option'(X0))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f2269,f95])).
fof(f1151,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X138))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,sK5(X137,bool))) )),
  inference(superposition,[],[f95,f149])).
fof(f1146,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(bool,bool),X118),s(bool,X117)))) | s(bool,X116) = s(bool,sK6(X116,X118,bool)) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f103,f149])).
fof(f940,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,f(s(X6,sK7(X7,X8,bool,X6)))) | s(fun(X6,bool),X7) = s(fun(X6,bool),X8) | s(bool,'T_0') = s(bool,g(s(X6,sK7(X7,X8,bool,X6))))) )),
  inference(superposition,[],[f110,f85])).
fof(f937,plain,(
  ( ! [X6,X8,X7] : (s(X8,f(s(bool,'F_0'))) != s(X8,g(s(bool,'F_0'))) | s(fun(bool,X8),X6) = s(fun(bool,X8),X7) | s(bool,'T_0') = s(bool,sK7(X6,X7,X8,bool))) )),
  inference(superposition,[],[f110,f85])).
fof(f845,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | s(bool,X15) != s(bool,X17) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X17)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f105,f85])).
fof(f807,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | s(bool,'T_0') = s(bool,sK6(X4,X5,bool))) )),
  inference(superposition,[],[f104,f85])).
fof(f455,plain,(
  ( ! [X1] : (s('type.option.option'(bool),X1) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X1) | s(bool,'T_0') = s(bool,sK5(X1,bool))) )),
  inference(superposition,[],[f95,f85])).
fof(f2267,plain,(
  ( ! [X2] : (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X2,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X2),s('type.option.option'(sK0),sK1)))) )),
  inference(subsumption_resolution,[],[f453,f84])).
fof(f453,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X2,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X2),s('type.option.option'(sK0),sK1)))) )),
  inference(superposition,[],[f407,f85])).
fof(f2269,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(X3,X5) = s(X3,sK6(X5,X4,X3))) )),
  inference(subsumption_resolution,[],[f406,f84])).
fof(f406,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | s(X3,X5) = s(X3,sK6(X5,X4,X3)) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) )),
  inference(superposition,[],[f103,f85])).
fof(f400,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | s(bool,X5) = s(bool,sK6(X5,X6,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f103,f85])).
fof(f265,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f140,f101])).
fof(f263,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f140,f101])).
fof(f658,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0')))) | s(bool,X6) = s(bool,sK6(X6,X7,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f103,f85])).
fof(f2262,plain,(
  ( ! [X6,X8,X7] : (s(X6,X8) = s(X6,sK6(X8,X7,X6)) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) )),
  inference(subsumption_resolution,[],[f667,f84])).
fof(f667,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(X6,X8) = s(X6,sK6(X8,X7,X6)) | s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) )),
  inference(superposition,[],[f103,f85])).
fof(f2260,plain,(
  ( ! [X2] : (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X2,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X2),s('type.option.option'(sK0),sK1)))) )),
  inference(subsumption_resolution,[],[f677,f84])).
fof(f677,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X2,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X2),s('type.option.option'(sK0),sK1)))) )),
  inference(superposition,[],[f407,f85])).
fof(f1515,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(condensation,[],[f1514])).
fof(f1514,plain,(
  ( ! [X61,X62,X60,X63] : (s(X60,X62) = s(X60,sK6(X62,X61,X60)) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),X61),s(X60,X62)))) )),
  inference(subsumption_resolution,[],[f1124,f87])).
fof(f1124,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X63)) | s(X60,X62) = s(X60,sK6(X62,X61,X60)) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),X61),s(X60,X62)))) )),
  inference(superposition,[],[f103,f149])).
fof(f2251,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X0),s('type.option.option'(sK0),sK1))) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X0,'type.option.option'(sK0)))) )),
  inference(condensation,[],[f2250])).
fof(f2250,plain,(
  ( ! [X76,X77] : (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X76,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X76),s('type.option.option'(sK0),sK1)))) )),
  inference(subsumption_resolution,[],[f1128,f87])).
fof(f1128,plain,(
  ( ! [X76,X77] : (p(s(bool,X77)) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X76,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X76),s('type.option.option'(sK0),sK1)))) )),
  inference(superposition,[],[f407,f149])).
fof(f1524,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(condensation,[],[f1523])).
fof(f1523,plain,(
  ( ! [X61,X62,X60,X63] : (s(X60,X62) = s(X60,sK6(X62,X61,X60)) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),X61),s(X60,X62))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1179,f87])).
fof(f1179,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X63)) | s(X60,X62) = s(X60,sK6(X62,X61,X60)) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),X61),s(X60,X62))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f103,f149])).
fof(f2245,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X0),s('type.option.option'(sK0),sK1))) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X0,'type.option.option'(sK0)))) )),
  inference(condensation,[],[f2244])).
fof(f2244,plain,(
  ( ! [X76,X77] : (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X76,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X76),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1183,f87])).
fof(f1183,plain,(
  ( ! [X76,X77] : (p(s(bool,X77)) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X76,'type.option.option'(sK0))) | s(bool,'T_0') = s(bool,happ(s(fun('type.option.option'(sK0),bool),X76),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f407,f149])).
fof(f1201,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,happ(s(fun(bool,bool),X118),s(bool,X117)))) | s(bool,X116) = s(bool,sK6(X116,X118,bool)) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f103,f149])).
fof(f1206,plain,(
  ( ! [X138,X137] : (s('type.option.option'(bool),X137) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X138))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X137) | s(bool,'T_0') = s(bool,sK5(X137,bool)) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f95,f149])).
fof(f2083,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | s(bool,X11) = s(bool,sK6(X11,X13,bool)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f103,f149])).
fof(f2084,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | s(bool,X14) = s(bool,sK6(X14,X16,bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f103,f149])).
fof(f2085,plain,(
  ( ! [X17,X18] : (p(s(bool,happ(s(fun(bool,bool),X18),s(bool,'F_0')))) | s(bool,X17) = s(bool,sK6(X17,X18,bool)) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f103,f85])).
fof(f2106,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(condensation,[],[f2105])).
fof(f2105,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X10) = s(X8,sK6(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2096,f87])).
fof(f2096,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X10) = s(X8,sK6(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f103,f149])).
fof(f2108,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(condensation,[],[f2107])).
fof(f2107,plain,(
  ( ! [X14,X12,X15,X13] : (s(X12,X14) = s(X12,sK6(X14,X13,X12)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) )),
  inference(subsumption_resolution,[],[f2097,f87])).
fof(f2097,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(X12,X14) = s(X12,sK6(X14,X13,X12)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) )),
  inference(superposition,[],[f103,f149])).
fof(f2233,plain,(
  ( ! [X17,X18,X16] : (s(X16,X18) = s(X16,sK6(X18,X17,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) )),
  inference(subsumption_resolution,[],[f2098,f84])).
fof(f2098,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(X16,X18) = s(X16,sK6(X18,X17,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) )),
  inference(superposition,[],[f103,f85])).
fof(f2192,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))) )),
  inference(superposition,[],[f85,f98])).
fof(f138,plain,(
  ( ! [X5] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X5))) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X5))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X5)))) )),
  inference(backward_demodulation,[],[f79,f80])).
fof(f407,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X0),s('type.option.option'(sK0),sK1)))) | s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6(sK2,X0,'type.option.option'(sK0)))) )),
  inference(forward_demodulation,[],[f397,f79])).
fof(f397,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X0),s('type.option.option'(sK0),sK1)))) | s('type.option.option'(sK0),sK2) = s('type.option.option'(sK0),sK6(sK2,X0,'type.option.option'(sK0)))) )),
  inference(superposition,[],[f103,f79])).
fof(f459,plain,(
  ( ! [X7] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),X7) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(X7,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(X7,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),X7)) )),
  inference(superposition,[],[f138,f95])).
fof(f1621,plain,(
  ( ! [X0] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),X0) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5('const.option.SOME_1'(s(sK0,sK5(X0,sK0))),sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5('const.option.SOME_1'(s(sK0,sK5(X0,sK0))),sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5(X0,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),X0)) )),
  inference(superposition,[],[f459,f95])).
fof(f84,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.misc.OPTION_EVERY_cong',aHL_FALSITY)).
fof(f91,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f63])).
fof(f268,plain,(
  ( ! [X2,X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X1))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f101,f85])).
fof(f269,plain,(
  ( ! [X4,X3] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f101,f85])).
fof(f305,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f102,f85])).
fof(f306,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X3))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X4))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f102,f85])).
fof(f398,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK6('F_0',X2,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f103,f85])).
fof(f846,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X18)))) | s(bool,'F_0') != s(bool,X20) | p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f105,f85])).
fof(f1274,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') != s(bool,f(s(X122,sK7(X123,X124,bool,X122)))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,X125) = s(bool,g(s(X122,sK7(X123,X124,bool,X122)))) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f110,f149])).
fof(f1150,plain,(
  ( ! [X134,X136,X135,X133] : (s(bool,X136) != s(bool,f(s(X133,sK7(X134,X135,bool,X133)))) | s(fun(X133,bool),X134) = s(fun(X133,bool),X135) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,g(s(X133,sK7(X134,X135,bool,X133))))) )),
  inference(superposition,[],[f110,f149])).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X1),s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X22,X21] : s(bool,happ(s(fun(X6,bool),X22),s(X6,X21))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X6,bool),X22),s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X21)))))),
  file('thm.misc.OPTION_EVERY_cong','thm.misc.OPTION_EVERY_def|split|1')).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK6(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X20,X21] : (! [X11] : (s(X6,X11) = s(X6,X21) => p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X11))))) <=> p(s(bool,happ(s(fun(X6,bool),X20),s(X6,X21)))))),
  file('thm.misc.OPTION_EVERY_cong','thm.bool.UNWIND_FORALL_THM2')).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f105,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f708,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X32))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))))) | s(bool,X32) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X29,bool),X30),s('type.option.option'(X29),'const.option.SOME_1'(s(X29,X31)))))) )),
  inference(superposition,[],[f101,f98])).
fof(f726,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,happ(s(fun(X45,bool),X46),s(X45,X47))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X45,bool),X46),s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X47)))),X48,bool)) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,happ(s(fun(X45,bool),X46),s(X45,X47))))))) )),
  inference(forward_demodulation,[],[f712,f98])).
fof(f712,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,happ(s(fun(bool,bool),X48),s(bool,happ(s(fun(X45,bool),X46),s(X45,X47)))))) | s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X45,bool),X46),s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X47))))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X45,bool),X46),s('type.option.option'(X45),'const.option.SOME_1'(s(X45,X47)))),X48,bool))) )),
  inference(superposition,[],[f103,f98])).
fof(f1042,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X32) = s(bool,happ(s(fun(X29,bool),X30),s(X29,X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X29,bool),X30),s('type.option.option'(X29),'const.option.SOME_1'(s(X29,X31)))))) )),
  inference(superposition,[],[f149,f98])).
fof(f1078,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(bool,X24) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X21,bool),X22),s('type.option.option'(X21),'const.option.SOME_1'(s(X21,X23))))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f149,f98])).
fof(f1536,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(X49,X51) != s(X49,X52) | p(s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))) | s(bool,X53) = s(bool,happ(s(fun(X49,bool),X50),s(X49,X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1251,f82])).
fof(f1251,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | s(X49,X51) != s(X49,X52) | p(s(bool,happ(s(fun(X49,bool),X50),s(X49,X52)))) | s(bool,X53) = s(bool,happ(s(fun(X49,bool),X50),s(X49,X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f105,f149])).
fof(f1537,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),X66),s(X65,X67)))) | s(bool,X68) = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK6(X67,X66,X65)))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f1254,f82])).
fof(f1254,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X65,bool),X66),s(X65,X67)))) | s(bool,X68) = s(bool,happ(s(fun(X65,bool),X66),s(X65,sK6(X67,X66,X65)))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f104,f149])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X0,bool),X2),s('type.option.option'(X0),X1))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f98,f95])).
fof(f1147,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,happ(s(fun(bool,bool),X121),s(bool,X120)))) | s(bool,X119) != s(bool,X122) | p(s(bool,happ(s(fun(bool,bool),X121),s(bool,X122)))) | s(bool,'T_0') = s(bool,X120) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f105,f149])).
fof(f1152,plain,(
  ( ! [X140,X141,X139] : (~p(s(bool,happ(s(fun(bool,bool),X140),s(bool,X141)))) | p(s(bool,happ(s(fun(bool,bool),X140),s(bool,X139)))) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,sK6(X139,X140,bool))) )),
  inference(superposition,[],[f104,f149])).
fof(f1223,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | s(bool,'T_0') = s(bool,sK6('T_0',X45,bool)) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f103,f149])).
fof(f1224,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X47)))) | s(bool,'T_0') != s(bool,X49) | p(s(bool,happ(s(fun(bool,bool),X48),s(bool,X49)))) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f105,f149])).
fof(f1271,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0')))) | s(bool,X109) != s(bool,X111) | p(s(bool,happ(s(fun(bool,bool),X110),s(bool,X111)))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f105,f149])).
fof(f1276,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,happ(s(fun(bool,bool),X129),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X129),s(bool,X128)))) | s(bool,X130) = s(bool,sK6(X128,X129,bool)) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f104,f149])).
fof(f409,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X7,bool),X8),s('type.option.option'(X7),'const.option.NONE_0')),X9,bool)) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f401,f94])).
fof(f401,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X7,bool),X8),s('type.option.option'(X7),'const.option.NONE_0'))) = s(bool,sK6('const.misc.OPTION_EVERY_2'(s(fun(X7,bool),X8),s('type.option.option'(X7),'const.option.NONE_0')),X9,bool))) )),
  inference(superposition,[],[f103,f94])).
fof(f1270,plain,(
  ( ! [X107,X108,X106] : (p(s(bool,happ(s(fun(bool,bool),X107),s(bool,'T_0')))) | s(bool,X106) = s(bool,sK6(X106,X107,bool)) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f103,f149])).
fof(f841,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.option.option'(X2),bool),X4),s('type.option.option'(X2),X3)))) | s('type.option.option'(X2),X5) != s('type.option.option'(X2),'const.option.SOME_1'(s(X2,sK5(X3,X2)))) | p(s(bool,happ(s(fun('type.option.option'(X2),bool),X4),s('type.option.option'(X2),X5)))) | s('type.option.option'(X2),'const.option.NONE_0') = s('type.option.option'(X2),X3)) )),
  inference(superposition,[],[f105,f95])).
fof(f461,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,happ(s(fun('type.option.option'(X11),bool),X13),s('type.option.option'(X11),X12)))) | s('type.option.option'(X11),'const.option.SOME_1'(s(X11,sK5(X12,X11)))) = s('type.option.option'(X11),sK6('const.option.SOME_1'(s(X11,sK5(X12,X11))),X13,'type.option.option'(X11))) | s('type.option.option'(X11),'const.option.NONE_0') = s('type.option.option'(X11),X12)) )),
  inference(superposition,[],[f103,f95])).
fof(f1632,plain,(
  ( ! [X5] : (s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X5))) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X5))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X5)))) ) | $spl14),
  inference(backward_demodulation,[],[f1628,f138])).
fof(f1631,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK2) | $spl14),
  inference(backward_demodulation,[],[f1628,f79])).
fof(f1277,plain,(
  ( ! [X134,X132,X133,X131] : (s(X133,f(s(bool,'T_0'))) != s(X133,g(s(bool,'T_0'))) | s(fun(bool,X133),X131) = s(fun(bool,X133),X132) | s(bool,X134) = s(bool,sK7(X131,X132,X133,bool)) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f110,f149])).
fof(f1275,plain,(
  ( ! [X127,X126] : (s('type.option.option'(bool),X126) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X126) | s(bool,X127) = s(bool,sK5(X126,bool)) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f95,f149])).
fof(f1268,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X100))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X101))) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f102,f149])).
fof(f1266,plain,(
  ( ! [X94,X95,X96] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X95))) | s(bool,X94) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f101,f149])).
fof(f1533,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1233,f82])).
fof(f1233,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f151,f149])).
fof(f1532,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | p(s(bool,X11)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1232,f82])).
fof(f1232,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | p(s(bool,X11)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f136,f149])).
fof(f1221,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,X39) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X39))) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f102,f149])).
fof(f1209,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f82,f149])).
fof(f1153,plain,(
  ( ! [X144,X142,X145,X143] : (s(X144,f(s(bool,X145))) != s(X144,g(s(bool,X145))) | s(fun(bool,X144),X142) = s(fun(bool,X144),X143) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,sK7(X142,X143,X144,bool))) )),
  inference(superposition,[],[f110,f149])).
fof(f1144,plain,(
  ( ! [X111,X112,X110] : (s(bool,X111) != s(bool,X112) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X110))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X112))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f102,f149])).
fof(f1142,plain,(
  ( ! [X105,X106,X104] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X105))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X106))) | s(bool,X104) = s(bool,X106) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f101,f149])).
fof(f1067,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f149,f149])).
fof(f464,plain,(
  ( ! [X21,X22,X20] : (s('type.option.option'('type.option.option'(X20)),'const.option.SOME_1'(s('type.option.option'(X20),X21))) != s('type.option.option'('type.option.option'(X20)),'const.option.SOME_1'(s('type.option.option'(X20),X22))) | s('type.option.option'(X20),X22) = s('type.option.option'(X20),'const.option.SOME_1'(s(X20,sK5(X21,X20)))) | s('type.option.option'(X20),'const.option.NONE_0') = s('type.option.option'(X20),X21)) )),
  inference(superposition,[],[f101,f95])).
fof(f462,plain,(
  ( ! [X14,X15,X16] : (s('type.option.option'(X14),X15) != s('type.option.option'(X14),X16) | s('type.option.option'('type.option.option'(X14)),'const.option.SOME_1'(s('type.option.option'(X14),X16))) = s('type.option.option'('type.option.option'(X14)),'const.option.SOME_1'(s('type.option.option'(X14),'const.option.SOME_1'(s(X14,sK5(X15,X14)))))) | s('type.option.option'(X14),'const.option.NONE_0') = s('type.option.option'(X14),X15)) )),
  inference(superposition,[],[f102,f95])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),X1) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X2) = s(X0,sK5(X1,X0)) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(superposition,[],[f101,f95])).
fof(f270,plain,(
  ( ! [X6,X7,X5] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X7))) | s(bool,X7) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(X5,bool),X6),s('type.option.option'(X5),'const.option.NONE_0')))) )),
  inference(superposition,[],[f101,f94])).
fof(f140,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f84,f86])).
fof(f136,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f135])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f124])).
fof(f124,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11,X18,X12,X19] : ((s(bool,X11) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X12) = s(bool,X19))) => ((p(s(bool,X11)) => p(s(bool,X12))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.misc.OPTION_EVERY_cong','thm.bool.IMP_CONG')).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.misc.OPTION_EVERY_cong',aHL_EXT)).
fof(f101,plain,(
  ( ! [X2,X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : ((s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) != s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X11,X12] : (s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X11))) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X12))) <=> s(X6,X11) = s(X6,X12))),
  file('thm.misc.OPTION_EVERY_cong','thm.option.SOME_|49||49|')).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f95,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK5(X1,X0)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,sK5(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : (s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1) | ? [X2] : s('type.option.option'(X0),X1) = s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X2))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X23] : (s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X23) | ? [X11] : s('type.option.option'(X6),X23) = s('type.option.option'(X6),'const.option.SOME_1'(s(X6,X11))))),
  file('thm.misc.OPTION_EVERY_cong','thm.option.option_nchotomy')).
fof(f86,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.misc.OPTION_EVERY_cong','thm.bool.EQ_CLAUSES|split|0')).
fof(f87,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f61])).
fof(f80,plain,(
  ( ! [X5] : (s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,X5))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X5))) | s('type.option.option'(sK0),sK2) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X5)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f4868,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl13 | $spl14)),
  inference(forward_demodulation,[],[f1625,f4643])).
fof(f1625,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) != s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f1625_D])).
fof(f1625_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) != s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f4858,plain,(
  $false | ($spl0 | $spl11 | $spl14)),
  inference(subsumption_resolution,[],[f4663,f4643])).
fof(f4663,plain,(
  s(bool,'T_0') != s(bool,sK5('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl0 | $spl11 | $spl14)),
  inference(backward_demodulation,[],[f4643,f925])).
fof(f925,plain,(
  s(bool,'T_0') != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | $spl11),
  inference(subsumption_resolution,[],[f904,f84])).
fof(f904,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | $spl11),
  inference(superposition,[],[f86,f868])).
fof(f868,plain,(
  s(bool,'F_0') = s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | $spl11),
  inference(unit_resulting_resolution,[],[f814,f101])).
fof(f814,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl11),
  inference(unit_resulting_resolution,[],[f804,f95])).
fof(f4854,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4853,f4643])).
fof(f4853,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4645,f86])).
fof(f4645,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f4643,f90])).
fof(f4852,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4644,f82])).
fof(f4644,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f4643,f84])).
fof(f4851,plain,(
  $false | ($spl0 | $spl14)),
  inference(trivial_inequality_removal,[],[f4647])).
fof(f4647,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f4643,f265])).
fof(f4642,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f4620,f94])).
fof(f4620,plain,(
  s(bool,'T_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f4611,f2277])).
fof(f4639,plain,(
  $false | ($spl2 | $spl14)),
  inference(subsumption_resolution,[],[f4638,f140])).
fof(f4638,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl14)),
  inference(forward_demodulation,[],[f4615,f94])).
fof(f4615,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl2 | $spl14)),
  inference(backward_demodulation,[],[f4611,f174])).
fof(f174,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | $spl2),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f4559,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f4559_D])).
fof(f4559_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4556,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f4556_D])).
fof(f4556_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f4593,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl12)),
  inference(forward_demodulation,[],[f4536,f173])).
fof(f4536,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl12)),
  inference(superposition,[],[f2351,f1626])).
fof(f2351,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f2328,f82])).
fof(f2328,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f151,f173])).
fof(f4561,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f4561_D])).
fof(f4561_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4592,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl2 | $spl12)),
  inference(forward_demodulation,[],[f4535,f174])).
fof(f4535,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl12)),
  inference(superposition,[],[f2302,f1626])).
fof(f4590,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl11 | $spl12)),
  inference(forward_demodulation,[],[f4533,f868])).
fof(f4533,plain,(
  s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl11 | $spl12)),
  inference(superposition,[],[f1513,f1626])).
fof(f1513,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X29)) ) | $spl11),
  inference(subsumption_resolution,[],[f1081,f140])).
fof(f1081,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s(bool,'T_0') = s(bool,X29)) ) | $spl11),
  inference(superposition,[],[f149,f868])).
fof(f4564,plain,(
  ( ! [X4] : (s(bool,sK5('const.option.SOME_1'(s(bool,X4)),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl38),
  inference(cnf_transformation,[],[f4564_D])).
fof(f4564_D,plain,(
  ( ! [X4] : (s(bool,sK5('const.option.SOME_1'(s(bool,X4)),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4532,plain,(
  ( ! [X15] : (s(bool,sK5('const.option.SOME_1'(s(bool,X15)),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f1458,f1626])).
fof(f1458,plain,(
  ( ! [X30,X31] : (s(bool,X31) = s(bool,sK5('const.option.SOME_1'(s(bool,X30)),bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl11),
  inference(duplicate_literal_removal,[],[f1082])).
fof(f1082,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,sK5('const.option.SOME_1'(s(bool,X30)),bool)) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl11),
  inference(superposition,[],[f149,f900])).
fof(f900,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK5('const.option.SOME_1'(s(bool,X0)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f868,f85])).
fof(f4579,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10)))) ) | $spl44),
  inference(cnf_transformation,[],[f4579_D])).
fof(f4579_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4521,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X10))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))) ) | $spl12),
  inference(superposition,[],[f305,f1626])).
fof(f4562,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl34),
  inference(cnf_transformation,[],[f4562_D])).
fof(f4562_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4557,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f4557_D])).
fof(f4557_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4516,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,X5) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl12),
  inference(superposition,[],[f149,f1626])).
fof(f4563,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl36),
  inference(cnf_transformation,[],[f4563_D])).
fof(f4563_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4515,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(superposition,[],[f149,f1626])).
fof(f4514,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))) ) | $spl12),
  inference(superposition,[],[f149,f1626])).
fof(f4513,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,X2) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(superposition,[],[f149,f1626])).
fof(f4507,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | $spl12),
  inference(superposition,[],[f85,f1626])).
fof(f4555,plain,(
  s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) | $spl26),
  inference(cnf_transformation,[],[f4555_D])).
fof(f4555_D,plain,(
  s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4568,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) | $spl42),
  inference(cnf_transformation,[],[f4568_D])).
fof(f4568_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4506,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) | s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) | $spl12),
  inference(superposition,[],[f103,f1626])).
fof(f4567,plain,(
  ( ! [X0] : (s(sK0,X0) != s(sK0,sK5(sK2,sK0)) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X0))))) ) | $spl40),
  inference(cnf_transformation,[],[f4567_D])).
fof(f4567_D,plain,(
  ( ! [X0] : (s(sK0,X0) != s(sK0,sK5(sK2,sK0)) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X0))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4569,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) | $spl43),
  inference(cnf_transformation,[],[f4569_D])).
fof(f4569_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f4505,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0))))) | s(sK0,X0) != s(sK0,sK5(sK2,sK0)) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X0))))) ) | $spl12),
  inference(superposition,[],[f105,f1626])).
fof(f4504,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) | $spl12),
  inference(superposition,[],[f2269,f1626])).
fof(f4560,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | $spl33),
  inference(cnf_transformation,[],[f4560_D])).
fof(f4560_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4503,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) | ($spl6 | $spl12)),
  inference(superposition,[],[f483,f1626])).
fof(f4566,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl2 | $spl12)),
  inference(forward_demodulation,[],[f4502,f174])).
fof(f4502,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl0 | $spl12)),
  inference(superposition,[],[f1626,f2302])).
fof(f4565,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl11 | $spl12)),
  inference(forward_demodulation,[],[f4501,f868])).
fof(f4501,plain,(
  s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | ($spl11 | $spl12)),
  inference(superposition,[],[f1626,f1513])).
fof(f4500,plain,(
  ( ! [X4] : (s(bool,sK5('const.option.SOME_1'(s(bool,X4)),bool)) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))) ) | ($spl11 | $spl12)),
  inference(superposition,[],[f1626,f1458])).
fof(f4499,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl12),
  inference(superposition,[],[f1626,f149])).
fof(f4498,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(superposition,[],[f1626,f149])).
fof(f4497,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))) ) | $spl12),
  inference(superposition,[],[f1626,f149])).
fof(f4496,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f1626,f149])).
fof(f4495,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | $spl12),
  inference(superposition,[],[f1626,f85])).
fof(f4494,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) | s(sK0,sK5(sK2,sK0)) = s(sK0,sK6(sK5(sK2,sK0),sK4,sK0)) | $spl12),
  inference(superposition,[],[f1626,f2269])).
fof(f482,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(cnf_transformation,[],[f482_D])).
fof(f482_D,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3950,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20)))) ) | $spl24),
  inference(cnf_transformation,[],[f3950_D])).
fof(f3950_D,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3938,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X19)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f3898])).
fof(f3898,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl0),
  inference(superposition,[],[f305,f2302])).
fof(f3949,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18)))) ) | $spl22),
  inference(cnf_transformation,[],[f3949_D])).
fof(f3949_D,plain,(
  ( ! [X18] : (s(bool,X18) != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3939,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) | s(bool,'T_0') = s(bool,X17)) ) | $spl11),
  inference(duplicate_literal_removal,[],[f3897])).
fof(f3897,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X18))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | $spl11),
  inference(superposition,[],[f305,f1513])).
fof(f2276,plain,(
  $false | ($spl0 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f173,f81,f80,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400,f263,f265,f139])).
fof(f2217,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2216,f2186])).
fof(f2216,plain,(
  s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) = s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK1,sK0)))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2186,f1630])).
fof(f2186,plain,(
  ( ! [X3] : (s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),X3),s('type.option.option'(sK0),sK1))) = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK5(sK1,sK0))))) ) | $spl15),
  inference(superposition,[],[f98,f1857])).
fof(f2099,plain,(
  ( ! [X3] : (s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK5(sK1,sK0))),X3,'type.option.option'(sK0))) | p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X3),s('type.option.option'(sK0),sK1))))) ) | $spl15),
  inference(forward_demodulation,[],[f2079,f1857])).
fof(f2079,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.option.option'(sK0),bool),X3),s('type.option.option'(sK0),sK1)))) | s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5(sK1,sK0)))) = s('type.option.option'(sK0),sK6('const.option.SOME_1'(s(sK0,sK5(sK1,sK0))),X3,'type.option.option'(sK0)))) ) | $spl15),
  inference(superposition,[],[f103,f1857])).
fof(f1981,plain,(
  ( ! [X0] : (s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,X0))) | s(sK0,X0) = s(sK0,sK5(sK1,sK0))) ) | $spl15),
  inference(superposition,[],[f101,f1857])).
fof(f1857,plain,(
  s('type.option.option'(sK0),sK1) = s('type.option.option'(sK0),'const.option.SOME_1'(s(sK0,sK5(sK1,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f1627,f95])).
fof(f1859,plain,(
  s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),'const.option.NONE_0'))) != s('type.option.option'('type.option.option'(sK0)),'const.option.SOME_1'(s('type.option.option'(sK0),sK1))) | $spl15),
  inference(unit_resulting_resolution,[],[f1627,f101])).
fof(f1627,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),sK1) | $spl15),
  inference(cnf_transformation,[],[f1627_D])).
fof(f1627_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') != s('type.option.option'(sK0),sK1) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f2275,plain,(
  $false | ($spl2 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f174,f81,f80,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400,f263,f265,f139])).
fof(f2274,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400,f263,f265,f139])).
fof(f2273,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400,f263,f265])).
fof(f2272,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400,f263])).
fof(f2271,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269,f400])).
fof(f2270,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267,f2269])).
fof(f2268,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455,f2267])).
fof(f2266,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658,f455])).
fof(f2265,plain,(
  $false | ($spl6 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f483,f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658])).
fof(f2264,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262,f658])).
fof(f2263,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260,f2262])).
fof(f2261,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807,f2260])).
fof(f2259,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845,f807])).
fof(f2258,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937,f845])).
fof(f2257,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940,f937])).
fof(f2256,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515,f940])).
fof(f2255,plain,(
  $false | ($spl11 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f958,f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515])).
fof(f958,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl11),
  inference(forward_demodulation,[],[f952,f868])).
fof(f952,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl11),
  inference(unit_resulting_resolution,[],[f925,f101])).
fof(f2254,plain,(
  $false | ($spl11 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f959,f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515])).
fof(f959,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl11),
  inference(forward_demodulation,[],[f951,f868])).
fof(f951,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl11),
  inference(unit_resulting_resolution,[],[f925,f101])).
fof(f2253,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251,f1515])).
fof(f2252,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146,f2251])).
fof(f2249,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151,f1146])).
fof(f2248,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524,f1151])).
fof(f2247,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245,f1524])).
fof(f2246,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201,f2245])).
fof(f2243,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206,f1201])).
fof(f2242,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083,f1206])).
fof(f2241,plain,(
  $false | ($spl11 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f1959,f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083])).
fof(f1959,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl11),
  inference(forward_demodulation,[],[f1943,f868])).
fof(f1943,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl11),
  inference(unit_resulting_resolution,[],[f925,f101])).
fof(f2240,plain,(
  $false | ($spl11 | $spl15 | $spl16)),
  inference(global_subsumption,[],[f1960,f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083])).
fof(f1960,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl11),
  inference(forward_demodulation,[],[f1941,f868])).
fof(f1941,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5('const.option.SOME_1'(s(bool,'F_0')),bool)))) | $spl11),
  inference(unit_resulting_resolution,[],[f925,f101])).
fof(f2239,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084,f2083])).
fof(f2238,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085,f2084])).
fof(f2237,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106,f2085])).
fof(f2236,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108,f2106])).
fof(f2235,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233,f2108])).
fof(f2234,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217,f2233])).
fof(f2232,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f81,f80,f139,f82,f85,f87,f86,f94,f95,f102,f101,f110,f136,f140,f149,f151,f165,f241,f270,f456,f462,f464,f1067,f1142,f1144,f1153,f1209,f1221,f1532,f1533,f1266,f1268,f1275,f1277,f1624,f1622,f461,f841,f1270,f409,f1276,f1271,f1224,f1223,f1152,f1147,f696,f1537,f1536,f1078,f1042,f726,f708,f105,f104,f103,f98,f1150,f1274,f846,f398,f306,f305,f269,f268,f91,f90,f84,f1621,f459,f407,f138,f79,f1630,f1627,f1859,f1857,f1981,f2099,f2186,f2192,f2217])).
fof(f2220,plain,(
  $false | ($spl3 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2219,f140])).
fof(f2219,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl15 | $spl16)),
  inference(forward_demodulation,[],[f2218,f180])).
fof(f180,plain,(
  s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f175,f85])).
fof(f2218,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | ($spl3 | $spl15 | $spl16)),
  inference(forward_demodulation,[],[f2217,f187])).
fof(f187,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | $spl3),
  inference(unit_resulting_resolution,[],[f186,f85])).
fof(f186,plain,(
  s(bool,'T_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | $spl3),
  inference(backward_demodulation,[],[f180,f139])).
fof(f1804,plain,(
  $false | ($spl3 | $spl11 | $spl14)),
  inference(subsumption_resolution,[],[f1718,f1659])).
fof(f1659,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1640,f241])).
fof(f1640,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),'const.option.NONE_0'))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1628,f334])).
fof(f334,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f329])).
fof(f329,plain,(
  ( ! [X1] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f271,f85])).
fof(f271,plain,(
  ( ! [X8] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X8))) | s(bool,X8) = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1)))) ) | $spl3),
  inference(superposition,[],[f101,f187])).
fof(f1718,plain,(
  s(bool,'T_0') != s(bool,sK5('const.option.SOME_1'(s(bool,'T_0')),bool)) | ($spl3 | $spl11 | $spl14)),
  inference(backward_demodulation,[],[f1659,f925])).
fof(f1801,plain,(
  ( ! [X22] : (s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) | $spl18),
  inference(cnf_transformation,[],[f1801_D])).
fof(f1801_D,plain,(
  ( ! [X22] : (s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1802,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))) ) | $spl20),
  inference(cnf_transformation,[],[f1802_D])).
fof(f1802_D,plain,(
  ( ! [X23] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1800,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) | s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) | ($spl3 | $spl14)),
  inference(forward_demodulation,[],[f1799,f1659])).
fof(f1799,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5(X22,bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1700,f1659])).
fof(f1700,plain,(
  ( ! [X23,X22] : (s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X23) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5(X22,bool)))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f792])).
fof(f792,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X23) | s('type.option.option'(bool),'const.option.SOME_1'(s(bool,X23))) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,sK5(X22,bool)))) | s('type.option.option'(bool),X22) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X22)) ) | $spl3),
  inference(superposition,[],[f102,f466])).
fof(f466,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool)) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0)) ) | $spl3),
  inference(forward_demodulation,[],[f454,f187])).
fof(f454,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))))) | s('type.option.option'(bool),'const.option.NONE_0') = s('type.option.option'(bool),X0) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) | $spl3),
  inference(superposition,[],[f95,f334])).
fof(f1795,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1794,f1659])).
fof(f1794,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1661,f86])).
fof(f1661,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f90])).
fof(f1793,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1660,f82])).
fof(f1660,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f84])).
fof(f1792,plain,(
  $false | ($spl3 | $spl14)),
  inference(trivial_inequality_removal,[],[f1665])).
fof(f1665,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f280])).
fof(f280,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl3),
  inference(forward_demodulation,[],[f266,f187])).
fof(f266,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))))) | $spl3),
  inference(unit_resulting_resolution,[],[f186,f101])).
fof(f1791,plain,(
  $false | ($spl3 | $spl14)),
  inference(trivial_inequality_removal,[],[f1666])).
fof(f1666,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f285])).
fof(f285,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f280,f101])).
fof(f1790,plain,(
  $false | ($spl3 | $spl14)),
  inference(trivial_inequality_removal,[],[f1667])).
fof(f1667,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1659,f289])).
fof(f289,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f285,f101])).
fof(f1658,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1657,f140])).
fof(f1657,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl14)),
  inference(forward_demodulation,[],[f1635,f94])).
fof(f1635,plain,(
  s(bool,'F_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1628,f187])).
fof(f1656,plain,(
  $false | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f1634,f94])).
fof(f1634,plain,(
  s(bool,'T_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),'const.option.NONE_0'))) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f1628,f186])).
fof(f1628,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | $spl14),
  inference(cnf_transformation,[],[f1628_D])).
fof(f1628_D,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1630,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK1,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f1630_D])).
fof(f1630_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK1,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1622,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK1,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1)),
  inference(equality_resolution,[],[f459])).
fof(f1626,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f1626_D])).
fof(f1626_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1624,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK1) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0))))),
  inference(forward_demodulation,[],[f1623,f79])).
fof(f1623,plain,(
  s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK2)),
  inference(trivial_inequality_removal,[],[f1620])).
fof(f1620,plain,(
  s('type.option.option'(sK0),sK1) != s('type.option.option'(sK0),sK1) | s(bool,happ(s(fun(sK0,bool),sK3),s(sK0,sK5(sK2,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK5(sK2,sK0)))) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK2)),
  inference(superposition,[],[f459,f79])).
fof(f802,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) | $spl8),
  inference(cnf_transformation,[],[f802_D])).
fof(f802_D,plain,(
  ( ! [X0] : (s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f804,plain,(
  s('type.option.option'(bool),'const.option.NONE_0') != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f804_D])).
fof(f804_D,plain,(
  s('type.option.option'(bool),'const.option.NONE_0') != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f797,plain,(
  ( ! [X0] : (s('type.option.option'(bool),'const.option.NONE_0') != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s('type.option.option'(bool),X0) = s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK5(X0,bool))) ) | $spl3),
  inference(equality_factoring,[],[f466])).
fof(f531,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f530,f482])).
fof(f530,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f489,f86])).
fof(f489,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f482,f90])).
fof(f529,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f488,f82])).
fof(f488,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f482,f84])).
fof(f528,plain,(
  $false | ($spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f495])).
fof(f495,plain,(
  s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) != s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f482,f280])).
fof(f527,plain,(
  $false | ($spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f496])).
fof(f496,plain,(
  s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) != s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f482,f285])).
fof(f526,plain,(
  $false | ($spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f497])).
fof(f497,plain,(
  s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) != s('type.option.option'('type.option.option'('type.option.option'(bool))),'const.option.SOME_1'(s('type.option.option'('type.option.option'(bool)),'const.option.SOME_1'(s('type.option.option'(bool),'const.option.SOME_1'(s(bool,'T_0'))))))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f482,f289])).
fof(f483,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | s(X9,X11) = s(X9,sK6(X11,X10,X9))) ) | $spl6),
  inference(cnf_transformation,[],[f483_D])).
fof(f483_D,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | s(X9,X11) = s(X9,sK6(X11,X10,X9))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f485,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))) | s(X24,X26) = s(X24,sK6(X26,X25,X24)) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(forward_demodulation,[],[f484,f187])).
fof(f484,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(X24,X26) = s(X24,sK6(X26,X25,X24)) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(subsumption_resolution,[],[f478,f87])).
fof(f478,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,X23)) | s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))) != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | s(X24,X26) = s(X24,sK6(X26,X25,X24)) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f449,f334])).
fof(f449,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X21)) | s(bool,X21) != s(bool,happ(s(fun(X18,bool),X19),s(X18,X20))) | s(X18,X20) = s(X18,sK6(X20,X19,X18))) ) | $spl3),
  inference(subsumption_resolution,[],[f432,f82])).
fof(f432,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(X18,bool),X19),s(X18,X20))) | p(s(bool,X21)) | s(X18,X20) = s(X18,sK6(X20,X19,X18))) ) | $spl3),
  inference(superposition,[],[f136,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(X0,X2) = s(X0,sK6(X2,X1,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f413,f84])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) ) | $spl3),
  inference(forward_demodulation,[],[f405,f187])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1)))) | s(X0,X2) = s(X0,sK6(X2,X1,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) ) | $spl3),
  inference(superposition,[],[f103,f334])).
fof(f481,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | s(X9,X11) = s(X9,sK6(X11,X10,X9)) | s(bool,'T_0') = s(bool,X8)) ) | $spl3),
  inference(subsumption_resolution,[],[f475,f87])).
fof(f475,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,happ(s(fun(X9,bool),X10),s(X9,X11))) | s(X9,X11) = s(X9,sK6(X11,X10,X9)) | s(bool,'T_0') = s(bool,X8)) ) | $spl3),
  inference(superposition,[],[f449,f85])).
fof(f173,plain,(
  s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) | $spl0),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f175,plain,(
  s(bool,'F_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | $spl3),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s(bool,'F_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f165,plain,(
  s(bool,'F_0') != s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK3),s('type.option.option'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.misc.OPTION_EVERY_2'(s(fun(sK0,bool),sK4),s('type.option.option'(sK0),sK1)))),
  inference(superposition,[],[f139,f85])).
% SZS output end Proof for thm.misc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 108
% Passive clauses: 598
% Generated clauses: 4983
% Final active clauses: 17
% Final passive clauses: 36
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 293
% Trivial inequalities: 76
% Fw subsumption resolutions: 247
% Fw demodulations: 605
% Bw demodulations: 394
% Condensations: 267
% Global subsumptions: 37
% 
% Simple tautologies: 97
% Equational tautologies: 70
% Forward subsumptions: 2554
% Fw demodulations to eq. taut.: 50
% Bw demodulations to eq. taut.: 25
% 
% Unit resulting resolution: 296
% Forward superposition: 1282
% Backward superposition: 1060
% Self superposition: 48
% Equality factoring: 36
% Equality resolution: 14
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 758
% SAT solver unit clauses: 19
% SAT solver binary clauses: 155
% 
% Sat splits: 32
% Sat splitting refutations: 57
% 
% Memory used [KB]: 3198
% Time elapsed: 0.183 s
% ------------------------------
% ------------------------------
% Success in time 1.55 s
