Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (136481)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 3538
% Generated clauses: 20161
% Final active clauses: 129
% Final passive clauses: 3108
% Input formulas: 68
% Initial clauses: 97
% 
% Duplicate literals: 590
% Trivial inequalities: 322
% Fw subsumption resolutions: 1523
% Fw demodulations: 3434
% Bw demodulations: 14
% 
% Simple tautologies: 266
% Equational tautologies: 62
% Forward subsumptions: 10047
% 
% Binary resolution: 168
% Factoring: 1
% Forward superposition: 5732
% Backward superposition: 7867
% Self superposition: 211
% Equality factoring: 114
% Equality resolution: 34
% 
% Split clauses: 150
% Split components: 302
% SAT solver clauses: 425
% SAT solver unit clauses: 7
% SAT solver binary clauses: 72
% 
% Sat splits: 150
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 423
% TWLsolver calls for satisfiability: 118
% 
% Memory used [KB]: 21364
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (136518)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 1489
% Generated clauses: 15920
% Final active clauses: 116
% Final passive clauses: 884
% Input formulas: 68
% Initial clauses: 97
% 
% Duplicate literals: 88
% Trivial inequalities: 216
% Fw subsumption resolutions: 553
% Fw demodulations: 2829
% Bw demodulations: 429
% Condensations: 120
% Global subsumptions: 8
% 
% Simple tautologies: 47
% Equational tautologies: 16
% Forward subsumptions: 8499
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 36
% 
% Binary resolution: 168
% Factoring: 1
% Forward superposition: 4142
% Backward superposition: 5282
% Self superposition: 153
% Equality factoring: 14
% Equality resolution: 20
% 
% Split clauses: 38
% Split components: 76
% SAT solver clauses: 1999
% SAT solver unit clauses: 148
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 17
% SAT solver learnt literals: 4
% 
% Sat splits: 38
% Sat splitting refutations: 44
% 
% TWLsolver clauses: 1966
% TWLsolver calls for satisfiability: 350
% 
% Memory used [KB]: 25330
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f7850,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f208,f7820])).
fof(f7820,plain,(
  ( ! [X350] : (~p(s(bool,X350))) )),
  inference(forward_demodulation,[],[f7819,f271])).
fof(f271,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X5) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X1),s(X0,X2),s(fun('type.num.num',X0),X3),s(fun(bool,X0),X4),s(X0,X5)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X15] : ! [X16,X31,X32,X33,X34] : s(X15,X34) = s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X15),X16),s(X15,X31),s(fun('type.num.num',X15),X32),s(fun(bool,X15),X33),s(X15,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|8','thm.bvp_to_wordProps.tag_case_def|split|4')).
fof(f7819,plain,(
  ( ! [X350,X346,X349,X347,X348] : (~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',bool),X346),s(bool,X347),s(fun('type.num.num',bool),X348),s(fun(bool,bool),X349),s(bool,X350))))) )),
  inference(subsumption_resolution,[],[f7713,f7670])).
fof(f7670,plain,(
  ( ! [X6,X11,X9] : (s(X6,X11) = s(X6,happ(s(fun('type.num.num',X6),X9),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f3774,f271])).
fof(f3774,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',sK0))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X1),s(fun(bool,X0),X4),s(X0,X5)))) )),
  inference(superposition,[],[f273,f207])).
fof(f207,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f169])).
fof(f169,plain,(
  ? [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|8','thm.bvp_to_wordProps.tag_distinct|split|8')).
fof(f273,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X15] : ! [X6,X16,X31,X32,X33,X34] : s(X15,happ(s(fun('type.num.num',X15),X32),s('type.num.num',X6))) = s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X6))),s(fun('type.num.num',X15),X16),s(X15,X31),s(fun('type.num.num',X15),X32),s(fun(bool,X15),X33),s(X15,X34)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|8','thm.bvp_to_wordProps.tag_case_def|split|2')).
fof(f7713,plain,(
  ( ! [X350,X346,X349,X347,X348] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),X348),s('type.num.num',sK0))) | ~p(s(bool,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'),s(fun('type.num.num',bool),X346),s(bool,X347),s(fun('type.num.num',bool),X348),s(fun(bool,bool),X349),s(bool,X350))))) )),
  inference(superposition,[],[f233,f3774])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X29] : (s(bool,'F_0') = s(bool,X29) <=> ~p(s(bool,X29)))),
  file('thm.bvp_to_wordProps.tag_distinct|split|8','thm.bool.EQ_CLAUSES|split|3')).
fof(f208,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvp_to_wordProps.tag_distinct|split|8',aHL_TRUTH)).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 123
% Passive clauses: 755
% Generated clauses: 7661
% Final active clauses: 106
% Final passive clauses: 603
% Input formulas: 68
% Initial clauses: 97
% 
% Duplicate literals: 9
% Trivial inequalities: 15
% Fw subsumption resolutions: 85
% Fw demodulations: 1322
% Bw demodulations: 43
% Condensations: 10
% 
% Simple tautologies: 27
% Equational tautologies: 59
% Forward subsumptions: 5182
% Fw demodulations to eq. taut.: 142
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 769
% Factoring: 1
% Forward superposition: 2264
% Backward superposition: 2973
% Self superposition: 101
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 785
% SAT solver unit clauses: 268
% SAT solver binary clauses: 417
% 
% Sat splits: 14
% 
% Memory used [KB]: 15479
% Time elapsed: 0.568 s
% ------------------------------
% ------------------------------
% Success in time 1.901 s
