Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exh_to_patProof
% (112034)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1198
% Passive clauses: 8590
% Generated clauses: 25523
% Final active clauses: 1177
% Final passive clauses: 7016
% Input formulas: 133
% Initial clauses: 447
% 
% Duplicate literals: 258
% Trivial inequalities: 57
% Fw subsumption resolutions: 1831
% Fw demodulations: 1842
% Bw demodulations: 11
% 
% Simple tautologies: 289
% Equational tautologies: 54
% Forward subsumptions: 12197
% 
% Binary resolution: 8665
% Factoring: 1193
% Forward superposition: 4084
% Backward superposition: 7044
% Self superposition: 29
% Equality factoring: 2
% Equality resolution: 39
% 
% Split clauses: 335
% Split components: 680
% SAT solver clauses: 612
% SAT solver unit clauses: 10
% SAT solver binary clauses: 99
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 335
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 611
% TWLsolver calls for satisfiability: 543
% 
% Memory used [KB]: 15991
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exh_to_patProof
% (112048)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 428
% Passive clauses: 2341
% Generated clauses: 10873
% Final active clauses: 309
% Final passive clauses: 1202
% Input formulas: 133
% Initial clauses: 447
% 
% Duplicate literals: 76
% Trivial inequalities: 21
% Fw subsumption resolutions: 600
% Fw demodulations: 740
% Bw demodulations: 704
% Condensations: 67
% Global subsumptions: 22
% 
% Simple tautologies: 254
% Equational tautologies: 8
% Forward subsumptions: 3867
% Bw demodulations to eq. taut.: 12
% 
% Binary resolution: 1303
% Factoring: 8
% Forward superposition: 1816
% Backward superposition: 1947
% Self superposition: 40
% Equality factoring: 2
% Equality resolution: 4
% 
% Split clauses: 92
% Split components: 184
% SAT solver clauses: 2946
% SAT solver unit clauses: 481
% SAT solver binary clauses: 210
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 13
% 
% Sat splits: 92
% Sat splitting refutations: 79
% 
% TWLsolver clauses: 2857
% TWLsolver calls for satisfiability: 6466
% 
% Memory used [KB]: 12537
% Time elapsed: 0.708 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exh_to_patProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exh_to_patProof
% SZS output start Proof for thm.exh_to_patProof
fof(f11100,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f647,f1121,f1119,f646,f1120,f1118,f2254,f4704,f4714,f4713,f5326,f5329,f5331,f5333,f5335,f5339,f5341,f5345,f5347,f5359,f5361,f5373,f5375,f5385,f5395,f5411,f5437,f5439,f5448,f5453,f5463,f5472,f5483,f5486,f5491,f5493,f5508,f5513,f5566,f5568,f5577,f5582,f5592,f5601,f5612,f5615,f5620,f5622,f5637,f5642,f5682,f5684,f5686,f5696,f5698,f5700,f5716,f5718,f5726,f5728,f5744,f5746,f5754,f5756,f5772,f5774,f5782,f5784,f5797,f5799,f6455,f6456,f6457,f6458,f5818,f5820,f5830,f5832,f5848,f5856,f5866,f5868,f5870,f5880,f5882,f5892,f5894,f5904,f5906,f5916,f5918,f5928,f5943,f5958,f6001,f6003,f6013,f6015,f6017,f6019,f6037,f6039,f6041,f6049,f6051,f6053,f6055,f6057,f6069,f6078,f6088,f6097,f6099,f6113,f6115,f6122,f6124,f6138,f6140,f6147,f6149,f6162,f6164,f6171,f6173,f6186,f6201,f6227,f6229,f6241,f6243,f6255,f6257,f6269,f6289,f6291,f6301,f6303,f6305,f6317,f6319,f6326,f6332,f6338,f6350,f6353,f6355,f6357,f6359,f6361,f6363,f6365,f6367,f6369,f6388,f6390,f6392,f6394,f6397,f6399,f6410,f6412,f6422,f6424,f6429,f6431,f6434,f6436,f6438,f6471,f6472,f6473,f6474,f6475,f6476,f6477,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6485,f6486,f6487,f6488,f6489,f6490,f6491,f6492,f6493,f6494,f6495,f6496,f6497,f6498,f6499,f6500,f6501,f6502,f6503,f6504,f6505,f6506,f6507,f6508,f6509,f6510,f6511,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6524,f6525,f6526,f6527,f6528,f6529,f6530,f6531,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6545,f6546,f6547,f6548,f6549,f6550,f6551,f6552,f6553,f6554,f6555,f6556,f6557,f6558,f6559,f6560,f6561,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6574,f6575,f6576,f6577,f6578,f6579,f6580,f6581,f6582,f6583,f6584,f6585,f6586,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6594,f6595,f6596,f6597,f6598,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6607,f6608,f6609,f6610,f6611,f6612,f6613,f6614,f6615,f6616,f6617,f6618,f6619,f6620,f6623,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6633,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6720,f6721,f6722,f7073,f7075,f7866,f7868,f8430,f8431,f8482,f8483,f8532,f8533,f8593,f8594,f5327,f11099])).
fof(f11099,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f11098,f1121])).
fof(f11098,plain,(
  s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) | $spl0),
  inference(forward_demodulation,[],[f11055,f10966])).
fof(f10966,plain,(
  s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(sK27,'const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2363,f2255,f3735,f3705,f7515,f912,f861])).
fof(f861,plain,(
  ( ! [X0,X1] : (~sP14(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK142(X1,X0)))) | sP12(X0,X1) | sP11(X0,X1) | sP10(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(X1,X0)))) | sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(X1,X0))))) | sP10(X0,X1) | sP11(X0,X1) | sP12(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK142(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK142(X1,X0))))) | sP13(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP10(X0,X1) & ~sP11(X0,X1) & ~sP12(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP13(X0,X1)) | sP14(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK141,sK142])],[f550])).
fof(f550,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | ? [X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5)))) | sP10(X0,X1) | sP11(X0,X1) | sP12(X0,X1) | ? [X4] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP13(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP10(X0,X1) & ~sP11(X0,X1) & ~sP12(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP13(X0,X1)) | sP14(X0,X1)))),
  inference(rectify,[],[f549])).
fof(f549,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP10(X1,X0) | sP11(X1,X0) | sP12(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP13(X1,X0)) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP10(X1,X0) & ~sP11(X1,X0) & ~sP12(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP13(X1,X0)) | sP14(X1,X0)))),
  inference(flattening,[],[f548])).
fof(f548,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP10(X1,X0) | sP11(X1,X0) | sP12(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP13(X1,X0))) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP10(X1,X0) & ~sP11(X1,X0) & ~sP12(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP13(X1,X0)) | sP14(X1,X0)))),
  inference(nnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X1,X0] : (sP14(X1,X0) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP10(X1,X0) | sP11(X1,X0) | sP12(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP13(X1,X0)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f912,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X99,X55] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X55))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X99)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.patSem.v_distinct|split|3')).
fof(f7515,plain,(
  ( ! [X0,X1] : (~sP12('const.patSem.Litv_1'(s('type.ast.lit',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1078,f874])).
fof(f874,plain,(
  ( ! [X0,X1] : (~sP12(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK146(X1,X0)),s('type.list.list'('type.patLang.exp'),sK148(X1,X0)),s('type.num.num',sK149(X1,X0))))) )),
  inference(cnf_transformation,[],[f557])).
fof(f557,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK145(X1,X0)),s('type.list.list'('type.patLang.exp'),sK147(X1,X0)),s('type.num.num',sK149(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK146(X1,X0)),s('type.list.list'('type.patLang.exp'),sK148(X1,X0)),s('type.num.num',sK149(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK147(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK145(X1,X0)))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK148(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK146(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK147(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK145(X1,X0)),s('type.list.list'('type.patSem.v'),sK146(X1,X0)))))))),s('type.list.list'('type.patLang.exp'),sK147(X1,X0)),s('type.list.list'('type.patLang.exp'),sK148(X1,X0))))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP12(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK145,sK146,sK147,sK148,sK149])],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ? [X7,X8,X9,X10,X11] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patLang.exp'),X9),s('type.num.num',X11))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X11))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patSem.v'),X8))))))),s('type.list.list'('type.patLang.exp'),X9),s('type.list.list'('type.patLang.exp'),X10)))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP12(X0,X1)))),
  inference(rectify,[],[f555])).
fof(f555,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8)))))) & (! [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | sP12(X1,X0)))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X1,X0] : (sP12(X1,X0) <=> ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1078,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2,X3] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2,X3] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X97,X29,X28,X55] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X55))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X28),s('type.list.list'('type.patLang.exp'),X29),s('type.num.num',X97)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.patSem.v_distinct|split|2')).
fof(f3705,plain,(
  ( ! [X0,X1] : (~sP11('const.patSem.Litv_1'(s('type.ast.lit',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1037,f878])).
fof(f878,plain,(
  ( ! [X0,X1] : (~sP11(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK151(X1,X0)),s('type.patLang.exp',sK153(X1,X0))))) )),
  inference(cnf_transformation,[],[f560])).
fof(f560,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK150(X1,X0)),s('type.patLang.exp',sK152(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK151(X1,X0)),s('type.patLang.exp',sK153(X1,X0)))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK150(X1,X0)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK151(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK150(X1,X0)),s('type.list.list'('type.patSem.v'),sK151(X1,X0)))))))),s('type.patLang.exp',sK152(X1,X0)))),s('type.patLang.exp',sK153(X1,X0))))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP11(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK150,sK151,sK152,sK153])],[f559])).
fof(f559,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ? [X6,X7,X8,X9] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X6),s('type.patLang.exp',X8))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X7),s('type.patLang.exp',X9))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7))))))),s('type.patLang.exp',X8))),s('type.patLang.exp',X9)))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP11(X0,X1)))),
  inference(rectify,[],[f558])).
fof(f558,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13)))))) & (! [X10,X11,X12,X13] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1037,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X29,X28,X55] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X55))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X28),s('type.patLang.exp',X29)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.patSem.v_distinct|split|1')).
fof(f3735,plain,(
  ( ! [X0,X1] : (~sP10('const.patSem.Litv_1'(s('type.ast.lit',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1038,f882])).
fof(f882,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK154(X1,X0)),s('type.list.list'('type.patSem.v'),sK156(X1,X0))))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK154(X1,X0)),s('type.list.list'('type.patSem.v'),sK155(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK154(X1,X0)),s('type.list.list'('type.patSem.v'),sK156(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK155(X1,X0)),s('type.list.list'('type.patSem.v'),sK156(X1,X0))))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP10(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK154,sK155,sK156])],[f562])).
fof(f562,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ? [X5,X6,X7] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X6))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7)))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP10(X0,X1)))),
  inference(rectify,[],[f561])).
fof(f561,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16)))))) & (! [X14,X15,X16] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f1038,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X2))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X29,X28,X55] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X55))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X28),s('type.list.list'('type.patSem.v'),X29)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.patSem.v_distinct|split|0')).
fof(f2255,plain,(
  sP14('const.patSem.Litv_1'(s('type.ast.lit',sK28)),sK27) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f1124])).
fof(f1124,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP14(X1,X0)) )),
  inference(forward_demodulation,[],[f885,f816])).
fof(f816,plain,(
  ( ! [X0,X1] : (s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X9] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X8),s('type.patSem.v',X9))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X8))),s('type.patSem.v',X9)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','ARITY_0_2_const.exh_to_patProof.v_rel')).
fof(f885,plain,(
  ( ! [X0,X1] : (sP14(X1,X0) | ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | sP14(X1,X0)) & (~sP14(X1,X0) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> sP14(X1,X0))),
  inference(definition_folding,[],[f237,f418,f417,f416,f415,f414])).
fof(f417,plain,(
  ! [X1,X0] : (sP13(X1,X0) <=> ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f237,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X28,X29] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X28))),s('type.patSem.v',X29)))) <=> (? [X7] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X17,X18,X19] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X20,X21,X22,X23] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X20),s('type.patLang.exp',X22))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X20))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patSem.v'),X21))))))),s('type.patLang.exp',X22))),s('type.patLang.exp',X23))))) | ? [X20,X21,X24,X25,X26] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patLang.exp'),X24),s('type.num.num',X26))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X26))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X24))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X20))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X24))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patSem.v'),X21))))))),s('type.list.list'('type.patLang.exp'),X24),s('type.list.list'('type.patLang.exp'),X25))))) | ? [X26] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X26))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X26)))) | ? [X18,X19] : (s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X29) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19)))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.exh_to_patProof.v_rel_cases')).
fof(f2363,plain,(
  ( ! [X0,X1] : (~sP13('const.patSem.Litv_1'(s('type.ast.lit',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f914,f870])).
fof(f870,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK144(X1,X0))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK143(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK144(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK143(X1,X0)),s('type.list.list'('type.patSem.v'),sK144(X1,X0))))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP13(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK143,sK144])],[f553])).
fof(f553,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | ? [X4,X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X4),s('type.list.list'('type.patSem.v'),X5)))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP13(X0,X1)))),
  inference(rectify,[],[f552])).
fof(f552,plain,(
  ! [X1,X0] : ((~sP13(X1,X0) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))) & (! [X2,X3] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP13(X1,X0)))),
  inference(nnf_transformation,[],[f417])).
fof(f914,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X99,X55] : ~s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X55))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X99)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.patSem.v_distinct|split|4')).
fof(f11055,plain,(
  s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(sK27,'const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2363,f2255,f3735,f3705,f7515,f912,f862])).
fof(f862,plain,(
  ( ! [X0,X1] : (~sP14(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK142(X1,X0)))) | sP12(X0,X1) | sP11(X0,X1) | sP10(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK141(X1,X0)))) | sP13(X0,X1)) )),
  inference(cnf_transformation,[],[f551])).
fof(f5327,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f5327_D])).
fof(f5327_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f8594,plain,(
  ( ! [X5] : (p(s(bool,sK105(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f8594_D])).
fof(f8594_D,plain,(
  ( ! [X5] : (p(s(bool,sK105(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f8593,plain,(
  ( ! [X4,X5] : (p(s(bool,sK105(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f8581,f755])).
fof(f755,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X118] : ! [X120] : s(X118,X120) = s(X118,happ(s(fun(X118,X118),'_i_0_0'),s(X118,X120)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','_i0_ax')).
fof(f8581,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK105(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f780,f755])).
fof(f780,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK105(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK104(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK105(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK104,sK105])],[f505])).
fof(f505,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f504])).
fof(f504,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f503])).
fof(f503,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X33] : ! [X35,X34] : ((? [X36] : p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36)))) | p(s(bool,X34))) <=> ? [X36] : (p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36)))) | p(s(bool,X34))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f8533,plain,(
  ( ! [X5] : (p(s(bool,sK104(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl20),
  inference(cnf_transformation,[],[f8533_D])).
fof(f8533_D,plain,(
  ( ! [X5] : (p(s(bool,sK104(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f8532,plain,(
  ( ! [X4,X5] : (p(s(bool,sK104(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f8520,f755])).
fof(f8520,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK104(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f778,f755])).
fof(f778,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK104(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f506])).
fof(f8483,plain,(
  ( ! [X5] : (p(s(bool,sK103('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f8483_D])).
fof(f8483_D,plain,(
  ( ! [X5] : (p(s(bool,sK103('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f8482,plain,(
  ( ! [X4,X5] : (p(s(bool,sK103('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f8470,f755])).
fof(f8470,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK103('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f777,f755])).
fof(f777,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK103(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK102(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK103(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK102,sK103])],[f501])).
fof(f501,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f500])).
fof(f500,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f499])).
fof(f499,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X33] : ! [X35,X34] : ((p(s(bool,X35)) | ? [X36] : p(s(bool,happ(s(fun(X33,bool),X34),s(X33,X36))))) <=> ? [X36] : (p(s(bool,X35)) | p(s(bool,happ(s(fun(X33,bool),X34),s(X33,X36))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f8431,plain,(
  ( ! [X5] : (p(s(bool,sK102('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl16),
  inference(cnf_transformation,[],[f8431_D])).
fof(f8431_D,plain,(
  ( ! [X5] : (p(s(bool,sK102('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8430,plain,(
  ( ! [X4,X5] : (p(s(bool,sK102('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f8418,f755])).
fof(f8418,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK102('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f775,f755])).
fof(f775,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK102(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f502])).
fof(f7868,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl14),
  inference(cnf_transformation,[],[f7868_D])).
fof(f7868_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f7866,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f7855,f755])).
fof(f7855,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f767,f755])).
fof(f767,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X33] : ! [X35,X36] : (p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36)))) => p(s(bool,happ(s(fun(X33,bool),X35),s(X33,'const.min.@_1'(s(fun(X33,bool),X35)))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.SELECT_AX')).
fof(f7075,plain,(
  p(s(bool,sK98('_i_0_0',bool))) | $spl12),
  inference(cnf_transformation,[],[f7075_D])).
fof(f7075_D,plain,(
  p(s(bool,sK98('_i_0_0',bool))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f7073,plain,(
  ( ! [X3] : (p(s(bool,sK98('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f7062,f755])).
fof(f7062,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK98('_i_0_0',bool)))))) )),
  inference(superposition,[],[f761,f755])).
fof(f761,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK98(X1,X0)))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK97(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK98(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK97,sK98])],[f489])).
fof(f489,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f488])).
fof(f488,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X33] : ! [X35] : (~? [X36] : p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36)))) <=> ! [X36] : ~p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.NOT_EXISTS_THM')).
fof(f6722,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5340,f1122])).
fof(f1122,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f651,f648])).
fof(f648,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_lit|split|1',aHL_TRUTH)).
fof(f651,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f85])).
fof(f85,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f5340,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1169,f5327,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X31,X32] : ((p(s(bool,X31)) => p(s(bool,X32))) => ((p(s(bool,X32)) => p(s(bool,X31))) => s(bool,X31) = s(bool,X32)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1169,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1122,f700])).
fof(f700,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X30] : (s(bool,'T_0') = s(bool,X30) <=> p(s(bool,X30)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f6721,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5346,f1122])).
fof(f5346,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1169,f5327,f817])).
fof(f6720,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5348,f1122])).
fof(f5348,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1207,f5327,f902])).
fof(f902,plain,(
  ( ! [X0,X1] : (~sP16(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1] : ((~sP16(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP16(X0,X1)))),
  inference(rectify,[],[f572])).
fof(f572,plain,(
  ! [X1,X0] : ((~sP16(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP16(X1,X0)))),
  inference(flattening,[],[f571])).
fof(f571,plain,(
  ! [X1,X0] : ((~sP16(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP16(X1,X0)))),
  inference(nnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X1,X0] : (sP16(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f1207,plain,(
  sP16('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1122,f648,f904])).
fof(f904,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP16(X0,X1)) )),
  inference(cnf_transformation,[],[f573])).
fof(f6719,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5349,f5327])).
fof(f5349,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4267,f5327,f902])).
fof(f4267,plain,(
  sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | $spl3),
  inference(unit_resulting_resolution,[],[f648,f4218,f904])).
fof(f4218,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | $spl3),
  inference(unit_resulting_resolution,[],[f1121,f855])).
fof(f855,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X27,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))),s('type.patSem.v',X27)))) <=> s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.exh_to_patProof.v_rel_lit|split|0')).
fof(f6718,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5351,f5327])).
fof(f5351,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2270,f5327,f902])).
fof(f2270,plain,(
  sP16('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f907])).
fof(f907,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP16(X0,X1)) )),
  inference(cnf_transformation,[],[f573])).
fof(f6717,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5352,f5327])).
fof(f5352,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f5327,f902])).
fof(f1239,plain,(
  ( ! [X0] : (sP16('F_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1122,f1219])).
fof(f1219,plain,(
  ( ! [X0,X1] : (sP16(X1,'const.bool.BOUNDED_1'(s(bool,X0))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1217,f648])).
fof(f1217,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP16(X1,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(superposition,[],[f907,f655])).
fof(f655,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X60] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X60)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.BOUNDED_THM')).
fof(f6716,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5353,f5327])).
fof(f5353,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1431,f5327,f902])).
fof(f1431,plain,(
  ( ! [X0] : (sP16('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f1122,f656,f907])).
fof(f656,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X33] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X33,X33),'_i_0_0'),s(fun(X33,X33),'_i_0_0'))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.ind_type.ISO_REFL')).
fof(f6715,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5354,f1122])).
fof(f5354,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2269,f5327,f902])).
fof(f2269,plain,(
  sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f904])).
fof(f6714,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5356,f1122])).
fof(f5356,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1222,f5327,f902])).
fof(f1222,plain,(
  ( ! [X0] : (sP16('const.bool.BOUNDED_1'(s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1122,f1210])).
fof(f1210,plain,(
  ( ! [X0,X1] : (sP16('const.bool.BOUNDED_1'(s(bool,X0)),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1208,f648])).
fof(f1208,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP16('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f904,f655])).
fof(f6713,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5357,f1122])).
fof(f5357,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1430,f5327,f902])).
fof(f1430,plain,(
  ( ! [X0] : (sP16('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1122,f656,f904])).
fof(f6712,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5364,f1122])).
fof(f5364,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1216,f5327,f902])).
fof(f1216,plain,(
  sP16('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1122,f648,f907])).
fof(f6711,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5365,f1122])).
fof(f5365,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2270,f5327,f902])).
fof(f6710,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5366,f1122])).
fof(f5366,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1239,f5327,f902])).
fof(f6709,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5367,f1122])).
fof(f5367,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1431,f5327,f902])).
fof(f6708,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5368,f5327])).
fof(f5368,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2269,f5327,f902])).
fof(f6707,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5369,f5327])).
fof(f5369,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4268,f5327,f902])).
fof(f4268,plain,(
  sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f648,f4218,f907])).
fof(f6706,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5370,f5327])).
fof(f5370,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1222,f5327,f902])).
fof(f6705,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5371,f5327])).
fof(f5371,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1430,f5327,f902])).
fof(f6704,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5410,f5562])).
fof(f5562,plain,(
  ( ! [X0] : (sP20('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1137,f1142,f5327,f967])).
fof(f967,plain,(
  ( ! [X2,X0,X1] : (~sP18(X1,X2,X0) | p(s(bool,X0)) | ~sP19(X0,X2,X1) | sP20(X0,X1,X2) | ~sP17(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f602])).
fof(f602,plain,(
  ! [X0,X1,X2] : ((~sP20(X0,X1,X2) | (sP17(X2,X1,X0) & sP18(X1,X2,X0) & sP19(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP17(X2,X1,X0) | ~sP18(X1,X2,X0) | ~sP19(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP20(X0,X1,X2)))),
  inference(rectify,[],[f601])).
fof(f601,plain,(
  ! [X2,X1,X0] : ((~sP20(X2,X1,X0) | (sP17(X0,X1,X2) & sP18(X1,X0,X2) & sP19(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP17(X0,X1,X2) | ~sP18(X1,X0,X2) | ~sP19(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP20(X2,X1,X0)))),
  inference(flattening,[],[f600])).
fof(f600,plain,(
  ! [X2,X1,X0] : ((~sP20(X2,X1,X0) | (sP17(X0,X1,X2) & sP18(X1,X0,X2) & sP19(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP17(X0,X1,X2) | ~sP18(X1,X0,X2) | ~sP19(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP20(X2,X1,X0)))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X2,X1,X0] : (sP20(X2,X1,X0) <=> (sP17(X0,X1,X2) & sP18(X1,X0,X2) & sP19(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP20])])).
fof(f1142,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f973])).
fof(f973,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f608])).
fof(f608,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(rectify,[],[f607])).
fof(f607,plain,(
  ! [X1,X0,X2] : ((~sP18(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP18(X1,X0,X2)))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X1,X0,X2] : ((~sP18(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP18(X1,X0,X2)))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X1,X0,X2] : (sP18(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f1137,plain,(
  ( ! [X0,X1] : (sP19(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f969])).
fof(f969,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((~sP19(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP19(X0,X1,X2)))),
  inference(rectify,[],[f604])).
fof(f604,plain,(
  ! [X2,X0,X1] : ((~sP19(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP19(X2,X0,X1)))),
  inference(flattening,[],[f603])).
fof(f603,plain,(
  ! [X2,X0,X1] : ((~sP19(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP19(X2,X0,X1)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X2,X0,X1] : (sP19(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP19])])).
fof(f1147,plain,(
  ( ! [X0,X1] : (sP17(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f977])).
fof(f977,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP17(X0,X1,X2)))),
  inference(flattening,[],[f609])).
fof(f609,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP17(X0,X1,X2)))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : (sP17(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f5410,plain,(
  ( ! [X0] : (~sP20('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f648,f5327,f964])).
fof(f964,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f602])).
fof(f6703,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5412,f5699])).
fof(f5699,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f975])).
fof(f975,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f608])).
fof(f5412,plain,(
  ~sP18('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f5327,f966])).
fof(f966,plain,(
  ( ! [X2,X0,X1] : (~sP18(X1,X2,X0) | p(s(bool,X1)) | ~sP19(X0,X2,X1) | sP20(X0,X1,X2) | ~sP17(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f602])).
fof(f1979,plain,(
  ~sP20('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f648,f1169,f982])).
fof(f982,plain,(
  ( ! [X2,X0,X1] : (~sP20(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f611])).
fof(f611,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP20(X2,X1,X0)) & (~sP20(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP20(X2,X1,X0))),
  inference(definition_folding,[],[f293,f427,f426,f425,f424])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X115,X113,X114] : ((p(s(bool,X114)) <=> s(bool,X113) = s(bool,X115)) <=> ((p(s(bool,X114)) | p(s(bool,X113)) | p(s(bool,X115))) & (p(s(bool,X114)) | ~p(s(bool,X115)) | ~p(s(bool,X113))) & (p(s(bool,X113)) | ~p(s(bool,X115)) | ~p(s(bool,X114))) & (p(s(bool,X115)) | ~p(s(bool,X113)) | ~p(s(bool,X114)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f1148,plain,(
  ( ! [X0,X1] : (sP17(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f648,f978])).
fof(f978,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f610])).
fof(f6702,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5413,f5699])).
fof(f5413,plain,(
  ~sP18('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f5327,f966])).
fof(f4269,plain,(
  ~sP20('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | $spl3),
  inference(unit_resulting_resolution,[],[f648,f648,f4218,f964])).
fof(f6701,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5417,f5699])).
fof(f5417,plain,(
  ~sP18('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f5327,f966])).
fof(f2284,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1169,f1118,f982])).
fof(f6700,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5419,f5699])).
fof(f5419,plain,(
  ~sP18('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f5327,f966])).
fof(f1138,plain,(
  ( ! [X0,X1] : (sP19(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f970])).
fof(f970,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f605])).
fof(f6699,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5420,f5699])).
fof(f5420,plain,(
  ~sP18('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f5327,f966])).
fof(f6698,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5421,f5699])).
fof(f5421,plain,(
  ~sP18('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f5327,f966])).
fof(f4281,plain,(
  ( ! [X0,X1] : (sP19(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4218,f970])).
fof(f6697,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5428,f5685])).
fof(f5685,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f971])).
fof(f971,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f605])).
fof(f5428,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1142,f5327,f966])).
fof(f6696,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5429,f5685])).
fof(f5429,plain,(
  ~sP19('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1142,f5327,f966])).
fof(f6695,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5454,f5685])).
fof(f5454,plain,(
  ~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f2272,f5327,f966])).
fof(f2272,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f973])).
fof(f6694,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5476,f5685])).
fof(f5476,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1143,f5327,f966])).
fof(f1143,plain,(
  ( ! [X0,X1] : (sP18(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f974])).
fof(f974,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f608])).
fof(f6693,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5477,f5685])).
fof(f5477,plain,(
  ~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1143,f5327,f966])).
fof(f6692,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5502,f5685])).
fof(f5502,plain,(
  ~sP19('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4286,f5327,f966])).
fof(f4286,plain,(
  ( ! [X0,X1] : (sP18(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4218,f974])).
fof(f6691,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5541,f5699])).
fof(f5541,plain,(
  ~sP18('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f5327,f967])).
fof(f6690,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5542,f5699])).
fof(f5542,plain,(
  ~sP18('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f5327,f967])).
fof(f6689,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5546,f5699])).
fof(f5546,plain,(
  ~sP18('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f5327,f967])).
fof(f6688,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5548,f5699])).
fof(f5548,plain,(
  ~sP18('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f5327,f967])).
fof(f6687,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5549,f5699])).
fof(f5549,plain,(
  ~sP18('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f5327,f967])).
fof(f6686,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5550,f5699])).
fof(f5550,plain,(
  ~sP18('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f5327,f967])).
fof(f6685,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5557,f5685])).
fof(f5557,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1142,f5327,f967])).
fof(f6684,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5558,f5685])).
fof(f5558,plain,(
  ~sP19('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1142,f5327,f967])).
fof(f6683,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5583,f5685])).
fof(f5583,plain,(
  ~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f2272,f5327,f967])).
fof(f6682,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5590,f5831])).
fof(f5831,plain,(
  ( ! [X0] : (~sP20(X0,'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2261,f5327,f983])).
fof(f983,plain,(
  ( ! [X2,X0,X1] : (~sP20(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f611])).
fof(f2261,plain,(
  s(bool,'T_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f820])).
fof(f820,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f371])).
fof(f5590,plain,(
  sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f1137,f2272,f5327,f967])).
fof(f2275,plain,(
  ( ! [X0,X1] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f979])).
fof(f979,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f610])).
fof(f6681,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5605,f5685])).
fof(f5605,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1143,f5327,f967])).
fof(f6680,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5606,f5685])).
fof(f5606,plain,(
  ~sP19('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1143,f5327,f967])).
fof(f6679,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5631,f5685])).
fof(f5631,plain,(
  ~sP19('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0') | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4286,f5327,f967])).
fof(f6678,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5639,f4269])).
fof(f5639,plain,(
  sP20('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1147,f1137,f4286,f5327,f967])).
fof(f6677,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5681,f5685])).
fof(f5681,plain,(
  ( ! [X0] : (~sP19('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f648,f5327,f968])).
fof(f968,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f605])).
fof(f6676,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5695,f5699])).
fof(f5695,plain,(
  ( ! [X0] : (~sP18('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f648,f5327,f972])).
fof(f972,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f608])).
fof(f6675,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5702,f5327])).
fof(f5702,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f1122,f5327,f976])).
fof(f976,plain,(
  ( ! [X2,X0,X1] : (~sP17(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f610])).
fof(f2274,plain,(
  ( ! [X0,X1] : (sP17(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f978])).
fof(f6674,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5703,f5327])).
fof(f5703,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1122,f5327,f976])).
fof(f1149,plain,(
  ( ! [X0,X1] : (sP17('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f648,f979])).
fof(f6673,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5704,f5327])).
fof(f5704,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f1122,f5327,f976])).
fof(f6672,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5705,f5327])).
fof(f5705,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f976])).
fof(f6671,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5706,f5327])).
fof(f5706,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f2273,plain,(
  ( ! [X0,X1] : (sP17(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f977])).
fof(f6670,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5708,f5327])).
fof(f5708,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6669,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5709,f5327])).
fof(f5709,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f5327,f5327,f976])).
fof(f6668,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5710,f5327])).
fof(f5710,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6667,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5712,f5327])).
fof(f5712,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f1122,f5327,f976])).
fof(f6666,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5713,f5327])).
fof(f5713,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1122,f5327,f976])).
fof(f6665,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5714,f5327])).
fof(f5714,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f1122,f5327,f976])).
fof(f6664,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5720,f5327])).
fof(f5720,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f6663,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5721,f5327])).
fof(f5721,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f5327,f5327,f976])).
fof(f6662,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5722,f5327])).
fof(f5722,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6661,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5723,f5327])).
fof(f5723,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f5327,f5327,f976])).
fof(f6660,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5724,f5327])).
fof(f5724,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6659,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5730,f5327])).
fof(f5730,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f1122,f5327,f976])).
fof(f6658,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5731,f5327])).
fof(f5731,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1122,f5327,f976])).
fof(f6657,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5732,f5327])).
fof(f5732,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f1122,f5327,f976])).
fof(f6656,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5734,f5327])).
fof(f5734,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f6655,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5735,f5327])).
fof(f5735,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f5327,f5327,f976])).
fof(f6654,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5736,f5327])).
fof(f5736,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6653,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5737,f5327])).
fof(f5737,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f5327,f5327,f976])).
fof(f6652,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5738,f5327])).
fof(f5738,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6651,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5739,f5327])).
fof(f5739,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1122,f5327,f976])).
fof(f6650,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5740,f5327])).
fof(f5740,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f1122,f5327,f976])).
fof(f6649,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5742,f5327])).
fof(f5742,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f1122,f5327,f976])).
fof(f6648,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5747,f5327])).
fof(f5747,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f976])).
fof(f6647,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5748,f5327])).
fof(f5748,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f6646,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5749,f5327])).
fof(f5749,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f5327,f5327,f976])).
fof(f6645,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5750,f5327])).
fof(f5750,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6644,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5752,f5327])).
fof(f5752,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6643,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5757,f5327])).
fof(f5757,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1122,f5327,f976])).
fof(f6642,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5758,f5327])).
fof(f5758,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f1122,f5327,f976])).
fof(f6641,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5760,f5327])).
fof(f5760,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f1122,f5327,f976])).
fof(f6640,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5761,f5327])).
fof(f5761,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f976])).
fof(f6639,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5762,f5327])).
fof(f5762,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f6638,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5764,f5327])).
fof(f5764,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6637,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5765,f5327])).
fof(f5765,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f5327,f5327,f976])).
fof(f6636,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5766,f5327])).
fof(f5766,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6635,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5767,f5327])).
fof(f5767,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1122,f5327,f976])).
fof(f6634,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5768,f5327])).
fof(f5768,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f1122,f5327,f976])).
fof(f6633,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5770,f5327])).
fof(f5770,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f1122,f5327,f976])).
fof(f6632,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5775,f5327])).
fof(f5775,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f976])).
fof(f6631,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5776,f5327])).
fof(f5776,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f5327,f5327,f976])).
fof(f6630,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5777,f5327])).
fof(f5777,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f5327,f5327,f976])).
fof(f6629,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5778,f5327])).
fof(f5778,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f5327,f5327,f976])).
fof(f6628,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5780,f5327])).
fof(f5780,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2275,f5327,f5327,f976])).
fof(f6627,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5785,f1169])).
fof(f5785,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1975,f5327,f980])).
fof(f980,plain,(
  ( ! [X2,X0,X1] : (sP20(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f611])).
fof(f1975,plain,(
  ~sP20('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f648,f1169,f982])).
fof(f6626,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5786,f1169])).
fof(f5786,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1979,f5327,f980])).
fof(f6625,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5787,f4445])).
fof(f4445,plain,(
  s(bool,'T_0') != s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | $spl3),
  inference(unit_resulting_resolution,[],[f648,f4218,f1104])).
fof(f1104,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1103])).
fof(f1103,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1051])).
fof(f1051,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f635])).
fof(f635,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X36,X130,X39,X131] : ((s(bool,X36) = s(bool,X130) & (p(s(bool,X130)) => s(bool,X39) = s(bool,X131))) => ((p(s(bool,X36)) => p(s(bool,X39))) <=> (p(s(bool,X130)) => p(s(bool,X131)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.IMP_CONG')).
fof(f5787,plain,(
  s(bool,'T_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4269,f5327,f980])).
fof(f6624,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5794,f1169])).
fof(f5794,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2284,f5327,f980])).
fof(f6623,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5795,f1169])).
fof(f5795,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2285,f5327,f980])).
fof(f2285,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1169,f1118,f982])).
fof(f6620,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5809,f4261])).
fof(f4261,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | $spl3),
  inference(unit_resulting_resolution,[],[f1122,f4218,f817])).
fof(f5809,plain,(
  s(bool,'F_0') != s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4272,f5327,f983])).
fof(f4272,plain,(
  sP20('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f1147,f1138,f1143,f4218,f966])).
fof(f6619,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5847,f6381])).
fof(f6381,plain,(
  ( ! [X0,X1] : (sP22('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1150,f5327,f1107])).
fof(f1107,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X0)) | ~sP21(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f990])).
fof(f990,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP21(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | (sP21(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP21(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP22(X0,X1,X2)))),
  inference(rectify,[],[f613])).
fof(f613,plain,(
  ! [X2,X0,X1] : ((~sP22(X2,X0,X1) | (sP21(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP21(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP22(X2,X0,X1)))),
  inference(flattening,[],[f612])).
fof(f612,plain,(
  ! [X2,X0,X1] : ((~sP22(X2,X0,X1) | (sP21(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP21(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP22(X2,X0,X1)))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X2,X0,X1] : (sP22(X2,X0,X1) <=> (sP21(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f1150,plain,(
  ( ! [X0,X1] : (sP21(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f992])).
fof(f992,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f616])).
fof(f616,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP21(X0,X1,X2)))),
  inference(flattening,[],[f615])).
fof(f615,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP21(X0,X1,X2)))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2] : (sP21(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP21])])).
fof(f5847,plain,(
  ( ! [X0,X1] : (~sP22('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f985])).
fof(f985,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f614])).
fof(f6618,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5855,f6381])).
fof(f5855,plain,(
  ( ! [X0,X1] : (~sP22('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f986])).
fof(f986,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f614])).
fof(f6617,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5865,f5869])).
fof(f5869,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f994])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f616])).
fof(f5865,plain,(
  ( ! [X0] : (~sP21('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f648,f5327,f991])).
fof(f991,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f616])).
fof(f6616,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5875,f1122])).
fof(f5875,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2421,f5327,f995])).
fof(f995,plain,(
  ( ! [X2,X0,X1] : (sP22(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP22(X2,X0,X1)) & (~sP22(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f617])).
fof(f617,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP22(X2,X0,X1)) & (~sP22(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP22(X2,X0,X1))),
  inference(definition_folding,[],[f295,f430,f429])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X115,X113,X114] : ((p(s(bool,X114)) <=> (p(s(bool,X113)) & p(s(bool,X115)))) <=> ((p(s(bool,X114)) | ~p(s(bool,X113)) | ~p(s(bool,X115))) & (p(s(bool,X113)) | ~p(s(bool,X114))) & (p(s(bool,X115)) | ~p(s(bool,X114)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sat.sat_23_prove_ltb')).
fof(f2421,plain,(
  ~sP22('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2416,f984])).
fof(f984,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | sP21(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f614])).
fof(f2416,plain,(
  ~sP21('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f648,f1122,f648,f991])).
fof(f6615,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5876,f5327])).
fof(f5876,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2286,f5327,f995])).
fof(f2286,plain,(
  ( ! [X0] : (~sP22('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f985])).
fof(f6614,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5877,f5327])).
fof(f5877,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2287,f5327,f995])).
fof(f2287,plain,(
  ( ! [X0] : (~sP22('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f986])).
fof(f6613,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5878,f5327])).
fof(f5878,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4353,f5327,f995])).
fof(f4353,plain,(
  ~sP22('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f648,f648,f4218,f1000])).
fof(f1000,plain,(
  ( ! [X2,X0,X1] : (~sP22(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f618])).
fof(f6612,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5883,f1122])).
fof(f5883,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f995])).
fof(f1246,plain,(
  ( ! [X0] : (~sP22('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f648,f1122,f985])).
fof(f6611,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5884,f5327])).
fof(f5884,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4332,f5327,f995])).
fof(f4332,plain,(
  ( ! [X0] : (~sP22('T_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f648,f4218,f985])).
fof(f6610,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5885,f5327])).
fof(f5885,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1247,f5327,f995])).
fof(f1247,plain,(
  ( ! [X0] : (~sP22('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f1122,f986])).
fof(f6609,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5886,f5327])).
fof(f5886,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4334,f5327,f995])).
fof(f4334,plain,(
  ( ! [X0] : (~sP22('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f648,f4218,f986])).
fof(f6608,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5888,f1122])).
fof(f5888,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2286,f5327,f995])).
fof(f6607,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5889,f5327])).
fof(f5889,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2287,f5327,f995])).
fof(f6606,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5899,f1122])).
fof(f5899,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2421,f5327,f996])).
fof(f996,plain,(
  ( ! [X2,X0,X1] : (sP22(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f618])).
fof(f6605,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5900,f5327])).
fof(f5900,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2286,f5327,f996])).
fof(f6604,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5901,f5327])).
fof(f5901,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2287,f5327,f996])).
fof(f6603,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5902,f5327])).
fof(f5902,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4353,f5327,f996])).
fof(f6602,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5907,f5327])).
fof(f5907,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f996])).
fof(f6601,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5908,f5327])).
fof(f5908,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4332,f5327,f996])).
fof(f6600,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5909,f1122])).
fof(f5909,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1247,f5327,f996])).
fof(f6599,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f5910,f5327])).
fof(f5910,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4334,f5327,f996])).
fof(f6598,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5912,f5327])).
fof(f5912,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2286,f5327,f996])).
fof(f6597,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5913,f1122])).
fof(f5913,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2287,f5327,f996])).
fof(f6596,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5942,f6172])).
fof(f6172,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f5327,f5327,f1012])).
fof(f1012,plain,(
  ( ! [X2,X0,X1] : (sP24(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP24(X2,X0,X1)) & (~sP24(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f625])).
fof(f625,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP24(X2,X0,X1)) & (~sP24(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP24(X2,X0,X1))),
  inference(definition_folding,[],[f297,f433,f432])).
fof(f432,plain,(
  ! [X2,X0,X1] : (sP23(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f433,plain,(
  ! [X2,X0,X1] : (sP24(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP23(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP24])])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X115,X113,X114] : ((p(s(bool,X114)) <=> (p(s(bool,X113)) | p(s(bool,X115)))) <=> ((p(s(bool,X114)) | ~p(s(bool,X113))) & (p(s(bool,X114)) | ~p(s(bool,X115))) & (p(s(bool,X113)) | p(s(bool,X115)) | ~p(s(bool,X114)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sat.sat_25_prove_ltb')).
fof(f5942,plain,(
  ( ! [X0,X1] : (~sP24(X0,X1,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f1001])).
fof(f1001,plain,(
  ( ! [X2,X0,X1] : (~sP24(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP23(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP23(X0,X1,X2) | sP24(X0,X1,X2)))),
  inference(rectify,[],[f620])).
fof(f620,plain,(
  ! [X2,X0,X1] : ((~sP24(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP23(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP23(X2,X0,X1) | sP24(X2,X0,X1)))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X2,X0,X1] : ((~sP24(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP23(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP23(X2,X0,X1)) | sP24(X2,X0,X1)))),
  inference(nnf_transformation,[],[f433])).
fof(f6595,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5957,f6172])).
fof(f5957,plain,(
  ( ! [X0,X1] : (~sP24(X0,'T_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f1002])).
fof(f1002,plain,(
  ( ! [X2,X0,X1] : (~sP24(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f621])).
fof(f6594,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5987,f5327])).
fof(f5987,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1157,f2289,f5327,f1007])).
fof(f1007,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | ~sP23(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f621])).
fof(f2289,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f1001])).
fof(f1157,plain,(
  ( ! [X0,X1] : (sP23('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f1011])).
fof(f1011,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP23(X0,X1,X2)))),
  inference(rectify,[],[f623])).
fof(f623,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP23(X2,X0,X1)))),
  inference(flattening,[],[f622])).
fof(f622,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP23(X2,X0,X1)))),
  inference(nnf_transformation,[],[f432])).
fof(f6593,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5988,f5327])).
fof(f5988,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1157,f1261,f5327,f1007])).
fof(f1261,plain,(
  ( ! [X0] : (~sP24('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f1122,f1002])).
fof(f6592,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5989,f5327])).
fof(f5989,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1156,f1261,f5327,f1007])).
fof(f1156,plain,(
  ( ! [X0,X1] : (sP23(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f648,f1010])).
fof(f1010,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f624])).
fof(f6591,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5991,f5327])).
fof(f5991,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2297,f1261,f5327,f1007])).
fof(f2297,plain,(
  ( ! [X0,X1] : (sP23(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f1009])).
fof(f1009,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f624])).
fof(f6590,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5992,f5327])).
fof(f5992,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1157,f2293,f5327,f1007])).
fof(f2293,plain,(
  ( ! [X0] : (~sP24('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f1002])).
fof(f6589,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5993,f5327])).
fof(f5993,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2298,f2293,f5327,f1007])).
fof(f2298,plain,(
  ( ! [X0,X1] : (sP23(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f1010])).
fof(f6588,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f5995,f5327])).
fof(f5995,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2297,f2293,f5327,f1007])).
fof(f6587,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6000,f6098])).
fof(f6098,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f1011])).
fof(f6000,plain,(
  ~sP23('F_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1122,f5327,f1007])).
fof(f6586,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6002,f6098])).
fof(f6002,plain,(
  ~sP23('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1122,f5327,f1007])).
fof(f6585,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6012,f6098])).
fof(f6012,plain,(
  ( ! [X0] : (~sP23('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1007])).
fof(f1257,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f1122,f1001])).
fof(f6584,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6014,f6098])).
fof(f6014,plain,(
  ( ! [X0] : (~sP23('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1007])).
fof(f6583,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6016,f6098])).
fof(f6016,plain,(
  ( ! [X0] : (~sP23('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1007])).
fof(f6582,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6018,f6098])).
fof(f6018,plain,(
  ( ! [X0] : (~sP23('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1007])).
fof(f6581,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6021,f5327])).
fof(f6021,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1155,f5327,f1007])).
fof(f1155,plain,(
  ( ! [X0,X1] : (sP23(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f1009])).
fof(f6580,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6023,f5327])).
fof(f6023,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1155,f5327,f1007])).
fof(f6579,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6024,f5327])).
fof(f6024,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f2297,f5327,f1007])).
fof(f6578,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6026,f5327])).
fof(f6026,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f2297,f5327,f1007])).
fof(f6577,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6028,f5327])).
fof(f6028,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f2298,f5327,f1007])).
fof(f6576,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6029,f5327])).
fof(f6029,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1157,f5327,f1007])).
fof(f6575,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6030,f5327])).
fof(f6030,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1157,f5327,f1007])).
fof(f6574,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6032,f5327])).
fof(f6032,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1157,f5327,f1007])).
fof(f6573,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6033,f6098])).
fof(f6033,plain,(
  ~sP23('F_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f5327,f1007])).
fof(f6572,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6034,f6098])).
fof(f6034,plain,(
  ~sP23('F_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1122,f5327,f1007])).
fof(f6571,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6043,f6098])).
fof(f6043,plain,(
  ( ! [X0] : (~sP23('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1007])).
fof(f6570,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6044,f6098])).
fof(f6044,plain,(
  ( ! [X0] : (~sP23('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1007])).
fof(f6569,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6045,f6098])).
fof(f6045,plain,(
  ( ! [X0] : (~sP23('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1007])).
fof(f6568,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6046,f6098])).
fof(f6046,plain,(
  ( ! [X0] : (~sP23('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1007])).
fof(f6567,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6061,f5327])).
fof(f6061,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2297,f648,f5327,f1008])).
fof(f1008,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f624])).
fof(f6566,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6062,f5327])).
fof(f6062,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1156,f648,f5327,f1008])).
fof(f6565,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6063,f5327])).
fof(f6063,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2298,f648,f5327,f1008])).
fof(f6564,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6068,f6098])).
fof(f6068,plain,(
  ( ! [X0] : (~sP23('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f1122,f5327,f1008])).
fof(f6563,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6077,f6098])).
fof(f6077,plain,(
  ( ! [X0,X1] : (~sP23('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f5327,f1008])).
fof(f6562,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6079,f5327])).
fof(f6079,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1155,f648,f5327,f1008])).
fof(f6561,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6080,f5327])).
fof(f6080,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2297,f648,f5327,f1008])).
fof(f6560,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6082,f5327])).
fof(f6082,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2298,f648,f5327,f1008])).
fof(f6559,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6087,f6098])).
fof(f6087,plain,(
  ( ! [X0] : (~sP23('T_0','F_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f1122,f5327,f1008])).
fof(f6558,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6096,f6098])).
fof(f6096,plain,(
  ( ! [X0,X1] : (~sP23('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327,f5327,f1008])).
fof(f6557,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6101,f5327])).
fof(f6101,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1122,f5327,f1012])).
fof(f6556,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6102,f5327])).
fof(f6102,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1122,f5327,f1012])).
fof(f6555,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6103,f5327])).
fof(f6103,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1122,f5327,f1012])).
fof(f6554,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6105,f5327])).
fof(f6105,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6553,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6106,f5327])).
fof(f6106,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1012])).
fof(f6552,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6107,f5327])).
fof(f6107,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6551,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6108,f1122])).
fof(f6108,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f5327,f5327,f1012])).
fof(f2487,plain,(
  ~sP24('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2482,f1003])).
fof(f1003,plain,(
  ( ! [X2,X0,X1] : (~sP24(X0,X1,X2) | sP23(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f621])).
fof(f2482,plain,(
  ~sP23('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f648,f1122,f1122,f1008])).
fof(f6550,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6109,f1122])).
fof(f6109,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1122,f5327,f1012])).
fof(f6549,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6110,f1122])).
fof(f6110,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1122,f5327,f1012])).
fof(f6548,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6116,f1122])).
fof(f6116,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1012])).
fof(f6547,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6117,f1122])).
fof(f6117,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6546,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6118,f1122])).
fof(f6118,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1012])).
fof(f6545,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6119,f1122])).
fof(f6119,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6544,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6125,f5327])).
fof(f6125,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f5327,f1012])).
fof(f6543,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6126,f5327])).
fof(f6126,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1122,f5327,f1012])).
fof(f6542,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6128,f5327])).
fof(f6128,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1122,f5327,f1012])).
fof(f6541,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6129,f5327])).
fof(f6129,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1012])).
fof(f6540,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6130,f5327])).
fof(f6130,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6539,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6132,f5327])).
fof(f6132,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6538,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6133,f1122])).
fof(f6133,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f5327,f5327,f1012])).
fof(f6537,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6134,f1122])).
fof(f6134,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f5327,f1012])).
fof(f6536,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6135,f1122])).
fof(f6135,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1122,f5327,f1012])).
fof(f6535,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6141,f1122])).
fof(f6141,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1012])).
fof(f6534,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6142,f1122])).
fof(f6142,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6533,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6143,f1122])).
fof(f6143,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1012])).
fof(f6532,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6144,f1122])).
fof(f6144,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6531,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6151,f5327])).
fof(f6151,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f1122,f5327,f1012])).
fof(f6530,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6152,f1122])).
fof(f6152,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f1122,f5327,f1012])).
fof(f6529,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6153,f5327])).
fof(f6153,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f1012])).
fof(f6528,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6154,f5327])).
fof(f6154,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6527,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6156,f5327])).
fof(f6156,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6526,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6157,f1122])).
fof(f6157,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f5327,f5327,f1012])).
fof(f6525,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6159,f5327])).
fof(f6159,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f1122,f5327,f1012])).
fof(f6524,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6160,f1122])).
fof(f6160,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f1122,f5327,f1012])).
fof(f6523,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6166,f5327])).
fof(f6166,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f5327,f5327,f1012])).
fof(f6522,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6167,f5327])).
fof(f6167,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f5327,f1012])).
fof(f6521,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6168,f5327])).
fof(f6168,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f5327,f5327,f1012])).
fof(f6520,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6169,f1122])).
fof(f6169,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f5327,f5327,f1012])).
fof(f6519,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6175,f5327])).
fof(f6175,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2299,f648,f5327,f1015])).
fof(f1015,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP24(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f626])).
fof(f2299,plain,(
  ( ! [X0] : (sP24('T_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f1013])).
fof(f1013,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP24(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f626])).
fof(f6518,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6176,f5327])).
fof(f6176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1270,f648,f5327,f1015])).
fof(f1270,plain,(
  ( ! [X0] : (sP24('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f648,f1014])).
fof(f1014,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP24(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f626])).
fof(f6517,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6177,f5327])).
fof(f6177,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2304,f648,f5327,f1015])).
fof(f2304,plain,(
  ( ! [X0] : (sP24('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f1014])).
fof(f6516,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6222,f5327])).
fof(f6222,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1264,f648,f5327,f1015])).
fof(f1264,plain,(
  ( ! [X0] : (sP24('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f648,f1013])).
fof(f6515,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6223,f5327])).
fof(f6223,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2299,f648,f5327,f1015])).
fof(f6514,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6225,f5327])).
fof(f6225,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2304,f648,f5327,f1015])).
fof(f6513,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6230,f1122])).
fof(f6230,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1295,f5327,f1018])).
fof(f1018,plain,(
  ( ! [X2,X0,X1] : (~sP26(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0,X1,X2] : ((~sP26(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP25(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP25(X0,X1,X2) | sP26(X0,X1,X2)))),
  inference(rectify,[],[f628])).
fof(f628,plain,(
  ! [X2,X0,X1] : ((~sP26(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP25(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP25(X2,X0,X1) | sP26(X2,X0,X1)))),
  inference(flattening,[],[f627])).
fof(f627,plain,(
  ! [X2,X0,X1] : ((~sP26(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP25(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP25(X2,X0,X1)) | sP26(X2,X0,X1)))),
  inference(nnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : (sP26(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP25(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP26])])).
fof(f1295,plain,(
  ( ! [X0] : (sP26('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1122,f648,f1030])).
fof(f1030,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP26(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP26(X2,X0,X1)) & (~sP26(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f633])).
fof(f633,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP26(X2,X0,X1)) & (~sP26(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP26(X2,X0,X1))),
  inference(definition_folding,[],[f385,f436,f435])).
fof(f435,plain,(
  ! [X2,X0,X1] : (sP25(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP25])])).
fof(f385,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X115,X113,X114] : ((p(s(bool,X114)) <=> (p(s(bool,X113)) => p(s(bool,X115)))) <=> ((p(s(bool,X114)) | p(s(bool,X113))) & (p(s(bool,X114)) | ~p(s(bool,X115))) & (~p(s(bool,X113)) | p(s(bool,X115)) | ~p(s(bool,X114)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sat.sat_27_prove_ltb')).
fof(f6512,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6231,f5327])).
fof(f6231,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1311,f5327,f1018])).
fof(f1311,plain,(
  ( ! [X0] : (sP26('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f648,f1031])).
fof(f1031,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP26(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f634])).
fof(f6511,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6232,f5327])).
fof(f6232,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2314,f5327,f1018])).
fof(f2314,plain,(
  ( ! [X0] : (sP26('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f1031])).
fof(f6510,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6234,f1122])).
fof(f6234,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2323,f5327,f1018])).
fof(f2323,plain,(
  ( ! [X0] : (sP26('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1159,f1118,f1105])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP25(X0,X1,X2) | sP26(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1021])).
fof(f1021,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2) | ~sP25(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f629])).
fof(f1159,plain,(
  ( ! [X0,X1] : (sP25(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1122,f1026])).
fof(f1026,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f632])).
fof(f632,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP25(X0,X1,X2)))),
  inference(rectify,[],[f631])).
fof(f631,plain,(
  ! [X2,X0,X1] : ((~sP25(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP25(X2,X0,X1)))),
  inference(flattening,[],[f630])).
fof(f630,plain,(
  ! [X2,X0,X1] : ((~sP25(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP25(X2,X0,X1)))),
  inference(nnf_transformation,[],[f435])).
fof(f6509,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6235,f5327])).
fof(f6235,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2324,f5327,f1018])).
fof(f2324,plain,(
  ( ! [X0] : (sP26('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1161,f1118,f1105])).
fof(f1161,plain,(
  ( ! [X0,X1] : (sP25(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f648,f1027])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f632])).
fof(f6508,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6236,f1122])).
fof(f6236,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1378,f5327,f1018])).
fof(f1378,plain,(
  ( ! [X0,X1] : (sP26('const.bool.BOUNDED_1'(s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1122,f1298])).
fof(f1298,plain,(
  ( ! [X2,X0,X1] : (sP26('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1296,f648])).
fof(f1296,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP26('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1)) )),
  inference(superposition,[],[f1030,f655])).
fof(f6507,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6237,f5327])).
fof(f6237,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1397,f5327,f1018])).
fof(f1397,plain,(
  ( ! [X0,X1] : (sP26('const.bool.BOUNDED_1'(s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f648,f1314])).
fof(f1314,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP26('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(subsumption_resolution,[],[f1312,f648])).
fof(f1312,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP26('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(superposition,[],[f1031,f655])).
fof(f6506,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6238,f5327])).
fof(f6238,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f5327,f1018])).
fof(f2329,plain,(
  ( ! [X0,X1] : (sP26('const.bool.BOUNDED_1'(s(bool,X0)),'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f1314])).
fof(f6505,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6239,f5327])).
fof(f6239,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1400,f5327,f1018])).
fof(f1400,plain,(
  ( ! [X2,X0,X1] : (sP26('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(subsumption_resolution,[],[f1398,f648])).
fof(f1398,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP26('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(superposition,[],[f1314,f655])).
fof(f6504,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6247,f1122])).
fof(f6247,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2563,f5327,f1018])).
fof(f2563,plain,(
  sP26('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1162,f648,f1122,f1024])).
fof(f1024,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2) | ~sP25(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f629])).
fof(f1162,plain,(
  ( ! [X0,X1] : (sP25('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f1028])).
fof(f1028,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f632])).
fof(f6503,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6248,f5327])).
fof(f6248,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2323,f5327,f1018])).
fof(f6502,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6249,f5327])).
fof(f6249,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2324,f5327,f1018])).
fof(f6501,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6250,f5327])).
fof(f6250,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1378,f5327,f1018])).
fof(f6500,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6251,f5327])).
fof(f6251,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1397,f5327,f1018])).
fof(f6499,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6252,f5327])).
fof(f6252,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2329,f5327,f1018])).
fof(f6498,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6253,f5327])).
fof(f6253,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1400,f5327,f1018])).
fof(f6497,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6288,f6304])).
fof(f6304,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f1028])).
fof(f6288,plain,(
  ~sP25('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f648,f5327,f1024])).
fof(f1285,plain,(
  ( ! [X0] : (~sP26('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f1122,f1019])).
fof(f1019,plain,(
  ( ! [X2,X0,X1] : (~sP26(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f629])).
fof(f6496,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6290,f6304])).
fof(f6290,plain,(
  ~sP25('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2309,f648,f5327,f1024])).
fof(f2309,plain,(
  ( ! [X0] : (~sP26('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1122,f1118,f1019])).
fof(f6495,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6300,f6304])).
fof(f6300,plain,(
  ( ! [X0] : (~sP25('T_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f648,f648,f5327,f1025])).
fof(f1025,plain,(
  ( ! [X2,X0,X1] : (~sP25(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f632])).
fof(f6494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6313,f1122])).
fof(f6313,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f648,f5327,f1029])).
fof(f1029,plain,(
  ( ! [X2,X0,X1] : (sP26(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f634])).
fof(f6493,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6314,f1122])).
fof(f6314,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2309,f648,f5327,f1029])).
fof(f6492,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6321,f5327])).
fof(f6321,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2309,f648,f5327,f1029])).
fof(f6491,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6322,f1122])).
fof(f6322,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2579,f648,f5327,f1029])).
fof(f2579,plain,(
  ~sP26('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2573,f1020])).
fof(f1020,plain,(
  ( ! [X2,X0,X1] : (~sP26(X0,X1,X2) | sP25(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f629])).
fof(f2573,plain,(
  ~sP25('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f648,f648,f1122,f1025])).
fof(f6490,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6387,f1150])).
fof(f6387,plain,(
  ( ! [X0] : (~sP21(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f1107])).
fof(f6489,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6389,f1150])).
fof(f6389,plain,(
  ( ! [X0] : (~sP21(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4332,f5327,f1107])).
fof(f6488,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6391,f1150])).
fof(f6391,plain,(
  ( ! [X0] : (~sP21('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1247,f5327,f1107])).
fof(f6487,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f6393,f1150])).
fof(f6393,plain,(
  ( ! [X0] : (~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X0,'T_0')) ) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4334,f5327,f1107])).
fof(f6486,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6396,f2288])).
fof(f2288,plain,(
  ( ! [X0,X1] : (sP21(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f992])).
fof(f6396,plain,(
  ( ! [X0] : (~sP21(X0,'F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2286,f5327,f1107])).
fof(f6485,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6398,f2288])).
fof(f6398,plain,(
  ( ! [X0] : (~sP21('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2287,f5327,f1107])).
fof(f6484,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6402,f5327])).
fof(f6402,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f5327,f1125])).
fof(f1125,plain,(
  ( ! [X0,X1] : (~sP15(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f887,f1122])).
fof(f887,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP15(X0,X1)) )),
  inference(cnf_transformation,[],[f566])).
fof(f566,plain,(
  ! [X0,X1] : ((~sP15(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP15(X0,X1)))),
  inference(flattening,[],[f565])).
fof(f565,plain,(
  ! [X0,X1] : ((~sP15(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP15(X0,X1)))),
  inference(nnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1] : (sP15(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f2263,plain,(
  ( ! [X0] : (sP15(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f888])).
fof(f888,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP15(X0,X1)) )),
  inference(cnf_transformation,[],[f566])).
fof(f6483,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6403,f5327])).
fof(f6403,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1168,f5327,f1125])).
fof(f1168,plain,(
  ( ! [X2,X3] : (sP15(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(subsumption_resolution,[],[f1166,f648])).
fof(f1166,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP15(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(superposition,[],[f888,f655])).
fof(f6482,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6404,f5327])).
fof(f6404,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f5327,f1125])).
fof(f1424,plain,(
  ( ! [X0,X1] : (sP15(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f656,f888])).
fof(f6481,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6405,f5327])).
fof(f6405,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1136,f5327,f1125])).
fof(f1136,plain,(
  ( ! [X0] : (sP15('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f648,f889])).
fof(f889,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP15(X0,X1)) )),
  inference(cnf_transformation,[],[f566])).
fof(f6480,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6406,f5327])).
fof(f6406,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2264,f5327,f1125])).
fof(f2264,plain,(
  ( ! [X0] : (sP15('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f889])).
fof(f6479,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6407,f5327])).
fof(f6407,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1167,f5327,f1125])).
fof(f1167,plain,(
  ( ! [X0,X1] : (sP15('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(subsumption_resolution,[],[f1165,f648])).
fof(f1165,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP15('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f889,f655])).
fof(f6478,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6408,f5327])).
fof(f6408,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f5327,f1125])).
fof(f1425,plain,(
  ( ! [X0,X1] : (sP15('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f656,f889])).
fof(f6477,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6413,f5327])).
fof(f6413,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1135,f5327,f1125])).
fof(f1135,plain,(
  ( ! [X0] : (sP15(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f888])).
fof(f6476,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6414,f5327])).
fof(f6414,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f5327,f1125])).
fof(f6475,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6415,f5327])).
fof(f6415,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1168,f5327,f1125])).
fof(f6474,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6416,f5327])).
fof(f6416,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1424,f5327,f1125])).
fof(f6473,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6418,f5327])).
fof(f6418,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2264,f5327,f1125])).
fof(f6472,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6419,f5327])).
fof(f6419,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1167,f5327,f1125])).
fof(f6471,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6420,f5327])).
fof(f6420,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1425,f5327,f1125])).
fof(f6438,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f659,f5327,f856])).
fof(f856,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f546])).
fof(f659,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X0)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X8] : s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X8))) = s('type.patSem.v',happ(s(fun('type.ast.lit','type.patSem.v'),'const.patSem.Litv_0'),s('type.ast.lit',X8)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','ARITY_0_1_const.patSem.Litv')).
fof(f6436,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2256,f5327,f850])).
fof(f850,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP9(X0,X1)) & (~sP9(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> sP9(X0,X1))),
  inference(definition_folding,[],[f372,f412,f411,f410,f409,f408])).
fof(f408,plain,(
  ! [X2,X4,X3] : (sP5(X2,X4,X3) <=> ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f409,plain,(
  ! [X2,X4,X3] : (sP6(X2,X4,X3) <=> ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f410,plain,(
  ! [X2,X4,X3] : (sP7(X2,X4,X3) <=> ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f411,plain,(
  ! [X2,X4,X3] : (sP8(X2,X4,X3) <=> ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f412,plain,(
  ! [X0,X1] : (sP9(X0,X1) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | sP5(X2,X4,X3) | sP6(X2,X4,X3) | sP7(X2,X4,X3) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | sP8(X2,X4,X3)) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f372,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(ennf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (! [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X5,X13] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X5),s('type.patSem.v',X13)))) <=> ! [X14] : (! [X15,X16] : ((? [X7] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X17,X18,X19] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X20,X21,X22,X23] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X20),s('type.patLang.exp',X22))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X20))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patSem.v'),X21))))))),s('type.patLang.exp',X22))),s('type.patLang.exp',X23))))) | ? [X20,X21,X24,X25,X26] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patLang.exp'),X24),s('type.num.num',X26))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X26))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X24))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X20))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X24))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.list.list'('type.patSem.v'),X20),s('type.list.list'('type.patSem.v'),X21))))))),s('type.list.list'('type.patLang.exp'),X24),s('type.list.list'('type.patLang.exp'),X25))))) | ? [X26] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X26))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X26)))) | ? [X18,X19] : (s('type.patSem.v',X15) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.patSem.v',X15))),s('type.patSem.v',X16))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X14),s('type.patSem.v',X5))),s('type.patSem.v',X13))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.exh_to_patProof.v_rel_def')).
fof(f2256,plain,(
  sP9(sK27,'const.patSem.Litv_1'(s('type.ast.lit',sK28))) | $spl0),
  inference(unit_resulting_resolution,[],[f1118,f849])).
fof(f849,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP9(X0,X1)) )),
  inference(cnf_transformation,[],[f543])).
fof(f6434,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2255,f5327,f1123])).
fof(f1123,plain,(
  ( ! [X0,X1] : (~sP14(X1,X0) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(forward_demodulation,[],[f886,f816])).
fof(f886,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | ~sP14(X1,X0)) )),
  inference(cnf_transformation,[],[f564])).
fof(f6431,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2267,f5327,f2788])).
fof(f2788,plain,(
  ( ! [X6] : (sP16(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | p(s(bool,X6))) ) | $spl0),
  inference(subsumption_resolution,[],[f2735,f648])).
fof(f2735,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP16(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) ) | $spl0),
  inference(superposition,[],[f907,f2261])).
fof(f2267,plain,(
  ~sP16('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f903])).
fof(f903,plain,(
  ( ! [X0,X1] : (~sP16(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f573])).
fof(f6429,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2265,f5327,f2787])).
fof(f2787,plain,(
  ( ! [X5] : (sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X5) | p(s(bool,X5))) ) | $spl0),
  inference(subsumption_resolution,[],[f2734,f648])).
fof(f2734,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),X5)) ) | $spl0),
  inference(superposition,[],[f904,f2261])).
fof(f2265,plain,(
  ~sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f648,f1118,f903])).
fof(f6424,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1135,f5327,f5327,f1125])).
fof(f6422,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1135,f1122,f5327,f1125])).
fof(f6412,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1135,f5327,f5327,f1125])).
fof(f6410,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1136,f1122,f5327,f1125])).
fof(f6399,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1153,f2287,f5327,f1107])).
fof(f1153,plain,(
  ( ! [X0,X1] : (sP21('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f994])).
fof(f6397,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1151,f2286,f5327,f1107])).
fof(f1151,plain,(
  ( ! [X0,X1] : (sP21(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1122,f993])).
fof(f993,plain,(
  ( ! [X2,X0,X1] : (sP21(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f616])).
fof(f6394,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4340,f4334,f5327,f1107])).
fof(f4340,plain,(
  ( ! [X0,X1] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4218,f994])).
fof(f6392,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1153,f1247,f5327,f1107])).
fof(f6390,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f4339,f4332,f5327,f1107])).
fof(f4339,plain,(
  ( ! [X0,X1] : (sP21(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))),s('type.patSem.v',sK27)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f4218,f993])).
fof(f6388,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1151,f1246,f5327,f1107])).
fof(f6369,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f920,f5327,f1104])).
fof(f920,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X37,X33] : ! [X36,X39] : s(X33,X36) = s(X33,'const.pair.FST_1'(s('type.pair.prod'(X33,X37),'const.pair.,_2'(s(X33,X36),s(X37,X39)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.pair.FST')).
fof(f6367,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f919,f5327,f1104])).
fof(f919,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X33,X37] : ! [X36,X39] : s(X37,X39) = s(X37,'const.pair.SND_1'(s('type.pair.prod'(X33,X37),'const.pair.,_2'(s(X33,X36),s(X37,X39)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.pair.SND')).
fof(f6365,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f916,f5327,f1104])).
fof(f916,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X37,X33] : ! [X36] : s(X33,X36) = s(X33,'const.sum.OUTL_1'(s('type.sum.sum'(X33,X37),'const.sum.INL_1'(s(X33,X36)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sum.OUTL')).
fof(f6363,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f915,f5327,f1104])).
fof(f915,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X33,X37] : ! [X36] : s(X37,X36) = s(X37,'const.sum.OUTR_1'(s('type.sum.sum'(X33,X37),'const.sum.INR_1'(s(X37,X36)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.sum.OUTR')).
fof(f6361,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f766,f5327,f1104])).
fof(f766,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X33] : ! [X55,X91] : s(X33,X55) = s(X33,'const.ind_type.FCONS_3'(s(X33,X55),s(fun('type.num.num',X33),X91),s('type.num.num','const.num.|48|_0')))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.ind_type.FCONS|split|0')).
fof(f6359,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f939,f5327,f1104])).
fof(f939,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK166(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK166(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK166(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK166(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK166])],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X37,X33] : ! [X124,X92] : ? [X125] : (s(X37,X124) = s(X37,happ(s(fun('type.list.list'(X33),X37),X125),s('type.list.list'(X33),'const.list.NIL_0'))) & ! [X28,X29] : s(X37,happ(s(fun('type.list.list'(X33),X37),X125),s('type.list.list'(X33),'const.list.CONS_2'(s(X33,X28),s('type.list.list'(X33),X29))))) = s(X37,happ(s(fun(X37,X37),happ(s(fun('type.list.list'(X33),fun(X37,X37)),happ(s(fun(X33,fun('type.list.list'(X33),fun(X37,X37))),X92),s(X33,X28))),s('type.list.list'(X33),X29))),s(X37,happ(s(fun('type.list.list'(X33),X37),X125),s('type.list.list'(X33),X29))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.list.list_Axiom')).
fof(f6357,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f755,f5327,f1104])).
fof(f6355,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f648,f2261,f5327,f1104])).
fof(f6353,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f655,f5327,f1104])).
fof(f6350,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1311,f648,f648,f5327,f1032])).
fof(f1032,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP26(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f634])).
fof(f6338,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2579,f648,f5327,f1030])).
fof(f6332,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f648,f5327,f5327,f1029])).
fof(f6326,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2579,f648,f1122,f5327,f1029])).
fof(f6319,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f5327,f648,f5327,f1029])).
fof(f6317,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1122,f648,f5327,f1029])).
fof(f6305,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2573,f5327,f1028])).
fof(f6303,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2573,f5327,f1026])).
fof(f6301,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1161,f648,f648,f5327,f1025])).
fof(f6291,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1162,f2309,f648,f5327,f1024])).
fof(f6289,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1162,f1285,f648,f5327,f1024])).
fof(f6269,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1311,f648,f5327,f1019])).
fof(f6257,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1295,f5327,f5327,f1018])).
fof(f6255,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2563,f1122,f5327,f1018])).
fof(f6243,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1295,f5327,f5327,f1018])).
fof(f6241,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1295,f1122,f5327,f1018])).
fof(f6229,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1264,f5327,f648,f5327,f1015])).
fof(f6227,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1264,f1122,f648,f5327,f1015])).
fof(f6201,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1264,f648,f5327,f5327,f1015])).
fof(f6186,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1270,f648,f1122,f5327,f1015])).
fof(f6173,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f5327,f1012])).
fof(f6171,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1122,f5327,f5327,f1012])).
fof(f6164,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f1122,f5327,f1012])).
fof(f6162,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2487,f1122,f1122,f5327,f1012])).
fof(f6149,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f5327,f1012])).
fof(f6147,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f5327,f5327,f1012])).
fof(f6140,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f1122,f5327,f1012])).
fof(f6138,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f1122,f5327,f1012])).
fof(f6124,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f5327,f5327,f5327,f1012])).
fof(f6122,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1122,f5327,f5327,f1012])).
fof(f6115,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f5327,f1122,f5327,f1012])).
fof(f6113,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1122,f1122,f5327,f1012])).
fof(f6099,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2482,f5327,f1011])).
fof(f6097,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1155,f648,f5327,f5327,f1008])).
fof(f6088,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1155,f648,f1122,f5327,f1008])).
fof(f6078,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1155,f648,f5327,f5327,f1008])).
fof(f6069,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1156,f648,f1122,f5327,f1008])).
fof(f6057,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1157,f5327,f5327,f1007])).
fof(f6055,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2293,f2298,f5327,f5327,f1007])).
fof(f6053,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1261,f1156,f5327,f5327,f1007])).
fof(f6051,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f2297,f5327,f5327,f1007])).
fof(f6049,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1155,f5327,f5327,f1007])).
fof(f6041,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1157,f1122,f5327,f1007])).
fof(f6039,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f2297,f1122,f5327,f1007])).
fof(f6037,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1257,f1155,f1122,f5327,f1007])).
fof(f6019,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1157,f2293,f5327,f5327,f1007])).
fof(f6017,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1157,f1261,f5327,f5327,f1007])).
fof(f6015,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1157,f2289,f5327,f5327,f1007])).
fof(f6013,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1157,f1257,f5327,f5327,f1007])).
fof(f6003,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1157,f2293,f1122,f5327,f1007])).
fof(f6001,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1157,f1261,f1122,f5327,f1007])).
fof(f5958,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1270,f648,f5327,f1002])).
fof(f5943,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1264,f648,f5327,f1001])).
fof(f5928,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2407,f648,f648,f5327,f1000])).
fof(f2407,plain,(
  sP22('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1150,f648,f648,f987])).
fof(f987,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP21(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f614])).
fof(f5918,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f5327,f996])).
fof(f5916,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1247,f1122,f5327,f996])).
fof(f5906,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f5327,f996])).
fof(f5904,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2421,f1122,f5327,f996])).
fof(f5894,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f5327,f995])).
fof(f5892,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f1122,f5327,f995])).
fof(f5882,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1246,f5327,f5327,f995])).
fof(f5880,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2421,f1122,f5327,f995])).
fof(f5870,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2416,f5327,f994])).
fof(f5868,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2416,f5327,f993])).
fof(f5866,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1150,f648,f648,f5327,f991])).
fof(f5856,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2407,f648,f5327,f986])).
fof(f5848,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2407,f648,f5327,f985])).
fof(f5832,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2719,f2261,f5327,f983])).
fof(f2719,plain,(
  sP20('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))) | $spl0),
  inference(unit_resulting_resolution,[],[f648,f2261,f981])).
fof(f981,plain,(
  ( ! [X2,X0,X1] : (sP20(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f611])).
fof(f5830,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1961,f655,f5327,f983])).
fof(f1961,plain,(
  ( ! [X0] : (sP20('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f648,f655,f981])).
fof(f5820,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2717,f2261,f5327,f983])).
fof(f2717,plain,(
  sP20('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f648,f2261,f981])).
fof(f5818,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1956,f655,f5327,f983])).
fof(f1956,plain,(
  ( ! [X0] : (sP20('T_0','const.bool.BOUNDED_1'(s(bool,X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f648,f655,f981])).
fof(f6458,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f5814])).
fof(f5814,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3635,f5327,f983])).
fof(f3635,plain,(
  sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2273,f1138,f1143,f1122,f966])).
fof(f6457,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f5813])).
fof(f5813,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3601,f5327,f983])).
fof(f3601,plain,(
  sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1137,f2272,f648,f965])).
fof(f965,plain,(
  ( ! [X2,X0,X1] : (~sP18(X1,X2,X0) | ~p(s(bool,X2)) | ~sP19(X0,X2,X1) | sP20(X0,X1,X2) | ~sP17(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f602])).
fof(f6456,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f5807])).
fof(f5807,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3634,f5327,f983])).
fof(f3634,plain,(
  sP20('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1147,f1138,f1143,f1122,f966])).
fof(f6455,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f5806])).
fof(f5806,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3596,f5327,f983])).
fof(f3596,plain,(
  sP20('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1149,f1137,f1142,f648,f965])).
fof(f5799,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1979,f1169,f5327,f980])).
fof(f5797,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1975,f1169,f5327,f980])).
fof(f5784,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f5327,f976])).
fof(f5782,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1122,f5327,f5327,f976])).
fof(f5774,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f1122,f5327,f976])).
fof(f5772,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1122,f1122,f5327,f976])).
fof(f5756,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f5327,f976])).
fof(f5754,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f1122,f5327,f5327,f976])).
fof(f5746,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f5327,f1122,f5327,f976])).
fof(f5744,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1122,f1122,f5327,f976])).
fof(f5728,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f5327,f5327,f976])).
fof(f5726,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1122,f5327,f5327,f976])).
fof(f5718,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1147,f5327,f1122,f5327,f976])).
fof(f5716,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1122,f1122,f5327,f976])).
fof(f5700,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2389,f5327,f975])).
fof(f2389,plain,(
  ~sP18('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f648,f1122,f648,f972])).
fof(f5698,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2389,f5327,f974])).
fof(f5696,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1142,f648,f648,f5327,f972])).
fof(f5686,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f5327,f971])).
fof(f2378,plain,(
  ~sP19('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f648,f1122,f648,f968])).
fof(f5684,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2378,f5327,f970])).
fof(f5682,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2650,f648,f648,f5327,f968])).
fof(f2650,plain,(
  ( ! [X0] : (sP19('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1956,f963])).
fof(f963,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,X2) | sP19(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f602])).
fof(f5642,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f4286,f5327,f967])).
fof(f5637,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f4286,f5327,f967])).
fof(f5622,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f1143,f5327,f967])).
fof(f5620,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f1143,f5327,f967])).
fof(f5615,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f1143,f5327,f967])).
fof(f5612,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f1143,f5327,f967])).
fof(f5601,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f2272,f5327,f967])).
fof(f5592,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f2272,f5327,f967])).
fof(f5582,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f1142,f5327,f967])).
fof(f5577,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f1142,f5327,f967])).
fof(f5568,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f1142,f5327,f967])).
fof(f5566,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f1142,f5327,f967])).
fof(f5513,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f4286,f5327,f966])).
fof(f5508,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f4286,f5327,f966])).
fof(f5493,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f1143,f5327,f966])).
fof(f5491,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f1143,f5327,f966])).
fof(f5486,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f1143,f5327,f966])).
fof(f5483,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f1143,f5327,f966])).
fof(f5472,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1138,f2272,f5327,f966])).
fof(f5463,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f2284,f1137,f2272,f5327,f966])).
fof(f5453,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f4281,f1142,f5327,f966])).
fof(f5448,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1138,f1142,f5327,f966])).
fof(f5439,plain,(
  $false | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1148,f4269,f1137,f1142,f5327,f966])).
fof(f5437,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1148,f1979,f1137,f1142,f5327,f966])).
fof(f5411,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1961,f648,f648,f5327,f964])).
fof(f5395,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1205,f648,f5327,f907])).
fof(f1205,plain,(
  ~sP16('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f648,f648,f903])).
fof(f5385,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1205,f648,f5327,f904])).
fof(f5375,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1207,f5327,f5327,f902])).
fof(f5373,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1216,f1122,f5327,f902])).
fof(f5361,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1207,f5327,f5327,f902])).
fof(f5359,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1207,f1122,f5327,f902])).
fof(f5347,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1122,f1169,f5327,f817])).
fof(f5345,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f1169,f5327,f817])).
fof(f5341,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1122,f1169,f5327,f817])).
fof(f5339,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f5327,f1169,f5327,f817])).
fof(f5335,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1169,f5327,f705])).
fof(f705,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X30] : (s(bool,'F_0') = s(bool,X30) <=> ~p(s(bool,X30)))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f5333,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f655,f5327,f700])).
fof(f5331,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f648,f5327])).
fof(f5329,plain,(
  p(s(bool,sK97('_i_0_0',bool))) | $spl10),
  inference(cnf_transformation,[],[f5329_D])).
fof(f5329_D,plain,(
  p(s(bool,sK97('_i_0_0',bool))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5326,plain,(
  ( ! [X3] : (p(s(bool,sK97('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f5315,f755])).
fof(f5315,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK97('_i_0_0',bool)))))) )),
  inference(superposition,[],[f760,f755])).
fof(f760,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK97(X1,X0)))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f4713,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f4713_D])).
fof(f4713_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4714,plain,(
  ( ! [X3] : (~p(s(bool,sK99(X3,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f4714_D])).
fof(f4714_D,plain,(
  ( ! [X3] : (~p(s(bool,sK99(X3,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f4704,plain,(
  ( ! [X3] : (~p(s(bool,sK99(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f769,f755])).
fof(f769,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK99(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f495])).
fof(f495,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK99(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK99(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK99])],[f365])).
fof(f365,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X33] : ! [X35,X30] : (! [X36] : (s(X33,X30) = s(X33,X36) => p(s(bool,happ(s(fun(X33,bool),X35),s(X33,X36))))) => p(s(bool,'|dollar|exists_1'(s(fun(X33,bool),X35)))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.bool.JRH_INDUCT_UTIL')).
fof(f2254,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2253,f1122])).
fof(f2253,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2234,f1188])).
fof(f1188,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v',sK27))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1134,f705])).
fof(f1134,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v',sK27)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1119,f1120])).
fof(f2234,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v',sK27)))) | $spl2),
  inference(superposition,[],[f658,f1120])).
fof(f658,plain,(
  ( ! [X0] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X0))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.exh_to_patProof.v_rel_rules|split|0')).
fof(f1118,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | $spl0),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1120,plain,(
  s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) | $spl2),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f646,plain,(
  s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))))),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | s('type.patSem.v',sK27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | s('type.patSem.v',sK27) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f438])).
fof(f438,plain,(
  ? [X0,X1] : ((p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ? [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) <~> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X6),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_lit|split|1','thm.exh_to_patProof.v_rel_lit|split|1')).
fof(f1119,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) | $spl1),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1121,plain,(
  s('type.patSem.v',sK27) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) | $spl3),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  s('type.patSem.v',sK27) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f647,plain,(
  s('type.patSem.v',sK27) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))) | ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK27),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK28))))))),
  inference(cnf_transformation,[],[f439])).
% SZS output end Proof for thm.exh_to_patProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 596
% Passive clauses: 2040
% Generated clauses: 10469
% Final active clauses: 548
% Final passive clauses: 1311
% Input formulas: 133
% Initial clauses: 447
% 
% Duplicate literals: 8
% Trivial inequalities: 8
% Fw subsumption resolutions: 296
% Fw demodulations: 1672
% Bw demodulations: 36
% Condensations: 19
% 
% Simple tautologies: 88
% Equational tautologies: 39
% Forward subsumptions: 5160
% Fw demodulations to eq. taut.: 720
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 4730
% Factoring: 1
% Forward superposition: 1902
% Backward superposition: 1278
% Self superposition: 16
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 2471
% SAT solver unit clauses: 900
% SAT solver binary clauses: 1097
% 
% Sat splits: 10
% Sat splitting refutations: 409
% 
% Memory used [KB]: 9850
% Time elapsed: 0.341 s
% ------------------------------
% ------------------------------
% Success in time 1.68 s
