Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.closProps
% (55637)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 107
% Passive clauses: 2665
% Generated clauses: 12535
% Final active clauses: 106
% Final passive clauses: 2484
% Input formulas: 39
% Initial clauses: 66
% 
% Duplicate literals: 455
% Trivial inequalities: 303
% Fw subsumption resolutions: 1160
% Fw demodulations: 1361
% 
% Simple tautologies: 328
% Equational tautologies: 23
% Forward subsumptions: 5848
% 
% Binary resolution: 61
% Forward superposition: 4636
% Backward superposition: 4239
% Self superposition: 193
% Equality factoring: 61
% Equality resolution: 42
% 
% Split clauses: 231
% Split components: 465
% SAT solver clauses: 463
% SAT solver unit clauses: 25
% SAT solver binary clauses: 103
% 
% Sat splits: 231
% 
% TWLsolver clauses: 463
% TWLsolver calls for satisfiability: 76
% 
% Memory used [KB]: 17526
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.closProps
% (55669)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 69
% Passive clauses: 1037
% Generated clauses: 8146
% Final active clauses: 69
% Final passive clauses: 942
% Input formulas: 39
% Initial clauses: 66
% 
% Duplicate literals: 290
% Trivial inequalities: 125
% Fw subsumption resolutions: 282
% Fw demodulations: 564
% Bw demodulations: 25
% Condensations: 622
% Global subsumptions: 4
% 
% Simple tautologies: 158
% Equational tautologies: 20
% Forward subsumptions: 3600
% 
% Binary resolution: 43
% Forward superposition: 2454
% Backward superposition: 2448
% Self superposition: 155
% Equality factoring: 46
% Equality resolution: 15
% 
% Split clauses: 35
% Split components: 72
% SAT solver clauses: 1514
% SAT solver unit clauses: 92
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 8
% 
% Sat splits: 35
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1498
% TWLsolver calls for satisfiability: 230
% 
% Memory used [KB]: 21492
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.closProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.closProps
% SZS output start Proof for thm.closProps
fof(f12259,plain,(
  $false),
  inference(global_subsumption,[],[f139,f138,f140,f143,f159,f160,f161,f162,f170,f171,f180,f181,f182,f183,f184,f215,f193,f194,f195,f196,f197,f198,f199,f200,f201,f202,f203,f223,f225,f228,f287,f288,f299,f309,f310,f347,f486,f489,f491,f541,f539,f576,f638,f678,f679,f693,f744,f791,f793,f898,f911,f918,f931,f1045,f1046,f1047,f1048,f1049,f1050,f1052,f1054,f1057,f1176,f1217,f1216,f1218,f1219,f1220,f1221,f1222,f1270,f1275,f1276,f1277,f1278,f1283,f1288,f1289,f1290,f1291,f1313,f1423,f1449,f1461,f1527,f1528,f1529,f1530,f1531,f1532,f1659,f1669,f1533,f2041,f2047,f2084,f2085,f2390,f2391,f2392,f2400,f2651,f2652,f2653,f2917,f2918,f2919,f2920,f2921,f2923,f2924,f2984,f2985,f3015,f3016,f3027,f2044,f2042,f222,f2656,f2654,f221,f1836,f1834,f1841,f1839,f2114,f2113,f2728,f2727,f2432,f2431,f1829,f1856,f1833,f1828,f1827,f1671,f1670,f1647,f1142,f1115,f3026,f3025,f2995,f2994,f168,f940,f941,f942,f169,f1853,f1855,f1859,f3000,f3002,f2915,f3033,f3035,f2916,f3032,f3034,f2999,f3001,f555,f661,f666,f663,f654,f651,f549,f546,f544,f649,f2104,f1661,f1663,f1525,f1851,f1126,f1526,f1043,f1044,f1665,f1667,f1122,f1124,f1093,f1095,f1660,f1662,f1664,f1666,f1121,f1123,f1092,f1094,f2102,f1588,f1586,f1267,f1265,f1097,f1098,f1166,f1165,f1136,f1133,f1128,f1109,f1106,f1101,f173,f172,f683,f680,f796,f798,f2721,f2698,f2696,f2695,f2693,f2689,f2658,f2107,f2083,f2081,f2080,f2078,f2074,f2046,f1843,f1838,f1643,f1641,f1640,f1638,f1634,f1633,f1590,f1424,f1372,f1346,f1344,f1343,f1310,f1308,f1307,f1302,f1301,f1269,f1199,f1198,f1197,f1196,f1141,f1140,f1139,f1138,f1135,f1134,f1114,f1113,f1112,f1111,f1108,f1107,f1087,f1056,f1051,f1018,f964,f956,f800,f795,f722,f721,f711,f710,f688,f684,f668,f665,f656,f653,f618,f617,f616,f614,f559,f554,f548,f543,f267,f235,f142,f4596,f4595,f4597,f4598,f4600,f4602,f4648,f4649,f4650,f4651,f4652,f4653,f4655,f4656,f4657,f4658,f4659,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4713,f4714,f4715,f4716,f4717,f4718,f4720,f4722,f4724,f4725,f4726,f4727,f4728,f4729,f4730,f4882,f4883,f4731,f4732,f4733,f4738,f4739,f4744,f4745,f4746,f4747,f4761,f4762,f4764,f4769,f4770,f4771,f4772,f4774,f4775,f4776,f4777,f4778,f4779,f4780,f4782,f4783,f4784,f4795,f4796,f4801,f4809,f4810,f4811,f4812,f4830,f4835,f4836,f4837,f4838,f4840,f4841,f4842,f4843,f4844,f4845,f4846,f4848,f4849,f4850,f4861,f4862,f5067,f5066,f5068,f5069,f5071,f5073,f5141,f5142,f5143,f5144,f5146,f5147,f5148,f5149,f5150,f5151,f5152,f5153,f5154,f5155,f5156,f5159,f5161,f5163,f5164,f5165,f5166,f5167,f5168,f5169,f5170,f5387,f5388,f5389,f5390,f5171,f5172,f5173,f5175,f5176,f5177,f5178,f5179,f5180,f5214,f5215,f5216,f5217,f5218,f5219,f5220,f5225,f5226,f5227,f5232,f5233,f5234,f5235,f5251,f5252,f5254,f5259,f5260,f5261,f5262,f5264,f5265,f5266,f5267,f5268,f5269,f5270,f5272,f5273,f5274,f5275,f5288,f5289,f5293,f5294,f5304,f5305,f5306,f5307,f5323,f5324,f5329,f5334,f5335,f5336,f5337,f5339,f5340,f5341,f5342,f5343,f5344,f5345,f5347,f5348,f5349,f5350,f5363,f5364,f5634,f5635,f5608,f5609,f5839,f5838,f5840,f5841,f5842,f5843,f5844,f5845,f5847,f5848,f5849,f5850,f5851,f5852,f5971,f5972,f5987,f5988,f5992,f5993,f5994,f5995,f6013,f6014,f6015,f6016,f6017,f6032,f6033,f6034,f6035,f6036,f6041,f6042,f6043,f6044,f6522,f6521,f6523,f6524,f6525,f6526,f6527,f6528,f6529,f6531,f6532,f6533,f6534,f6535,f6536,f6623,f6624,f6644,f6645,f6646,f6647,f6668,f6669,f6670,f6671,f6672,f6673,f6690,f6691,f6692,f6697,f6698,f6699,f6700,f7303,f7044,f7045,f7306,f7046,f7047,f7307,f7049,f7050,f7051,f7052,f7053,f7054,f7055,f7056,f7057,f7058,f7059,f7309,f7065,f7067,f7070,f7071,f7072,f7073,f7074,f7075,f7076,f7077,f7078,f7079,f7080,f7081,f7082,f7313,f7314,f7083,f7084,f7086,f7087,f7088,f7089,f7090,f7091,f7092,f7160,f7161,f7162,f7163,f7165,f7166,f7167,f7168,f7169,f7170,f7171,f7172,f7173,f7174,f7175,f7318,f7181,f7183,f7186,f7187,f7188,f7189,f7190,f7191,f7192,f7193,f7194,f7195,f7196,f7197,f7198,f7199,f7201,f7203,f7208,f7209,f7210,f7244,f7245,f7332,f7255,f7259,f7261,f7262,f7264,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7339,f7340,f7609,f7608,f7610,f7611,f7612,f7613,f7614,f7615,f7616,f7617,f7619,f7620,f7621,f7622,f7623,f7624,f7625,f7745,f7767,f7768,f7772,f7773,f7774,f7775,f7799,f7800,f7801,f7802,f7803,f7804,f7805,f7806,f7828,f7829,f7830,f7831,f7832,f7833,f7834,f7839,f7840,f7841,f7842,f7912,f7911,f8031,f8030,f8110,f8109,f8111,f8112,f8113,f8115,f8116,f8117,f8118,f8119,f8120,f8121,f8122,f8123,f8124,f8125,f8126,f8127,f8255,f8256,f8284,f8285,f8286,f8287,f8312,f8313,f8314,f8315,f8316,f8317,f8318,f8319,f8320,f8345,f8346,f8347,f8348,f8349,f8354,f8355,f8356,f8357,f8626,f8625,f8627,f8628,f8629,f8631,f8632,f8633,f8634,f8635,f8636,f8637,f8638,f8639,f8640,f8641,f8642,f8643,f8644,f8782,f8783,f8814,f8815,f8816,f8817,f8843,f8844,f8845,f8846,f8847,f8848,f8849,f8850,f8851,f8879,f8880,f8881,f8882,f8883,f8884,f8889,f8890,f8891,f8892,f9055,f9111,f9113,f9115,f9118,f9119,f9120,f9121,f9122,f9123,f9124,f9170,f9174,f9176,f9177,f9179,f9183,f9184,f9185,f9186,f9187,f9188,f9189,f9190,f9191,f9303,f9305,f11013,f11012,f11014,f11015,f11016,f11017,f11018,f11019,f11020,f11021,f11022,f11023,f11024,f11025,f11027,f11028,f11029,f11030,f11031,f11032,f11033,f11183,f11184,f11219,f11220,f11221,f11222,f11251,f11252,f11253,f11254,f11255,f11256,f11257,f11258,f11259,f11291,f11292,f11293,f11294,f11295,f11296,f11297,f11302,f11303,f11304,f11305,f12093,f12094,f12108,f12236])).
fof(f12236,plain,(
  p(s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))))),
  inference(forward_demodulation,[],[f12235,f5166])).
fof(f12235,plain,(
  p(s(bool,happ(s(fun('type.closSem.state'(sK0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(sK0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s('type.closSem.state'(sK0),sK5))))),
  inference(forward_demodulation,[],[f12234,f170])).
fof(f12234,plain,(
  p(s(bool,happ(s(fun('type.closSem.state'(sK0),bool),happ(s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(sK0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(sK0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(sK0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))))),s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6))),s('type.closSem.state'(sK0),sK5))))),
  inference(forward_demodulation,[],[f12155,f184])).
fof(f12155,plain,(
  p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(sK0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(sK0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(sK0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))))))))),
  inference(superposition,[],[f221,f138])).
fof(f12108,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X1),s('type.closSem.v',X2),s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),X4))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X4)))))))))) )),
  inference(unit_resulting_resolution,[],[f221,f225])).
fof(f12094,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f12002,f142])).
fof(f12002,plain,(
  ( ! [X12] : (p(s(bool,'F_0')) | s(bool,X12) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'))) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f616,f7911])).
fof(f12093,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f12000,f142])).
fof(f12000,plain,(
  ( ! [X8,X9] : (p(s(bool,'F_0')) | s(bool,X8) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f489,f7911])).
fof(f11305,plain,(
  ( ! [X749,X751,X745,X747,X741,X743,X748,X750,X744,X746,X742] : (p(s(bool,happ(s(fun(fun(X741,fun(X742,X743)),bool),happ(s(fun(fun(X741,fun(X742,X743)),fun(fun(X741,fun(X742,X743)),bool)),'=_0'),s(fun(X741,fun(X742,X743)),X751))),s(fun(X741,fun(X742,X743)),'_dst_x00x11x01_3'(s(fun(X745,fun(X746,X743)),X747),s(fun(X741,fun(X742,X745)),happ(s(fun(X744,fun(X741,fun(X742,X745))),X748),s(X744,X750))),s(fun(X742,X746),X749)))))) | s(fun(X741,fun(X742,X743)),X751) != s(fun(X741,fun(X742,X743)),happ(s(fun(X744,fun(X741,fun(X742,X743))),'_dst_x000x111x001_3'(s(fun(X745,fun(X746,X743)),X747),s(fun(X744,fun(X741,fun(X742,X745))),X748),s(fun(X742,X746),X749))),s(X744,X750)))) )),
  inference(superposition,[],[f173,f203])).
fof(f11304,plain,(
  ( ! [X736,X732,X734,X738,X730,X737,X733,X735,X739,X731,X740] : (p(s(bool,happ(s(fun(fun(X730,fun(X731,X732)),bool),happ(s(fun(fun(X730,fun(X731,X732)),fun(fun(X730,fun(X731,X732)),bool)),'=_0'),s(fun(X730,fun(X731,X732)),'_dst_x00x11x01_3'(s(fun(X734,fun(X735,X732)),X736),s(fun(X730,fun(X731,X734)),happ(s(fun(X733,fun(X730,fun(X731,X734))),X737),s(X733,X739))),s(fun(X731,X735),X738))))),s(fun(X730,fun(X731,X732)),X740)))) | s(fun(X730,fun(X731,X732)),X740) != s(fun(X730,fun(X731,X732)),happ(s(fun(X733,fun(X730,fun(X731,X732))),'_dst_x000x111x001_3'(s(fun(X734,fun(X735,X732)),X736),s(fun(X733,fun(X730,fun(X731,X734))),X737),s(fun(X731,X735),X738))),s(X733,X739)))) )),
  inference(superposition,[],[f173,f203])).
fof(f11303,plain,(
  ( ! [X728,X724,X726,X720,X722,X729,X725,X727,X721,X723,X719] : (~p(s(bool,happ(s(fun(fun(X719,fun(X720,X721)),bool),happ(s(fun(fun(X719,fun(X720,X721)),fun(fun(X719,fun(X720,X721)),bool)),'=_0'),s(fun(X719,fun(X720,X721)),X729))),s(fun(X719,fun(X720,X721)),'_dst_x00x11x01_3'(s(fun(X723,fun(X724,X721)),X725),s(fun(X719,fun(X720,X723)),happ(s(fun(X722,fun(X719,fun(X720,X723))),X726),s(X722,X728))),s(fun(X720,X724),X727)))))) | s(fun(X719,fun(X720,X721)),X729) = s(fun(X719,fun(X720,X721)),happ(s(fun(X722,fun(X719,fun(X720,X721))),'_dst_x000x111x001_3'(s(fun(X723,fun(X724,X721)),X725),s(fun(X722,fun(X719,fun(X720,X723))),X726),s(fun(X720,X724),X727))),s(X722,X728)))) )),
  inference(superposition,[],[f172,f203])).
fof(f11302,plain,(
  ( ! [X716,X718,X712,X714,X708,X710,X717,X713,X715,X709,X711] : (~p(s(bool,happ(s(fun(fun(X708,fun(X709,X710)),bool),happ(s(fun(fun(X708,fun(X709,X710)),fun(fun(X708,fun(X709,X710)),bool)),'=_0'),s(fun(X708,fun(X709,X710)),'_dst_x00x11x01_3'(s(fun(X712,fun(X713,X710)),X714),s(fun(X708,fun(X709,X712)),happ(s(fun(X711,fun(X708,fun(X709,X712))),X715),s(X711,X717))),s(fun(X709,X713),X716))))),s(fun(X708,fun(X709,X710)),X718)))) | s(fun(X708,fun(X709,X710)),X718) = s(fun(X708,fun(X709,X710)),happ(s(fun(X711,fun(X708,fun(X709,X710))),'_dst_x000x111x001_3'(s(fun(X712,fun(X713,X710)),X714),s(fun(X711,fun(X708,fun(X709,X712))),X715),s(fun(X709,X713),X716))),s(X711,X717)))) )),
  inference(superposition,[],[f172,f203])).
fof(f11297,plain,(
  ( ! [X660,X656,X658,X652,X654,X648,X650,X646,X661,X657,X659,X653,X655,X649,X651,X647] : (s(fun(X656,fun(X657,X658)),'_dst_x00x00x11_3'(s(fun(fun(X646,fun(X647,X648)),fun(X659,X658)),X660),s(fun(X646,fun(X647,X648)),'_dst_x00x11x01_3'(s(fun(X650,fun(X651,X648)),X652),s(fun(X646,fun(X647,X650)),happ(s(fun(X649,fun(X646,fun(X647,X650))),X653),s(X649,X655))),s(fun(X647,X651),X654))),s(fun(X656,fun(X657,X659)),X661))) = s(fun(X656,fun(X657,X658)),happ(s(fun(X649,fun(X656,fun(X657,X658))),'_dst_x000x100x011_3'(s(fun(fun(X646,fun(X647,X648)),fun(X659,X658)),X660),s(fun(X649,fun(X646,fun(X647,X648))),'_dst_x000x111x001_3'(s(fun(X650,fun(X651,X648)),X652),s(fun(X649,fun(X646,fun(X647,X650))),X653),s(fun(X647,X651),X654))),s(fun(X656,fun(X657,X659)),X661))),s(X649,X655)))) )),
  inference(superposition,[],[f202,f203])).
fof(f11296,plain,(
  ( ! [X639,X633,X635,X631,X644,X640,X642,X636,X638,X632,X634,X630,X645,X641,X643,X637] : (s(fun(X640,fun(X641,X642)),'_dst_x00x01x00_3'(s(fun(X643,fun(fun(X630,fun(X631,X632)),X642)),X644),s(fun(X641,X643),X645),s(fun(X630,fun(X631,X632)),'_dst_x00x11x01_3'(s(fun(X634,fun(X635,X632)),X636),s(fun(X630,fun(X631,X634)),happ(s(fun(X633,fun(X630,fun(X631,X634))),X637),s(X633,X639))),s(fun(X631,X635),X638))))) = s(fun(X640,fun(X641,X642)),happ(s(fun(X633,fun(X640,fun(X641,X642))),'_dst_x000x001x100_3'(s(fun(X643,fun(fun(X630,fun(X631,X632)),X642)),X644),s(fun(X641,X643),X645),s(fun(X633,fun(X630,fun(X631,X632))),'_dst_x000x111x001_3'(s(fun(X634,fun(X635,X632)),X636),s(fun(X633,fun(X630,fun(X631,X634))),X637),s(fun(X631,X635),X638))))),s(X633,X639)))) )),
  inference(superposition,[],[f201,f203])).
fof(f11295,plain,(
  ( ! [X629,X625,X627,X621,X623,X617,X619,X628,X624,X626,X620,X622,X616,X618] : (s(X626,happ(s(fun(X619,X626),'_dst_x0x1x1_3'(s(fun(fun(X616,fun(X617,X618)),fun(X627,X626)),X628),s(fun(X619,fun(X616,fun(X617,X618))),'_dst_x000x111x001_3'(s(fun(X620,fun(X621,X618)),X622),s(fun(X619,fun(X616,fun(X617,X620))),X623),s(fun(X617,X621),X624))),s(fun(X619,X627),X629))),s(X619,X625))) = s(X626,happ(s(fun(X627,X626),happ(s(fun(fun(X616,fun(X617,X618)),fun(X627,X626)),X628),s(fun(X616,fun(X617,X618)),'_dst_x00x11x01_3'(s(fun(X620,fun(X621,X618)),X622),s(fun(X616,fun(X617,X620)),happ(s(fun(X619,fun(X616,fun(X617,X620))),X623),s(X619,X625))),s(fun(X617,X621),X624))))),s(X627,happ(s(fun(X619,X627),X629),s(X619,X625)))))) )),
  inference(superposition,[],[f196,f203])).
fof(f11294,plain,(
  ( ! [X606,X610,X602,X613,X615,X609,X605,X607,X611,X603,X612,X614,X608,X604] : (s(X612,happ(s(fun(X605,X612),'_dst_x0x1x1_3'(s(fun(X613,fun(fun(X602,fun(X603,X604)),X612)),X614),s(fun(X605,X613),X615),s(fun(X605,fun(X602,fun(X603,X604))),'_dst_x000x111x001_3'(s(fun(X606,fun(X607,X604)),X608),s(fun(X605,fun(X602,fun(X603,X606))),X609),s(fun(X603,X607),X610))))),s(X605,X611))) = s(X612,happ(s(fun(fun(X602,fun(X603,X604)),X612),happ(s(fun(X613,fun(fun(X602,fun(X603,X604)),X612)),X614),s(X613,happ(s(fun(X605,X613),X615),s(X605,X611))))),s(fun(X602,fun(X603,X604)),'_dst_x00x11x01_3'(s(fun(X606,fun(X607,X604)),X608),s(fun(X602,fun(X603,X606)),happ(s(fun(X605,fun(X602,fun(X603,X606))),X609),s(X605,X611))),s(fun(X603,X607),X610)))))) )),
  inference(superposition,[],[f196,f203])).
fof(f11293,plain,(
  ( ! [X600,X596,X598,X592,X594,X588,X590,X601,X597,X599,X593,X595,X589,X591] : (s(X598,happ(s(fun(X591,X598),'_dst_x0x1x0_3'(s(fun(fun(X588,fun(X589,X590)),fun(X599,X598)),X600),s(fun(X591,fun(X588,fun(X589,X590))),'_dst_x000x111x001_3'(s(fun(X592,fun(X593,X590)),X594),s(fun(X591,fun(X588,fun(X589,X592))),X595),s(fun(X589,X593),X596))),s(X599,X601))),s(X591,X597))) = s(X598,happ(s(fun(X599,X598),happ(s(fun(fun(X588,fun(X589,X590)),fun(X599,X598)),X600),s(fun(X588,fun(X589,X590)),'_dst_x00x11x01_3'(s(fun(X592,fun(X593,X590)),X594),s(fun(X588,fun(X589,X592)),happ(s(fun(X591,fun(X588,fun(X589,X592))),X595),s(X591,X597))),s(fun(X589,X593),X596))))),s(X599,X601)))) )),
  inference(superposition,[],[f195,f203])).
fof(f11292,plain,(
  ( ! [X575,X584,X586,X580,X582,X576,X578,X574,X585,X587,X581,X583,X577,X579] : (s(X584,happ(s(fun(X577,X584),'_dst_x0x0x1_3'(s(fun(X585,fun(fun(X574,fun(X575,X576)),X584)),X586),s(X585,X587),s(fun(X577,fun(X574,fun(X575,X576))),'_dst_x000x111x001_3'(s(fun(X578,fun(X579,X576)),X580),s(fun(X577,fun(X574,fun(X575,X578))),X581),s(fun(X575,X579),X582))))),s(X577,X583))) = s(X584,happ(s(fun(fun(X574,fun(X575,X576)),X584),happ(s(fun(X585,fun(fun(X574,fun(X575,X576)),X584)),X586),s(X585,X587))),s(fun(X574,fun(X575,X576)),'_dst_x00x11x01_3'(s(fun(X578,fun(X579,X576)),X580),s(fun(X574,fun(X575,X578)),happ(s(fun(X577,fun(X574,fun(X575,X578))),X581),s(X577,X583))),s(fun(X575,X579),X582)))))) )),
  inference(superposition,[],[f194,f203])).
fof(f11291,plain,(
  ( ! [X573,X569,X571,X565,X567,X563,X572,X568,X570,X564,X566,X562] : (s(X572,happ(s(fun(X565,X572),'_dst_x0x1_2'(s(fun(fun(X562,fun(X563,X564)),X572),X573),s(fun(X565,fun(X562,fun(X563,X564))),'_dst_x000x111x001_3'(s(fun(X566,fun(X567,X564)),X568),s(fun(X565,fun(X562,fun(X563,X566))),X569),s(fun(X563,X567),X570))))),s(X565,X571))) = s(X572,happ(s(fun(fun(X562,fun(X563,X564)),X572),X573),s(fun(X562,fun(X563,X564)),'_dst_x00x11x01_3'(s(fun(X566,fun(X567,X564)),X568),s(fun(X562,fun(X563,X566)),happ(s(fun(X565,fun(X562,fun(X563,X566))),X569),s(X565,X571))),s(fun(X563,X567),X570)))))) )),
  inference(superposition,[],[f182,f203])).
fof(f11259,plain,(
  ( ! [X103,X101,X111,X109,X107,X105,X112,X102,X100,X110,X108,X106,X104] : (s(fun(X101,X102),happ(s(fun(X109,fun(X101,X102)),'_dst_x0x1x0_3'(s(fun(X103,fun(X100,fun(X101,X102))),'_dst_x000x111x001_3'(s(fun(X104,fun(X105,X102)),X106),s(fun(X103,fun(X100,fun(X101,X104))),X107),s(fun(X101,X105),X108))),s(fun(X109,X103),X110),s(X100,X112))),s(X109,X111))) = s(fun(X101,X102),happ(s(fun(X100,fun(X101,X102)),'_dst_x00x11x01_3'(s(fun(X104,fun(X105,X102)),X106),s(fun(X100,fun(X101,X104)),happ(s(fun(X103,fun(X100,fun(X101,X104))),X107),s(X103,happ(s(fun(X109,X103),X110),s(X109,X111))))),s(fun(X101,X105),X108))),s(X100,X112)))) )),
  inference(superposition,[],[f195,f203])).
fof(f11258,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X87,X95,X93,X91,X89,X98,X96] : (s(fun(X88,X89),happ(s(fun(X96,fun(X88,X89)),'_dst_x0x1x1_3'(s(fun(X90,fun(X87,fun(X88,X89))),'_dst_x000x111x001_3'(s(fun(X91,fun(X92,X89)),X93),s(fun(X90,fun(X87,fun(X88,X91))),X94),s(fun(X88,X92),X95))),s(fun(X96,X90),X97),s(fun(X96,X87),X99))),s(X96,X98))) = s(fun(X88,X89),happ(s(fun(X87,fun(X88,X89)),'_dst_x00x11x01_3'(s(fun(X91,fun(X92,X89)),X93),s(fun(X87,fun(X88,X91)),happ(s(fun(X90,fun(X87,fun(X88,X91))),X94),s(X90,happ(s(fun(X96,X90),X97),s(X96,X98))))),s(fun(X88,X92),X95))),s(X87,happ(s(fun(X96,X87),X99),s(X96,X98)))))) )),
  inference(superposition,[],[f196,f203])).
fof(f11257,plain,(
  ( ! [X80,X78,X76,X85,X83,X81,X79,X77,X86,X84,X82] : (s(fun(X77,X78),happ(s(fun(X76,fun(X77,X78)),'_dst_x00x11x01_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X76,fun(X77,X80)),happ(s(fun(X79,fun(X76,fun(X77,X80))),X83),s(X79,X85))),s(fun(X77,X81),X84))),s(X76,X86))) = s(fun(X77,X78),'const.pair.UNCURRY_2'(s(fun(X79,fun(X76,fun(X77,X78))),'_dst_x000x111x001_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X79,fun(X76,fun(X77,X80))),X83),s(fun(X77,X81),X84))),s('type.pair.prod'(X79,X76),'const.pair.,_2'(s(X79,X85),s(X76,X86)))))) )),
  inference(superposition,[],[f183,f203])).
fof(f11256,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X65,X75,X73] : (s(fun(X66,X67),happ(s(fun(X65,fun(X66,X67)),'_dst_x00x11x01_3'(s(fun(X69,fun(X70,X67)),X71),s(fun(X65,fun(X66,X69)),happ(s(fun(X68,fun(X65,fun(X66,X69))),X72),s(X68,X74))),s(fun(X66,X70),X73))),s(X65,X75))) = s(fun(X66,X67),'const.pair.pair_CASE_2'(s('type.pair.prod'(X68,X65),'const.pair.,_2'(s(X68,X74),s(X65,X75))),s(fun(X68,fun(X65,fun(X66,X67))),'_dst_x000x111x001_3'(s(fun(X69,fun(X70,X67)),X71),s(fun(X68,fun(X65,fun(X66,X69))),X72),s(fun(X66,X70),X73)))))) )),
  inference(superposition,[],[f184,f203])).
fof(f11255,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X53,X63] : (s(fun(X53,X63),'_dst_x0x1_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X53,fun(X54,X55)),'_dst_x00x11x01_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X53,fun(X54,X57)),happ(s(fun(X56,fun(X53,fun(X54,X57))),X60),s(X56,X62))),s(fun(X54,X58),X61))))) = s(fun(X53,X63),happ(s(fun(X56,fun(X53,X63)),'_dst_x00x11_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X56,fun(X53,fun(X54,X55))),'_dst_x000x111x001_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X56,fun(X53,fun(X54,X57))),X60),s(fun(X54,X58),X61))))),s(X56,X62)))) )),
  inference(superposition,[],[f193,f203])).
fof(f11254,plain,(
  ( ! [X47,X45,X43,X41,X52,X50,X48,X46,X44,X42,X40,X51,X49] : (s(fun(X41,X42),happ(s(fun(X50,fun(X41,X42)),'_dst_x0x0x1_3'(s(fun(X43,fun(X40,fun(X41,X42))),'_dst_x000x111x001_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X43,fun(X40,fun(X41,X44))),X47),s(fun(X41,X45),X48))),s(X43,X49),s(fun(X50,X40),X51))),s(X50,X52))) = s(fun(X41,X42),happ(s(fun(X40,fun(X41,X42)),'_dst_x00x11x01_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X40,fun(X41,X44)),happ(s(fun(X43,fun(X40,fun(X41,X44))),X47),s(X43,X49))),s(fun(X41,X45),X48))),s(X40,happ(s(fun(X50,X40),X51),s(X50,X52)))))) )),
  inference(superposition,[],[f194,f203])).
fof(f11253,plain,(
  ( ! [X30,X28,X26,X39,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s(fun(X26,X36),'_dst_x0x0x1_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X26,fun(X27,X28)),'_dst_x00x11x01_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X26,fun(X27,X30)),happ(s(fun(X29,fun(X26,fun(X27,X30))),X33),s(X29,X35))),s(fun(X27,X31),X34))))) = s(fun(X26,X36),happ(s(fun(X29,fun(X26,X36)),'_dst_x00x00x11_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X29,fun(X26,fun(X27,X28))),'_dst_x000x111x001_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X29,fun(X26,fun(X27,X30))),X33),s(fun(X27,X31),X34))))),s(X29,X35)))) )),
  inference(superposition,[],[f199,f203])).
fof(f11252,plain,(
  ( ! [X24,X14,X12,X23,X21,X19,X17,X25,X15,X13,X22,X20,X18,X16] : (s(fun(X12,X22),'_dst_x0x1x1_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X12,fun(X13,X14)),'_dst_x00x11x01_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X12,fun(X13,X16)),happ(s(fun(X15,fun(X12,fun(X13,X16))),X19),s(X15,X21))),s(fun(X13,X17),X20))),s(fun(X12,X23),X25))) = s(fun(X12,X22),happ(s(fun(X15,fun(X12,X22)),'_dst_x00x11x01_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X15,fun(X12,fun(X13,X14))),'_dst_x000x111x001_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X15,fun(X12,fun(X13,X16))),X19),s(fun(X13,X17),X20))),s(fun(X12,X23),X25))),s(X15,X21)))) )),
  inference(superposition,[],[f200,f203])).
fof(f11251,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,fun(X1,X10)),'_dst_x00x11_2'(s(fun(X2,X10),X11),s(fun(X0,fun(X1,X2)),'_dst_x00x11x01_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X0,fun(X1,X4)),happ(s(fun(X3,fun(X0,fun(X1,X4))),X7),s(X3,X9))),s(fun(X1,X5),X8))))) = s(fun(X0,fun(X1,X10)),happ(s(fun(X3,fun(X0,fun(X1,X10))),'_dst_x000x111_2'(s(fun(X2,X10),X11),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x111x001_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X3,fun(X0,fun(X1,X4))),X7),s(fun(X1,X5),X8))))),s(X3,X9)))) )),
  inference(superposition,[],[f197,f203])).
fof(f11222,plain,(
  ( ! [X536,X532,X534,X528,X530,X537,X533,X535,X529,X531,X527] : (p(s(bool,happ(s(fun(fun(X527,fun(X528,X529)),bool),happ(s(fun(fun(X527,fun(X528,X529)),fun(fun(X527,fun(X528,X529)),bool)),'=_0'),s(fun(X527,fun(X528,X529)),X537))),s(fun(X527,fun(X528,X529)),happ(s(fun(X533,fun(X527,fun(X528,X529))),'_dst_x000x111x001_3'(s(fun(X530,fun(X531,X529)),X532),s(fun(X533,fun(X527,fun(X528,X530))),X534),s(fun(X528,X531),X536))),s(X533,X535)))))) | s(fun(X527,fun(X528,X529)),X537) != s(fun(X527,fun(X528,X529)),'_dst_x00x11x01_3'(s(fun(X530,fun(X531,X529)),X532),s(fun(X527,fun(X528,X530)),happ(s(fun(X533,fun(X527,fun(X528,X530))),X534),s(X533,X535))),s(fun(X528,X531),X536)))) )),
  inference(superposition,[],[f173,f203])).
fof(f11221,plain,(
  ( ! [X524,X526,X520,X522,X516,X518,X525,X521,X523,X517,X519] : (p(s(bool,happ(s(fun(fun(X516,fun(X517,X518)),bool),happ(s(fun(fun(X516,fun(X517,X518)),fun(fun(X516,fun(X517,X518)),bool)),'=_0'),s(fun(X516,fun(X517,X518)),happ(s(fun(X522,fun(X516,fun(X517,X518))),'_dst_x000x111x001_3'(s(fun(X519,fun(X520,X518)),X521),s(fun(X522,fun(X516,fun(X517,X519))),X523),s(fun(X517,X520),X525))),s(X522,X524))))),s(fun(X516,fun(X517,X518)),X526)))) | s(fun(X516,fun(X517,X518)),X526) != s(fun(X516,fun(X517,X518)),'_dst_x00x11x01_3'(s(fun(X519,fun(X520,X518)),X521),s(fun(X516,fun(X517,X519)),happ(s(fun(X522,fun(X516,fun(X517,X519))),X523),s(X522,X524))),s(fun(X517,X520),X525)))) )),
  inference(superposition,[],[f173,f203])).
fof(f11220,plain,(
  ( ! [X509,X511,X505,X507,X512,X514,X508,X510,X506,X513,X515] : (~p(s(bool,happ(s(fun(fun(X505,fun(X506,X507)),bool),happ(s(fun(fun(X505,fun(X506,X507)),fun(fun(X505,fun(X506,X507)),bool)),'=_0'),s(fun(X505,fun(X506,X507)),X515))),s(fun(X505,fun(X506,X507)),happ(s(fun(X511,fun(X505,fun(X506,X507))),'_dst_x000x111x001_3'(s(fun(X508,fun(X509,X507)),X510),s(fun(X511,fun(X505,fun(X506,X508))),X512),s(fun(X506,X509),X514))),s(X511,X513)))))) | s(fun(X505,fun(X506,X507)),X515) = s(fun(X505,fun(X506,X507)),'_dst_x00x11x01_3'(s(fun(X508,fun(X509,X507)),X510),s(fun(X505,fun(X506,X508)),happ(s(fun(X511,fun(X505,fun(X506,X508))),X512),s(X511,X513))),s(fun(X506,X509),X514)))) )),
  inference(superposition,[],[f172,f203])).
fof(f11219,plain,(
  ( ! [X499,X494,X500,X502,X496,X498,X504,X495,X501,X503,X497] : (~p(s(bool,happ(s(fun(fun(X494,fun(X495,X496)),bool),happ(s(fun(fun(X494,fun(X495,X496)),fun(fun(X494,fun(X495,X496)),bool)),'=_0'),s(fun(X494,fun(X495,X496)),happ(s(fun(X500,fun(X494,fun(X495,X496))),'_dst_x000x111x001_3'(s(fun(X497,fun(X498,X496)),X499),s(fun(X500,fun(X494,fun(X495,X497))),X501),s(fun(X495,X498),X503))),s(X500,X502))))),s(fun(X494,fun(X495,X496)),X504)))) | s(fun(X494,fun(X495,X496)),X504) = s(fun(X494,fun(X495,X496)),'_dst_x00x11x01_3'(s(fun(X497,fun(X498,X496)),X499),s(fun(X494,fun(X495,X497)),happ(s(fun(X500,fun(X494,fun(X495,X497))),X501),s(X500,X502))),s(fun(X495,X498),X503)))) )),
  inference(superposition,[],[f172,f203])).
fof(f11184,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,X2),'_dst_x0x1x1_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X1,X3),happ(s(fun(X0,fun(X1,X3)),happ(s(fun(X6,fun(X0,fun(X1,X3))),X7),s(X6,X8))),s(X0,X10))),s(fun(X1,X4),X9))) = s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),happ(s(fun(X6,fun(X0,fun(X1,X2))),'_dst_x000x111x001_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X6,fun(X0,fun(X1,X3))),X7),s(fun(X1,X4),X9))),s(X6,X8))),s(X0,X10)))) )),
  inference(superposition,[],[f200,f203])).
fof(f11183,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (s(fun(X10,fun(X11,X12)),'_dst_x00x11x01_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X10,fun(X11,X14)),happ(s(fun(X13,fun(X10,fun(X11,X14))),X17),s(X13,happ(s(fun(X19,X13),X20),s(X19,X21))))),s(fun(X11,X15),X18))) = s(fun(X10,fun(X11,X12)),happ(s(fun(X19,fun(X10,fun(X11,X12))),'_dst_x0x1_2'(s(fun(X13,fun(X10,fun(X11,X12))),'_dst_x000x111x001_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X13,fun(X10,fun(X11,X14))),X17),s(fun(X11,X15),X18))),s(fun(X19,X13),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f203,f182])).
fof(f11033,plain,(
  ( ! [X269,X271,X267,X277,X273,X275,X268,X270,X276,X272,X274] : (s(fun(X267,fun(X268,X274)),'_dst_x00x11x01_3'(s(fun(X269,fun(X275,X274)),X276),s(fun(X267,fun(X268,X269)),'const.pair.UNCURRY_2'(s(fun(X270,fun(X271,fun(X267,fun(X268,X269)))),X272),s('type.pair.prod'(X270,X271),X273))),s(fun(X268,X275),X277))) = s(fun(X267,fun(X268,X274)),happ(s(fun('type.pair.prod'(X270,X271),fun(X267,fun(X268,X274))),'_dst_x000x111x001_3'(s(fun(X269,fun(X275,X274)),X276),s(fun('type.pair.prod'(X270,X271),fun(X267,fun(X268,X269))),happ(s(fun(fun(X270,fun(X271,fun(X267,fun(X268,X269)))),fun('type.pair.prod'(X270,X271),fun(X267,fun(X268,X269)))),'const.pair.UNCURRY_0'),s(fun(X270,fun(X271,fun(X267,fun(X268,X269)))),X272))),s(fun(X268,X275),X277))),s('type.pair.prod'(X270,X271),X273)))) )),
  inference(superposition,[],[f203,f180])).
fof(f11032,plain,(
  ( ! [X243,X255,X253,X251,X249,X246,X244,X242,X254,X252,X250,X248,X247,X245] : (s(fun(X242,fun(X243,X252)),happ(s(fun(X245,fun(X242,fun(X243,X252))),'_dst_x000x111x001_3'(s(fun(X244,fun(X253,X252)),X254),s(fun(X245,fun(X242,fun(X243,X244))),'_dst_x0x0x1_3'(s(fun(X246,fun(X247,fun(X242,fun(X243,X244)))),X248),s(X246,X249),s(fun(X245,X247),X250))),s(fun(X243,X253),X255))),s(X245,X251))) = s(fun(X242,fun(X243,X252)),'_dst_x00x11x01_3'(s(fun(X244,fun(X253,X252)),X254),s(fun(X242,fun(X243,X244)),happ(s(fun(X247,fun(X242,fun(X243,X244))),happ(s(fun(X246,fun(X247,fun(X242,fun(X243,X244)))),X248),s(X246,X249))),s(X247,happ(s(fun(X245,X247),X250),s(X245,X251))))),s(fun(X243,X253),X255)))) )),
  inference(superposition,[],[f203,f194])).
fof(f11031,plain,(
  ( ! [X241,X239,X237,X235,X233,X240,X238,X236,X234,X232] : (s(fun(X232,fun(X233,X238)),'_dst_x00x11x01_3'(s(fun(X234,fun(X239,X238)),X240),s(fun(X232,fun(X233,X234)),X236),s(fun(X233,X239),X241))) = s(fun(X232,fun(X233,X238)),happ(s(fun(X235,fun(X232,fun(X233,X238))),'_dst_x000x111x001_3'(s(fun(X234,fun(X239,X238)),X240),s(fun(X235,fun(X232,fun(X233,X234))),'_dst_x0_1'(s(fun(X232,fun(X233,X234)),X236))),s(fun(X233,X239),X241))),s(X235,X237)))) )),
  inference(superposition,[],[f203,f170])).
fof(f11030,plain,(
  ( ! [X218,X222,X220,X231,X229,X227,X225,X219,X223,X221,X230,X228,X226,X224] : (s(fun(X218,fun(X219,X228)),happ(s(fun(X221,fun(X218,fun(X219,X228))),'_dst_x000x111x001_3'(s(fun(X220,fun(X229,X228)),X230),s(fun(X221,fun(X218,fun(X219,X220))),'_dst_x0x1x0_3'(s(fun(X222,fun(X223,fun(X218,fun(X219,X220)))),X224),s(fun(X221,X222),X225),s(X223,X226))),s(fun(X219,X229),X231))),s(X221,X227))) = s(fun(X218,fun(X219,X228)),'_dst_x00x11x01_3'(s(fun(X220,fun(X229,X228)),X230),s(fun(X218,fun(X219,X220)),happ(s(fun(X223,fun(X218,fun(X219,X220))),happ(s(fun(X222,fun(X223,fun(X218,fun(X219,X220)))),X224),s(X222,happ(s(fun(X221,X222),X225),s(X221,X227))))),s(X223,X226))),s(fun(X219,X229),X231)))) )),
  inference(superposition,[],[f203,f195])).
fof(f11029,plain,(
  ( ! [X212,X210,X216,X208,X206,X204,X215,X213,X211,X209,X217,X207,X205,X214] : (s(fun(X204,fun(X205,X214)),'_dst_x00x11x01_3'(s(fun(X206,fun(X215,X214)),X216),s(fun(X204,fun(X205,X206)),happ(s(fun(X210,fun(X204,fun(X205,X206))),'_dst_x0x1x1_3'(s(fun(X208,fun(X207,fun(X204,fun(X205,X206)))),X209),s(fun(X210,X208),X211),s(fun(X210,X207),X213))),s(X210,X212))),s(fun(X205,X215),X217))) = s(fun(X204,fun(X205,X214)),happ(s(fun(X207,fun(X204,fun(X205,X214))),'_dst_x000x111x001_3'(s(fun(X206,fun(X215,X214)),X216),s(fun(X207,fun(X204,fun(X205,X206))),happ(s(fun(X208,fun(X207,fun(X204,fun(X205,X206)))),X209),s(X208,happ(s(fun(X210,X208),X211),s(X210,X212))))),s(fun(X205,X215),X217))),s(X207,happ(s(fun(X210,X207),X213),s(X210,X212)))))) )),
  inference(superposition,[],[f203,f196])).
fof(f11028,plain,(
  ( ! [X191,X198,X196,X194,X192,X202,X200,X190,X199,X197,X195,X193,X203,X201] : (s(fun(X190,fun(X191,X200)),'_dst_x00x11x01_3'(s(fun(X192,fun(X201,X200)),X202),s(fun(X190,fun(X191,X192)),happ(s(fun(X196,fun(X190,fun(X191,X192))),'_dst_x0x1x0_3'(s(fun(X194,fun(X193,fun(X190,fun(X191,X192)))),X195),s(fun(X196,X194),X197),s(X193,X199))),s(X196,X198))),s(fun(X191,X201),X203))) = s(fun(X190,fun(X191,X200)),happ(s(fun(X193,fun(X190,fun(X191,X200))),'_dst_x000x111x001_3'(s(fun(X192,fun(X201,X200)),X202),s(fun(X193,fun(X190,fun(X191,X192))),happ(s(fun(X194,fun(X193,fun(X190,fun(X191,X192)))),X195),s(X194,happ(s(fun(X196,X194),X197),s(X196,X198))))),s(fun(X191,X201),X203))),s(X193,X199)))) )),
  inference(superposition,[],[f203,f195])).
fof(f11027,plain,(
  ( ! [X177,X189,X187,X185,X182,X180,X178,X176,X188,X186,X184,X183,X181,X179] : (s(fun(X176,fun(X177,X186)),'_dst_x00x11x01_3'(s(fun(X178,fun(X187,X186)),X188),s(fun(X176,fun(X177,X178)),happ(s(fun(X183,fun(X176,fun(X177,X178))),'_dst_x0x0x1_3'(s(fun(X180,fun(X179,fun(X176,fun(X177,X178)))),X181),s(X180,X182),s(fun(X183,X179),X184))),s(X183,X185))),s(fun(X177,X187),X189))) = s(fun(X176,fun(X177,X186)),happ(s(fun(X179,fun(X176,fun(X177,X186))),'_dst_x000x111x001_3'(s(fun(X178,fun(X187,X186)),X188),s(fun(X179,fun(X176,fun(X177,X178))),happ(s(fun(X180,fun(X179,fun(X176,fun(X177,X178)))),X181),s(X180,X182))),s(fun(X177,X187),X189))),s(X179,happ(s(fun(X183,X179),X184),s(X183,X185)))))) )),
  inference(superposition,[],[f203,f194])).
fof(f11025,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X162,X160] : (s(fun(X152,fun(X153,X160)),'_dst_x00x11x01_3'(s(fun(X154,fun(X161,X160)),X162),s(fun(X152,fun(X153,X154)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X156,X155),'const.pair.,_2'(s(X156,X158),s(X155,X159))),s(fun(X156,fun(X155,fun(X152,fun(X153,X154)))),X157))),s(fun(X153,X161),X163))) = s(fun(X152,fun(X153,X160)),happ(s(fun(X155,fun(X152,fun(X153,X160))),'_dst_x000x111x001_3'(s(fun(X154,fun(X161,X160)),X162),s(fun(X155,fun(X152,fun(X153,X154))),happ(s(fun(X156,fun(X155,fun(X152,fun(X153,X154)))),X157),s(X156,X158))),s(fun(X153,X161),X163))),s(X155,X159)))) )),
  inference(superposition,[],[f203,f184])).
fof(f11024,plain,(
  ( ! [X146,X144,X142,X140,X151,X149,X147,X145,X143,X141,X150,X148] : (s(fun(X140,fun(X141,X148)),'_dst_x00x11x01_3'(s(fun(X142,fun(X149,X148)),X150),s(fun(X140,fun(X141,X142)),X145),s(fun(X141,X149),X151))) = s(fun(X140,fun(X141,X148)),happ(s(fun(X143,fun(X140,fun(X141,X148))),'_dst_x000x111x001_3'(s(fun(X142,fun(X149,X148)),X150),s(fun(X143,fun(X140,fun(X141,X142))),'_dst_x0x1_2'(s(fun(X144,fun(X140,fun(X141,X142))),'_dst_x0_1'(s(fun(X140,fun(X141,X142)),X145))),s(fun(X143,X144),X146))),s(fun(X141,X149),X151))),s(X143,X147)))) )),
  inference(superposition,[],[f203,f1049])).
fof(f11023,plain,(
  ( ! [X134,X132,X130,X128,X138,X136,X135,X133,X131,X129,X139,X137] : (s(fun(X128,fun(X129,X136)),'_dst_x00x11x01_3'(s(fun(X130,fun(X137,X136)),X138),s(fun(X128,fun(X129,X130)),happ(s(fun(X132,fun(X128,fun(X129,X130))),X133),s(X132,happ(s(fun(X131,X132),X134),s(X131,X135))))),s(fun(X129,X137),X139))) = s(fun(X128,fun(X129,X136)),happ(s(fun(X131,fun(X128,fun(X129,X136))),'_dst_x000x111x001_3'(s(fun(X130,fun(X137,X136)),X138),s(fun(X131,fun(X128,fun(X129,X130))),'_dst_x0x1_2'(s(fun(X132,fun(X128,fun(X129,X130))),X133),s(fun(X131,X132),X134))),s(fun(X129,X137),X139))),s(X131,X135)))) )),
  inference(superposition,[],[f203,f182])).
fof(f11022,plain,(
  ( ! [X127,X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (s(fun(X116,fun(X117,X124)),happ(s(fun(X119,fun(X116,fun(X117,X124))),'_dst_x000x111x001_3'(s(fun(X118,fun(X125,X124)),X126),s(fun(X119,fun(X116,fun(X117,X118))),X120),s(fun(X117,X125),X127))),s(X119,happ(s(fun(X121,X119),X122),s(X121,X123))))) = s(fun(X116,fun(X117,X124)),'_dst_x00x11x01_3'(s(fun(X118,fun(X125,X124)),X126),s(fun(X116,fun(X117,X118)),happ(s(fun(X121,fun(X116,fun(X117,X118))),'_dst_x0x1_2'(s(fun(X119,fun(X116,fun(X117,X118))),X120),s(fun(X121,X119),X122))),s(X121,X123))),s(fun(X117,X125),X127)))) )),
  inference(superposition,[],[f203,f182])).
fof(f11021,plain,(
  ( ! [X94,X103,X101,X99,X97,X107,X105,X95,X102,X100,X98,X96,X106,X104] : (s(fun(X94,fun(X95,X104)),'_dst_x00x11x01_3'(s(fun(X96,fun(X105,X104)),X106),s(fun(X94,fun(X95,X96)),'_dst_x0x0x1_3'(s(fun(X98,fun(X99,fun(X95,X96))),X100),s(X98,X101),s(fun(X94,X99),happ(s(fun(X97,fun(X94,X99)),X102),s(X97,X103))))),s(fun(X95,X105),X107))) = s(fun(X94,fun(X95,X104)),happ(s(fun(X97,fun(X94,fun(X95,X104))),'_dst_x000x111x001_3'(s(fun(X96,fun(X105,X104)),X106),s(fun(X97,fun(X94,fun(X95,X96))),'_dst_x00x00x11_3'(s(fun(X98,fun(X99,fun(X95,X96))),X100),s(X98,X101),s(fun(X97,fun(X94,X99)),X102))),s(fun(X95,X105),X107))),s(X97,X103)))) )),
  inference(superposition,[],[f203,f199])).
fof(f11020,plain,(
  ( ! [X80,X92,X90,X88,X87,X85,X83,X81,X93,X91,X89,X86,X84,X82] : (s(fun(X80,fun(X81,X90)),'_dst_x00x11x01_3'(s(fun(X82,fun(X91,X90)),X92),s(fun(X80,fun(X81,X82)),'_dst_x0x1x0_3'(s(fun(X84,fun(X85,fun(X81,X82))),X86),s(fun(X80,X84),X87),s(X85,X88))),s(fun(X81,X91),X93))) = s(fun(X80,fun(X81,X90)),happ(s(fun(X83,fun(X80,fun(X81,X90))),'_dst_x000x111x001_3'(s(fun(X82,fun(X91,X90)),X92),s(fun(X83,fun(X80,fun(X81,X82))),'_dst_x00x01x00_3'(s(fun(X84,fun(X85,fun(X81,X82))),X86),s(fun(X80,X84),X87),s(X85,X88))),s(fun(X81,X91),X93))),s(X83,X89)))) )),
  inference(superposition,[],[f203,f198])).
fof(f11019,plain,(
  ( ! [X70,X68,X78,X76,X74,X72,X71,X69,X79,X77,X75,X73] : (s(fun(X68,fun(X69,X76)),'_dst_x00x11x01_3'(s(fun(X70,fun(X77,X76)),X78),s(fun(X68,fun(X69,X70)),'_dst_x0x1_2'(s(fun(X72,fun(X69,X70)),X73),s(fun(X68,X72),happ(s(fun(X71,fun(X68,X72)),X74),s(X71,X75))))),s(fun(X69,X77),X79))) = s(fun(X68,fun(X69,X76)),happ(s(fun(X71,fun(X68,fun(X69,X76))),'_dst_x000x111x001_3'(s(fun(X70,fun(X77,X76)),X78),s(fun(X71,fun(X68,fun(X69,X70))),'_dst_x00x11_2'(s(fun(X72,fun(X69,X70)),X73),s(fun(X71,fun(X68,X72)),X74))),s(fun(X69,X77),X79))),s(X71,X75)))) )),
  inference(superposition,[],[f203,f193])).
fof(f11018,plain,(
  ( ! [X61,X59,X57,X66,X64,X54,X62,X60,X58,X56,X67,X65,X55,X63] : (s(fun(X54,fun(X55,X64)),'_dst_x00x11x01_3'(s(fun(X56,fun(X65,X64)),X66),s(fun(X54,fun(X55,X56)),'_dst_x0x1x1_3'(s(fun(X58,fun(X59,fun(X55,X56))),X60),s(fun(X54,X58),happ(s(fun(X57,fun(X54,X58)),X61),s(X57,X63))),s(fun(X54,X59),X62))),s(fun(X55,X65),X67))) = s(fun(X54,fun(X55,X64)),happ(s(fun(X57,fun(X54,fun(X55,X64))),'_dst_x000x111x001_3'(s(fun(X56,fun(X65,X64)),X66),s(fun(X57,fun(X54,fun(X55,X56))),'_dst_x00x11x01_3'(s(fun(X58,fun(X59,fun(X55,X56))),X60),s(fun(X57,fun(X54,X58)),X61),s(fun(X54,X59),X62))),s(fun(X55,X65),X67))),s(X57,X63)))) )),
  inference(superposition,[],[f203,f200])).
fof(f11017,plain,(
  ( ! [X47,X45,X43,X41,X52,X50,X48,X46,X44,X42,X40,X53,X51,X49] : (s(fun(X40,fun(X41,X50)),'_dst_x00x11x01_3'(s(fun(X42,fun(X51,X50)),X52),s(fun(X40,fun(X41,X42)),'_dst_x00x11x01_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X40,fun(X41,X44)),happ(s(fun(X43,fun(X40,fun(X41,X44))),X47),s(X43,X49))),s(fun(X41,X45),X48))),s(fun(X41,X51),X53))) = s(fun(X40,fun(X41,X50)),happ(s(fun(X43,fun(X40,fun(X41,X50))),'_dst_x000x111x001_3'(s(fun(X42,fun(X51,X50)),X52),s(fun(X43,fun(X40,fun(X41,X42))),'_dst_x000x111x001_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X43,fun(X40,fun(X41,X44))),X47),s(fun(X41,X45),X48))),s(fun(X41,X51),X53))),s(X43,X49)))) )),
  inference(superposition,[],[f203,f203])).
fof(f11016,plain,(
  ( ! [X30,X28,X26,X39,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s(fun(X26,fun(X27,X36)),'_dst_x00x11x01_3'(s(fun(X28,fun(X37,X36)),X38),s(fun(X26,fun(X27,X28)),'_dst_x00x01x00_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X27,X30),X33),s(X31,happ(s(fun(X29,X31),X34),s(X29,X35))))),s(fun(X27,X37),X39))) = s(fun(X26,fun(X27,X36)),happ(s(fun(X29,fun(X26,fun(X27,X36))),'_dst_x000x111x001_3'(s(fun(X28,fun(X37,X36)),X38),s(fun(X29,fun(X26,fun(X27,X28))),'_dst_x000x001x100_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X27,X30),X33),s(fun(X29,X31),X34))),s(fun(X27,X37),X39))),s(X29,X35)))) )),
  inference(superposition,[],[f203,f201])).
fof(f11015,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s(fun(X14,fun(X15,X22)),'_dst_x00x11x01_3'(s(fun(X16,fun(X23,X22)),X24),s(fun(X14,fun(X15,X16)),'_dst_x00x11_2'(s(fun(X18,X16),X19),s(fun(X14,fun(X15,X18)),happ(s(fun(X17,fun(X14,fun(X15,X18))),X20),s(X17,X21))))),s(fun(X15,X23),X25))) = s(fun(X14,fun(X15,X22)),happ(s(fun(X17,fun(X14,fun(X15,X22))),'_dst_x000x111x001_3'(s(fun(X16,fun(X23,X22)),X24),s(fun(X17,fun(X14,fun(X15,X16))),'_dst_x000x111_2'(s(fun(X18,X16),X19),s(fun(X17,fun(X14,fun(X15,X18))),X20))),s(fun(X15,X23),X25))),s(X17,X21)))) )),
  inference(superposition,[],[f203,f197])).
fof(f11014,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(fun(X0,fun(X1,X10)),'_dst_x00x11x01_3'(s(fun(X2,fun(X11,X10)),X12),s(fun(X0,fun(X1,X2)),'_dst_x00x00x11_3'(s(fun(X4,fun(X5,X2)),X6),s(X4,happ(s(fun(X3,X4),X7),s(X3,X9))),s(fun(X0,fun(X1,X5)),X8))),s(fun(X1,X11),X13))) = s(fun(X0,fun(X1,X10)),happ(s(fun(X3,fun(X0,fun(X1,X10))),'_dst_x000x111x001_3'(s(fun(X2,fun(X11,X10)),X12),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x100x011_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X3,X4),X7),s(fun(X0,fun(X1,X5)),X8))),s(fun(X1,X11),X13))),s(X3,X9)))) )),
  inference(superposition,[],[f203,f202])).
fof(f11012,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),'_dst_x00x11x01_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X0,fun(X1,X3)),happ(s(fun(X6,fun(X0,fun(X1,X3))),X7),s(X6,X8))),s(fun(X1,X4),X9))))),s(fun(X0,fun(X1,X2)),happ(s(fun(X6,fun(X0,fun(X1,X2))),'_dst_x000x111x001_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X6,fun(X0,fun(X1,X3))),X7),s(fun(X1,X4),X9))),s(X6,X8))))))) )),
  inference(unit_resulting_resolution,[],[f203,f173])).
fof(f11013,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),happ(s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x111x001_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X3,fun(X0,fun(X1,X4))),X7),s(fun(X1,X5),X8))),s(X3,X9))))),s(fun(X0,fun(X1,X2)),'_dst_x00x11x01_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X0,fun(X1,X4)),happ(s(fun(X3,fun(X0,fun(X1,X4))),X7),s(X3,X9))),s(fun(X1,X5),X8))))))) )),
  inference(unit_resulting_resolution,[],[f203,f173])).
fof(f9305,plain,(
  ( ! [X1011,X1009,X1015,X1013,X1010,X1008,X1014,X1012,X1007] : (p(s(bool,happ(s(fun('type.pair.prod'(X1007,X1008),bool),X1011),s('type.pair.prod'(X1007,X1008),'const.pair.,_2'(s(X1007,X1014),s(X1008,X1015)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X1007,X1008),bool),X1011),s('type.pair.prod'(X1007,X1008),sK7(happ(s(fun(X1009,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0x1_2'(s(fun(X1010,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X1007,X1008),bool),X1011))),s(fun(X1009,X1010),X1012))),s(X1009,X1013)),X1008,X1007)))))) )),
  inference(forward_demodulation,[],[f9304,f5177])).
fof(f9304,plain,(
  ( ! [X1011,X1009,X1015,X1013,X1010,X1008,X1014,X1012,X1007] : (p(s(bool,happ(s(fun(X1009,bool),'_dst_x0x1x0_3'(s(fun(X1010,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X1007,X1008),bool),X1011))),s(fun(X1009,X1010),X1012),s('type.pair.prod'(X1007,X1008),'const.pair.,_2'(s(X1007,X1014),s(X1008,X1015))))),s(X1009,X1013)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X1007,X1008),bool),X1011),s('type.pair.prod'(X1007,X1008),sK7(happ(s(fun(X1009,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0x1_2'(s(fun(X1010,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X1007,X1008),bool),X1011))),s(fun(X1009,X1010),X1012))),s(X1009,X1013)),X1008,X1007)))))) )),
  inference(forward_demodulation,[],[f9223,f5170])).
fof(f9223,plain,(
  ( ! [X1011,X1009,X1015,X1013,X1010,X1008,X1014,X1012,X1007] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1007,X1008),bool),X1011),s('type.pair.prod'(X1007,X1008),sK7(happ(s(fun(X1009,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0x1_2'(s(fun(X1010,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X1007,X1008),bool),X1011))),s(fun(X1009,X1010),X1012))),s(X1009,X1013)),X1008,X1007))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1007,X1008),bool),happ(s(fun(X1009,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0x1_2'(s(fun(X1010,fun('type.pair.prod'(X1007,X1008),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X1007,X1008),bool),X1011))),s(fun(X1009,X1010),X1012))),s(X1009,X1013))),s('type.pair.prod'(X1007,X1008),'const.pair.,_2'(s(X1007,X1014),s(X1008,X1015))))))) )),
  inference(superposition,[],[f168,f1049])).
fof(f9303,plain,(
  ( ! [X1002,X1000,X1006,X1004,X1003,X1001,X1005,X999] : (p(s(bool,happ(s(fun('type.pair.prod'(X999,X1000),bool),X1003),s('type.pair.prod'(X999,X1000),X1006)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X999,X1000),bool),X1003),s('type.pair.prod'(X999,X1000),'const.pair.,_2'(s(X999,sK8(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999)),s(X1000,sK9(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999)))))))) )),
  inference(forward_demodulation,[],[f9302,f5177])).
fof(f9302,plain,(
  ( ! [X1002,X1000,X1006,X1004,X1003,X1001,X1005,X999] : (p(s(bool,happ(s(fun(X1001,bool),'_dst_x0x1x0_3'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004),s('type.pair.prod'(X999,X1000),X1006))),s(X1001,X1005)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X999,X1000),bool),X1003),s('type.pair.prod'(X999,X1000),'const.pair.,_2'(s(X999,sK8(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999)),s(X1000,sK9(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999)))))))) )),
  inference(forward_demodulation,[],[f9222,f5170])).
fof(f9222,plain,(
  ( ! [X1002,X1000,X1006,X1004,X1003,X1001,X1005,X999] : (~p(s(bool,happ(s(fun('type.pair.prod'(X999,X1000),bool),X1003),s('type.pair.prod'(X999,X1000),'const.pair.,_2'(s(X999,sK8(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999)),s(X1000,sK9(happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005)),X1000,X999))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X999,X1000),bool),happ(s(fun(X1001,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0x1_2'(s(fun(X1002,fun('type.pair.prod'(X999,X1000),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X999,X1000),bool),X1003))),s(fun(X1001,X1002),X1004))),s(X1001,X1005))),s('type.pair.prod'(X999,X1000),X1006))))) )),
  inference(superposition,[],[f169,f1049])).
fof(f9191,plain,(
  ( ! [X633,X635,X629,X631,X627,X632,X634,X628,X630,X626] : (s(fun(X626,fun(X627,X634)),'_dst_x00x11_2'(s(fun(X628,X634),X635),s(fun(X626,fun(X627,X628)),X631))) = s(fun(X626,fun(X627,X634)),happ(s(fun(X629,fun(X626,fun(X627,X634))),'_dst_x000x111_2'(s(fun(X628,X634),X635),s(fun(X629,fun(X626,fun(X627,X628))),'_dst_x0x1_2'(s(fun(X630,fun(X626,fun(X627,X628))),'_dst_x0_1'(s(fun(X626,fun(X627,X628)),X631))),s(fun(X629,X630),X632))))),s(X629,X633)))) )),
  inference(superposition,[],[f197,f1049])).
fof(f9190,plain,(
  ( ! [X625,X621,X623,X617,X619,X624,X620,X622,X616,X618] : (s(X617,happ(s(fun(X616,X617),X620),s(X616,X625))) = s(X617,happ(s(fun(X622,X617),'_dst_x0x1x0_3'(s(fun(X618,fun(X616,X617)),'_dst_x0x1_2'(s(fun(X619,fun(X616,X617)),'_dst_x0_1'(s(fun(X616,X617),X620))),s(fun(X618,X619),X621))),s(fun(X622,X618),X623),s(X616,X625))),s(X622,X624)))) )),
  inference(superposition,[],[f195,f1049])).
fof(f9189,plain,(
  ( ! [X606,X613,X615,X609,X611,X607,X612,X614,X608,X610] : (s(X607,happ(s(fun(X606,X607),X610),s(X606,happ(s(fun(X612,X606),X615),s(X612,X614))))) = s(X607,happ(s(fun(X612,X607),'_dst_x0x1x1_3'(s(fun(X608,fun(X606,X607)),'_dst_x0x1_2'(s(fun(X609,fun(X606,X607)),'_dst_x0_1'(s(fun(X606,X607),X610))),s(fun(X608,X609),X611))),s(fun(X612,X608),X613),s(fun(X612,X606),X615))),s(X612,X614)))) )),
  inference(superposition,[],[f196,f1049])).
fof(f9188,plain,(
  ( ! [X600,X602,X598,X605,X601,X603,X599,X604] : (s(X599,happ(s(fun(X598,X599),X602),s(X598,X605))) = s(X599,'const.pair.UNCURRY_2'(s(fun(X600,fun(X598,X599)),'_dst_x0x1_2'(s(fun(X601,fun(X598,X599)),'_dst_x0_1'(s(fun(X598,X599),X602))),s(fun(X600,X601),X603))),s('type.pair.prod'(X600,X598),'const.pair.,_2'(s(X600,X604),s(X598,X605)))))) )),
  inference(superposition,[],[f183,f1049])).
fof(f9187,plain,(
  ( ! [X596,X592,X594,X590,X597,X593,X595,X591] : (s(X591,happ(s(fun(X590,X591),X594),s(X590,X597))) = s(X591,'const.pair.pair_CASE_2'(s('type.pair.prod'(X592,X590),'const.pair.,_2'(s(X592,X596),s(X590,X597))),s(fun(X592,fun(X590,X591)),'_dst_x0x1_2'(s(fun(X593,fun(X590,X591)),'_dst_x0_1'(s(fun(X590,X591),X594))),s(fun(X592,X593),X595)))))) )),
  inference(superposition,[],[f184,f1049])).
fof(f9186,plain,(
  ( ! [X588,X584,X586,X582,X589,X585,X587,X581,X583] : (s(fun(X581,X588),'_dst_x0x1_2'(s(fun(X582,X588),X589),s(fun(X581,X582),X585))) = s(fun(X581,X588),happ(s(fun(X583,fun(X581,X588)),'_dst_x00x11_2'(s(fun(X582,X588),X589),s(fun(X583,fun(X581,X582)),'_dst_x0x1_2'(s(fun(X584,fun(X581,X582)),'_dst_x0_1'(s(fun(X581,X582),X585))),s(fun(X583,X584),X586))))),s(X583,X587)))) )),
  inference(superposition,[],[f193,f1049])).
fof(f9185,plain,(
  ( ! [X573,X575,X571,X580,X576,X578,X572,X574,X577,X579] : (s(X572,happ(s(fun(X571,X572),X575),s(X571,happ(s(fun(X578,X571),X579),s(X578,X580))))) = s(X572,happ(s(fun(X578,X572),'_dst_x0x0x1_3'(s(fun(X573,fun(X571,X572)),'_dst_x0x1_2'(s(fun(X574,fun(X571,X572)),'_dst_x0_1'(s(fun(X571,X572),X575))),s(fun(X573,X574),X576))),s(X573,X577),s(fun(X578,X571),X579))),s(X578,X580)))) )),
  inference(superposition,[],[f194,f1049])).
fof(f9184,plain,(
  ( ! [X569,X565,X567,X561,X563,X568,X570,X564,X566,X560,X562] : (s(fun(X560,X567),'_dst_x0x0x1_3'(s(fun(X568,fun(X561,X567)),X569),s(X568,X570),s(fun(X560,X561),X564))) = s(fun(X560,X567),happ(s(fun(X562,fun(X560,X567)),'_dst_x00x00x11_3'(s(fun(X568,fun(X561,X567)),X569),s(X568,X570),s(fun(X562,fun(X560,X561)),'_dst_x0x1_2'(s(fun(X563,fun(X560,X561)),'_dst_x0_1'(s(fun(X560,X561),X564))),s(fun(X562,X563),X565))))),s(X562,X566)))) )),
  inference(superposition,[],[f199,f1049])).
fof(f9183,plain,(
  ( ! [X557,X559,X553,X555,X549,X551,X556,X558,X552,X554,X550] : (s(fun(X549,X556),'_dst_x0x1x1_3'(s(fun(X550,fun(X557,X556)),X558),s(fun(X549,X550),X553),s(fun(X549,X557),X559))) = s(fun(X549,X556),happ(s(fun(X551,fun(X549,X556)),'_dst_x00x11x01_3'(s(fun(X550,fun(X557,X556)),X558),s(fun(X551,fun(X549,X550)),'_dst_x0x1_2'(s(fun(X552,fun(X549,X550)),'_dst_x0_1'(s(fun(X549,X550),X553))),s(fun(X551,X552),X554))),s(fun(X549,X557),X559))),s(X551,X555)))) )),
  inference(superposition,[],[f200,f1049])).
fof(f9179,plain,(
  ( ! [X524,X526,X522,X525,X527,X523] : (s(bool,'F_0') != s(bool,X524) | s(bool,'T_0') = s(bool,X527) | s(bool,X527) = s(bool,happ(s(fun(X522,bool),'_dst_x0x1_2'(s(fun(X523,bool),'_dst_x0_1'(s(bool,X524))),s(fun(X522,X523),X525))),s(X522,X526)))) )),
  inference(superposition,[],[f1176,f1049])).
fof(f9177,plain,(
  ( ! [X511,X516,X512,X514,X513,X515] : (s(bool,X513) != s(bool,X516) | s(bool,'T_0') = s(bool,X516) | s(bool,'F_0') = s(bool,happ(s(fun(X511,bool),'_dst_x0x1_2'(s(fun(X512,bool),'_dst_x0_1'(s(bool,X513))),s(fun(X511,X512),X514))),s(X511,X515)))) )),
  inference(superposition,[],[f638,f1049])).
fof(f9176,plain,(
  ( ! [X509,X505,X507,X508,X510,X506] : (s(bool,X507) != s(bool,X510) | s(bool,'T_0') = s(bool,happ(s(fun(X505,bool),'_dst_x0x1_2'(s(fun(X506,bool),'_dst_x0_1'(s(bool,X507))),s(fun(X505,X506),X508))),s(X505,X509))) | s(bool,'F_0') = s(bool,X510)) )),
  inference(superposition,[],[f638,f1049])).
fof(f9174,plain,(
  ( ! [X492,X494,X496,X498,X493,X495,X497] : (s(bool,'F_0') = s(bool,happ(s(fun(X497,bool),'_dst_x0_1'(s(bool,X494))),s(X497,X498))) | s(bool,'T_0') = s(bool,happ(s(fun(X492,bool),'_dst_x0x1_2'(s(fun(X493,bool),'_dst_x0_1'(s(bool,X494))),s(fun(X492,X493),X495))),s(X492,X496)))) )),
  inference(superposition,[],[f614,f1049])).
fof(f9170,plain,(
  ( ! [X472,X469,X471,X467,X473,X468,X470] : (s(bool,'T_0') = s(bool,happ(s(fun(X472,bool),'_dst_x0_1'(s(bool,X469))),s(X472,X473))) | s(bool,'F_0') = s(bool,happ(s(fun(X467,bool),'_dst_x0x1_2'(s(fun(X468,bool),'_dst_x0_1'(s(bool,X469))),s(fun(X467,X468),X470))),s(X467,X471)))) )),
  inference(superposition,[],[f267,f1049])).
fof(f9124,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(fun(X66,fun(X67,X68)),'_dst_x00x00x11_3'(s(fun(X60,fun(X69,X68)),X70),s(X60,X63),s(fun(X66,fun(X67,X69)),X71))) = s(fun(X66,fun(X67,X68)),happ(s(fun(X61,fun(X66,fun(X67,X68))),'_dst_x000x100x011_3'(s(fun(X60,fun(X69,X68)),X70),s(fun(X61,X60),'_dst_x0x1_2'(s(fun(X62,X60),'_dst_x0_1'(s(X60,X63))),s(fun(X61,X62),X64))),s(fun(X66,fun(X67,X69)),X71))),s(X61,X65)))) )),
  inference(superposition,[],[f202,f1049])).
fof(f9123,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s(fun(X54,fun(X55,X56)),'_dst_x00x01x00_3'(s(fun(X57,fun(X48,X56)),X58),s(fun(X55,X57),X59),s(X48,X51))) = s(fun(X54,fun(X55,X56)),happ(s(fun(X49,fun(X54,fun(X55,X56))),'_dst_x000x001x100_3'(s(fun(X57,fun(X48,X56)),X58),s(fun(X55,X57),X59),s(fun(X49,X48),'_dst_x0x1_2'(s(fun(X50,X48),'_dst_x0_1'(s(X48,X51))),s(fun(X49,X50),X52))))),s(X49,X53)))) )),
  inference(superposition,[],[f201,f1049])).
fof(f9122,plain,(
  ( ! [X39,X47,X45,X43,X41,X38,X46,X44,X42,X40] : (s(X44,happ(s(fun(X45,X44),happ(s(fun(X38,fun(X45,X44)),X46),s(X38,X41))),s(X45,happ(s(fun(X39,X45),X47),s(X39,X43))))) = s(X44,happ(s(fun(X39,X44),'_dst_x0x1x1_3'(s(fun(X38,fun(X45,X44)),X46),s(fun(X39,X38),'_dst_x0x1_2'(s(fun(X40,X38),'_dst_x0_1'(s(X38,X41))),s(fun(X39,X40),X42))),s(fun(X39,X45),X47))),s(X39,X43)))) )),
  inference(superposition,[],[f196,f1049])).
fof(f9121,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X36,X34,X32] : (s(X34,happ(s(fun(X28,X34),happ(s(fun(X35,fun(X28,X34)),X36),s(X35,happ(s(fun(X29,X35),X37),s(X29,X33))))),s(X28,X31))) = s(X34,happ(s(fun(X29,X34),'_dst_x0x1x1_3'(s(fun(X35,fun(X28,X34)),X36),s(fun(X29,X35),X37),s(fun(X29,X28),'_dst_x0x1_2'(s(fun(X30,X28),'_dst_x0_1'(s(X28,X31))),s(fun(X29,X30),X32))))),s(X29,X33)))) )),
  inference(superposition,[],[f196,f1049])).
fof(f9120,plain,(
  ( ! [X26,X24,X23,X21,X19,X27,X25,X22,X20,X18] : (s(X24,happ(s(fun(X25,X24),happ(s(fun(X18,fun(X25,X24)),X26),s(X18,X21))),s(X25,X27))) = s(X24,happ(s(fun(X19,X24),'_dst_x0x1x0_3'(s(fun(X18,fun(X25,X24)),X26),s(fun(X19,X18),'_dst_x0x1_2'(s(fun(X20,X18),'_dst_x0_1'(s(X18,X21))),s(fun(X19,X20),X22))),s(X25,X27))),s(X19,X23)))) )),
  inference(superposition,[],[f195,f1049])).
fof(f9119,plain,(
  ( ! [X14,X12,X10,X8,X17,X15,X13,X11,X9,X16] : (s(X14,happ(s(fun(X8,X14),happ(s(fun(X15,fun(X8,X14)),X16),s(X15,X17))),s(X8,X11))) = s(X14,happ(s(fun(X9,X14),'_dst_x0x0x1_3'(s(fun(X15,fun(X8,X14)),X16),s(X15,X17),s(fun(X9,X8),'_dst_x0x1_2'(s(fun(X10,X8),'_dst_x0_1'(s(X8,X11))),s(fun(X9,X10),X12))))),s(X9,X13)))) )),
  inference(superposition,[],[f194,f1049])).
fof(f9118,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,X3))) = s(X6,happ(s(fun(X1,X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),'_dst_x0_1'(s(X0,X3))),s(fun(X1,X2),X4))))),s(X1,X5)))) )),
  inference(superposition,[],[f182,f1049])).
fof(f9115,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X40,bool),'_dst_x0_1'(s(bool,X41))),s(fun(X39,X40),X42))),s(X39,X43))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f1049,f223])).
fof(f9113,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(bool,'T_0') = s(bool,X29) | s(bool,X32) = s(bool,happ(s(fun(X27,bool),'_dst_x0x1_2'(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29))),s(fun(X27,X28),X30))),s(X27,X31))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f1049,f223])).
fof(f9111,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X12,X15) = s(X12,happ(s(fun(X17,X12),'_dst_x0x1_2'(s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),'_dst_x0_1'(s(X12,X15))),s(fun(X13,X14),X16))),s(fun(X17,X13),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f1049,f182])).
fof(f9055,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,X3) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X4,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X2,X1),'_dst_x0_1'(s(X1,X3))),s(fun(X4,fun(X0,X2)),X5))),s(X4,X6))),s(X0,X7)))) )),
  inference(superposition,[],[f1049,f193])).
fof(f8892,plain,(
  ( ! [X666,X668,X670,X664,X673,X667,X669,X663,X665,X671,X672] : (p(s(bool,happ(s(fun(fun(X663,fun(X664,X665)),bool),happ(s(fun(fun(X663,fun(X664,X665)),fun(fun(X663,fun(X664,X665)),bool)),'=_0'),s(fun(X663,fun(X664,X665)),X673))),s(fun(X663,fun(X664,X665)),'_dst_x00x00x11_3'(s(fun(X667,fun(X668,X665)),X669),s(X667,happ(s(fun(X666,X667),X670),s(X666,X672))),s(fun(X663,fun(X664,X668)),X671)))))) | s(fun(X663,fun(X664,X665)),X673) != s(fun(X663,fun(X664,X665)),happ(s(fun(X666,fun(X663,fun(X664,X665))),'_dst_x000x100x011_3'(s(fun(X667,fun(X668,X665)),X669),s(fun(X666,X667),X670),s(fun(X663,fun(X664,X668)),X671))),s(X666,X672)))) )),
  inference(superposition,[],[f173,f202])).
fof(f8891,plain,(
  ( ! [X660,X662,X656,X658,X652,X654,X661,X657,X659,X653,X655] : (p(s(bool,happ(s(fun(fun(X652,fun(X653,X654)),bool),happ(s(fun(fun(X652,fun(X653,X654)),fun(fun(X652,fun(X653,X654)),bool)),'=_0'),s(fun(X652,fun(X653,X654)),'_dst_x00x00x11_3'(s(fun(X656,fun(X657,X654)),X658),s(X656,happ(s(fun(X655,X656),X659),s(X655,X661))),s(fun(X652,fun(X653,X657)),X660))))),s(fun(X652,fun(X653,X654)),X662)))) | s(fun(X652,fun(X653,X654)),X662) != s(fun(X652,fun(X653,X654)),happ(s(fun(X655,fun(X652,fun(X653,X654))),'_dst_x000x100x011_3'(s(fun(X656,fun(X657,X654)),X658),s(fun(X655,X656),X659),s(fun(X652,fun(X653,X657)),X660))),s(X655,X661)))) )),
  inference(superposition,[],[f173,f202])).
fof(f8890,plain,(
  ( ! [X648,X650,X644,X646,X642,X649,X651,X645,X647,X641,X643] : (~p(s(bool,happ(s(fun(fun(X641,fun(X642,X643)),bool),happ(s(fun(fun(X641,fun(X642,X643)),fun(fun(X641,fun(X642,X643)),bool)),'=_0'),s(fun(X641,fun(X642,X643)),X651))),s(fun(X641,fun(X642,X643)),'_dst_x00x00x11_3'(s(fun(X645,fun(X646,X643)),X647),s(X645,happ(s(fun(X644,X645),X648),s(X644,X650))),s(fun(X641,fun(X642,X646)),X649)))))) | s(fun(X641,fun(X642,X643)),X651) = s(fun(X641,fun(X642,X643)),happ(s(fun(X644,fun(X641,fun(X642,X643))),'_dst_x000x100x011_3'(s(fun(X645,fun(X646,X643)),X647),s(fun(X644,X645),X648),s(fun(X641,fun(X642,X646)),X649))),s(X644,X650)))) )),
  inference(superposition,[],[f172,f202])).
fof(f8889,plain,(
  ( ! [X639,X633,X635,X631,X640,X636,X638,X632,X634,X630,X637] : (~p(s(bool,happ(s(fun(fun(X630,fun(X631,X632)),bool),happ(s(fun(fun(X630,fun(X631,X632)),fun(fun(X630,fun(X631,X632)),bool)),'=_0'),s(fun(X630,fun(X631,X632)),'_dst_x00x00x11_3'(s(fun(X634,fun(X635,X632)),X636),s(X634,happ(s(fun(X633,X634),X637),s(X633,X639))),s(fun(X630,fun(X631,X635)),X638))))),s(fun(X630,fun(X631,X632)),X640)))) | s(fun(X630,fun(X631,X632)),X640) = s(fun(X630,fun(X631,X632)),happ(s(fun(X633,fun(X630,fun(X631,X632))),'_dst_x000x100x011_3'(s(fun(X634,fun(X635,X632)),X636),s(fun(X633,X634),X637),s(fun(X630,fun(X631,X635)),X638))),s(X633,X639)))) )),
  inference(superposition,[],[f172,f202])).
fof(f8884,plain,(
  ( ! [X573,X575,X569,X571,X580,X582,X576,X578,X572,X574,X568,X570,X581,X583,X577,X579] : (s(fun(X578,fun(X579,X580)),'_dst_x00x01x00_3'(s(fun(X581,fun(fun(X568,fun(X569,X570)),X580)),X582),s(fun(X579,X581),X583),s(fun(X568,fun(X569,X570)),'_dst_x00x00x11_3'(s(fun(X572,fun(X573,X570)),X574),s(X572,happ(s(fun(X571,X572),X575),s(X571,X577))),s(fun(X568,fun(X569,X573)),X576))))) = s(fun(X578,fun(X579,X580)),happ(s(fun(X571,fun(X578,fun(X579,X580))),'_dst_x000x001x100_3'(s(fun(X581,fun(fun(X568,fun(X569,X570)),X580)),X582),s(fun(X579,X581),X583),s(fun(X571,fun(X568,fun(X569,X570))),'_dst_x000x100x011_3'(s(fun(X572,fun(X573,X570)),X574),s(fun(X571,X572),X575),s(fun(X568,fun(X569,X573)),X576))))),s(X571,X577)))) )),
  inference(superposition,[],[f201,f202])).
fof(f8883,plain,(
  ( ! [X565,X567,X561,X563,X557,X559,X555,X564,X566,X560,X562,X556,X558,X554] : (s(X564,happ(s(fun(X557,X564),'_dst_x0x1x1_3'(s(fun(fun(X554,fun(X555,X556)),fun(X565,X564)),X566),s(fun(X557,fun(X554,fun(X555,X556))),'_dst_x000x100x011_3'(s(fun(X558,fun(X559,X556)),X560),s(fun(X557,X558),X561),s(fun(X554,fun(X555,X559)),X562))),s(fun(X557,X565),X567))),s(X557,X563))) = s(X564,happ(s(fun(X565,X564),happ(s(fun(fun(X554,fun(X555,X556)),fun(X565,X564)),X566),s(fun(X554,fun(X555,X556)),'_dst_x00x00x11_3'(s(fun(X558,fun(X559,X556)),X560),s(X558,happ(s(fun(X557,X558),X561),s(X557,X563))),s(fun(X554,fun(X555,X559)),X562))))),s(X565,happ(s(fun(X557,X565),X567),s(X557,X563)))))) )),
  inference(superposition,[],[f196,f202])).
fof(f8882,plain,(
  ( ! [X540,X542,X546,X553,X549,X551,X545,X541,X543,X547,X552,X548,X550,X544] : (s(X550,happ(s(fun(X543,X550),'_dst_x0x1x1_3'(s(fun(X551,fun(fun(X540,fun(X541,X542)),X550)),X552),s(fun(X543,X551),X553),s(fun(X543,fun(X540,fun(X541,X542))),'_dst_x000x100x011_3'(s(fun(X544,fun(X545,X542)),X546),s(fun(X543,X544),X547),s(fun(X540,fun(X541,X545)),X548))))),s(X543,X549))) = s(X550,happ(s(fun(fun(X540,fun(X541,X542)),X550),happ(s(fun(X551,fun(fun(X540,fun(X541,X542)),X550)),X552),s(X551,happ(s(fun(X543,X551),X553),s(X543,X549))))),s(fun(X540,fun(X541,X542)),'_dst_x00x00x11_3'(s(fun(X544,fun(X545,X542)),X546),s(X544,happ(s(fun(X543,X544),X547),s(X543,X549))),s(fun(X540,fun(X541,X545)),X548)))))) )),
  inference(superposition,[],[f196,f202])).
fof(f8881,plain,(
  ( ! [X536,X538,X532,X534,X528,X530,X526,X537,X539,X533,X535,X529,X531,X527] : (s(X536,happ(s(fun(X529,X536),'_dst_x0x1x0_3'(s(fun(fun(X526,fun(X527,X528)),fun(X537,X536)),X538),s(fun(X529,fun(X526,fun(X527,X528))),'_dst_x000x100x011_3'(s(fun(X530,fun(X531,X528)),X532),s(fun(X529,X530),X533),s(fun(X526,fun(X527,X531)),X534))),s(X537,X539))),s(X529,X535))) = s(X536,happ(s(fun(X537,X536),happ(s(fun(fun(X526,fun(X527,X528)),fun(X537,X536)),X538),s(fun(X526,fun(X527,X528)),'_dst_x00x00x11_3'(s(fun(X530,fun(X531,X528)),X532),s(X530,happ(s(fun(X529,X530),X533),s(X529,X535))),s(fun(X526,fun(X527,X531)),X534))))),s(X537,X539)))) )),
  inference(superposition,[],[f195,f202])).
fof(f8880,plain,(
  ( ! [X524,X520,X522,X516,X518,X512,X514,X525,X521,X523,X517,X519,X513,X515] : (s(X522,happ(s(fun(X515,X522),'_dst_x0x0x1_3'(s(fun(X523,fun(fun(X512,fun(X513,X514)),X522)),X524),s(X523,X525),s(fun(X515,fun(X512,fun(X513,X514))),'_dst_x000x100x011_3'(s(fun(X516,fun(X517,X514)),X518),s(fun(X515,X516),X519),s(fun(X512,fun(X513,X517)),X520))))),s(X515,X521))) = s(X522,happ(s(fun(fun(X512,fun(X513,X514)),X522),happ(s(fun(X523,fun(fun(X512,fun(X513,X514)),X522)),X524),s(X523,X525))),s(fun(X512,fun(X513,X514)),'_dst_x00x00x11_3'(s(fun(X516,fun(X517,X514)),X518),s(X516,happ(s(fun(X515,X516),X519),s(X515,X521))),s(fun(X512,fun(X513,X517)),X520)))))) )),
  inference(superposition,[],[f194,f202])).
fof(f8879,plain,(
  ( ! [X509,X511,X505,X507,X500,X502,X508,X510,X504,X506,X501,X503] : (s(X510,happ(s(fun(fun(X500,fun(X501,X502)),X510),X511),s(fun(X500,fun(X501,X502)),'_dst_x00x00x11_3'(s(fun(X504,fun(X505,X502)),X506),s(X504,happ(s(fun(X503,X504),X507),s(X503,X509))),s(fun(X500,fun(X501,X505)),X508))))) = s(X510,happ(s(fun(X503,X510),'_dst_x0x1_2'(s(fun(fun(X500,fun(X501,X502)),X510),X511),s(fun(X503,fun(X500,fun(X501,X502))),'_dst_x000x100x011_3'(s(fun(X504,fun(X505,X502)),X506),s(fun(X503,X504),X507),s(fun(X500,fun(X501,X505)),X508))))),s(X503,X509)))) )),
  inference(superposition,[],[f182,f202])).
fof(f8851,plain,(
  ( ! [X103,X101,X111,X109,X107,X105,X112,X102,X100,X110,X108,X106,X104] : (s(fun(X101,X102),happ(s(fun(X100,fun(X101,X102)),'_dst_x00x00x11_3'(s(fun(X104,fun(X105,X102)),X106),s(X104,happ(s(fun(X103,X104),X107),s(X103,happ(s(fun(X109,X103),X110),s(X109,X111))))),s(fun(X100,fun(X101,X105)),X108))),s(X100,X112))) = s(fun(X101,X102),happ(s(fun(X109,fun(X101,X102)),'_dst_x0x1x0_3'(s(fun(X103,fun(X100,fun(X101,X102))),'_dst_x000x100x011_3'(s(fun(X104,fun(X105,X102)),X106),s(fun(X103,X104),X107),s(fun(X100,fun(X101,X105)),X108))),s(fun(X109,X103),X110),s(X100,X112))),s(X109,X111)))) )),
  inference(superposition,[],[f195,f202])).
fof(f8850,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X87,X95,X93,X91,X89,X98,X96] : (s(fun(X88,X89),happ(s(fun(X96,fun(X88,X89)),'_dst_x0x1x1_3'(s(fun(X90,fun(X87,fun(X88,X89))),'_dst_x000x100x011_3'(s(fun(X91,fun(X92,X89)),X93),s(fun(X90,X91),X94),s(fun(X87,fun(X88,X92)),X95))),s(fun(X96,X90),X97),s(fun(X96,X87),X99))),s(X96,X98))) = s(fun(X88,X89),happ(s(fun(X87,fun(X88,X89)),'_dst_x00x00x11_3'(s(fun(X91,fun(X92,X89)),X93),s(X91,happ(s(fun(X90,X91),X94),s(X90,happ(s(fun(X96,X90),X97),s(X96,X98))))),s(fun(X87,fun(X88,X92)),X95))),s(X87,happ(s(fun(X96,X87),X99),s(X96,X98)))))) )),
  inference(superposition,[],[f196,f202])).
fof(f8849,plain,(
  ( ! [X80,X78,X76,X85,X83,X81,X79,X77,X86,X84,X82] : (s(fun(X77,X78),happ(s(fun(X76,fun(X77,X78)),'_dst_x00x00x11_3'(s(fun(X80,fun(X81,X78)),X82),s(X80,happ(s(fun(X79,X80),X83),s(X79,X85))),s(fun(X76,fun(X77,X81)),X84))),s(X76,X86))) = s(fun(X77,X78),'const.pair.UNCURRY_2'(s(fun(X79,fun(X76,fun(X77,X78))),'_dst_x000x100x011_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X79,X80),X83),s(fun(X76,fun(X77,X81)),X84))),s('type.pair.prod'(X79,X76),'const.pair.,_2'(s(X79,X85),s(X76,X86)))))) )),
  inference(superposition,[],[f183,f202])).
fof(f8848,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X65,X75,X73] : (s(fun(X66,X67),happ(s(fun(X65,fun(X66,X67)),'_dst_x00x00x11_3'(s(fun(X69,fun(X70,X67)),X71),s(X69,happ(s(fun(X68,X69),X72),s(X68,X74))),s(fun(X65,fun(X66,X70)),X73))),s(X65,X75))) = s(fun(X66,X67),'const.pair.pair_CASE_2'(s('type.pair.prod'(X68,X65),'const.pair.,_2'(s(X68,X74),s(X65,X75))),s(fun(X68,fun(X65,fun(X66,X67))),'_dst_x000x100x011_3'(s(fun(X69,fun(X70,X67)),X71),s(fun(X68,X69),X72),s(fun(X65,fun(X66,X70)),X73)))))) )),
  inference(superposition,[],[f184,f202])).
fof(f8847,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X53,X63] : (s(fun(X53,X63),'_dst_x0x1_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X53,fun(X54,X55)),'_dst_x00x00x11_3'(s(fun(X57,fun(X58,X55)),X59),s(X57,happ(s(fun(X56,X57),X60),s(X56,X62))),s(fun(X53,fun(X54,X58)),X61))))) = s(fun(X53,X63),happ(s(fun(X56,fun(X53,X63)),'_dst_x00x11_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X56,fun(X53,fun(X54,X55))),'_dst_x000x100x011_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X56,X57),X60),s(fun(X53,fun(X54,X58)),X61))))),s(X56,X62)))) )),
  inference(superposition,[],[f193,f202])).
fof(f8846,plain,(
  ( ! [X47,X45,X43,X41,X52,X50,X48,X46,X44,X42,X40,X51,X49] : (s(fun(X41,X42),happ(s(fun(X40,fun(X41,X42)),'_dst_x00x00x11_3'(s(fun(X44,fun(X45,X42)),X46),s(X44,happ(s(fun(X43,X44),X47),s(X43,X49))),s(fun(X40,fun(X41,X45)),X48))),s(X40,happ(s(fun(X50,X40),X51),s(X50,X52))))) = s(fun(X41,X42),happ(s(fun(X50,fun(X41,X42)),'_dst_x0x0x1_3'(s(fun(X43,fun(X40,fun(X41,X42))),'_dst_x000x100x011_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X43,X44),X47),s(fun(X40,fun(X41,X45)),X48))),s(X43,X49),s(fun(X50,X40),X51))),s(X50,X52)))) )),
  inference(superposition,[],[f194,f202])).
fof(f8845,plain,(
  ( ! [X30,X28,X26,X39,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s(fun(X26,X36),'_dst_x0x0x1_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X26,fun(X27,X28)),'_dst_x00x00x11_3'(s(fun(X30,fun(X31,X28)),X32),s(X30,happ(s(fun(X29,X30),X33),s(X29,X35))),s(fun(X26,fun(X27,X31)),X34))))) = s(fun(X26,X36),happ(s(fun(X29,fun(X26,X36)),'_dst_x00x00x11_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X29,fun(X26,fun(X27,X28))),'_dst_x000x100x011_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X29,X30),X33),s(fun(X26,fun(X27,X31)),X34))))),s(X29,X35)))) )),
  inference(superposition,[],[f199,f202])).
fof(f8844,plain,(
  ( ! [X24,X14,X12,X23,X21,X19,X17,X25,X15,X13,X22,X20,X18,X16] : (s(fun(X12,X22),'_dst_x0x1x1_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X12,fun(X13,X14)),'_dst_x00x00x11_3'(s(fun(X16,fun(X17,X14)),X18),s(X16,happ(s(fun(X15,X16),X19),s(X15,X21))),s(fun(X12,fun(X13,X17)),X20))),s(fun(X12,X23),X25))) = s(fun(X12,X22),happ(s(fun(X15,fun(X12,X22)),'_dst_x00x11x01_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X15,fun(X12,fun(X13,X14))),'_dst_x000x100x011_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X15,X16),X19),s(fun(X12,fun(X13,X17)),X20))),s(fun(X12,X23),X25))),s(X15,X21)))) )),
  inference(superposition,[],[f200,f202])).
fof(f8843,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,fun(X1,X10)),'_dst_x00x11_2'(s(fun(X2,X10),X11),s(fun(X0,fun(X1,X2)),'_dst_x00x00x11_3'(s(fun(X4,fun(X5,X2)),X6),s(X4,happ(s(fun(X3,X4),X7),s(X3,X9))),s(fun(X0,fun(X1,X5)),X8))))) = s(fun(X0,fun(X1,X10)),happ(s(fun(X3,fun(X0,fun(X1,X10))),'_dst_x000x111_2'(s(fun(X2,X10),X11),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x100x011_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X3,X4),X7),s(fun(X0,fun(X1,X5)),X8))))),s(X3,X9)))) )),
  inference(superposition,[],[f197,f202])).
fof(f8817,plain,(
  ( ! [X466,X472,X474,X469,X471,X465,X467,X475,X473,X468,X470] : (p(s(bool,happ(s(fun(fun(X465,fun(X466,X467)),bool),happ(s(fun(fun(X465,fun(X466,X467)),fun(fun(X465,fun(X466,X467)),bool)),'=_0'),s(fun(X465,fun(X466,X467)),X475))),s(fun(X465,fun(X466,X467)),happ(s(fun(X471,fun(X465,fun(X466,X467))),'_dst_x000x100x011_3'(s(fun(X468,fun(X469,X467)),X470),s(fun(X471,X468),X472),s(fun(X465,fun(X466,X469)),X474))),s(X471,X473)))))) | s(fun(X465,fun(X466,X467)),X475) != s(fun(X465,fun(X466,X467)),'_dst_x00x00x11_3'(s(fun(X468,fun(X469,X467)),X470),s(X468,happ(s(fun(X471,X468),X472),s(X471,X473))),s(fun(X465,fun(X466,X469)),X474)))) )),
  inference(superposition,[],[f173,f202])).
fof(f8816,plain,(
  ( ! [X455,X461,X463,X457,X459,X454,X460,X462,X456,X458,X464] : (p(s(bool,happ(s(fun(fun(X454,fun(X455,X456)),bool),happ(s(fun(fun(X454,fun(X455,X456)),fun(fun(X454,fun(X455,X456)),bool)),'=_0'),s(fun(X454,fun(X455,X456)),happ(s(fun(X460,fun(X454,fun(X455,X456))),'_dst_x000x100x011_3'(s(fun(X457,fun(X458,X456)),X459),s(fun(X460,X457),X461),s(fun(X454,fun(X455,X458)),X463))),s(X460,X462))))),s(fun(X454,fun(X455,X456)),X464)))) | s(fun(X454,fun(X455,X456)),X464) != s(fun(X454,fun(X455,X456)),'_dst_x00x00x11_3'(s(fun(X457,fun(X458,X456)),X459),s(X457,happ(s(fun(X460,X457),X461),s(X460,X462))),s(fun(X454,fun(X455,X458)),X463)))) )),
  inference(superposition,[],[f173,f202])).
fof(f8815,plain,(
  ( ! [X445,X447,X443,X453,X449,X451,X444,X446,X452,X448,X450] : (~p(s(bool,happ(s(fun(fun(X443,fun(X444,X445)),bool),happ(s(fun(fun(X443,fun(X444,X445)),fun(fun(X443,fun(X444,X445)),bool)),'=_0'),s(fun(X443,fun(X444,X445)),X453))),s(fun(X443,fun(X444,X445)),happ(s(fun(X449,fun(X443,fun(X444,X445))),'_dst_x000x100x011_3'(s(fun(X446,fun(X447,X445)),X448),s(fun(X449,X446),X450),s(fun(X443,fun(X444,X447)),X452))),s(X449,X451)))))) | s(fun(X443,fun(X444,X445)),X453) = s(fun(X443,fun(X444,X445)),'_dst_x00x00x11_3'(s(fun(X446,fun(X447,X445)),X448),s(X446,happ(s(fun(X449,X446),X450),s(X449,X451))),s(fun(X443,fun(X444,X447)),X452)))) )),
  inference(superposition,[],[f172,f202])).
fof(f8814,plain,(
  ( ! [X441,X436,X438,X432,X434,X440,X442,X437,X439,X433,X435] : (~p(s(bool,happ(s(fun(fun(X432,fun(X433,X434)),bool),happ(s(fun(fun(X432,fun(X433,X434)),fun(fun(X432,fun(X433,X434)),bool)),'=_0'),s(fun(X432,fun(X433,X434)),happ(s(fun(X438,fun(X432,fun(X433,X434))),'_dst_x000x100x011_3'(s(fun(X435,fun(X436,X434)),X437),s(fun(X438,X435),X439),s(fun(X432,fun(X433,X436)),X441))),s(X438,X440))))),s(fun(X432,fun(X433,X434)),X442)))) | s(fun(X432,fun(X433,X434)),X442) = s(fun(X432,fun(X433,X434)),'_dst_x00x00x11_3'(s(fun(X435,fun(X436,X434)),X437),s(X435,happ(s(fun(X438,X435),X439),s(X438,X440))),s(fun(X432,fun(X433,X436)),X441)))) )),
  inference(superposition,[],[f172,f202])).
fof(f8783,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,X2),'_dst_x0x0x1_3'(s(fun(X3,fun(X4,X2)),X5),s(X3,happ(s(fun(X6,X3),X7),s(X6,X8))),s(fun(X1,X4),happ(s(fun(X0,fun(X1,X4)),X9),s(X0,X10))))) = s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),happ(s(fun(X6,fun(X0,fun(X1,X2))),'_dst_x000x100x011_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X6,X3),X7),s(fun(X0,fun(X1,X4)),X9))),s(X6,X8))),s(X0,X10)))) )),
  inference(superposition,[],[f199,f202])).
fof(f8782,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (s(fun(X10,fun(X11,X12)),'_dst_x00x00x11_3'(s(fun(X14,fun(X15,X12)),X16),s(X14,happ(s(fun(X13,X14),X17),s(X13,happ(s(fun(X19,X13),X20),s(X19,X21))))),s(fun(X10,fun(X11,X15)),X18))) = s(fun(X10,fun(X11,X12)),happ(s(fun(X19,fun(X10,fun(X11,X12))),'_dst_x0x1_2'(s(fun(X13,fun(X10,fun(X11,X12))),'_dst_x000x100x011_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X13,X14),X17),s(fun(X10,fun(X11,X15)),X18))),s(fun(X19,X13),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f202,f182])).
fof(f8644,plain,(
  ( ! [X383,X377,X379,X385,X372,X374,X370,X380,X382,X376,X378,X384,X373,X375,X371,X381] : (s(fun(X380,fun(X381,X382)),'_dst_x00x00x11_3'(s(fun(fun(X370,fun(X371,X372)),fun(X383,X382)),X384),s(fun(X370,fun(X371,X372)),'_dst_x00x00x11_3'(s(fun(X374,fun(X375,X372)),X376),s(X374,happ(s(fun(X373,X374),X377),s(X373,X379))),s(fun(X370,fun(X371,X375)),X378))),s(fun(X380,fun(X381,X383)),X385))) = s(fun(X380,fun(X381,X382)),happ(s(fun(X373,fun(X380,fun(X381,X382))),'_dst_x000x100x011_3'(s(fun(fun(X370,fun(X371,X372)),fun(X383,X382)),X384),s(fun(X373,fun(X370,fun(X371,X372))),'_dst_x000x100x011_3'(s(fun(X374,fun(X375,X372)),X376),s(fun(X373,X374),X377),s(fun(X370,fun(X371,X375)),X378))),s(fun(X380,fun(X381,X383)),X385))),s(X373,X379)))) )),
  inference(superposition,[],[f202,f202])).
fof(f8643,plain,(
  ( ! [X356,X358,X354,X364,X366,X360,X362,X368,X357,X359,X355,X365,X367,X361,X363,X369] : (s(fun(X364,fun(X365,X366)),'_dst_x00x00x11_3'(s(fun(fun(X354,fun(X355,X356)),fun(X367,X366)),X368),s(fun(X354,fun(X355,X356)),'_dst_x00x01x00_3'(s(fun(X358,fun(X359,X356)),X360),s(fun(X355,X358),X361),s(X359,happ(s(fun(X357,X359),X362),s(X357,X363))))),s(fun(X364,fun(X365,X367)),X369))) = s(fun(X364,fun(X365,X366)),happ(s(fun(X357,fun(X364,fun(X365,X366))),'_dst_x000x100x011_3'(s(fun(fun(X354,fun(X355,X356)),fun(X367,X366)),X368),s(fun(X357,fun(X354,fun(X355,X356))),'_dst_x000x001x100_3'(s(fun(X358,fun(X359,X356)),X360),s(fun(X355,X358),X361),s(fun(X357,X359),X362))),s(fun(X364,fun(X365,X367)),X369))),s(X357,X363)))) )),
  inference(superposition,[],[f202,f201])).
fof(f8642,plain,(
  ( ! [X350,X344,X346,X352,X341,X343,X349,X351,X345,X347,X353,X340,X342,X348] : (s(fun(X348,fun(X349,X350)),'_dst_x00x00x11_3'(s(fun(fun(X340,fun(X341,X342)),fun(X351,X350)),X352),s(fun(X340,fun(X341,X342)),'_dst_x00x11_2'(s(fun(X344,X342),X345),s(fun(X340,fun(X341,X344)),happ(s(fun(X343,fun(X340,fun(X341,X344))),X346),s(X343,X347))))),s(fun(X348,fun(X349,X351)),X353))) = s(fun(X348,fun(X349,X350)),happ(s(fun(X343,fun(X348,fun(X349,X350))),'_dst_x000x100x011_3'(s(fun(fun(X340,fun(X341,X342)),fun(X351,X350)),X352),s(fun(X343,fun(X340,fun(X341,X342))),'_dst_x000x111_2'(s(fun(X344,X342),X345),s(fun(X343,fun(X340,fun(X341,X344))),X346))),s(fun(X348,fun(X349,X351)),X353))),s(X343,X347)))) )),
  inference(superposition,[],[f202,f197])).
fof(f8641,plain,(
  ( ! [X325,X327,X333,X335,X329,X331,X337,X339,X326,X332,X334,X328,X330,X336,X338] : (s(fun(X334,fun(X335,X336)),'_dst_x00x00x11_3'(s(fun(fun(X325,X326),fun(X337,X336)),X338),s(fun(X325,X326),'_dst_x0x0x1_3'(s(fun(X328,fun(X329,X326)),X330),s(X328,X331),s(fun(X325,X329),happ(s(fun(X327,fun(X325,X329)),X332),s(X327,X333))))),s(fun(X334,fun(X335,X337)),X339))) = s(fun(X334,fun(X335,X336)),happ(s(fun(X327,fun(X334,fun(X335,X336))),'_dst_x000x100x011_3'(s(fun(fun(X325,X326),fun(X337,X336)),X338),s(fun(X327,fun(X325,X326)),'_dst_x00x00x11_3'(s(fun(X328,fun(X329,X326)),X330),s(X328,X331),s(fun(X327,fun(X325,X329)),X332))),s(fun(X334,fun(X335,X337)),X339))),s(X327,X333)))) )),
  inference(superposition,[],[f202,f199])).
fof(f8640,plain,(
  ( ! [X313,X315,X321,X323,X310,X316,X318,X312,X314,X324,X320,X322,X311,X317,X319] : (s(fun(X319,fun(X320,X321)),'_dst_x00x00x11_3'(s(fun(fun(X310,X311),fun(X322,X321)),X323),s(fun(X310,X311),'_dst_x0x1x0_3'(s(fun(X313,fun(X314,X311)),X315),s(fun(X310,X313),X316),s(X314,X317))),s(fun(X319,fun(X320,X322)),X324))) = s(fun(X319,fun(X320,X321)),happ(s(fun(X312,fun(X319,fun(X320,X321))),'_dst_x000x100x011_3'(s(fun(fun(X310,X311),fun(X322,X321)),X323),s(fun(X312,fun(X310,X311)),'_dst_x00x01x00_3'(s(fun(X313,fun(X314,X311)),X315),s(fun(X310,X313),X316),s(X314,X317))),s(fun(X319,fun(X320,X322)),X324))),s(X312,X318)))) )),
  inference(superposition,[],[f202,f198])).
fof(f8639,plain,(
  ( ! [X300,X302,X298,X308,X304,X306,X301,X303,X297,X299,X309,X305,X307] : (s(fun(X304,fun(X305,X306)),'_dst_x00x00x11_3'(s(fun(fun(X297,X298),fun(X307,X306)),X308),s(fun(X297,X298),'_dst_x0x1_2'(s(fun(X300,X298),X301),s(fun(X297,X300),happ(s(fun(X299,fun(X297,X300)),X302),s(X299,X303))))),s(fun(X304,fun(X305,X307)),X309))) = s(fun(X304,fun(X305,X306)),happ(s(fun(X299,fun(X304,fun(X305,X306))),'_dst_x000x100x011_3'(s(fun(fun(X297,X298),fun(X307,X306)),X308),s(fun(X299,fun(X297,X298)),'_dst_x00x11_2'(s(fun(X300,X298),X301),s(fun(X299,fun(X297,X300)),X302))),s(fun(X304,fun(X305,X307)),X309))),s(X299,X303)))) )),
  inference(superposition,[],[f202,f193])).
fof(f8638,plain,(
  ( ! [X292,X282,X294,X288,X290,X296,X285,X287,X293,X283,X295,X291,X289,X284,X286] : (s(fun(X291,fun(X292,X293)),'_dst_x00x00x11_3'(s(fun(fun(X282,X283),fun(X294,X293)),X295),s(fun(X282,X283),'_dst_x0x1x1_3'(s(fun(X285,fun(X286,X283)),X287),s(fun(X282,X285),happ(s(fun(X284,fun(X282,X285)),X288),s(X284,X290))),s(fun(X282,X286),X289))),s(fun(X291,fun(X292,X294)),X296))) = s(fun(X291,fun(X292,X293)),happ(s(fun(X284,fun(X291,fun(X292,X293))),'_dst_x000x100x011_3'(s(fun(fun(X282,X283),fun(X294,X293)),X295),s(fun(X284,fun(X282,X283)),'_dst_x00x11x01_3'(s(fun(X285,fun(X286,X283)),X287),s(fun(X284,fun(X282,X285)),X288),s(fun(X282,X286),X289))),s(fun(X291,fun(X292,X294)),X296))),s(X284,X290)))) )),
  inference(superposition,[],[f202,f200])).
fof(f8637,plain,(
  ( ! [X156,X154,X152,X151,X149,X147,X157,X155,X153,X150,X148] : (s(fun(X152,fun(X153,X154)),'_dst_x00x00x11_3'(s(fun(X147,fun(X155,X154)),X156),s(X147,'const.pair.UNCURRY_2'(s(fun(X148,fun(X149,X147)),X150),s('type.pair.prod'(X148,X149),X151))),s(fun(X152,fun(X153,X155)),X157))) = s(fun(X152,fun(X153,X154)),happ(s(fun('type.pair.prod'(X148,X149),fun(X152,fun(X153,X154))),'_dst_x000x100x011_3'(s(fun(X147,fun(X155,X154)),X156),s(fun('type.pair.prod'(X148,X149),X147),happ(s(fun(fun(X148,fun(X149,X147)),fun('type.pair.prod'(X148,X149),X147)),'const.pair.UNCURRY_0'),s(fun(X148,fun(X149,X147)),X150))),s(fun(X152,fun(X153,X155)),X157))),s('type.pair.prod'(X148,X149),X151)))) )),
  inference(superposition,[],[f202,f180])).
fof(f8636,plain,(
  ( ! [X127,X125,X123,X134,X132,X130,X128,X126,X124,X122,X135,X133,X131,X129] : (s(fun(X130,fun(X131,X132)),'_dst_x00x00x11_3'(s(fun(X122,fun(X133,X132)),X134),s(X122,happ(s(fun(X125,X122),happ(s(fun(X124,fun(X125,X122)),X126),s(X124,X127))),s(X125,happ(s(fun(X123,X125),X128),s(X123,X129))))),s(fun(X130,fun(X131,X133)),X135))) = s(fun(X130,fun(X131,X132)),happ(s(fun(X123,fun(X130,fun(X131,X132))),'_dst_x000x100x011_3'(s(fun(X122,fun(X133,X132)),X134),s(fun(X123,X122),'_dst_x0x0x1_3'(s(fun(X124,fun(X125,X122)),X126),s(X124,X127),s(fun(X123,X125),X128))),s(fun(X130,fun(X131,X133)),X135))),s(X123,X129)))) )),
  inference(superposition,[],[f202,f194])).
fof(f8635,plain,(
  ( ! [X121,X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s(fun(X116,fun(X117,X118)),'_dst_x00x00x11_3'(s(fun(X112,fun(X119,X118)),X120),s(X112,X114),s(fun(X116,fun(X117,X119)),X121))) = s(fun(X116,fun(X117,X118)),happ(s(fun(X113,fun(X116,fun(X117,X118))),'_dst_x000x100x011_3'(s(fun(X112,fun(X119,X118)),X120),s(fun(X113,X112),'_dst_x0_1'(s(X112,X114))),s(fun(X116,fun(X117,X119)),X121))),s(X113,X115)))) )),
  inference(superposition,[],[f202,f170])).
fof(f8634,plain,(
  ( ! [X103,X101,X99,X111,X109,X107,X105,X102,X100,X98,X110,X108,X106,X104] : (s(fun(X106,fun(X107,X108)),'_dst_x00x00x11_3'(s(fun(X98,fun(X109,X108)),X110),s(X98,happ(s(fun(X101,X98),happ(s(fun(X100,fun(X101,X98)),X102),s(X100,happ(s(fun(X99,X100),X103),s(X99,X105))))),s(X101,X104))),s(fun(X106,fun(X107,X109)),X111))) = s(fun(X106,fun(X107,X108)),happ(s(fun(X99,fun(X106,fun(X107,X108))),'_dst_x000x100x011_3'(s(fun(X98,fun(X109,X108)),X110),s(fun(X99,X98),'_dst_x0x1x0_3'(s(fun(X100,fun(X101,X98)),X102),s(fun(X99,X100),X103),s(X101,X104))),s(fun(X106,fun(X107,X109)),X111))),s(X99,X105)))) )),
  inference(superposition,[],[f202,f195])).
fof(f8633,plain,(
  ( ! [X94,X92,X90,X88,X97,X87,X85,X95,X93,X91,X89,X96,X86,X84] : (s(fun(X92,fun(X93,X94)),'_dst_x00x00x11_3'(s(fun(X84,fun(X95,X94)),X96),s(X84,happ(s(fun(X88,X84),'_dst_x0x1x1_3'(s(fun(X86,fun(X85,X84)),X87),s(fun(X88,X86),X89),s(fun(X88,X85),X91))),s(X88,X90))),s(fun(X92,fun(X93,X95)),X97))) = s(fun(X92,fun(X93,X94)),happ(s(fun(X85,fun(X92,fun(X93,X94))),'_dst_x000x100x011_3'(s(fun(X84,fun(X95,X94)),X96),s(fun(X85,X84),happ(s(fun(X86,fun(X85,X84)),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90))))),s(fun(X92,fun(X93,X95)),X97))),s(X85,happ(s(fun(X88,X85),X91),s(X88,X90)))))) )),
  inference(superposition,[],[f202,f196])).
fof(f8632,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X83,X81,X71,X79,X77,X75,X73,X82] : (s(fun(X78,fun(X79,X80)),'_dst_x00x00x11_3'(s(fun(X70,fun(X81,X80)),X82),s(X70,happ(s(fun(X74,X70),'_dst_x0x1x0_3'(s(fun(X72,fun(X71,X70)),X73),s(fun(X74,X72),X75),s(X71,X77))),s(X74,X76))),s(fun(X78,fun(X79,X81)),X83))) = s(fun(X78,fun(X79,X80)),happ(s(fun(X71,fun(X78,fun(X79,X80))),'_dst_x000x100x011_3'(s(fun(X70,fun(X81,X80)),X82),s(fun(X71,X70),happ(s(fun(X72,fun(X71,X70)),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76))))),s(fun(X78,fun(X79,X81)),X83))),s(X71,X77)))) )),
  inference(superposition,[],[f202,f195])).
fof(f8631,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X56,X69,X67,X65,X63] : (s(fun(X64,fun(X65,X66)),'_dst_x00x00x11_3'(s(fun(X56,fun(X67,X66)),X68),s(X56,happ(s(fun(X61,X56),'_dst_x0x0x1_3'(s(fun(X58,fun(X57,X56)),X59),s(X58,X60),s(fun(X61,X57),X62))),s(X61,X63))),s(fun(X64,fun(X65,X67)),X69))) = s(fun(X64,fun(X65,X66)),happ(s(fun(X57,fun(X64,fun(X65,X66))),'_dst_x000x100x011_3'(s(fun(X56,fun(X67,X66)),X68),s(fun(X57,X56),happ(s(fun(X58,fun(X57,X56)),X59),s(X58,X60))),s(fun(X64,fun(X65,X67)),X69))),s(X57,happ(s(fun(X61,X57),X62),s(X61,X63)))))) )),
  inference(superposition,[],[f202,f194])).
fof(f8629,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X32,X42,X40] : (s(fun(X38,fun(X39,X40)),'_dst_x00x00x11_3'(s(fun(X32,fun(X41,X40)),X42),s(X32,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X33),'const.pair.,_2'(s(X34,X36),s(X33,X37))),s(fun(X34,fun(X33,X32)),X35))),s(fun(X38,fun(X39,X41)),X43))) = s(fun(X38,fun(X39,X40)),happ(s(fun(X33,fun(X38,fun(X39,X40))),'_dst_x000x100x011_3'(s(fun(X32,fun(X41,X40)),X42),s(fun(X33,X32),happ(s(fun(X34,fun(X33,X32)),X35),s(X34,X36))),s(fun(X38,fun(X39,X41)),X43))),s(X33,X37)))) )),
  inference(superposition,[],[f202,f184])).
fof(f8628,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (s(fun(X26,fun(X27,X28)),'_dst_x00x00x11_3'(s(fun(X20,fun(X29,X28)),X30),s(X20,happ(s(fun(X22,X20),X23),s(X22,happ(s(fun(X21,X22),X24),s(X21,X25))))),s(fun(X26,fun(X27,X29)),X31))) = s(fun(X26,fun(X27,X28)),happ(s(fun(X21,fun(X26,fun(X27,X28))),'_dst_x000x100x011_3'(s(fun(X20,fun(X29,X28)),X30),s(fun(X21,X20),'_dst_x0x1_2'(s(fun(X22,X20),X23),s(fun(X21,X22),X24))),s(fun(X26,fun(X27,X29)),X31))),s(X21,X25)))) )),
  inference(superposition,[],[f202,f182])).
fof(f8627,plain,(
  ( ! [X14,X12,X10,X8,X19,X17,X15,X13,X11,X9,X18,X16] : (s(fun(X14,fun(X15,X16)),'_dst_x00x00x11_3'(s(fun(X8,fun(X17,X16)),X18),s(X8,happ(s(fun(X11,X8),'_dst_x0x1_2'(s(fun(X9,X8),X10),s(fun(X11,X9),X12))),s(X11,X13))),s(fun(X14,fun(X15,X17)),X19))) = s(fun(X14,fun(X15,X16)),happ(s(fun(X9,fun(X14,fun(X15,X16))),'_dst_x000x100x011_3'(s(fun(X8,fun(X17,X16)),X18),s(fun(X9,X8),X10),s(fun(X14,fun(X15,X17)),X19))),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))))) )),
  inference(superposition,[],[f202,f182])).
fof(f8625,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),'_dst_x00x00x11_3'(s(fun(X3,fun(X4,X2)),X5),s(X3,happ(s(fun(X6,X3),X7),s(X6,X8))),s(fun(X0,fun(X1,X4)),X9))))),s(fun(X0,fun(X1,X2)),happ(s(fun(X6,fun(X0,fun(X1,X2))),'_dst_x000x100x011_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X6,X3),X7),s(fun(X0,fun(X1,X4)),X9))),s(X6,X8))))))) )),
  inference(unit_resulting_resolution,[],[f202,f173])).
fof(f8626,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),happ(s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x100x011_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X3,X4),X7),s(fun(X0,fun(X1,X5)),X8))),s(X3,X9))))),s(fun(X0,fun(X1,X2)),'_dst_x00x00x11_3'(s(fun(X4,fun(X5,X2)),X6),s(X4,happ(s(fun(X3,X4),X7),s(X3,X9))),s(fun(X0,fun(X1,X5)),X8))))))) )),
  inference(unit_resulting_resolution,[],[f202,f173])).
fof(f8357,plain,(
  ( ! [X606,X600,X602,X598,X605,X607,X601,X603,X597,X599,X604] : (p(s(bool,happ(s(fun(fun(X597,fun(X598,X599)),bool),happ(s(fun(fun(X597,fun(X598,X599)),fun(fun(X597,fun(X598,X599)),bool)),'=_0'),s(fun(X597,fun(X598,X599)),X607))),s(fun(X597,fun(X598,X599)),'_dst_x00x01x00_3'(s(fun(X601,fun(X602,X599)),X603),s(fun(X598,X601),X604),s(X602,happ(s(fun(X600,X602),X605),s(X600,X606)))))))) | s(fun(X597,fun(X598,X599)),X607) != s(fun(X597,fun(X598,X599)),happ(s(fun(X600,fun(X597,fun(X598,X599))),'_dst_x000x001x100_3'(s(fun(X601,fun(X602,X599)),X603),s(fun(X598,X601),X604),s(fun(X600,X602),X605))),s(X600,X606)))) )),
  inference(superposition,[],[f173,f201])).
fof(f8356,plain,(
  ( ! [X596,X592,X594,X588,X590,X586,X593,X595,X589,X591,X587] : (p(s(bool,happ(s(fun(fun(X586,fun(X587,X588)),bool),happ(s(fun(fun(X586,fun(X587,X588)),fun(fun(X586,fun(X587,X588)),bool)),'=_0'),s(fun(X586,fun(X587,X588)),'_dst_x00x01x00_3'(s(fun(X590,fun(X591,X588)),X592),s(fun(X587,X590),X593),s(X591,happ(s(fun(X589,X591),X594),s(X589,X595))))))),s(fun(X586,fun(X587,X588)),X596)))) | s(fun(X586,fun(X587,X588)),X596) != s(fun(X586,fun(X587,X588)),happ(s(fun(X589,fun(X586,fun(X587,X588))),'_dst_x000x001x100_3'(s(fun(X590,fun(X591,X588)),X592),s(fun(X587,X590),X593),s(fun(X589,X591),X594))),s(X589,X595)))) )),
  inference(superposition,[],[f173,f201])).
fof(f8355,plain,(
  ( ! [X575,X584,X580,X582,X576,X578,X585,X581,X583,X577,X579] : (~p(s(bool,happ(s(fun(fun(X575,fun(X576,X577)),bool),happ(s(fun(fun(X575,fun(X576,X577)),fun(fun(X575,fun(X576,X577)),bool)),'=_0'),s(fun(X575,fun(X576,X577)),X585))),s(fun(X575,fun(X576,X577)),'_dst_x00x01x00_3'(s(fun(X579,fun(X580,X577)),X581),s(fun(X576,X579),X582),s(X580,happ(s(fun(X578,X580),X583),s(X578,X584)))))))) | s(fun(X575,fun(X576,X577)),X585) = s(fun(X575,fun(X576,X577)),happ(s(fun(X578,fun(X575,fun(X576,X577))),'_dst_x000x001x100_3'(s(fun(X579,fun(X580,X577)),X581),s(fun(X576,X579),X582),s(fun(X578,X580),X583))),s(X578,X584)))) )),
  inference(superposition,[],[f172,f201])).
fof(f8354,plain,(
  ( ! [X573,X569,X571,X565,X567,X572,X574,X568,X570,X564,X566] : (~p(s(bool,happ(s(fun(fun(X564,fun(X565,X566)),bool),happ(s(fun(fun(X564,fun(X565,X566)),fun(fun(X564,fun(X565,X566)),bool)),'=_0'),s(fun(X564,fun(X565,X566)),'_dst_x00x01x00_3'(s(fun(X568,fun(X569,X566)),X570),s(fun(X565,X568),X571),s(X569,happ(s(fun(X567,X569),X572),s(X567,X573))))))),s(fun(X564,fun(X565,X566)),X574)))) | s(fun(X564,fun(X565,X566)),X574) = s(fun(X564,fun(X565,X566)),happ(s(fun(X567,fun(X564,fun(X565,X566))),'_dst_x000x001x100_3'(s(fun(X568,fun(X569,X566)),X570),s(fun(X565,X568),X571),s(fun(X567,X569),X572))),s(X567,X573)))) )),
  inference(superposition,[],[f172,f201])).
fof(f8349,plain,(
  ( ! [X509,X511,X507,X516,X512,X514,X508,X510,X506,X517,X513,X515] : (s(X516,happ(s(fun(fun(X506,fun(X507,X508)),X516),X517),s(fun(X506,fun(X507,X508)),'_dst_x00x01x00_3'(s(fun(X510,fun(X511,X508)),X512),s(fun(X507,X510),X513),s(X511,happ(s(fun(X509,X511),X514),s(X509,X515))))))) = s(X516,happ(s(fun(X509,X516),'_dst_x0x1_2'(s(fun(fun(X506,fun(X507,X508)),X516),X517),s(fun(X509,fun(X506,fun(X507,X508))),'_dst_x000x001x100_3'(s(fun(X510,fun(X511,X508)),X512),s(fun(X507,X510),X513),s(fun(X509,X511),X514))))),s(X509,X515)))) )),
  inference(superposition,[],[f182,f201])).
fof(f8348,plain,(
  ( ! [X499,X505,X492,X494,X500,X502,X496,X498,X504,X493,X495,X501,X503,X497] : (s(X502,happ(s(fun(X495,X502),'_dst_x0x0x1_3'(s(fun(X503,fun(fun(X492,fun(X493,X494)),X502)),X504),s(X503,X505),s(fun(X495,fun(X492,fun(X493,X494))),'_dst_x000x001x100_3'(s(fun(X496,fun(X497,X494)),X498),s(fun(X493,X496),X499),s(fun(X495,X497),X500))))),s(X495,X501))) = s(X502,happ(s(fun(fun(X492,fun(X493,X494)),X502),happ(s(fun(X503,fun(fun(X492,fun(X493,X494)),X502)),X504),s(X503,X505))),s(fun(X492,fun(X493,X494)),'_dst_x00x01x00_3'(s(fun(X496,fun(X497,X494)),X498),s(fun(X493,X496),X499),s(X497,happ(s(fun(X495,X497),X500),s(X495,X501)))))))) )),
  inference(superposition,[],[f194,f201])).
fof(f8347,plain,(
  ( ! [X478,X484,X486,X480,X482,X488,X490,X479,X485,X487,X481,X483,X489,X491] : (s(X488,happ(s(fun(X481,X488),'_dst_x0x1x0_3'(s(fun(fun(X478,fun(X479,X480)),fun(X489,X488)),X490),s(fun(X481,fun(X478,fun(X479,X480))),'_dst_x000x001x100_3'(s(fun(X482,fun(X483,X480)),X484),s(fun(X479,X482),X485),s(fun(X481,X483),X486))),s(X489,X491))),s(X481,X487))) = s(X488,happ(s(fun(X489,X488),happ(s(fun(fun(X478,fun(X479,X480)),fun(X489,X488)),X490),s(fun(X478,fun(X479,X480)),'_dst_x00x01x00_3'(s(fun(X482,fun(X483,X480)),X484),s(fun(X479,X482),X485),s(X483,happ(s(fun(X481,X483),X486),s(X481,X487))))))),s(X489,X491)))) )),
  inference(superposition,[],[f195,f201])).
fof(f8346,plain,(
  ( ! [X466,X476,X472,X474,X469,X471,X465,X467,X475,X477,X473,X468,X470,X464] : (s(X474,happ(s(fun(X467,X474),'_dst_x0x1x1_3'(s(fun(fun(X464,fun(X465,X466)),fun(X475,X474)),X476),s(fun(X467,fun(X464,fun(X465,X466))),'_dst_x000x001x100_3'(s(fun(X468,fun(X469,X466)),X470),s(fun(X465,X468),X471),s(fun(X467,X469),X472))),s(fun(X467,X475),X477))),s(X467,X473))) = s(X474,happ(s(fun(X475,X474),happ(s(fun(fun(X464,fun(X465,X466)),fun(X475,X474)),X476),s(fun(X464,fun(X465,X466)),'_dst_x00x01x00_3'(s(fun(X468,fun(X469,X466)),X470),s(fun(X465,X468),X471),s(X469,happ(s(fun(X467,X469),X472),s(X467,X473))))))),s(X475,happ(s(fun(X467,X475),X477),s(X467,X473)))))) )),
  inference(superposition,[],[f196,f201])).
fof(f8345,plain,(
  ( ! [X453,X455,X451,X461,X463,X457,X459,X452,X454,X450,X460,X462,X456,X458] : (s(X460,happ(s(fun(X453,X460),'_dst_x0x1x1_3'(s(fun(X461,fun(fun(X450,fun(X451,X452)),X460)),X462),s(fun(X453,X461),X463),s(fun(X453,fun(X450,fun(X451,X452))),'_dst_x000x001x100_3'(s(fun(X454,fun(X455,X452)),X456),s(fun(X451,X454),X457),s(fun(X453,X455),X458))))),s(X453,X459))) = s(X460,happ(s(fun(fun(X450,fun(X451,X452)),X460),happ(s(fun(X461,fun(fun(X450,fun(X451,X452)),X460)),X462),s(X461,happ(s(fun(X453,X461),X463),s(X453,X459))))),s(fun(X450,fun(X451,X452)),'_dst_x00x01x00_3'(s(fun(X454,fun(X455,X452)),X456),s(fun(X451,X454),X457),s(X455,happ(s(fun(X453,X455),X458),s(X453,X459)))))))) )),
  inference(superposition,[],[f196,f201])).
fof(f8320,plain,(
  ( ! [X103,X101,X111,X109,X107,X105,X112,X102,X100,X110,X108,X106,X104] : (s(fun(X101,X102),happ(s(fun(X100,fun(X101,X102)),'_dst_x00x01x00_3'(s(fun(X104,fun(X105,X102)),X106),s(fun(X101,X104),X107),s(X105,happ(s(fun(X103,X105),X108),s(X103,happ(s(fun(X109,X103),X110),s(X109,X111))))))),s(X100,X112))) = s(fun(X101,X102),happ(s(fun(X109,fun(X101,X102)),'_dst_x0x1x0_3'(s(fun(X103,fun(X100,fun(X101,X102))),'_dst_x000x001x100_3'(s(fun(X104,fun(X105,X102)),X106),s(fun(X101,X104),X107),s(fun(X103,X105),X108))),s(fun(X109,X103),X110),s(X100,X112))),s(X109,X111)))) )),
  inference(superposition,[],[f195,f201])).
fof(f8319,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X87,X95,X93,X91,X89,X98,X96] : (s(fun(X88,X89),happ(s(fun(X96,fun(X88,X89)),'_dst_x0x1x1_3'(s(fun(X90,fun(X87,fun(X88,X89))),'_dst_x000x001x100_3'(s(fun(X91,fun(X92,X89)),X93),s(fun(X88,X91),X94),s(fun(X90,X92),X95))),s(fun(X96,X90),X97),s(fun(X96,X87),X99))),s(X96,X98))) = s(fun(X88,X89),happ(s(fun(X87,fun(X88,X89)),'_dst_x00x01x00_3'(s(fun(X91,fun(X92,X89)),X93),s(fun(X88,X91),X94),s(X92,happ(s(fun(X90,X92),X95),s(X90,happ(s(fun(X96,X90),X97),s(X96,X98))))))),s(X87,happ(s(fun(X96,X87),X99),s(X96,X98)))))) )),
  inference(superposition,[],[f196,f201])).
fof(f8318,plain,(
  ( ! [X80,X78,X76,X85,X83,X81,X79,X77,X86,X84,X82] : (s(fun(X77,X78),happ(s(fun(X76,fun(X77,X78)),'_dst_x00x01x00_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X77,X80),X83),s(X81,happ(s(fun(X79,X81),X84),s(X79,X85))))),s(X76,X86))) = s(fun(X77,X78),'const.pair.UNCURRY_2'(s(fun(X79,fun(X76,fun(X77,X78))),'_dst_x000x001x100_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X77,X80),X83),s(fun(X79,X81),X84))),s('type.pair.prod'(X79,X76),'const.pair.,_2'(s(X79,X85),s(X76,X86)))))) )),
  inference(superposition,[],[f183,f201])).
fof(f8317,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X65,X75,X73] : (s(fun(X66,X67),happ(s(fun(X65,fun(X66,X67)),'_dst_x00x01x00_3'(s(fun(X69,fun(X70,X67)),X71),s(fun(X66,X69),X72),s(X70,happ(s(fun(X68,X70),X73),s(X68,X74))))),s(X65,X75))) = s(fun(X66,X67),'const.pair.pair_CASE_2'(s('type.pair.prod'(X68,X65),'const.pair.,_2'(s(X68,X74),s(X65,X75))),s(fun(X68,fun(X65,fun(X66,X67))),'_dst_x000x001x100_3'(s(fun(X69,fun(X70,X67)),X71),s(fun(X66,X69),X72),s(fun(X68,X70),X73)))))) )),
  inference(superposition,[],[f184,f201])).
fof(f8316,plain,(
  ( ! [X61,X59,X57,X64,X54,X62,X60,X58,X56,X55,X53,X63] : (s(fun(X53,X63),'_dst_x0x1_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X53,fun(X54,X55)),'_dst_x00x01x00_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X54,X57),X60),s(X58,happ(s(fun(X56,X58),X61),s(X56,X62))))))) = s(fun(X53,X63),happ(s(fun(X56,fun(X53,X63)),'_dst_x00x11_2'(s(fun(fun(X54,X55),X63),X64),s(fun(X56,fun(X53,fun(X54,X55))),'_dst_x000x001x100_3'(s(fun(X57,fun(X58,X55)),X59),s(fun(X54,X57),X60),s(fun(X56,X58),X61))))),s(X56,X62)))) )),
  inference(superposition,[],[f193,f201])).
fof(f8315,plain,(
  ( ! [X47,X45,X43,X41,X52,X50,X48,X46,X44,X42,X40,X51,X49] : (s(fun(X41,X42),happ(s(fun(X40,fun(X41,X42)),'_dst_x00x01x00_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X41,X44),X47),s(X45,happ(s(fun(X43,X45),X48),s(X43,X49))))),s(X40,happ(s(fun(X50,X40),X51),s(X50,X52))))) = s(fun(X41,X42),happ(s(fun(X50,fun(X41,X42)),'_dst_x0x0x1_3'(s(fun(X43,fun(X40,fun(X41,X42))),'_dst_x000x001x100_3'(s(fun(X44,fun(X45,X42)),X46),s(fun(X41,X44),X47),s(fun(X43,X45),X48))),s(X43,X49),s(fun(X50,X40),X51))),s(X50,X52)))) )),
  inference(superposition,[],[f194,f201])).
fof(f8314,plain,(
  ( ! [X30,X28,X26,X39,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s(fun(X26,X36),'_dst_x0x0x1_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X26,fun(X27,X28)),'_dst_x00x01x00_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X27,X30),X33),s(X31,happ(s(fun(X29,X31),X34),s(X29,X35))))))) = s(fun(X26,X36),happ(s(fun(X29,fun(X26,X36)),'_dst_x00x00x11_3'(s(fun(X37,fun(fun(X27,X28),X36)),X38),s(X37,X39),s(fun(X29,fun(X26,fun(X27,X28))),'_dst_x000x001x100_3'(s(fun(X30,fun(X31,X28)),X32),s(fun(X27,X30),X33),s(fun(X29,X31),X34))))),s(X29,X35)))) )),
  inference(superposition,[],[f199,f201])).
fof(f8313,plain,(
  ( ! [X24,X14,X12,X23,X21,X19,X17,X25,X15,X13,X22,X20,X18,X16] : (s(fun(X12,X22),'_dst_x0x1x1_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X12,fun(X13,X14)),'_dst_x00x01x00_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X13,X16),X19),s(X17,happ(s(fun(X15,X17),X20),s(X15,X21))))),s(fun(X12,X23),X25))) = s(fun(X12,X22),happ(s(fun(X15,fun(X12,X22)),'_dst_x00x11x01_3'(s(fun(fun(X13,X14),fun(X23,X22)),X24),s(fun(X15,fun(X12,fun(X13,X14))),'_dst_x000x001x100_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X13,X16),X19),s(fun(X15,X17),X20))),s(fun(X12,X23),X25))),s(X15,X21)))) )),
  inference(superposition,[],[f200,f201])).
fof(f8312,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,fun(X1,X10)),'_dst_x00x11_2'(s(fun(X2,X10),X11),s(fun(X0,fun(X1,X2)),'_dst_x00x01x00_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X1,X4),X7),s(X5,happ(s(fun(X3,X5),X8),s(X3,X9))))))) = s(fun(X0,fun(X1,X10)),happ(s(fun(X3,fun(X0,fun(X1,X10))),'_dst_x000x111_2'(s(fun(X2,X10),X11),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x001x100_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X1,X4),X7),s(fun(X3,X5),X8))))),s(X3,X9)))) )),
  inference(superposition,[],[f197,f201])).
fof(f8287,plain,(
  ( ! [X420,X422,X416,X418,X424,X415,X421,X423,X417,X419,X425] : (p(s(bool,happ(s(fun(fun(X415,fun(X416,X417)),bool),happ(s(fun(fun(X415,fun(X416,X417)),fun(fun(X415,fun(X416,X417)),bool)),'=_0'),s(fun(X415,fun(X416,X417)),X425))),s(fun(X415,fun(X416,X417)),happ(s(fun(X422,fun(X415,fun(X416,X417))),'_dst_x000x001x100_3'(s(fun(X418,fun(X419,X417)),X420),s(fun(X416,X418),X421),s(fun(X422,X419),X423))),s(X422,X424)))))) | s(fun(X415,fun(X416,X417)),X425) != s(fun(X415,fun(X416,X417)),'_dst_x00x01x00_3'(s(fun(X418,fun(X419,X417)),X420),s(fun(X416,X418),X421),s(X419,happ(s(fun(X422,X419),X423),s(X422,X424)))))) )),
  inference(superposition,[],[f173,f201])).
fof(f8286,plain,(
  ( ! [X412,X414,X408,X410,X405,X407,X413,X409,X411,X404,X406] : (p(s(bool,happ(s(fun(fun(X404,fun(X405,X406)),bool),happ(s(fun(fun(X404,fun(X405,X406)),fun(fun(X404,fun(X405,X406)),bool)),'=_0'),s(fun(X404,fun(X405,X406)),happ(s(fun(X411,fun(X404,fun(X405,X406))),'_dst_x000x001x100_3'(s(fun(X407,fun(X408,X406)),X409),s(fun(X405,X407),X410),s(fun(X411,X408),X412))),s(X411,X413))))),s(fun(X404,fun(X405,X406)),X414)))) | s(fun(X404,fun(X405,X406)),X414) != s(fun(X404,fun(X405,X406)),'_dst_x00x01x00_3'(s(fun(X407,fun(X408,X406)),X409),s(fun(X405,X407),X410),s(X408,happ(s(fun(X411,X408),X412),s(X411,X413)))))) )),
  inference(superposition,[],[f173,f201])).
fof(f8285,plain,(
  ( ! [X397,X399,X393,X395,X401,X403,X396,X398,X394,X400,X402] : (~p(s(bool,happ(s(fun(fun(X393,fun(X394,X395)),bool),happ(s(fun(fun(X393,fun(X394,X395)),fun(fun(X393,fun(X394,X395)),bool)),'=_0'),s(fun(X393,fun(X394,X395)),X403))),s(fun(X393,fun(X394,X395)),happ(s(fun(X400,fun(X393,fun(X394,X395))),'_dst_x000x001x100_3'(s(fun(X396,fun(X397,X395)),X398),s(fun(X394,X396),X399),s(fun(X400,X397),X401))),s(X400,X402)))))) | s(fun(X393,fun(X394,X395)),X403) = s(fun(X393,fun(X394,X395)),'_dst_x00x01x00_3'(s(fun(X396,fun(X397,X395)),X398),s(fun(X394,X396),X399),s(X397,happ(s(fun(X400,X397),X401),s(X400,X402)))))) )),
  inference(superposition,[],[f172,f201])).
fof(f8284,plain,(
  ( ! [X383,X389,X391,X385,X387,X382,X388,X390,X384,X386,X392] : (~p(s(bool,happ(s(fun(fun(X382,fun(X383,X384)),bool),happ(s(fun(fun(X382,fun(X383,X384)),fun(fun(X382,fun(X383,X384)),bool)),'=_0'),s(fun(X382,fun(X383,X384)),happ(s(fun(X389,fun(X382,fun(X383,X384))),'_dst_x000x001x100_3'(s(fun(X385,fun(X386,X384)),X387),s(fun(X383,X385),X388),s(fun(X389,X386),X390))),s(X389,X391))))),s(fun(X382,fun(X383,X384)),X392)))) | s(fun(X382,fun(X383,X384)),X392) = s(fun(X382,fun(X383,X384)),'_dst_x00x01x00_3'(s(fun(X385,fun(X386,X384)),X387),s(fun(X383,X385),X388),s(X386,happ(s(fun(X389,X386),X390),s(X389,X391)))))) )),
  inference(superposition,[],[f172,f201])).
fof(f8256,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X1,X2),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X1,X3),X6),s(X4,happ(s(fun(X7,X4),X8),s(X7,X9))))) = s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),happ(s(fun(X7,fun(X0,fun(X1,X2))),'_dst_x000x001x100_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X1,X3),X6),s(fun(X7,X4),X8))),s(X7,X9))),s(X0,X10)))) )),
  inference(superposition,[],[f198,f201])).
fof(f8255,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (s(fun(X10,fun(X11,X12)),'_dst_x00x01x00_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X11,X14),X17),s(X15,happ(s(fun(X13,X15),X18),s(X13,happ(s(fun(X19,X13),X20),s(X19,X21))))))) = s(fun(X10,fun(X11,X12)),happ(s(fun(X19,fun(X10,fun(X11,X12))),'_dst_x0x1_2'(s(fun(X13,fun(X10,fun(X11,X12))),'_dst_x000x001x100_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X11,X14),X17),s(fun(X13,X15),X18))),s(fun(X19,X13),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f201,f182])).
fof(f8127,plain,(
  ( ! [X350,X344,X356,X346,X352,X354,X343,X349,X351,X345,X357,X347,X353,X355,X342,X348] : (s(fun(X352,fun(X353,X354)),'_dst_x00x01x00_3'(s(fun(X355,fun(fun(X342,fun(X343,X344)),X354)),X356),s(fun(X353,X355),X357),s(fun(X342,fun(X343,X344)),'_dst_x00x01x00_3'(s(fun(X346,fun(X347,X344)),X348),s(fun(X343,X346),X349),s(X347,happ(s(fun(X345,X347),X350),s(X345,X351))))))) = s(fun(X352,fun(X353,X354)),happ(s(fun(X345,fun(X352,fun(X353,X354))),'_dst_x000x001x100_3'(s(fun(X355,fun(fun(X342,fun(X343,X344)),X354)),X356),s(fun(X353,X355),X357),s(fun(X345,fun(X342,fun(X343,X344))),'_dst_x000x001x100_3'(s(fun(X346,fun(X347,X344)),X348),s(fun(X343,X346),X349),s(fun(X345,X347),X350))))),s(X345,X351)))) )),
  inference(superposition,[],[f201,f201])).
fof(f8126,plain,(
  ( ! [X333,X335,X329,X331,X341,X337,X339,X332,X334,X328,X330,X340,X336,X338] : (s(fun(X336,fun(X337,X338)),'_dst_x00x01x00_3'(s(fun(X339,fun(fun(X328,fun(X329,X330)),X338)),X340),s(fun(X337,X339),X341),s(fun(X328,fun(X329,X330)),'_dst_x00x11_2'(s(fun(X332,X330),X333),s(fun(X328,fun(X329,X332)),happ(s(fun(X331,fun(X328,fun(X329,X332))),X334),s(X331,X335))))))) = s(fun(X336,fun(X337,X338)),happ(s(fun(X331,fun(X336,fun(X337,X338))),'_dst_x000x001x100_3'(s(fun(X339,fun(fun(X328,fun(X329,X330)),X338)),X340),s(fun(X337,X339),X341),s(fun(X331,fun(X328,fun(X329,X330))),'_dst_x000x111_2'(s(fun(X332,X330),X333),s(fun(X331,fun(X328,fun(X329,X332))),X334))))),s(X331,X335)))) )),
  inference(superposition,[],[f201,f197])).
fof(f8125,plain,(
  ( ! [X313,X315,X325,X327,X321,X323,X316,X318,X314,X324,X326,X320,X322,X317,X319] : (s(fun(X322,fun(X323,X324)),'_dst_x00x01x00_3'(s(fun(X325,fun(fun(X313,X314),X324)),X326),s(fun(X323,X325),X327),s(fun(X313,X314),'_dst_x0x0x1_3'(s(fun(X316,fun(X317,X314)),X318),s(X316,X319),s(fun(X313,X317),happ(s(fun(X315,fun(X313,X317)),X320),s(X315,X321))))))) = s(fun(X322,fun(X323,X324)),happ(s(fun(X315,fun(X322,fun(X323,X324))),'_dst_x000x001x100_3'(s(fun(X325,fun(fun(X313,X314),X324)),X326),s(fun(X323,X325),X327),s(fun(X315,fun(X313,X314)),'_dst_x00x00x11_3'(s(fun(X316,fun(X317,X314)),X318),s(X316,X319),s(fun(X315,fun(X313,X317)),X320))))),s(X315,X321)))) )),
  inference(superposition,[],[f201,f199])).
fof(f8124,plain,(
  ( ! [X300,X302,X298,X308,X310,X304,X306,X312,X301,X303,X299,X309,X311,X305,X307] : (s(fun(X307,fun(X308,X309)),'_dst_x00x01x00_3'(s(fun(X310,fun(fun(X298,X299),X309)),X311),s(fun(X308,X310),X312),s(fun(X298,X299),'_dst_x0x1x0_3'(s(fun(X301,fun(X302,X299)),X303),s(fun(X298,X301),X304),s(X302,X305))))) = s(fun(X307,fun(X308,X309)),happ(s(fun(X300,fun(X307,fun(X308,X309))),'_dst_x000x001x100_3'(s(fun(X310,fun(fun(X298,X299),X309)),X311),s(fun(X308,X310),X312),s(fun(X300,fun(X298,X299)),'_dst_x00x01x00_3'(s(fun(X301,fun(X302,X299)),X303),s(fun(X298,X301),X304),s(X302,X305))))),s(X300,X306)))) )),
  inference(superposition,[],[f201,f198])).
fof(f8123,plain,(
  ( ! [X292,X294,X288,X290,X296,X285,X287,X293,X295,X289,X291,X297,X286] : (s(fun(X292,fun(X293,X294)),'_dst_x00x01x00_3'(s(fun(X295,fun(fun(X285,X286),X294)),X296),s(fun(X293,X295),X297),s(fun(X285,X286),'_dst_x0x1_2'(s(fun(X288,X286),X289),s(fun(X285,X288),happ(s(fun(X287,fun(X285,X288)),X290),s(X287,X291))))))) = s(fun(X292,fun(X293,X294)),happ(s(fun(X287,fun(X292,fun(X293,X294))),'_dst_x000x001x100_3'(s(fun(X295,fun(fun(X285,X286),X294)),X296),s(fun(X293,X295),X297),s(fun(X287,fun(X285,X286)),'_dst_x00x11_2'(s(fun(X288,X286),X289),s(fun(X287,fun(X285,X288)),X290))))),s(X287,X291)))) )),
  inference(superposition,[],[f201,f193])).
fof(f8122,plain,(
  ( ! [X280,X282,X271,X277,X279,X273,X275,X281,X283,X270,X276,X278,X272,X274,X284] : (s(fun(X279,fun(X280,X281)),'_dst_x00x01x00_3'(s(fun(X282,fun(fun(X270,X271),X281)),X283),s(fun(X280,X282),X284),s(fun(X270,X271),'_dst_x0x1x1_3'(s(fun(X273,fun(X274,X271)),X275),s(fun(X270,X273),happ(s(fun(X272,fun(X270,X273)),X276),s(X272,X278))),s(fun(X270,X274),X277))))) = s(fun(X279,fun(X280,X281)),happ(s(fun(X272,fun(X279,fun(X280,X281))),'_dst_x000x001x100_3'(s(fun(X282,fun(fun(X270,X271),X281)),X283),s(fun(X280,X282),X284),s(fun(X272,fun(X270,X271)),'_dst_x00x11x01_3'(s(fun(X273,fun(X274,X271)),X275),s(fun(X272,fun(X270,X273)),X276),s(fun(X270,X274),X277))))),s(X272,X278)))) )),
  inference(superposition,[],[f201,f200])).
fof(f8121,plain,(
  ( ! [X156,X154,X152,X151,X149,X147,X157,X155,X153,X150,X148] : (s(fun(X152,fun(X153,X154)),'_dst_x00x01x00_3'(s(fun(X155,fun(X147,X154)),X156),s(fun(X153,X155),X157),s(X147,'const.pair.UNCURRY_2'(s(fun(X148,fun(X149,X147)),X150),s('type.pair.prod'(X148,X149),X151))))) = s(fun(X152,fun(X153,X154)),happ(s(fun('type.pair.prod'(X148,X149),fun(X152,fun(X153,X154))),'_dst_x000x001x100_3'(s(fun(X155,fun(X147,X154)),X156),s(fun(X153,X155),X157),s(fun('type.pair.prod'(X148,X149),X147),happ(s(fun(fun(X148,fun(X149,X147)),fun('type.pair.prod'(X148,X149),X147)),'const.pair.UNCURRY_0'),s(fun(X148,fun(X149,X147)),X150))))),s('type.pair.prod'(X148,X149),X151)))) )),
  inference(superposition,[],[f201,f180])).
fof(f8120,plain,(
  ( ! [X127,X125,X123,X134,X132,X130,X128,X126,X124,X122,X135,X133,X131,X129] : (s(fun(X130,fun(X131,X132)),'_dst_x00x01x00_3'(s(fun(X133,fun(X122,X132)),X134),s(fun(X131,X133),X135),s(X122,happ(s(fun(X125,X122),happ(s(fun(X124,fun(X125,X122)),X126),s(X124,X127))),s(X125,happ(s(fun(X123,X125),X128),s(X123,X129))))))) = s(fun(X130,fun(X131,X132)),happ(s(fun(X123,fun(X130,fun(X131,X132))),'_dst_x000x001x100_3'(s(fun(X133,fun(X122,X132)),X134),s(fun(X131,X133),X135),s(fun(X123,X122),'_dst_x0x0x1_3'(s(fun(X124,fun(X125,X122)),X126),s(X124,X127),s(fun(X123,X125),X128))))),s(X123,X129)))) )),
  inference(superposition,[],[f201,f194])).
fof(f8119,plain,(
  ( ! [X121,X118,X116,X114,X112,X120,X119,X117,X115,X113] : (s(fun(X116,fun(X117,X118)),'_dst_x00x01x00_3'(s(fun(X119,fun(X112,X118)),X120),s(fun(X117,X119),X121),s(X112,X114))) = s(fun(X116,fun(X117,X118)),happ(s(fun(X113,fun(X116,fun(X117,X118))),'_dst_x000x001x100_3'(s(fun(X119,fun(X112,X118)),X120),s(fun(X117,X119),X121),s(fun(X113,X112),'_dst_x0_1'(s(X112,X114))))),s(X113,X115)))) )),
  inference(superposition,[],[f201,f170])).
fof(f8118,plain,(
  ( ! [X103,X101,X99,X111,X109,X107,X105,X102,X100,X98,X110,X108,X106,X104] : (s(fun(X106,fun(X107,X108)),'_dst_x00x01x00_3'(s(fun(X109,fun(X98,X108)),X110),s(fun(X107,X109),X111),s(X98,happ(s(fun(X101,X98),happ(s(fun(X100,fun(X101,X98)),X102),s(X100,happ(s(fun(X99,X100),X103),s(X99,X105))))),s(X101,X104))))) = s(fun(X106,fun(X107,X108)),happ(s(fun(X99,fun(X106,fun(X107,X108))),'_dst_x000x001x100_3'(s(fun(X109,fun(X98,X108)),X110),s(fun(X107,X109),X111),s(fun(X99,X98),'_dst_x0x1x0_3'(s(fun(X100,fun(X101,X98)),X102),s(fun(X99,X100),X103),s(X101,X104))))),s(X99,X105)))) )),
  inference(superposition,[],[f201,f195])).
fof(f8117,plain,(
  ( ! [X94,X92,X90,X88,X97,X87,X85,X95,X93,X91,X89,X96,X86,X84] : (s(fun(X92,fun(X93,X94)),'_dst_x00x01x00_3'(s(fun(X95,fun(X84,X94)),X96),s(fun(X93,X95),X97),s(X84,happ(s(fun(X88,X84),'_dst_x0x1x1_3'(s(fun(X86,fun(X85,X84)),X87),s(fun(X88,X86),X89),s(fun(X88,X85),X91))),s(X88,X90))))) = s(fun(X92,fun(X93,X94)),happ(s(fun(X85,fun(X92,fun(X93,X94))),'_dst_x000x001x100_3'(s(fun(X95,fun(X84,X94)),X96),s(fun(X93,X95),X97),s(fun(X85,X84),happ(s(fun(X86,fun(X85,X84)),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90))))))),s(X85,happ(s(fun(X88,X85),X91),s(X88,X90)))))) )),
  inference(superposition,[],[f201,f196])).
fof(f8116,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X83,X81,X71,X79,X77,X75,X73,X82] : (s(fun(X78,fun(X79,X80)),'_dst_x00x01x00_3'(s(fun(X81,fun(X70,X80)),X82),s(fun(X79,X81),X83),s(X70,happ(s(fun(X74,X70),'_dst_x0x1x0_3'(s(fun(X72,fun(X71,X70)),X73),s(fun(X74,X72),X75),s(X71,X77))),s(X74,X76))))) = s(fun(X78,fun(X79,X80)),happ(s(fun(X71,fun(X78,fun(X79,X80))),'_dst_x000x001x100_3'(s(fun(X81,fun(X70,X80)),X82),s(fun(X79,X81),X83),s(fun(X71,X70),happ(s(fun(X72,fun(X71,X70)),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76))))))),s(X71,X77)))) )),
  inference(superposition,[],[f201,f195])).
fof(f8115,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X56,X69,X67,X65,X63] : (s(fun(X64,fun(X65,X66)),'_dst_x00x01x00_3'(s(fun(X67,fun(X56,X66)),X68),s(fun(X65,X67),X69),s(X56,happ(s(fun(X61,X56),'_dst_x0x0x1_3'(s(fun(X58,fun(X57,X56)),X59),s(X58,X60),s(fun(X61,X57),X62))),s(X61,X63))))) = s(fun(X64,fun(X65,X66)),happ(s(fun(X57,fun(X64,fun(X65,X66))),'_dst_x000x001x100_3'(s(fun(X67,fun(X56,X66)),X68),s(fun(X65,X67),X69),s(fun(X57,X56),happ(s(fun(X58,fun(X57,X56)),X59),s(X58,X60))))),s(X57,happ(s(fun(X61,X57),X62),s(X61,X63)))))) )),
  inference(superposition,[],[f201,f194])).
fof(f8113,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X32,X42,X40] : (s(fun(X38,fun(X39,X40)),'_dst_x00x01x00_3'(s(fun(X41,fun(X32,X40)),X42),s(fun(X39,X41),X43),s(X32,'const.pair.pair_CASE_2'(s('type.pair.prod'(X34,X33),'const.pair.,_2'(s(X34,X36),s(X33,X37))),s(fun(X34,fun(X33,X32)),X35))))) = s(fun(X38,fun(X39,X40)),happ(s(fun(X33,fun(X38,fun(X39,X40))),'_dst_x000x001x100_3'(s(fun(X41,fun(X32,X40)),X42),s(fun(X39,X41),X43),s(fun(X33,X32),happ(s(fun(X34,fun(X33,X32)),X35),s(X34,X36))))),s(X33,X37)))) )),
  inference(superposition,[],[f201,f184])).
fof(f8112,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X22,X20] : (s(fun(X26,fun(X27,X28)),'_dst_x00x01x00_3'(s(fun(X29,fun(X20,X28)),X30),s(fun(X27,X29),X31),s(X20,happ(s(fun(X22,X20),X23),s(X22,happ(s(fun(X21,X22),X24),s(X21,X25))))))) = s(fun(X26,fun(X27,X28)),happ(s(fun(X21,fun(X26,fun(X27,X28))),'_dst_x000x001x100_3'(s(fun(X29,fun(X20,X28)),X30),s(fun(X27,X29),X31),s(fun(X21,X20),'_dst_x0x1_2'(s(fun(X22,X20),X23),s(fun(X21,X22),X24))))),s(X21,X25)))) )),
  inference(superposition,[],[f201,f182])).
fof(f8111,plain,(
  ( ! [X14,X12,X10,X8,X19,X17,X15,X13,X11,X9,X18,X16] : (s(fun(X14,fun(X15,X16)),'_dst_x00x01x00_3'(s(fun(X17,fun(X8,X16)),X18),s(fun(X15,X17),X19),s(X8,happ(s(fun(X11,X8),'_dst_x0x1_2'(s(fun(X9,X8),X10),s(fun(X11,X9),X12))),s(X11,X13))))) = s(fun(X14,fun(X15,X16)),happ(s(fun(X9,fun(X14,fun(X15,X16))),'_dst_x000x001x100_3'(s(fun(X17,fun(X8,X16)),X18),s(fun(X15,X17),X19),s(fun(X9,X8),X10))),s(X9,happ(s(fun(X11,X9),X12),s(X11,X13)))))) )),
  inference(superposition,[],[f201,f182])).
fof(f8109,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X1,X3),X6),s(X4,happ(s(fun(X7,X4),X8),s(X7,X9))))))),s(fun(X0,fun(X1,X2)),happ(s(fun(X7,fun(X0,fun(X1,X2))),'_dst_x000x001x100_3'(s(fun(X3,fun(X4,X2)),X5),s(fun(X1,X3),X6),s(fun(X7,X4),X8))),s(X7,X9))))))) )),
  inference(unit_resulting_resolution,[],[f201,f173])).
fof(f8110,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),happ(s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x001x100_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X1,X4),X7),s(fun(X3,X5),X8))),s(X3,X9))))),s(fun(X0,fun(X1,X2)),'_dst_x00x01x00_3'(s(fun(X4,fun(X5,X2)),X6),s(fun(X1,X4),X7),s(X5,happ(s(fun(X3,X5),X8),s(X3,X9))))))))) )),
  inference(unit_resulting_resolution,[],[f201,f173])).
fof(f8030,plain,(
  s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0')))),
  inference(unit_resulting_resolution,[],[f184,f918,f638])).
fof(f8031,plain,(
  s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1)))))))),
  inference(unit_resulting_resolution,[],[f183,f918,f638])).
fof(f7911,plain,(
  s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0')))),
  inference(unit_resulting_resolution,[],[f184,f898,f638])).
fof(f7912,plain,(
  s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.pair.prod'('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.pair.,_2'(s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5)))))))),
  inference(unit_resulting_resolution,[],[f183,f898,f638])).
fof(f7842,plain,(
  ( ! [X536,X538,X532,X534,X530,X537,X539,X533,X535,X531] : (p(s(bool,happ(s(fun(fun(X530,X531),bool),happ(s(fun(fun(X530,X531),fun(fun(X530,X531),bool)),'=_0'),s(fun(X530,X531),X539))),s(fun(X530,X531),'_dst_x0x1x1_3'(s(fun(X533,fun(X534,X531)),X535),s(fun(X530,X533),happ(s(fun(X532,fun(X530,X533)),X536),s(X532,X538))),s(fun(X530,X534),X537)))))) | s(fun(X530,X531),X539) != s(fun(X530,X531),happ(s(fun(X532,fun(X530,X531)),'_dst_x00x11x01_3'(s(fun(X533,fun(X534,X531)),X535),s(fun(X532,fun(X530,X533)),X536),s(fun(X530,X534),X537))),s(X532,X538)))) )),
  inference(superposition,[],[f173,f200])).
fof(f7841,plain,(
  ( ! [X528,X524,X526,X520,X522,X529,X525,X527,X521,X523] : (p(s(bool,happ(s(fun(fun(X520,X521),bool),happ(s(fun(fun(X520,X521),fun(fun(X520,X521),bool)),'=_0'),s(fun(X520,X521),'_dst_x0x1x1_3'(s(fun(X523,fun(X524,X521)),X525),s(fun(X520,X523),happ(s(fun(X522,fun(X520,X523)),X526),s(X522,X528))),s(fun(X520,X524),X527))))),s(fun(X520,X521),X529)))) | s(fun(X520,X521),X529) != s(fun(X520,X521),happ(s(fun(X522,fun(X520,X521)),'_dst_x00x11x01_3'(s(fun(X523,fun(X524,X521)),X525),s(fun(X522,fun(X520,X523)),X526),s(fun(X520,X524),X527))),s(X522,X528)))) )),
  inference(superposition,[],[f173,f200])).
fof(f7840,plain,(
  ( ! [X511,X516,X518,X512,X514,X510,X517,X519,X513,X515] : (~p(s(bool,happ(s(fun(fun(X510,X511),bool),happ(s(fun(fun(X510,X511),fun(fun(X510,X511),bool)),'=_0'),s(fun(X510,X511),X519))),s(fun(X510,X511),'_dst_x0x1x1_3'(s(fun(X513,fun(X514,X511)),X515),s(fun(X510,X513),happ(s(fun(X512,fun(X510,X513)),X516),s(X512,X518))),s(fun(X510,X514),X517)))))) | s(fun(X510,X511),X519) = s(fun(X510,X511),happ(s(fun(X512,fun(X510,X511)),'_dst_x00x11x01_3'(s(fun(X513,fun(X514,X511)),X515),s(fun(X512,fun(X510,X513)),X516),s(fun(X510,X514),X517))),s(X512,X518)))) )),
  inference(superposition,[],[f172,f200])).
fof(f7839,plain,(
  ( ! [X509,X505,X507,X500,X502,X508,X504,X506,X501,X503] : (~p(s(bool,happ(s(fun(fun(X500,X501),bool),happ(s(fun(fun(X500,X501),fun(fun(X500,X501),bool)),'=_0'),s(fun(X500,X501),'_dst_x0x1x1_3'(s(fun(X503,fun(X504,X501)),X505),s(fun(X500,X503),happ(s(fun(X502,fun(X500,X503)),X506),s(X502,X508))),s(fun(X500,X504),X507))))),s(fun(X500,X501),X509)))) | s(fun(X500,X501),X509) = s(fun(X500,X501),happ(s(fun(X502,fun(X500,X501)),'_dst_x00x11x01_3'(s(fun(X503,fun(X504,X501)),X505),s(fun(X502,fun(X500,X503)),X506),s(fun(X500,X504),X507))),s(X502,X508)))) )),
  inference(superposition,[],[f172,f200])).
fof(f7834,plain,(
  ( ! [X447,X453,X455,X449,X451,X457,X452,X454,X448,X450,X456] : (s(X456,happ(s(fun(fun(X447,X448),X456),X457),s(fun(X447,X448),'_dst_x0x1x1_3'(s(fun(X450,fun(X451,X448)),X452),s(fun(X447,X450),happ(s(fun(X449,fun(X447,X450)),X453),s(X449,X455))),s(fun(X447,X451),X454))))) = s(X456,happ(s(fun(X449,X456),'_dst_x0x1_2'(s(fun(fun(X447,X448),X456),X457),s(fun(X449,fun(X447,X448)),'_dst_x00x11x01_3'(s(fun(X450,fun(X451,X448)),X452),s(fun(X449,fun(X447,X450)),X453),s(fun(X447,X451),X454))))),s(X449,X455)))) )),
  inference(superposition,[],[f182,f200])).
fof(f7833,plain,(
  ( ! [X445,X441,X443,X436,X438,X434,X444,X446,X440,X442,X437,X439,X435] : (s(X443,happ(s(fun(X436,X443),'_dst_x0x0x1_3'(s(fun(X444,fun(fun(X434,X435),X443)),X445),s(X444,X446),s(fun(X436,fun(X434,X435)),'_dst_x00x11x01_3'(s(fun(X437,fun(X438,X435)),X439),s(fun(X436,fun(X434,X437)),X440),s(fun(X434,X438),X441))))),s(X436,X442))) = s(X443,happ(s(fun(fun(X434,X435),X443),happ(s(fun(X444,fun(fun(X434,X435),X443)),X445),s(X444,X446))),s(fun(X434,X435),'_dst_x0x1x1_3'(s(fun(X437,fun(X438,X435)),X439),s(fun(X434,X437),happ(s(fun(X436,fun(X434,X437)),X440),s(X436,X442))),s(fun(X434,X438),X441)))))) )),
  inference(superposition,[],[f194,f200])).
fof(f7832,plain,(
  ( ! [X422,X428,X430,X424,X426,X432,X421,X423,X429,X431,X425,X427,X433] : (s(X430,happ(s(fun(X423,X430),'_dst_x0x1x0_3'(s(fun(fun(X421,X422),fun(X431,X430)),X432),s(fun(X423,fun(X421,X422)),'_dst_x00x11x01_3'(s(fun(X424,fun(X425,X422)),X426),s(fun(X423,fun(X421,X424)),X427),s(fun(X421,X425),X428))),s(X431,X433))),s(X423,X429))) = s(X430,happ(s(fun(X431,X430),happ(s(fun(fun(X421,X422),fun(X431,X430)),X432),s(fun(X421,X422),'_dst_x0x1x1_3'(s(fun(X424,fun(X425,X422)),X426),s(fun(X421,X424),happ(s(fun(X423,fun(X421,X424)),X427),s(X423,X429))),s(fun(X421,X425),X428))))),s(X431,X433)))) )),
  inference(superposition,[],[f195,f200])).
fof(f7831,plain,(
  ( ! [X412,X414,X408,X410,X420,X416,X418,X413,X415,X409,X411,X417,X419] : (s(X417,happ(s(fun(X410,X417),'_dst_x0x1x1_3'(s(fun(fun(X408,X409),fun(X418,X417)),X419),s(fun(X410,fun(X408,X409)),'_dst_x00x11x01_3'(s(fun(X411,fun(X412,X409)),X413),s(fun(X410,fun(X408,X411)),X414),s(fun(X408,X412),X415))),s(fun(X410,X418),X420))),s(X410,X416))) = s(X417,happ(s(fun(X418,X417),happ(s(fun(fun(X408,X409),fun(X418,X417)),X419),s(fun(X408,X409),'_dst_x0x1x1_3'(s(fun(X411,fun(X412,X409)),X413),s(fun(X408,X411),happ(s(fun(X410,fun(X408,X411)),X414),s(X410,X416))),s(fun(X408,X412),X415))))),s(X418,happ(s(fun(X410,X418),X420),s(X410,X416)))))) )),
  inference(superposition,[],[f196,f200])).
fof(f7830,plain,(
  ( ! [X397,X399,X395,X405,X407,X401,X403,X396,X398,X404,X406,X400,X402] : (s(X404,happ(s(fun(X397,X404),'_dst_x0x1x1_3'(s(fun(X405,fun(fun(X395,X396),X404)),X406),s(fun(X397,X405),X407),s(fun(X397,fun(X395,X396)),'_dst_x00x11x01_3'(s(fun(X398,fun(X399,X396)),X400),s(fun(X397,fun(X395,X398)),X401),s(fun(X395,X399),X402))))),s(X397,X403))) = s(X404,happ(s(fun(fun(X395,X396),X404),happ(s(fun(X405,fun(fun(X395,X396),X404)),X406),s(X405,happ(s(fun(X397,X405),X407),s(X397,X403))))),s(fun(X395,X396),'_dst_x0x1x1_3'(s(fun(X398,fun(X399,X396)),X400),s(fun(X395,X398),happ(s(fun(X397,fun(X395,X398)),X401),s(X397,X403))),s(fun(X395,X399),X402)))))) )),
  inference(superposition,[],[f196,f200])).
fof(f7829,plain,(
  ( ! [X389,X391,X385,X387,X393,X388,X390,X384,X386,X392,X394] : (~p(s(bool,happ(s(fun('type.pair.prod'(X384,X385),bool),'_dst_x0x1x1_3'(s(fun(X387,fun(X388,bool)),X389),s(fun('type.pair.prod'(X384,X385),X387),happ(s(fun(X386,fun('type.pair.prod'(X384,X385),X387)),X390),s(X386,X392))),s(fun('type.pair.prod'(X384,X385),X388),X391))),s('type.pair.prod'(X384,X385),sK7(happ(s(fun(X386,fun('type.pair.prod'(X384,X385),bool)),'_dst_x00x11x01_3'(s(fun(X387,fun(X388,bool)),X389),s(fun(X386,fun('type.pair.prod'(X384,X385),X387)),X390),s(fun('type.pair.prod'(X384,X385),X388),X391))),s(X386,X392)),X385,X384))))) | p(s(bool,happ(s(fun('type.pair.prod'(X384,X385),bool),happ(s(fun(X386,fun('type.pair.prod'(X384,X385),bool)),'_dst_x00x11x01_3'(s(fun(X387,fun(X388,bool)),X389),s(fun(X386,fun('type.pair.prod'(X384,X385),X387)),X390),s(fun('type.pair.prod'(X384,X385),X388),X391))),s(X386,X392))),s('type.pair.prod'(X384,X385),'const.pair.,_2'(s(X384,X393),s(X385,X394))))))) )),
  inference(superposition,[],[f168,f200])).
fof(f7828,plain,(
  ( ! [X383,X377,X379,X374,X380,X382,X376,X378,X375,X381] : (~p(s(bool,happ(s(fun('type.pair.prod'(X374,X375),bool),'_dst_x0x1x1_3'(s(fun(X377,fun(X378,bool)),X379),s(fun('type.pair.prod'(X374,X375),X377),happ(s(fun(X376,fun('type.pair.prod'(X374,X375),X377)),X380),s(X376,X382))),s(fun('type.pair.prod'(X374,X375),X378),X381))),s('type.pair.prod'(X374,X375),'const.pair.,_2'(s(X374,sK8(happ(s(fun(X376,fun('type.pair.prod'(X374,X375),bool)),'_dst_x00x11x01_3'(s(fun(X377,fun(X378,bool)),X379),s(fun(X376,fun('type.pair.prod'(X374,X375),X377)),X380),s(fun('type.pair.prod'(X374,X375),X378),X381))),s(X376,X382)),X375,X374)),s(X375,sK9(happ(s(fun(X376,fun('type.pair.prod'(X374,X375),bool)),'_dst_x00x11x01_3'(s(fun(X377,fun(X378,bool)),X379),s(fun(X376,fun('type.pair.prod'(X374,X375),X377)),X380),s(fun('type.pair.prod'(X374,X375),X378),X381))),s(X376,X382)),X375,X374))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X374,X375),bool),happ(s(fun(X376,fun('type.pair.prod'(X374,X375),bool)),'_dst_x00x11x01_3'(s(fun(X377,fun(X378,bool)),X379),s(fun(X376,fun('type.pair.prod'(X374,X375),X377)),X380),s(fun('type.pair.prod'(X374,X375),X378),X381))),s(X376,X382))),s('type.pair.prod'(X374,X375),X383))))) )),
  inference(superposition,[],[f169,f200])).
fof(f7806,plain,(
  ( ! [X80,X90,X88,X87,X85,X83,X81,X91,X89,X86,X84,X82] : (s(fun(X80,fun(X81,X90)),'_dst_x00x11_2'(s(fun(X82,X90),X91),s(fun(X80,fun(X81,X82)),'_dst_x0x1x1_3'(s(fun(X84,fun(X85,fun(X81,X82))),X86),s(fun(X80,X84),happ(s(fun(X83,fun(X80,X84)),X87),s(X83,X89))),s(fun(X80,X85),X88))))) = s(fun(X80,fun(X81,X90)),happ(s(fun(X83,fun(X80,fun(X81,X90))),'_dst_x000x111_2'(s(fun(X82,X90),X91),s(fun(X83,fun(X80,fun(X81,X82))),'_dst_x00x11x01_3'(s(fun(X84,fun(X85,fun(X81,X82))),X86),s(fun(X83,fun(X80,X84)),X87),s(fun(X80,X85),X88))))),s(X83,X89)))) )),
  inference(superposition,[],[f197,f200])).
fof(f7805,plain,(
  ( ! [X70,X68,X78,X76,X74,X72,X71,X69,X79,X77,X75,X73] : (s(X69,happ(s(fun(X76,X69),'_dst_x0x1x0_3'(s(fun(X70,fun(X68,X69)),'_dst_x00x11x01_3'(s(fun(X71,fun(X72,X69)),X73),s(fun(X70,fun(X68,X71)),X74),s(fun(X68,X72),X75))),s(fun(X76,X70),X77),s(X68,X79))),s(X76,X78))) = s(X69,happ(s(fun(X68,X69),'_dst_x0x1x1_3'(s(fun(X71,fun(X72,X69)),X73),s(fun(X68,X71),happ(s(fun(X70,fun(X68,X71)),X74),s(X70,happ(s(fun(X76,X70),X77),s(X76,X78))))),s(fun(X68,X72),X75))),s(X68,X79)))) )),
  inference(superposition,[],[f195,f200])).
fof(f7804,plain,(
  ( ! [X61,X59,X57,X66,X64,X62,X60,X58,X56,X67,X65,X63] : (s(X57,happ(s(fun(X64,X57),'_dst_x0x1x1_3'(s(fun(X58,fun(X56,X57)),'_dst_x00x11x01_3'(s(fun(X59,fun(X60,X57)),X61),s(fun(X58,fun(X56,X59)),X62),s(fun(X56,X60),X63))),s(fun(X64,X58),X65),s(fun(X64,X56),X67))),s(X64,X66))) = s(X57,happ(s(fun(X56,X57),'_dst_x0x1x1_3'(s(fun(X59,fun(X60,X57)),X61),s(fun(X56,X59),happ(s(fun(X58,fun(X56,X59)),X62),s(X58,happ(s(fun(X64,X58),X65),s(X64,X66))))),s(fun(X56,X60),X63))),s(X56,happ(s(fun(X64,X56),X67),s(X64,X66)))))) )),
  inference(superposition,[],[f196,f200])).
fof(f7803,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),'_dst_x0x1x1_3'(s(fun(X49,fun(X50,X47)),X51),s(fun(X46,X49),happ(s(fun(X48,fun(X46,X49)),X52),s(X48,X54))),s(fun(X46,X50),X53))),s(X46,X55))) = s(X47,'const.pair.UNCURRY_2'(s(fun(X48,fun(X46,X47)),'_dst_x00x11x01_3'(s(fun(X49,fun(X50,X47)),X51),s(fun(X48,fun(X46,X49)),X52),s(fun(X46,X50),X53))),s('type.pair.prod'(X48,X46),'const.pair.,_2'(s(X48,X54),s(X46,X55)))))) )),
  inference(superposition,[],[f183,f200])).
fof(f7802,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (s(X37,happ(s(fun(X36,X37),'_dst_x0x1x1_3'(s(fun(X39,fun(X40,X37)),X41),s(fun(X36,X39),happ(s(fun(X38,fun(X36,X39)),X42),s(X38,X44))),s(fun(X36,X40),X43))),s(X36,X45))) = s(X37,'const.pair.pair_CASE_2'(s('type.pair.prod'(X38,X36),'const.pair.,_2'(s(X38,X44),s(X36,X45))),s(fun(X38,fun(X36,X37)),'_dst_x00x11x01_3'(s(fun(X39,fun(X40,X37)),X41),s(fun(X38,fun(X36,X39)),X42),s(fun(X36,X40),X43)))))) )),
  inference(superposition,[],[f184,f200])).
fof(f7801,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X25,X34,X32] : (s(fun(X25,X34),'_dst_x0x1_2'(s(fun(X26,X34),X35),s(fun(X25,X26),'_dst_x0x1x1_3'(s(fun(X28,fun(X29,X26)),X30),s(fun(X25,X28),happ(s(fun(X27,fun(X25,X28)),X31),s(X27,X33))),s(fun(X25,X29),X32))))) = s(fun(X25,X34),happ(s(fun(X27,fun(X25,X34)),'_dst_x00x11_2'(s(fun(X26,X34),X35),s(fun(X27,fun(X25,X26)),'_dst_x00x11x01_3'(s(fun(X28,fun(X29,X26)),X30),s(fun(X27,fun(X25,X28)),X31),s(fun(X25,X29),X32))))),s(X27,X33)))) )),
  inference(superposition,[],[f193,f200])).
fof(f7800,plain,(
  ( ! [X24,X14,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X14,happ(s(fun(X22,X14),'_dst_x0x0x1_3'(s(fun(X15,fun(X13,X14)),'_dst_x00x11x01_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X15,fun(X13,X16)),X19),s(fun(X13,X17),X20))),s(X15,X21),s(fun(X22,X13),X23))),s(X22,X24))) = s(X14,happ(s(fun(X13,X14),'_dst_x0x1x1_3'(s(fun(X16,fun(X17,X14)),X18),s(fun(X13,X16),happ(s(fun(X15,fun(X13,X16)),X19),s(X15,X21))),s(fun(X13,X17),X20))),s(X13,happ(s(fun(X22,X13),X23),s(X22,X24)))))) )),
  inference(superposition,[],[f194,f200])).
fof(f7799,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,X9),'_dst_x0x0x1_3'(s(fun(X10,fun(X1,X9)),X11),s(X10,X12),s(fun(X0,X1),'_dst_x0x1x1_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),happ(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(fun(X0,X4),X7))))) = s(fun(X0,X9),happ(s(fun(X2,fun(X0,X9)),'_dst_x00x00x11_3'(s(fun(X10,fun(X1,X9)),X11),s(X10,X12),s(fun(X2,fun(X0,X1)),'_dst_x00x11x01_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(fun(X0,X4),X7))))),s(X2,X8)))) )),
  inference(superposition,[],[f199,f200])).
fof(f7775,plain,(
  ( ! [X364,X366,X372,X368,X370,X365,X367,X373,X369,X371] : (p(s(bool,happ(s(fun(fun(X364,X365),bool),happ(s(fun(fun(X364,X365),fun(fun(X364,X365),bool)),'=_0'),s(fun(X364,X365),X373))),s(fun(X364,X365),happ(s(fun(X369,fun(X364,X365)),'_dst_x00x11x01_3'(s(fun(X366,fun(X367,X365)),X368),s(fun(X369,fun(X364,X366)),X370),s(fun(X364,X367),X372))),s(X369,X371)))))) | s(fun(X364,X365),X373) != s(fun(X364,X365),'_dst_x0x1x1_3'(s(fun(X366,fun(X367,X365)),X368),s(fun(X364,X366),happ(s(fun(X369,fun(X364,X366)),X370),s(X369,X371))),s(fun(X364,X367),X372)))) )),
  inference(superposition,[],[f173,f200])).
fof(f7774,plain,(
  ( ! [X356,X358,X354,X360,X362,X357,X359,X355,X361,X363] : (p(s(bool,happ(s(fun(fun(X354,X355),bool),happ(s(fun(fun(X354,X355),fun(fun(X354,X355),bool)),'=_0'),s(fun(X354,X355),happ(s(fun(X359,fun(X354,X355)),'_dst_x00x11x01_3'(s(fun(X356,fun(X357,X355)),X358),s(fun(X359,fun(X354,X356)),X360),s(fun(X354,X357),X362))),s(X359,X361))))),s(fun(X354,X355),X363)))) | s(fun(X354,X355),X363) != s(fun(X354,X355),'_dst_x0x1x1_3'(s(fun(X356,fun(X357,X355)),X358),s(fun(X354,X356),happ(s(fun(X359,fun(X354,X356)),X360),s(X359,X361))),s(fun(X354,X357),X362)))) )),
  inference(superposition,[],[f173,f200])).
fof(f7773,plain,(
  ( ! [X350,X344,X346,X352,X349,X351,X345,X347,X353,X348] : (~p(s(bool,happ(s(fun(fun(X344,X345),bool),happ(s(fun(fun(X344,X345),fun(fun(X344,X345),bool)),'=_0'),s(fun(X344,X345),X353))),s(fun(X344,X345),happ(s(fun(X349,fun(X344,X345)),'_dst_x00x11x01_3'(s(fun(X346,fun(X347,X345)),X348),s(fun(X349,fun(X344,X346)),X350),s(fun(X344,X347),X352))),s(X349,X351)))))) | s(fun(X344,X345),X353) = s(fun(X344,X345),'_dst_x0x1x1_3'(s(fun(X346,fun(X347,X345)),X348),s(fun(X344,X346),happ(s(fun(X349,fun(X344,X346)),X350),s(X349,X351))),s(fun(X344,X347),X352)))) )),
  inference(superposition,[],[f172,f200])).
fof(f7772,plain,(
  ( ! [X335,X341,X343,X337,X339,X334,X340,X342,X336,X338] : (~p(s(bool,happ(s(fun(fun(X334,X335),bool),happ(s(fun(fun(X334,X335),fun(fun(X334,X335),bool)),'=_0'),s(fun(X334,X335),happ(s(fun(X339,fun(X334,X335)),'_dst_x00x11x01_3'(s(fun(X336,fun(X337,X335)),X338),s(fun(X339,fun(X334,X336)),X340),s(fun(X334,X337),X342))),s(X339,X341))))),s(fun(X334,X335),X343)))) | s(fun(X334,X335),X343) = s(fun(X334,X335),'_dst_x0x1x1_3'(s(fun(X336,fun(X337,X335)),X338),s(fun(X334,X336),happ(s(fun(X339,fun(X334,X336)),X340),s(X339,X341))),s(fun(X334,X337),X342)))) )),
  inference(superposition,[],[f172,f200])).
fof(f7768,plain,(
  ( ! [X292,X294,X300,X302,X296,X298,X293,X295,X301,X297,X299] : (~p(s(bool,happ(s(fun('type.pair.prod'(X292,X293),bool),happ(s(fun(X297,fun('type.pair.prod'(X292,X293),bool)),'_dst_x00x11x01_3'(s(fun(X294,fun(X295,bool)),X296),s(fun(X297,fun('type.pair.prod'(X292,X293),X294)),X298),s(fun('type.pair.prod'(X292,X293),X295),X300))),s(X297,X299))),s('type.pair.prod'(X292,X293),sK7('_dst_x0x1x1_3'(s(fun(X294,fun(X295,bool)),X296),s(fun('type.pair.prod'(X292,X293),X294),happ(s(fun(X297,fun('type.pair.prod'(X292,X293),X294)),X298),s(X297,X299))),s(fun('type.pair.prod'(X292,X293),X295),X300)),X293,X292))))) | p(s(bool,happ(s(fun('type.pair.prod'(X292,X293),bool),'_dst_x0x1x1_3'(s(fun(X294,fun(X295,bool)),X296),s(fun('type.pair.prod'(X292,X293),X294),happ(s(fun(X297,fun('type.pair.prod'(X292,X293),X294)),X298),s(X297,X299))),s(fun('type.pair.prod'(X292,X293),X295),X300))),s('type.pair.prod'(X292,X293),'const.pair.,_2'(s(X292,X301),s(X293,X302))))))) )),
  inference(superposition,[],[f168,f200])).
fof(f7767,plain,(
  ( ! [X282,X288,X290,X285,X287,X283,X289,X291,X284,X286] : (~p(s(bool,happ(s(fun('type.pair.prod'(X282,X283),bool),happ(s(fun(X287,fun('type.pair.prod'(X282,X283),bool)),'_dst_x00x11x01_3'(s(fun(X284,fun(X285,bool)),X286),s(fun(X287,fun('type.pair.prod'(X282,X283),X284)),X288),s(fun('type.pair.prod'(X282,X283),X285),X290))),s(X287,X289))),s('type.pair.prod'(X282,X283),'const.pair.,_2'(s(X282,sK8('_dst_x0x1x1_3'(s(fun(X284,fun(X285,bool)),X286),s(fun('type.pair.prod'(X282,X283),X284),happ(s(fun(X287,fun('type.pair.prod'(X282,X283),X284)),X288),s(X287,X289))),s(fun('type.pair.prod'(X282,X283),X285),X290)),X283,X282)),s(X283,sK9('_dst_x0x1x1_3'(s(fun(X284,fun(X285,bool)),X286),s(fun('type.pair.prod'(X282,X283),X284),happ(s(fun(X287,fun('type.pair.prod'(X282,X283),X284)),X288),s(X287,X289))),s(fun('type.pair.prod'(X282,X283),X285),X290)),X283,X282))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X282,X283),bool),'_dst_x0x1x1_3'(s(fun(X284,fun(X285,bool)),X286),s(fun('type.pair.prod'(X282,X283),X284),happ(s(fun(X287,fun('type.pair.prod'(X282,X283),X284)),X288),s(X287,X289))),s(fun('type.pair.prod'(X282,X283),X285),X290))),s('type.pair.prod'(X282,X283),X291))))) )),
  inference(superposition,[],[f169,f200])).
fof(f7745,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X9,X18,X16] : (s(fun(X9,X10),'_dst_x0x1x1_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),happ(s(fun(X11,fun(X9,X12)),X15),s(X11,happ(s(fun(X17,X11),X18),s(X17,X19))))),s(fun(X9,X13),X16))) = s(fun(X9,X10),happ(s(fun(X17,fun(X9,X10)),'_dst_x0x1_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x11x01_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X11,fun(X9,X12)),X15),s(fun(X9,X13),X16))),s(fun(X17,X11),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f200,f182])).
fof(f7625,plain,(
  ( ! [X198,X206,X204,X202,X200,X199,X197,X205,X203,X201] : (s(fun(X197,X203),'_dst_x0x1x1_3'(s(fun(X198,fun(X204,X203)),X205),s(fun(X197,X198),'const.pair.UNCURRY_2'(s(fun(X199,fun(X200,fun(X197,X198))),X201),s('type.pair.prod'(X199,X200),X202))),s(fun(X197,X204),X206))) = s(fun(X197,X203),happ(s(fun('type.pair.prod'(X199,X200),fun(X197,X203)),'_dst_x00x11x01_3'(s(fun(X198,fun(X204,X203)),X205),s(fun('type.pair.prod'(X199,X200),fun(X197,X198)),happ(s(fun(fun(X199,fun(X200,fun(X197,X198))),fun('type.pair.prod'(X199,X200),fun(X197,X198))),'const.pair.UNCURRY_0'),s(fun(X199,fun(X200,fun(X197,X198))),X201))),s(fun(X197,X204),X206))),s('type.pair.prod'(X199,X200),X202)))) )),
  inference(superposition,[],[f200,f180])).
fof(f7624,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X186,X184,X174,X183,X181,X179] : (s(fun(X174,X183),happ(s(fun(X176,fun(X174,X183)),'_dst_x00x11x01_3'(s(fun(X175,fun(X184,X183)),X185),s(fun(X176,fun(X174,X175)),'_dst_x0x0x1_3'(s(fun(X177,fun(X178,fun(X174,X175))),X179),s(X177,X180),s(fun(X176,X178),X181))),s(fun(X174,X184),X186))),s(X176,X182))) = s(fun(X174,X183),'_dst_x0x1x1_3'(s(fun(X175,fun(X184,X183)),X185),s(fun(X174,X175),happ(s(fun(X178,fun(X174,X175)),happ(s(fun(X177,fun(X178,fun(X174,X175))),X179),s(X177,X180))),s(X178,happ(s(fun(X176,X178),X181),s(X176,X182))))),s(fun(X174,X184),X186)))) )),
  inference(superposition,[],[f200,f194])).
fof(f7623,plain,(
  ( ! [X167,X165,X173,X171,X169,X166,X172,X170,X168] : (s(fun(X165,X170),'_dst_x0x1x1_3'(s(fun(X166,fun(X171,X170)),X172),s(fun(X165,X166),X168),s(fun(X165,X171),X173))) = s(fun(X165,X170),happ(s(fun(X167,fun(X165,X170)),'_dst_x00x11x01_3'(s(fun(X166,fun(X171,X170)),X172),s(fun(X167,fun(X165,X166)),'_dst_x0_1'(s(fun(X165,X166),X168))),s(fun(X165,X171),X173))),s(X167,X169)))) )),
  inference(superposition,[],[f200,f170])).
fof(f7622,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X164,X162,X160] : (s(fun(X152,X161),happ(s(fun(X154,fun(X152,X161)),'_dst_x00x11x01_3'(s(fun(X153,fun(X162,X161)),X163),s(fun(X154,fun(X152,X153)),'_dst_x0x1x0_3'(s(fun(X155,fun(X156,fun(X152,X153))),X157),s(fun(X154,X155),X158),s(X156,X159))),s(fun(X152,X162),X164))),s(X154,X160))) = s(fun(X152,X161),'_dst_x0x1x1_3'(s(fun(X153,fun(X162,X161)),X163),s(fun(X152,X153),happ(s(fun(X156,fun(X152,X153)),happ(s(fun(X155,fun(X156,fun(X152,X153))),X157),s(X155,happ(s(fun(X154,X155),X158),s(X154,X160))))),s(X156,X159))),s(fun(X152,X162),X164)))) )),
  inference(superposition,[],[f200,f195])).
fof(f7621,plain,(
  ( ! [X146,X144,X142,X140,X151,X149,X147,X145,X143,X141,X139,X150,X148] : (s(fun(X139,X148),'_dst_x0x1x1_3'(s(fun(X140,fun(X149,X148)),X150),s(fun(X139,X140),happ(s(fun(X144,fun(X139,X140)),'_dst_x0x1x1_3'(s(fun(X142,fun(X141,fun(X139,X140))),X143),s(fun(X144,X142),X145),s(fun(X144,X141),X147))),s(X144,X146))),s(fun(X139,X149),X151))) = s(fun(X139,X148),happ(s(fun(X141,fun(X139,X148)),'_dst_x00x11x01_3'(s(fun(X140,fun(X149,X148)),X150),s(fun(X141,fun(X139,X140)),happ(s(fun(X142,fun(X141,fun(X139,X140))),X143),s(X142,happ(s(fun(X144,X142),X145),s(X144,X146))))),s(fun(X139,X149),X151))),s(X141,happ(s(fun(X144,X141),X147),s(X144,X146)))))) )),
  inference(superposition,[],[f200,f196])).
fof(f7620,plain,(
  ( ! [X127,X134,X132,X130,X128,X138,X136,X126,X135,X133,X131,X129,X137] : (s(fun(X126,X135),'_dst_x0x1x1_3'(s(fun(X127,fun(X136,X135)),X137),s(fun(X126,X127),happ(s(fun(X131,fun(X126,X127)),'_dst_x0x1x0_3'(s(fun(X129,fun(X128,fun(X126,X127))),X130),s(fun(X131,X129),X132),s(X128,X134))),s(X131,X133))),s(fun(X126,X136),X138))) = s(fun(X126,X135),happ(s(fun(X128,fun(X126,X135)),'_dst_x00x11x01_3'(s(fun(X127,fun(X136,X135)),X137),s(fun(X128,fun(X126,X127)),happ(s(fun(X129,fun(X128,fun(X126,X127))),X130),s(X129,happ(s(fun(X131,X129),X132),s(X131,X133))))),s(fun(X126,X136),X138))),s(X128,X134)))) )),
  inference(superposition,[],[f200,f195])).
fof(f7619,plain,(
  ( ! [X125,X123,X121,X118,X116,X114,X124,X122,X120,X119,X117,X115,X113] : (s(fun(X113,X122),'_dst_x0x1x1_3'(s(fun(X114,fun(X123,X122)),X124),s(fun(X113,X114),happ(s(fun(X119,fun(X113,X114)),'_dst_x0x0x1_3'(s(fun(X116,fun(X115,fun(X113,X114))),X117),s(X116,X118),s(fun(X119,X115),X120))),s(X119,X121))),s(fun(X113,X123),X125))) = s(fun(X113,X122),happ(s(fun(X115,fun(X113,X122)),'_dst_x00x11x01_3'(s(fun(X114,fun(X123,X122)),X124),s(fun(X115,fun(X113,X114)),happ(s(fun(X116,fun(X115,fun(X113,X114))),X117),s(X116,X118))),s(fun(X113,X123),X125))),s(X115,happ(s(fun(X119,X115),X120),s(X119,X121)))))) )),
  inference(superposition,[],[f200,f194])).
fof(f7617,plain,(
  ( ! [X94,X92,X101,X99,X97,X95,X93,X91,X100,X98,X96] : (s(fun(X91,X98),'_dst_x0x1x1_3'(s(fun(X92,fun(X99,X98)),X100),s(fun(X91,X92),'const.pair.pair_CASE_2'(s('type.pair.prod'(X94,X93),'const.pair.,_2'(s(X94,X96),s(X93,X97))),s(fun(X94,fun(X93,fun(X91,X92))),X95))),s(fun(X91,X99),X101))) = s(fun(X91,X98),happ(s(fun(X93,fun(X91,X98)),'_dst_x00x11x01_3'(s(fun(X92,fun(X99,X98)),X100),s(fun(X93,fun(X91,X92)),happ(s(fun(X94,fun(X93,fun(X91,X92))),X95),s(X94,X96))),s(fun(X91,X99),X101))),s(X93,X97)))) )),
  inference(superposition,[],[f200,f184])).
fof(f7616,plain,(
  ( ! [X80,X90,X88,X87,X85,X83,X81,X89,X86,X84,X82] : (s(fun(X80,X87),'_dst_x0x1x1_3'(s(fun(X81,fun(X88,X87)),X89),s(fun(X80,X81),happ(s(fun(X83,fun(X80,X81)),X84),s(X83,happ(s(fun(X82,X83),X85),s(X82,X86))))),s(fun(X80,X88),X90))) = s(fun(X80,X87),happ(s(fun(X82,fun(X80,X87)),'_dst_x00x11x01_3'(s(fun(X81,fun(X88,X87)),X89),s(fun(X82,fun(X80,X81)),'_dst_x0x1_2'(s(fun(X83,fun(X80,X81)),X84),s(fun(X82,X83),X85))),s(fun(X80,X88),X90))),s(X82,X86)))) )),
  inference(superposition,[],[f200,f182])).
fof(f7615,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X69,X79,X77,X75,X73] : (s(fun(X69,X76),happ(s(fun(X71,fun(X69,X76)),'_dst_x00x11x01_3'(s(fun(X70,fun(X77,X76)),X78),s(fun(X71,fun(X69,X70)),X72),s(fun(X69,X77),X79))),s(X71,happ(s(fun(X73,X71),X74),s(X73,X75))))) = s(fun(X69,X76),'_dst_x0x1x1_3'(s(fun(X70,fun(X77,X76)),X78),s(fun(X69,X70),happ(s(fun(X73,fun(X69,X70)),'_dst_x0x1_2'(s(fun(X71,fun(X69,X70)),X72),s(fun(X73,X71),X74))),s(X73,X75))),s(fun(X69,X77),X79)))) )),
  inference(superposition,[],[f200,f182])).
fof(f7614,plain,(
  ( ! [X61,X59,X57,X54,X52,X50,X60,X58,X56,X55,X53,X51] : (s(fun(X50,X58),'_dst_x0x1x1_3'(s(fun(fun(X51,X52),fun(X59,X58)),X60),s(fun(X50,fun(X51,X52)),'_dst_x00x11_2'(s(fun(X54,X52),X55),s(fun(X50,fun(X51,X54)),happ(s(fun(X53,fun(X50,fun(X51,X54))),X56),s(X53,X57))))),s(fun(X50,X59),X61))) = s(fun(X50,X58),happ(s(fun(X53,fun(X50,X58)),'_dst_x00x11x01_3'(s(fun(fun(X51,X52),fun(X59,X58)),X60),s(fun(X53,fun(X50,fun(X51,X52))),'_dst_x000x111_2'(s(fun(X54,X52),X55),s(fun(X53,fun(X50,fun(X51,X54))),X56))),s(fun(X50,X59),X61))),s(X53,X57)))) )),
  inference(superposition,[],[f200,f197])).
fof(f7613,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40,X49] : (s(fun(X37,X46),'_dst_x0x1x1_3'(s(fun(X38,fun(X47,X46)),X48),s(fun(X37,X38),'_dst_x0x0x1_3'(s(fun(X40,fun(X41,X38)),X42),s(X40,X43),s(fun(X37,X41),happ(s(fun(X39,fun(X37,X41)),X44),s(X39,X45))))),s(fun(X37,X47),X49))) = s(fun(X37,X46),happ(s(fun(X39,fun(X37,X46)),'_dst_x00x11x01_3'(s(fun(X38,fun(X47,X46)),X48),s(fun(X39,fun(X37,X38)),'_dst_x00x00x11_3'(s(fun(X40,fun(X41,X38)),X42),s(X40,X43),s(fun(X39,fun(X37,X41)),X44))),s(fun(X37,X47),X49))),s(X39,X45)))) )),
  inference(superposition,[],[f200,f199])).
fof(f7612,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X36,X34,X32] : (s(fun(X24,X33),'_dst_x0x1x1_3'(s(fun(X25,fun(X34,X33)),X35),s(fun(X24,X25),'_dst_x0x1x0_3'(s(fun(X27,fun(X28,X25)),X29),s(fun(X24,X27),X30),s(X28,X31))),s(fun(X24,X34),X36))) = s(fun(X24,X33),happ(s(fun(X26,fun(X24,X33)),'_dst_x00x11x01_3'(s(fun(X25,fun(X34,X33)),X35),s(fun(X26,fun(X24,X25)),'_dst_x00x01x00_3'(s(fun(X27,fun(X28,X25)),X29),s(fun(X24,X27),X30),s(X28,X31))),s(fun(X24,X34),X36))),s(X26,X32)))) )),
  inference(superposition,[],[f200,f198])).
fof(f7611,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(fun(X13,X20),'_dst_x0x1x1_3'(s(fun(X14,fun(X21,X20)),X22),s(fun(X13,X14),'_dst_x0x1_2'(s(fun(X16,X14),X17),s(fun(X13,X16),happ(s(fun(X15,fun(X13,X16)),X18),s(X15,X19))))),s(fun(X13,X21),X23))) = s(fun(X13,X20),happ(s(fun(X15,fun(X13,X20)),'_dst_x00x11x01_3'(s(fun(X14,fun(X21,X20)),X22),s(fun(X15,fun(X13,X14)),'_dst_x00x11_2'(s(fun(X16,X14),X17),s(fun(X15,fun(X13,X16)),X18))),s(fun(X13,X21),X23))),s(X15,X19)))) )),
  inference(superposition,[],[f200,f193])).
fof(f7610,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,X9),'_dst_x0x1x1_3'(s(fun(X1,fun(X10,X9)),X11),s(fun(X0,X1),'_dst_x0x1x1_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),happ(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(fun(X0,X4),X7))),s(fun(X0,X10),X12))) = s(fun(X0,X9),happ(s(fun(X2,fun(X0,X9)),'_dst_x00x11x01_3'(s(fun(X1,fun(X10,X9)),X11),s(fun(X2,fun(X0,X1)),'_dst_x00x11x01_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(fun(X0,X4),X7))),s(fun(X0,X10),X12))),s(X2,X8)))) )),
  inference(superposition,[],[f200,f200])).
fof(f7608,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1x1_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X0,X2),happ(s(fun(X5,fun(X0,X2)),X6),s(X5,X7))),s(fun(X0,X3),X8))))),s(fun(X0,X1),happ(s(fun(X5,fun(X0,X1)),'_dst_x00x11x01_3'(s(fun(X2,fun(X3,X1)),X4),s(fun(X5,fun(X0,X2)),X6),s(fun(X0,X3),X8))),s(X5,X7))))))) )),
  inference(unit_resulting_resolution,[],[f200,f173])).
fof(f7609,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x00x11x01_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,fun(X0,X3)),X6),s(fun(X0,X4),X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x1x1_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),happ(s(fun(X2,fun(X0,X3)),X6),s(X2,X8))),s(fun(X0,X4),X7))))))) )),
  inference(unit_resulting_resolution,[],[f200,f173])).
fof(f7340,plain,(
  ( ! [X912,X918,X916,X914,X911,X909,X915,X913,X919,X917,X910] : (p(s(bool,happ(s(fun('type.pair.prod'(X909,X910),bool),happ(s(fun(X914,fun('type.pair.prod'(X909,X910),bool)),'_dst_x0x1x1_3'(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(fun(X914,X912),X915),s(fun(X914,X911),X917))),s(X914,X916))),s('type.pair.prod'(X909,X910),'const.pair.,_2'(s(X909,X918),s(X910,X919)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X909,X910),bool),happ(s(fun(X914,fun('type.pair.prod'(X909,X910),bool)),'_dst_x0x1x1_3'(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(fun(X914,X912),X915),s(fun(X914,X911),X917))),s(X914,X916))),s('type.pair.prod'(X909,X910),sK7(happ(s(fun(X911,fun('type.pair.prod'(X909,X910),bool)),happ(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(X912,happ(s(fun(X914,X912),X915),s(X914,X916))))),s(X911,happ(s(fun(X914,X911),X917),s(X914,X916)))),X910,X909)))))) )),
  inference(forward_demodulation,[],[f7293,f196])).
fof(f7293,plain,(
  ( ! [X912,X918,X916,X914,X911,X909,X915,X913,X919,X917,X910] : (~p(s(bool,happ(s(fun('type.pair.prod'(X909,X910),bool),happ(s(fun(X914,fun('type.pair.prod'(X909,X910),bool)),'_dst_x0x1x1_3'(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(fun(X914,X912),X915),s(fun(X914,X911),X917))),s(X914,X916))),s('type.pair.prod'(X909,X910),sK7(happ(s(fun(X911,fun('type.pair.prod'(X909,X910),bool)),happ(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(X912,happ(s(fun(X914,X912),X915),s(X914,X916))))),s(X911,happ(s(fun(X914,X911),X917),s(X914,X916)))),X910,X909))))) | p(s(bool,happ(s(fun('type.pair.prod'(X909,X910),bool),happ(s(fun(X911,fun('type.pair.prod'(X909,X910),bool)),happ(s(fun(X912,fun(X911,fun('type.pair.prod'(X909,X910),bool))),X913),s(X912,happ(s(fun(X914,X912),X915),s(X914,X916))))),s(X911,happ(s(fun(X914,X911),X917),s(X914,X916))))),s('type.pair.prod'(X909,X910),'const.pair.,_2'(s(X909,X918),s(X910,X919))))))) )),
  inference(superposition,[],[f168,f196])).
fof(f7339,plain,(
  ( ! [X907,X905,X899,X903,X901,X906,X904,X908,X902,X900] : (p(s(bool,happ(s(fun('type.pair.prod'(X899,X900),bool),happ(s(fun(X904,fun('type.pair.prod'(X899,X900),bool)),'_dst_x0x1x1_3'(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(fun(X904,X902),X905),s(fun(X904,X901),X907))),s(X904,X906))),s('type.pair.prod'(X899,X900),X908)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X899,X900),bool),happ(s(fun(X904,fun('type.pair.prod'(X899,X900),bool)),'_dst_x0x1x1_3'(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(fun(X904,X902),X905),s(fun(X904,X901),X907))),s(X904,X906))),s('type.pair.prod'(X899,X900),'const.pair.,_2'(s(X899,sK8(happ(s(fun(X901,fun('type.pair.prod'(X899,X900),bool)),happ(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(X902,happ(s(fun(X904,X902),X905),s(X904,X906))))),s(X901,happ(s(fun(X904,X901),X907),s(X904,X906)))),X900,X899)),s(X900,sK9(happ(s(fun(X901,fun('type.pair.prod'(X899,X900),bool)),happ(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(X902,happ(s(fun(X904,X902),X905),s(X904,X906))))),s(X901,happ(s(fun(X904,X901),X907),s(X904,X906)))),X900,X899)))))))) )),
  inference(forward_demodulation,[],[f7292,f196])).
fof(f7292,plain,(
  ( ! [X907,X905,X899,X903,X901,X906,X904,X908,X902,X900] : (~p(s(bool,happ(s(fun('type.pair.prod'(X899,X900),bool),happ(s(fun(X904,fun('type.pair.prod'(X899,X900),bool)),'_dst_x0x1x1_3'(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(fun(X904,X902),X905),s(fun(X904,X901),X907))),s(X904,X906))),s('type.pair.prod'(X899,X900),'const.pair.,_2'(s(X899,sK8(happ(s(fun(X901,fun('type.pair.prod'(X899,X900),bool)),happ(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(X902,happ(s(fun(X904,X902),X905),s(X904,X906))))),s(X901,happ(s(fun(X904,X901),X907),s(X904,X906)))),X900,X899)),s(X900,sK9(happ(s(fun(X901,fun('type.pair.prod'(X899,X900),bool)),happ(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(X902,happ(s(fun(X904,X902),X905),s(X904,X906))))),s(X901,happ(s(fun(X904,X901),X907),s(X904,X906)))),X900,X899))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X899,X900),bool),happ(s(fun(X901,fun('type.pair.prod'(X899,X900),bool)),happ(s(fun(X902,fun(X901,fun('type.pair.prod'(X899,X900),bool))),X903),s(X902,happ(s(fun(X904,X902),X905),s(X904,X906))))),s(X901,happ(s(fun(X904,X901),X907),s(X904,X906))))),s('type.pair.prod'(X899,X900),X908))))) )),
  inference(superposition,[],[f169,f196])).
fof(f7273,plain,(
  ( ! [X656,X658,X652,X654,X648,X650,X657,X653,X655,X649,X651,X647] : (s(fun(X647,fun(X648,X657)),'_dst_x00x11_2'(s(fun(X649,X657),X658),s(fun(X647,fun(X648,X649)),happ(s(fun(X653,fun(X647,fun(X648,X649))),'_dst_x0x1x1_3'(s(fun(X651,fun(X650,fun(X647,fun(X648,X649)))),X652),s(fun(X653,X651),X654),s(fun(X653,X650),X656))),s(X653,X655))))) = s(fun(X647,fun(X648,X657)),happ(s(fun(X650,fun(X647,fun(X648,X657))),'_dst_x000x111_2'(s(fun(X649,X657),X658),s(fun(X650,fun(X647,fun(X648,X649))),happ(s(fun(X651,fun(X650,fun(X647,fun(X648,X649)))),X652),s(X651,happ(s(fun(X653,X651),X654),s(X653,X655))))))),s(X650,happ(s(fun(X653,X650),X656),s(X653,X655)))))) )),
  inference(superposition,[],[f197,f196])).
fof(f7272,plain,(
  ( ! [X639,X644,X646,X640,X642,X638,X645,X641,X643,X637] : (s(X638,happ(s(fun(X637,X638),happ(s(fun(X642,fun(X637,X638)),'_dst_x0x1x1_3'(s(fun(X640,fun(X639,fun(X637,X638))),X641),s(fun(X642,X640),X643),s(fun(X642,X639),X645))),s(X642,X644))),s(X637,X646))) = s(X638,'const.pair.UNCURRY_2'(s(fun(X639,fun(X637,X638)),happ(s(fun(X640,fun(X639,fun(X637,X638))),X641),s(X640,happ(s(fun(X642,X640),X643),s(X642,X644))))),s('type.pair.prod'(X639,X637),'const.pair.,_2'(s(X639,happ(s(fun(X642,X639),X645),s(X642,X644))),s(X637,X646)))))) )),
  inference(superposition,[],[f183,f196])).
fof(f7271,plain,(
  ( ! [X633,X635,X629,X631,X627,X636,X632,X634,X628,X630] : (s(X628,happ(s(fun(X627,X628),happ(s(fun(X632,fun(X627,X628)),'_dst_x0x1x1_3'(s(fun(X630,fun(X629,fun(X627,X628))),X631),s(fun(X632,X630),X633),s(fun(X632,X629),X635))),s(X632,X634))),s(X627,X636))) = s(X628,'const.pair.pair_CASE_2'(s('type.pair.prod'(X629,X627),'const.pair.,_2'(s(X629,happ(s(fun(X632,X629),X635),s(X632,X634))),s(X627,X636))),s(fun(X629,fun(X627,X628)),happ(s(fun(X630,fun(X629,fun(X627,X628))),X631),s(X630,happ(s(fun(X632,X630),X633),s(X632,X634)))))))) )),
  inference(superposition,[],[f184,f196])).
fof(f7270,plain,(
  ( ! [X625,X621,X623,X617,X619,X624,X626,X620,X622,X616,X618] : (s(fun(X616,X625),'_dst_x0x1_2'(s(fun(X617,X625),X626),s(fun(X616,X617),happ(s(fun(X621,fun(X616,X617)),'_dst_x0x1x1_3'(s(fun(X619,fun(X618,fun(X616,X617))),X620),s(fun(X621,X619),X622),s(fun(X621,X618),X624))),s(X621,X623))))) = s(fun(X616,X625),happ(s(fun(X618,fun(X616,X625)),'_dst_x00x11_2'(s(fun(X617,X625),X626),s(fun(X618,fun(X616,X617)),happ(s(fun(X619,fun(X618,fun(X616,X617))),X620),s(X619,happ(s(fun(X621,X619),X622),s(X621,X623))))))),s(X618,happ(s(fun(X621,X618),X624),s(X621,X623)))))) )),
  inference(superposition,[],[f193,f196])).
fof(f7269,plain,(
  ( ! [X606,X610,X613,X615,X609,X605,X607,X611,X612,X614,X608,X604] : (s(X605,happ(s(fun(X604,X605),happ(s(fun(X609,fun(X604,X605)),'_dst_x0x1x1_3'(s(fun(X607,fun(X606,fun(X604,X605))),X608),s(fun(X609,X607),X610),s(fun(X609,X606),X612))),s(X609,X611))),s(X604,happ(s(fun(X613,X604),X614),s(X613,X615))))) = s(X605,happ(s(fun(X613,X605),'_dst_x0x0x1_3'(s(fun(X606,fun(X604,X605)),happ(s(fun(X607,fun(X606,fun(X604,X605))),X608),s(X607,happ(s(fun(X609,X607),X610),s(X609,X611))))),s(X606,happ(s(fun(X609,X606),X612),s(X609,X611))),s(fun(X613,X604),X614))),s(X613,X615)))) )),
  inference(superposition,[],[f194,f196])).
fof(f7268,plain,(
  ( ! [X600,X602,X596,X598,X592,X594,X601,X603,X597,X599,X593,X595,X591] : (s(fun(X591,X600),'_dst_x0x0x1_3'(s(fun(X601,fun(X592,X600)),X602),s(X601,X603),s(fun(X591,X592),happ(s(fun(X596,fun(X591,X592)),'_dst_x0x1x1_3'(s(fun(X594,fun(X593,fun(X591,X592))),X595),s(fun(X596,X594),X597),s(fun(X596,X593),X599))),s(X596,X598))))) = s(fun(X591,X600),happ(s(fun(X593,fun(X591,X600)),'_dst_x00x00x11_3'(s(fun(X601,fun(X592,X600)),X602),s(X601,X603),s(fun(X593,fun(X591,X592)),happ(s(fun(X594,fun(X593,fun(X591,X592))),X595),s(X594,happ(s(fun(X596,X594),X597),s(X596,X598))))))),s(X593,happ(s(fun(X596,X593),X599),s(X596,X598)))))) )),
  inference(superposition,[],[f199,f196])).
fof(f7267,plain,(
  ( ! [X588,X590,X584,X586,X582,X589,X585,X587,X581,X583] : (s(X582,happ(s(fun(X581,X582),happ(s(fun(X586,fun(X581,X582)),'_dst_x0x1x1_3'(s(fun(X584,fun(X583,fun(X581,X582))),X585),s(fun(X586,X584),X587),s(fun(X586,X583),X589))),s(X586,X588))),s(X581,X590))) = s(X582,happ(s(fun(X586,X582),'_dst_x0x1x0_3'(s(fun(X583,fun(X581,X582)),happ(s(fun(X584,fun(X583,fun(X581,X582))),X585),s(X584,happ(s(fun(X586,X584),X587),s(X586,X588))))),s(fun(X586,X583),X589),s(X581,X590))),s(X586,X588)))) )),
  inference(superposition,[],[f195,f196])).
fof(f7264,plain,(
  ( ! [X561,X557,X559,X555,X560,X562,X556,X558] : (s(bool,'F_0') != s(bool,happ(s(fun(X558,bool),'_dst_x0x1x1_3'(s(fun(X556,fun(X555,bool)),X557),s(fun(X558,X556),X559),s(fun(X558,X555),X561))),s(X558,X560))) | s(bool,'T_0') = s(bool,X562) | s(bool,X562) = s(bool,happ(s(fun(X555,bool),happ(s(fun(X556,fun(X555,bool)),X557),s(X556,happ(s(fun(X558,X556),X559),s(X558,X560))))),s(X555,happ(s(fun(X558,X555),X561),s(X558,X560)))))) )),
  inference(superposition,[],[f1176,f196])).
fof(f7262,plain,(
  ( ! [X540,X542,X546,X545,X547,X543,X541,X544] : (s(bool,X547) != s(bool,happ(s(fun(X543,bool),'_dst_x0x1x1_3'(s(fun(X541,fun(X540,bool)),X542),s(fun(X543,X541),X544),s(fun(X543,X540),X546))),s(X543,X545))) | s(bool,'T_0') = s(bool,X547) | s(bool,'F_0') = s(bool,happ(s(fun(X540,bool),happ(s(fun(X541,fun(X540,bool)),X542),s(X541,happ(s(fun(X543,X541),X544),s(X543,X545))))),s(X540,happ(s(fun(X543,X540),X546),s(X543,X545)))))) )),
  inference(superposition,[],[f638,f196])).
fof(f7261,plain,(
  ( ! [X536,X538,X532,X534,X537,X539,X533,X535] : (s(bool,X539) != s(bool,happ(s(fun(X535,bool),'_dst_x0x1x1_3'(s(fun(X533,fun(X532,bool)),X534),s(fun(X535,X533),X536),s(fun(X535,X532),X538))),s(X535,X537))) | s(bool,'T_0') = s(bool,happ(s(fun(X532,bool),happ(s(fun(X533,fun(X532,bool)),X534),s(X533,happ(s(fun(X535,X533),X536),s(X535,X537))))),s(X532,happ(s(fun(X535,X532),X538),s(X535,X537))))) | s(bool,'F_0') = s(bool,X539)) )),
  inference(superposition,[],[f638,f196])).
fof(f7259,plain,(
  ( ! [X520,X522,X516,X518,X521,X523,X517,X519,X515] : (s(bool,'F_0') = s(bool,happ(s(fun(X522,bool),'_dst_x0_1'(s(bool,happ(s(fun(X518,bool),'_dst_x0x1x1_3'(s(fun(X516,fun(X515,bool)),X517),s(fun(X518,X516),X519),s(fun(X518,X515),X521))),s(X518,X520))))),s(X522,X523))) | s(bool,'T_0') = s(bool,happ(s(fun(X515,bool),happ(s(fun(X516,fun(X515,bool)),X517),s(X516,happ(s(fun(X518,X516),X519),s(X518,X520))))),s(X515,happ(s(fun(X518,X515),X521),s(X518,X520)))))) )),
  inference(superposition,[],[f614,f196])).
fof(f7255,plain,(
  ( ! [X484,X486,X482,X488,X490,X485,X487,X483,X489] : (s(bool,'T_0') = s(bool,happ(s(fun(X489,bool),'_dst_x0_1'(s(bool,happ(s(fun(X485,bool),'_dst_x0x1x1_3'(s(fun(X483,fun(X482,bool)),X484),s(fun(X485,X483),X486),s(fun(X485,X482),X488))),s(X485,X487))))),s(X489,X490))) | s(bool,'F_0') = s(bool,happ(s(fun(X482,bool),happ(s(fun(X483,fun(X482,bool)),X484),s(X483,happ(s(fun(X485,X483),X486),s(X485,X487))))),s(X482,happ(s(fun(X485,X482),X488),s(X485,X487)))))) )),
  inference(superposition,[],[f267,f196])).
fof(f7332,plain,(
  ( ! [X476,X478,X480,X475,X477,X479,X481] : (s(bool,'T_0') = s(bool,happ(s(fun(X478,bool),'_dst_x0x1x1_3'(s(fun(X476,fun(X475,bool)),X477),s(fun(X478,X476),X479),s(fun(X478,X475),X481))),s(X478,X480))) | s(bool,'F_0') = s(bool,happ(s(fun(X475,bool),happ(s(fun(X476,fun(X475,bool)),X477),s(X476,happ(s(fun(X478,X476),X479),s(X478,X480))))),s(X475,happ(s(fun(X478,X475),X481),s(X478,X480)))))) )),
  inference(forward_demodulation,[],[f7254,f159])).
fof(f7254,plain,(
  ( ! [X476,X478,X480,X475,X477,X479,X481] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X478,bool),'_dst_x0x1x1_3'(s(fun(X476,fun(X475,bool)),X477),s(fun(X478,X476),X479),s(fun(X478,X475),X481))),s(X478,X480))))) | s(bool,'F_0') = s(bool,happ(s(fun(X475,bool),happ(s(fun(X476,fun(X475,bool)),X477),s(X476,happ(s(fun(X478,X476),X479),s(X478,X480))))),s(X475,happ(s(fun(X478,X475),X481),s(X478,X480)))))) )),
  inference(superposition,[],[f235,f196])).
fof(f7245,plain,(
  ( ! [X412,X414,X416,X413,X415] : (~p(s(bool,happ(s(fun(X413,bool),'_dst_x0x1x1_3'(s(fun(X412,fun(X412,bool)),'=_0'),s(fun(X413,X412),X414),s(fun(X413,X412),X416))),s(X413,X415)))) | s(X412,happ(s(fun(X413,X412),X414),s(X413,X415))) = s(X412,happ(s(fun(X413,X412),X416),s(X413,X415)))) )),
  inference(superposition,[],[f172,f196])).
fof(f7244,plain,(
  ( ! [X408,X410,X407,X409,X411] : (p(s(bool,happ(s(fun(X408,bool),'_dst_x0x1x1_3'(s(fun(X407,fun(X407,bool)),'=_0'),s(fun(X408,X407),X409),s(fun(X408,X407),X411))),s(X408,X410)))) | s(X407,happ(s(fun(X408,X407),X409),s(X408,X410))) != s(X407,happ(s(fun(X408,X407),X411),s(X408,X410)))) )),
  inference(superposition,[],[f173,f196])).
fof(f7210,plain,(
  ( ! [X47,X45,X43,X41,X48,X46,X44,X42,X40,X49] : (s(X48,happ(s(fun(X40,X48),X49),s(X40,happ(s(fun(X44,X40),'_dst_x0x1x1_3'(s(fun(X42,fun(X41,X40)),X43),s(fun(X44,X42),X45),s(fun(X44,X41),X47))),s(X44,X46))))) = s(X48,happ(s(fun(X41,X48),'_dst_x0x1_2'(s(fun(X40,X48),X49),s(fun(X41,X40),happ(s(fun(X42,fun(X41,X40)),X43),s(X42,happ(s(fun(X44,X42),X45),s(X44,X46))))))),s(X41,happ(s(fun(X44,X41),X47),s(X44,X46)))))) )),
  inference(superposition,[],[f182,f196])).
fof(f7209,plain,(
  ( ! [X30,X28,X39,X37,X35,X33,X31,X29,X38,X36,X34,X32] : (s(X36,happ(s(fun(X28,X36),happ(s(fun(X37,fun(X28,X36)),X38),s(X37,X39))),s(X28,happ(s(fun(X32,X28),'_dst_x0x1x1_3'(s(fun(X30,fun(X29,X28)),X31),s(fun(X32,X30),X33),s(fun(X32,X29),X35))),s(X32,X34))))) = s(X36,happ(s(fun(X29,X36),'_dst_x0x0x1_3'(s(fun(X37,fun(X28,X36)),X38),s(X37,X39),s(fun(X29,X28),happ(s(fun(X30,fun(X29,X28)),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34))))))),s(X29,happ(s(fun(X32,X29),X35),s(X32,X34)))))) )),
  inference(superposition,[],[f194,f196])).
fof(f7208,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X27,X25,X22,X20,X18,X16] : (s(X24,happ(s(fun(X25,X24),happ(s(fun(X16,fun(X25,X24)),X26),s(X16,happ(s(fun(X20,X16),'_dst_x0x1x1_3'(s(fun(X18,fun(X17,X16)),X19),s(fun(X20,X18),X21),s(fun(X20,X17),X23))),s(X20,X22))))),s(X25,X27))) = s(X24,happ(s(fun(X17,X24),'_dst_x0x1x0_3'(s(fun(X16,fun(X25,X24)),X26),s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),X19),s(X18,happ(s(fun(X20,X18),X21),s(X20,X22))))),s(X25,X27))),s(X17,happ(s(fun(X20,X17),X23),s(X20,X22)))))) )),
  inference(superposition,[],[f195,f196])).
fof(f7203,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X98,X96] : (s(bool,X98) = s(bool,happ(s(fun(X94,bool),'_dst_x0x1x1_3'(s(fun(X92,fun(X91,bool)),X93),s(fun(X94,X92),X95),s(fun(X94,X91),X97))),s(X94,X96))) | s(bool,'T_0') = s(bool,happ(s(fun(X91,bool),happ(s(fun(X92,fun(X91,bool)),X93),s(X92,happ(s(fun(X94,X92),X95),s(X94,X96))))),s(X91,happ(s(fun(X94,X91),X97),s(X94,X96))))) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f196,f223])).
fof(f7201,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),'_dst_x0x1x1_3'(s(fun(X76,fun(X75,bool)),X77),s(fun(X78,X76),X79),s(fun(X78,X75),X81))),s(X78,X80))) | s(bool,X82) = s(bool,happ(s(fun(X75,bool),happ(s(fun(X76,fun(X75,bool)),X77),s(X76,happ(s(fun(X78,X76),X79),s(X78,X80))))),s(X75,happ(s(fun(X78,X75),X81),s(X78,X80))))) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f196,f223])).
fof(f7199,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X63,bool),'_dst_x0x1x1_3'(s(fun(X62,fun(X62,bool)),'=_0'),s(fun(X63,X62),X64),s(fun(X63,X62),X64))),s(X63,X65)))) )),
  inference(superposition,[],[f196,f744])).
fof(f7198,plain,(
  ( ! [X61,X59,X57,X60,X58,X56,X55] : (s(X55,'const.pair.UNCURRY_2'(s(fun(X56,fun(X57,X55)),happ(s(fun(X58,fun(X56,fun(X57,X55))),X59),s(X58,X60))),s('type.pair.prod'(X56,X57),happ(s(fun(X58,'type.pair.prod'(X56,X57)),X61),s(X58,X60))))) = s(X55,happ(s(fun(X58,X55),'_dst_x0x1x1_3'(s(fun(fun(X56,fun(X57,X55)),fun('type.pair.prod'(X56,X57),X55)),'const.pair.UNCURRY_0'),s(fun(X58,fun(X56,fun(X57,X55))),X59),s(fun(X58,'type.pair.prod'(X56,X57)),X61))),s(X58,X60)))) )),
  inference(superposition,[],[f196,f180])).
fof(f7197,plain,(
  ( ! [X54,X52,X50,X48,X53,X51,X49] : (s(X48,'const.pair.pair_CASE_2'(s('type.pair.prod'(X49,X50),happ(s(fun(X51,'type.pair.prod'(X49,X50)),X52),s(X51,X53))),s(fun(X49,fun(X50,X48)),happ(s(fun(X51,fun(X49,fun(X50,X48))),X54),s(X51,X53))))) = s(X48,happ(s(fun(X51,X48),'_dst_x0x1x1_3'(s(fun('type.pair.prod'(X49,X50),fun(fun(X49,fun(X50,X48)),X48)),'const.pair.pair_CASE_0'),s(fun(X51,'type.pair.prod'(X49,X50)),X52),s(fun(X51,fun(X49,fun(X50,X48))),X54))),s(X51,X53)))) )),
  inference(superposition,[],[f196,f181])).
fof(f7196,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(X40,happ(s(fun(X44,X40),'_dst_x0x1x1_3'(s(fun(X42,fun(X41,X40)),X43),s(fun(X44,X42),X45),s(fun(X44,X41),X47))),s(X44,X46))) = s(X40,happ(s(fun(X44,X40),'_dst_x0x1_2'(s(fun(X41,X40),happ(s(fun(X42,fun(X41,X40)),X43),s(X42,happ(s(fun(X44,X42),X45),s(X44,X46))))),s(fun(X44,X41),X47))),s(X44,X46)))) )),
  inference(superposition,[],[f196,f182])).
fof(f7195,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(X32,happ(s(fun(X36,X32),'_dst_x0x1x1_3'(s(fun(X34,fun(X33,X32)),X35),s(fun(X36,X34),X37),s(fun(X36,X33),X39))),s(X36,X38))) = s(X32,'const.pair.UNCURRY_2'(s(fun(X34,fun(X33,X32)),X35),s('type.pair.prod'(X34,X33),'const.pair.,_2'(s(X34,happ(s(fun(X36,X34),X37),s(X36,X38))),s(X33,happ(s(fun(X36,X33),X39),s(X36,X38)))))))) )),
  inference(superposition,[],[f196,f183])).
fof(f7194,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(X24,happ(s(fun(X28,X24),'_dst_x0x1x1_3'(s(fun(X26,fun(X25,X24)),X27),s(fun(X28,X26),X29),s(fun(X28,X25),X31))),s(X28,X30))) = s(X24,'const.pair.pair_CASE_2'(s('type.pair.prod'(X26,X25),'const.pair.,_2'(s(X26,happ(s(fun(X28,X26),X29),s(X28,X30))),s(X25,happ(s(fun(X28,X25),X31),s(X28,X30))))),s(fun(X26,fun(X25,X24)),X27)))) )),
  inference(superposition,[],[f196,f184])).
fof(f7193,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X20,X16),'_dst_x0x1x1_3'(s(fun(X18,fun(X17,X16)),X19),s(fun(X20,X18),X21),s(fun(X20,X17),X23))),s(X20,X22))) = s(X16,happ(s(fun(X20,X16),'_dst_x0x0x1_3'(s(fun(X18,fun(X17,X16)),X19),s(X18,happ(s(fun(X20,X18),X21),s(X20,X22))),s(fun(X20,X17),X23))),s(X20,X22)))) )),
  inference(superposition,[],[f196,f194])).
fof(f7192,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X12,X8),'_dst_x0x1x1_3'(s(fun(X10,fun(X9,X8)),X11),s(fun(X12,X10),X13),s(fun(X12,X9),X15))),s(X12,X14))) = s(X8,happ(s(fun(X12,X8),'_dst_x0x1x0_3'(s(fun(X10,fun(X9,X8)),X11),s(fun(X12,X10),X13),s(X9,happ(s(fun(X12,X9),X15),s(X12,X14))))),s(X12,X14)))) )),
  inference(superposition,[],[f196,f195])).
fof(f7191,plain,(
  ( ! [X280,X277,X279,X275,X276,X278] : (s(X277,happ(s(fun('type.closSem.state'(X275),X277),'_dst_x0x1x1_3'(s(fun(X278,fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X277)),X279),s(fun('type.closSem.state'(X275),X278),X280),s(fun('type.closSem.state'(X275),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))),s('type.closSem.state'(X275),X276))) = s(X277,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X277),happ(s(fun(X278,fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X277)),X279),s(X278,happ(s(fun('type.closSem.state'(X275),X278),X280),s('type.closSem.state'(X275),X276))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X275),X276)))))) )),
  inference(superposition,[],[f196,f160])).
fof(f7190,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (s(X271,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X271),'_dst_x0x1x1_3'(s(fun(X272,fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),X271)),X273),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X272),X274),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269))),'const.closSem.evaluate_0'))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270))) = s(X271,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),X271),happ(s(fun(X272,fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),X271)),X273),s(X272,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X272),X274),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270)))))) )),
  inference(superposition,[],[f196,f161])).
fof(f7189,plain,(
  ( ! [X261,X263,X257,X259,X265,X267,X260,X262,X258,X268,X264,X266] : (s(X265,happ(s(fun(X260,X265),'_dst_x0x1x1_3'(s(fun(X266,fun(fun(X257,fun(X258,X259)),X265)),X267),s(fun(X260,X266),X268),s(fun(X260,fun(X257,fun(X258,X259))),'_dst_x000x111_2'(s(fun(X261,X259),X262),s(fun(X260,fun(X257,fun(X258,X261))),X263))))),s(X260,X264))) = s(X265,happ(s(fun(fun(X257,fun(X258,X259)),X265),happ(s(fun(X266,fun(fun(X257,fun(X258,X259)),X265)),X267),s(X266,happ(s(fun(X260,X266),X268),s(X260,X264))))),s(fun(X257,fun(X258,X259)),'_dst_x00x11_2'(s(fun(X261,X259),X262),s(fun(X257,fun(X258,X261)),happ(s(fun(X260,fun(X257,fun(X258,X261))),X263),s(X260,X264)))))))) )),
  inference(superposition,[],[f196,f197])).
fof(f7188,plain,(
  ( ! [X255,X253,X251,X249,X246,X244,X254,X256,X250,X248,X252,X247,X245] : (s(X253,happ(s(fun(X246,X253),'_dst_x0x1x1_3'(s(fun(X254,fun(fun(X244,X245),X253)),X255),s(fun(X246,X254),X256),s(fun(X246,fun(X244,X245)),'_dst_x00x00x11_3'(s(fun(X247,fun(X248,X245)),X249),s(X247,X250),s(fun(X246,fun(X244,X248)),X251))))),s(X246,X252))) = s(X253,happ(s(fun(fun(X244,X245),X253),happ(s(fun(X254,fun(fun(X244,X245),X253)),X255),s(X254,happ(s(fun(X246,X254),X256),s(X246,X252))))),s(fun(X244,X245),'_dst_x0x0x1_3'(s(fun(X247,fun(X248,X245)),X249),s(X247,X250),s(fun(X244,X248),happ(s(fun(X246,fun(X244,X248)),X251),s(X246,X252)))))))) )),
  inference(superposition,[],[f196,f199])).
fof(f7187,plain,(
  ( ! [X243,X241,X239,X237,X235,X233,X242,X240,X238,X236,X234] : (s(X240,happ(s(fun(X235,X240),'_dst_x0x1x1_3'(s(fun(X241,fun(fun(X233,X234),X240)),X242),s(fun(X235,X241),X243),s(fun(X235,fun(X233,X234)),'_dst_x00x11_2'(s(fun(X236,X234),X237),s(fun(X235,fun(X233,X236)),X238))))),s(X235,X239))) = s(X240,happ(s(fun(fun(X233,X234),X240),happ(s(fun(X241,fun(fun(X233,X234),X240)),X242),s(X241,happ(s(fun(X235,X241),X243),s(X235,X239))))),s(fun(X233,X234),'_dst_x0x1_2'(s(fun(X236,X234),X237),s(fun(X233,X236),happ(s(fun(X235,fun(X233,X236)),X238),s(X235,X239)))))))) )),
  inference(superposition,[],[f196,f193])).
fof(f7186,plain,(
  ( ! [X222,X220,X231,X229,X227,X225,X223,X221,X230,X228,X226,X224,X232] : (s(X229,happ(s(fun(X222,X229),'_dst_x0x1x1_3'(s(fun(X230,fun(fun(X220,X221),X229)),X231),s(fun(X222,X230),X232),s(fun(X222,fun(X220,X221)),'_dst_x00x01x00_3'(s(fun(X223,fun(X224,X221)),X225),s(fun(X220,X223),X226),s(X224,X227))))),s(X222,X228))) = s(X229,happ(s(fun(fun(X220,X221),X229),happ(s(fun(X230,fun(fun(X220,X221),X229)),X231),s(X230,happ(s(fun(X222,X230),X232),s(X222,X228))))),s(fun(X220,X221),'_dst_x0x1x0_3'(s(fun(X223,fun(X224,X221)),X225),s(fun(X220,X223),X226),s(X224,X227)))))) )),
  inference(superposition,[],[f196,f198])).
fof(f7183,plain,(
  ( ! [X198,X204,X202,X200,X199,X197,X203,X201] : (s(X201,happ(s(fun(X197,X201),'_dst_x0x1x1_3'(s(fun(X202,fun(bool,X201)),X203),s(fun(X197,X202),X204),s(fun(X197,bool),X198))),s(X197,X199))) = s(X201,happ(s(fun(bool,X201),happ(s(fun(X202,fun(bool,X201)),X203),s(X202,happ(s(fun(X197,X202),X204),s(X197,X199))))),s(bool,X200))) | s(bool,'T_0') = s(bool,happ(s(fun(X197,bool),X198),s(X197,X199))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f196,f223])).
fof(f7181,plain,(
  ( ! [X187,X185,X182,X188,X186,X184,X183,X181] : (s(X184,happ(s(fun(X181,X184),'_dst_x0x1x1_3'(s(fun(X185,fun(bool,X184)),X186),s(fun(X181,X185),X187),s(fun(X181,bool),X182))),s(X181,X183))) = s(X184,happ(s(fun(bool,X184),happ(s(fun(X185,fun(bool,X184)),X186),s(X185,happ(s(fun(X181,X185),X187),s(X181,X183))))),s(bool,'T_0'))) | s(bool,X188) = s(bool,happ(s(fun(X181,bool),X182),s(X181,X183))) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f196,f223])).
fof(f7318,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(X160,happ(s(fun(bool,X160),'_dst_x0x1x0_3'(s(fun(X161,fun(bool,X160)),X162),s(fun(bool,X161),X163),s(bool,'T_0'))),s(bool,X159))) = s(X160,happ(s(fun(bool,X160),'_dst_x0x1x1_3'(s(fun(X161,fun(bool,X160)),X162),s(fun(bool,X161),X163),s(fun(bool,bool),'_i_0_0'))),s(bool,X159))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(forward_demodulation,[],[f7177,f195])).
fof(f7177,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(X160,happ(s(fun(bool,X160),'_dst_x0x1x1_3'(s(fun(X161,fun(bool,X160)),X162),s(fun(bool,X161),X163),s(fun(bool,bool),'_i_0_0'))),s(bool,X159))) = s(X160,happ(s(fun(bool,X160),happ(s(fun(X161,fun(bool,X160)),X162),s(X161,happ(s(fun(bool,X161),X163),s(bool,X159))))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f196,f235])).
fof(f7175,plain,(
  ( ! [X146,X151,X149,X147,X145,X150,X148] : (s(X148,happ(s(fun(bool,X148),happ(s(fun(X149,fun(bool,X148)),X150),s(X149,happ(s(fun(X145,X149),X151),s(X145,X147))))),s(bool,'T_0'))) = s(X148,happ(s(fun(X145,X148),'_dst_x0x1x1_3'(s(fun(X149,fun(bool,X148)),X150),s(fun(X145,X149),X151),s(fun(X145,bool),'_dst_x0_1'(s(bool,X146))))),s(X145,X147))) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f196,f267])).
fof(f7174,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139] : (s(X141,happ(s(fun(bool,X141),happ(s(fun(X142,fun(bool,X141)),X143),s(X142,happ(s(fun(X138,X142),X144),s(X138,X140))))),s(bool,'F_0'))) = s(X141,happ(s(fun(X138,X141),'_dst_x0x1x1_3'(s(fun(X142,fun(bool,X141)),X143),s(fun(X138,X142),X144),s(fun(X138,bool),'_dst_x0_1'(s(bool,X139))))),s(X138,X140))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f196,f614])).
fof(f7173,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(X134,happ(s(fun(bool,X134),happ(s(fun(X135,fun(bool,X134)),X136),s(X135,happ(s(fun(X132,X135),X137),s(X132,X133))))),s(bool,'T_0'))) = s(X134,happ(s(fun(X132,X134),'_dst_x0x1x1_3'(s(fun(X135,fun(bool,X134)),X136),s(fun(X132,X135),X137),s(fun(X132,bool),happ(s(fun(X132,fun(X132,bool)),'=_0'),s(X132,X133))))),s(X132,X133)))) )),
  inference(superposition,[],[f196,f744])).
fof(f7172,plain,(
  ( ! [X127,X125,X123,X130,X128,X126,X124,X131,X129] : (s(X128,happ(s(fun(X123,X128),happ(s(fun(X129,fun(X123,X128)),X130),s(X129,happ(s(fun('type.pair.prod'(X124,X125),X129),X131),s('type.pair.prod'(X124,X125),X127))))),s(X123,'const.pair.UNCURRY_2'(s(fun(X124,fun(X125,X123)),X126),s('type.pair.prod'(X124,X125),X127))))) = s(X128,happ(s(fun('type.pair.prod'(X124,X125),X128),'_dst_x0x1x1_3'(s(fun(X129,fun(X123,X128)),X130),s(fun('type.pair.prod'(X124,X125),X129),X131),s(fun('type.pair.prod'(X124,X125),X123),happ(s(fun(fun(X124,fun(X125,X123)),fun('type.pair.prod'(X124,X125),X123)),'const.pair.UNCURRY_0'),s(fun(X124,fun(X125,X123)),X126))))),s('type.pair.prod'(X124,X125),X127)))) )),
  inference(superposition,[],[f196,f180])).
fof(f7171,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (s(X119,happ(s(fun(X114,X119),happ(s(fun(X120,fun(X114,X119)),X121),s(X120,happ(s(fun(fun(X115,fun(X116,X114)),X120),X122),s(fun(X115,fun(X116,X114)),X118))))),s(X114,'const.pair.pair_CASE_2'(s('type.pair.prod'(X115,X116),X117),s(fun(X115,fun(X116,X114)),X118))))) = s(X119,happ(s(fun(fun(X115,fun(X116,X114)),X119),'_dst_x0x1x1_3'(s(fun(X120,fun(X114,X119)),X121),s(fun(fun(X115,fun(X116,X114)),X120),X122),s(fun(fun(X115,fun(X116,X114)),X114),happ(s(fun('type.pair.prod'(X115,X116),fun(fun(X115,fun(X116,X114)),X114)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X115,X116),X117))))),s(fun(X115,fun(X116,X114)),X118)))) )),
  inference(superposition,[],[f196,f181])).
fof(f7170,plain,(
  ( ! [X103,X111,X109,X107,X105,X112,X102,X110,X108,X106,X104,X113] : (s(X110,happ(s(fun(X103,X110),'_dst_x0x1x1_3'(s(fun(X111,fun(X102,X110)),X112),s(fun(X103,X111),X113),s(fun(X103,X102),'_dst_x0x0x1_3'(s(fun(X104,fun(X105,X102)),X106),s(X104,X107),s(fun(X103,X105),X108))))),s(X103,X109))) = s(X110,happ(s(fun(X102,X110),happ(s(fun(X111,fun(X102,X110)),X112),s(X111,happ(s(fun(X103,X111),X113),s(X103,X109))))),s(X102,happ(s(fun(X105,X102),happ(s(fun(X104,fun(X105,X102)),X106),s(X104,X107))),s(X105,happ(s(fun(X103,X105),X108),s(X103,X109)))))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7169,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (s(X98,happ(s(fun(X94,X98),happ(s(fun(X99,fun(X94,X98)),X100),s(X99,happ(s(fun(X95,X99),X101),s(X95,X97))))),s(X94,X96))) = s(X98,happ(s(fun(X95,X98),'_dst_x0x1x1_3'(s(fun(X99,fun(X94,X98)),X100),s(fun(X95,X99),X101),s(fun(X95,X94),'_dst_x0_1'(s(X94,X96))))),s(X95,X97)))) )),
  inference(superposition,[],[f196,f170])).
fof(f7168,plain,(
  ( ! [X92,X90,X88,X87,X85,X83,X93,X91,X89,X86,X84,X82] : (s(X90,happ(s(fun(X83,X90),'_dst_x0x1x1_3'(s(fun(X91,fun(X82,X90)),X92),s(fun(X83,X91),X93),s(fun(X83,X82),'_dst_x0x1x0_3'(s(fun(X84,fun(X85,X82)),X86),s(fun(X83,X84),X87),s(X85,X88))))),s(X83,X89))) = s(X90,happ(s(fun(X82,X90),happ(s(fun(X91,fun(X82,X90)),X92),s(X91,happ(s(fun(X83,X91),X93),s(X83,X89))))),s(X82,happ(s(fun(X85,X82),happ(s(fun(X84,fun(X85,X82)),X86),s(X84,happ(s(fun(X83,X84),X87),s(X83,X89))))),s(X85,X88)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7167,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73] : (s(X78,happ(s(fun(X71,X78),'_dst_x0x1x1_3'(s(fun(X79,fun(X70,X78)),X80),s(fun(X71,X79),X81),s(fun(X71,X70),happ(s(fun(X72,fun(X71,X70)),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76))))))),s(X71,happ(s(fun(X74,X71),X77),s(X74,X76))))) = s(X78,happ(s(fun(X70,X78),happ(s(fun(X79,fun(X70,X78)),X80),s(X79,happ(s(fun(X71,X79),X81),s(X71,happ(s(fun(X74,X71),X77),s(X74,X76))))))),s(X70,happ(s(fun(X74,X70),'_dst_x0x1x1_3'(s(fun(X72,fun(X71,X70)),X73),s(fun(X74,X72),X75),s(fun(X74,X71),X77))),s(X74,X76)))))) )),
  inference(superposition,[],[f196,f196])).
fof(f7166,plain,(
  ( ! [X61,X59,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s(X66,happ(s(fun(X59,X66),'_dst_x0x1x1_3'(s(fun(X67,fun(X58,X66)),X68),s(fun(X59,X67),X69),s(fun(X59,X58),happ(s(fun(X60,fun(X59,X58)),X61),s(X60,happ(s(fun(X62,X60),X63),s(X62,X64))))))),s(X59,X65))) = s(X66,happ(s(fun(X58,X66),happ(s(fun(X67,fun(X58,X66)),X68),s(X67,happ(s(fun(X59,X67),X69),s(X59,X65))))),s(X58,happ(s(fun(X62,X58),'_dst_x0x1x0_3'(s(fun(X60,fun(X59,X58)),X61),s(fun(X62,X60),X63),s(X59,X65))),s(X62,X64)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7165,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (s(X54,happ(s(fun(X47,X54),'_dst_x0x1x1_3'(s(fun(X55,fun(X46,X54)),X56),s(fun(X47,X55),X57),s(fun(X47,X46),happ(s(fun(X48,fun(X47,X46)),X49),s(X48,X50))))),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53))))) = s(X54,happ(s(fun(X46,X54),happ(s(fun(X55,fun(X46,X54)),X56),s(X55,happ(s(fun(X47,X55),X57),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53))))))),s(X46,happ(s(fun(X51,X46),'_dst_x0x0x1_3'(s(fun(X48,fun(X47,X46)),X49),s(X48,X50),s(fun(X51,X47),X52))),s(X51,X53)))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7163,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (s(X32,happ(s(fun(X27,X32),'_dst_x0x1x1_3'(s(fun(X33,fun(X26,X32)),X34),s(fun(X27,X33),X35),s(fun(X27,X26),happ(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))))),s(X27,X31))) = s(X32,happ(s(fun(X26,X32),happ(s(fun(X33,fun(X26,X32)),X34),s(X33,happ(s(fun(X27,X33),X35),s(X27,X31))))),s(X26,'const.pair.pair_CASE_2'(s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))),s(fun(X28,fun(X27,X26)),X29)))))) )),
  inference(superposition,[],[f196,f184])).
fof(f7162,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X22,happ(s(fun(X17,X22),'_dst_x0x1x1_3'(s(fun(X23,fun(X16,X22)),X24),s(fun(X17,X23),X25),s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X19),s(fun(X17,X18),X20))))),s(X17,X21))) = s(X22,happ(s(fun(X16,X22),happ(s(fun(X23,fun(X16,X22)),X24),s(X23,happ(s(fun(X17,X23),X25),s(X17,X21))))),s(X16,happ(s(fun(X18,X16),X19),s(X18,happ(s(fun(X17,X18),X20),s(X17,X21)))))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7161,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(X12,happ(s(fun(X7,X12),'_dst_x0x1x1_3'(s(fun(X13,fun(X6,X12)),X14),s(fun(X7,X13),X15),s(fun(X7,X6),X8))),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))) = s(X12,happ(s(fun(X6,X12),happ(s(fun(X13,fun(X6,X12)),X14),s(X13,happ(s(fun(X7,X13),X15),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))))),s(X6,happ(s(fun(X9,X6),'_dst_x0x1_2'(s(fun(X7,X6),X8),s(fun(X9,X7),X10))),s(X9,X11)))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7160,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,happ(s(fun(X0,X2),'_dst_x0x1x1_3'(s(fun(X3,fun(X0,X2)),X4),s(fun(X0,X3),X5),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))) = s(X2,happ(s(fun(X0,X2),happ(s(fun(X3,fun(X0,X2)),X4),s(X3,happ(s(fun(X0,X3),X5),s(X0,X1))))),s(X0,X1)))) )),
  inference(superposition,[],[f196,f159])).
fof(f7092,plain,(
  ( ! [X167,X165,X163,X169,X166,X164,X162,X170,X168] : (s(X163,happ(s(fun(X162,X163),'const.pair.UNCURRY_2'(s(fun(X164,fun(X165,fun(X162,X163))),X166),s('type.pair.prod'(X164,X165),happ(s(fun(X167,'type.pair.prod'(X164,X165)),X168),s(X167,X169))))),s(X162,happ(s(fun(X167,X162),X170),s(X167,X169))))) = s(X163,happ(s(fun(X167,X163),'_dst_x0x1x1_3'(s(fun('type.pair.prod'(X164,X165),fun(X162,X163)),happ(s(fun(fun(X164,fun(X165,fun(X162,X163))),fun('type.pair.prod'(X164,X165),fun(X162,X163))),'const.pair.UNCURRY_0'),s(fun(X164,fun(X165,fun(X162,X163))),X166))),s(fun(X167,'type.pair.prod'(X164,X165)),X168),s(fun(X167,X162),X170))),s(X167,X169)))) )),
  inference(superposition,[],[f196,f180])).
fof(f7091,plain,(
  ( ! [X158,X156,X154,X161,X159,X157,X155,X153,X160] : (s(X154,happ(s(fun(X153,X154),'const.pair.pair_CASE_2'(s('type.pair.prod'(X155,X156),X157),s(fun(X155,fun(X156,fun(X153,X154))),happ(s(fun(X158,fun(X155,fun(X156,fun(X153,X154)))),X159),s(X158,X160))))),s(X153,happ(s(fun(X158,X153),X161),s(X158,X160))))) = s(X154,happ(s(fun(X158,X154),'_dst_x0x1x1_3'(s(fun(fun(X155,fun(X156,fun(X153,X154))),fun(X153,X154)),happ(s(fun('type.pair.prod'(X155,X156),fun(fun(X155,fun(X156,fun(X153,X154))),fun(X153,X154))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X155,X156),X157))),s(fun(X158,fun(X155,fun(X156,fun(X153,X154)))),X159),s(fun(X158,X153),X161))),s(X158,X160)))) )),
  inference(superposition,[],[f196,f181])).
fof(f7090,plain,(
  ( ! [X146,X144,X152,X142,X151,X149,X147,X145,X143,X141,X150,X148] : (s(X142,happ(s(fun(X149,X142),'_dst_x0x1x1_3'(s(fun(X143,fun(X141,X142)),'_dst_x0x0x1_3'(s(fun(X144,fun(X145,fun(X141,X142))),X146),s(X144,X147),s(fun(X143,X145),X148))),s(fun(X149,X143),X150),s(fun(X149,X141),X152))),s(X149,X151))) = s(X142,happ(s(fun(X141,X142),happ(s(fun(X145,fun(X141,X142)),happ(s(fun(X144,fun(X145,fun(X141,X142))),X146),s(X144,X147))),s(X145,happ(s(fun(X143,X145),X148),s(X143,happ(s(fun(X149,X143),X150),s(X149,X151))))))),s(X141,happ(s(fun(X149,X141),X152),s(X149,X151)))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7089,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (s(X134,happ(s(fun(X133,X134),X136),s(X133,happ(s(fun(X137,X133),X140),s(X137,X139))))) = s(X134,happ(s(fun(X137,X134),'_dst_x0x1x1_3'(s(fun(X135,fun(X133,X134)),'_dst_x0_1'(s(fun(X133,X134),X136))),s(fun(X137,X135),X138),s(fun(X137,X133),X140))),s(X137,X139)))) )),
  inference(superposition,[],[f196,f170])).
fof(f7088,plain,(
  ( ! [X127,X125,X123,X121,X132,X130,X128,X126,X124,X122,X131,X129] : (s(X122,happ(s(fun(X129,X122),'_dst_x0x1x1_3'(s(fun(X123,fun(X121,X122)),'_dst_x0x1x0_3'(s(fun(X124,fun(X125,fun(X121,X122))),X126),s(fun(X123,X124),X127),s(X125,X128))),s(fun(X129,X123),X130),s(fun(X129,X121),X132))),s(X129,X131))) = s(X122,happ(s(fun(X121,X122),happ(s(fun(X125,fun(X121,X122)),happ(s(fun(X124,fun(X125,fun(X121,X122))),X126),s(X124,happ(s(fun(X123,X124),X127),s(X123,happ(s(fun(X129,X123),X130),s(X129,X131))))))),s(X125,X128))),s(X121,happ(s(fun(X129,X121),X132),s(X129,X131)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7087,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X120,X110,X119,X117,X115,X113] : (s(X110,happ(s(fun(X117,X110),'_dst_x0x1x1_3'(s(fun(X111,fun(X109,X110)),happ(s(fun(X112,fun(X111,fun(X109,X110))),X113),s(X112,happ(s(fun(X114,X112),X115),s(X114,X116))))),s(fun(X117,X111),X118),s(fun(X117,X109),X120))),s(X117,X119))) = s(X110,happ(s(fun(X109,X110),happ(s(fun(X114,fun(X109,X110)),'_dst_x0x1x0_3'(s(fun(X112,fun(X111,fun(X109,X110))),X113),s(fun(X114,X112),X115),s(X111,happ(s(fun(X117,X111),X118),s(X117,X119))))),s(X114,X116))),s(X109,happ(s(fun(X117,X109),X120),s(X117,X119)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7086,plain,(
  ( ! [X103,X101,X99,X107,X105,X102,X100,X108,X106,X104] : (s(X100,happ(s(fun(X99,X100),happ(s(fun(X104,fun(X99,X100)),'_dst_x0x1x1_3'(s(fun(X102,fun(X101,fun(X99,X100))),X103),s(fun(X104,X102),X105),s(fun(X104,X101),X107))),s(X104,X106))),s(X99,happ(s(fun(X104,X99),X108),s(X104,X106))))) = s(X100,happ(s(fun(X104,X100),'_dst_x0x1x1_3'(s(fun(X101,fun(X99,X100)),happ(s(fun(X102,fun(X101,fun(X99,X100))),X103),s(X102,happ(s(fun(X104,X102),X105),s(X104,X106))))),s(fun(X104,X101),X107),s(fun(X104,X99),X108))),s(X104,X106)))) )),
  inference(superposition,[],[f196,f196])).
fof(f7084,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X79,X86,X84,X82] : (s(X80,happ(s(fun(X85,X80),'_dst_x0x1x1_3'(s(fun(X81,fun(X79,X80)),happ(s(fun(X82,fun(X81,fun(X79,X80))),X83),s(X82,X84))),s(fun(X85,X81),X86),s(fun(X85,X79),X88))),s(X85,X87))) = s(X80,happ(s(fun(X79,X80),'const.pair.pair_CASE_2'(s('type.pair.prod'(X82,X81),'const.pair.,_2'(s(X82,X84),s(X81,happ(s(fun(X85,X81),X86),s(X85,X87))))),s(fun(X82,fun(X81,fun(X79,X80))),X83))),s(X79,happ(s(fun(X85,X79),X88),s(X85,X87)))))) )),
  inference(superposition,[],[f196,f184])).
fof(f7083,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X69,X77,X75,X73] : (s(X70,happ(s(fun(X75,X70),'_dst_x0x1x1_3'(s(fun(X71,fun(X69,X70)),happ(s(fun(X72,fun(X71,fun(X69,X70))),X73),s(X72,X74))),s(fun(X75,X71),X76),s(fun(X75,X69),X78))),s(X75,X77))) = s(X70,happ(s(fun(X69,X70),happ(s(fun(X75,fun(X69,X70)),'_dst_x0x0x1_3'(s(fun(X72,fun(X71,fun(X69,X70))),X73),s(X72,X74),s(fun(X75,X71),X76))),s(X75,X77))),s(X69,happ(s(fun(X75,X69),X78),s(X75,X77)))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7314,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X5,bool),'_dst_x0x1x1_3'(s(fun(X1,fun(X0,bool)),'_dst_x0x1_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))))),s(fun(X5,X1),X6),s(fun(X5,X0),'_dst_x0x0x1_3'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4),s(fun(X5,X1),X6))))),s(X5,X7))))) )),
  inference(backward_demodulation,[],[f7082,f4596])).
fof(f7313,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X3,bool),'_dst_x0x1x1_3'(s(fun(X1,fun(X0,bool)),'_dst_x0x1_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X1,X0),X2))),s(fun(X3,X1),X4),s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))))),s(X3,X5))))) )),
  inference(backward_demodulation,[],[f7082,f1043])).
fof(f7082,plain,(
  ( ! [X61,X59,X68,X66,X64,X62,X60,X67,X65,X63] : (s(X60,happ(s(fun(X65,X60),'_dst_x0x1x1_3'(s(fun(X61,fun(X59,X60)),'_dst_x0x1_2'(s(fun(X62,fun(X59,X60)),X63),s(fun(X61,X62),X64))),s(fun(X65,X61),X66),s(fun(X65,X59),X68))),s(X65,X67))) = s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,happ(s(fun(X61,X62),X64),s(X61,happ(s(fun(X65,X61),X66),s(X65,X67))))))),s(X59,happ(s(fun(X65,X59),X68),s(X65,X67)))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7081,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(X52,happ(s(fun(X55,X52),'_dst_x0x1x1_3'(s(fun(X53,fun(X51,X52)),X54),s(fun(X55,X53),X56),s(fun(X55,X51),X58))),s(X55,X57))) = s(X52,happ(s(fun(X51,X52),happ(s(fun(X55,fun(X51,X52)),'_dst_x0x1_2'(s(fun(X53,fun(X51,X52)),X54),s(fun(X55,X53),X56))),s(X55,X57))),s(X51,happ(s(fun(X55,X51),X58),s(X55,X57)))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7080,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(X46,happ(s(fun(X45,X46),happ(s(fun(X47,fun(X45,X46)),X48),s(X47,X49))),s(X45,happ(s(fun(X47,X45),X50),s(X47,X49))))) = s(X46,happ(s(fun(X47,X46),'_dst_x0x1x1_3'(s(fun(fun(X45,X46),fun(X45,X46)),'_i_0_0'),s(fun(X47,fun(X45,X46)),X48),s(fun(X47,X45),X50))),s(X47,X49)))) )),
  inference(superposition,[],[f196,f159])).
fof(f7079,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X44,X42,X40] : (s(fun(X35,X36),happ(s(fun(X41,fun(X35,X36)),'_dst_x0x1x1_3'(s(fun(X37,fun(X34,fun(X35,X36))),'_dst_x000x111_2'(s(fun(X38,X36),X39),s(fun(X37,fun(X34,fun(X35,X38))),X40))),s(fun(X41,X37),X42),s(fun(X41,X34),X44))),s(X41,X43))) = s(fun(X35,X36),happ(s(fun(X34,fun(X35,X36)),'_dst_x00x11_2'(s(fun(X38,X36),X39),s(fun(X34,fun(X35,X38)),happ(s(fun(X37,fun(X34,fun(X35,X38))),X40),s(X37,happ(s(fun(X41,X37),X42),s(X41,X43))))))),s(X34,happ(s(fun(X41,X34),X44),s(X41,X43)))))) )),
  inference(superposition,[],[f196,f197])).
fof(f7078,plain,(
  ( ! [X30,X28,X26,X24,X33,X23,X31,X29,X27,X25,X32,X22] : (s(X23,happ(s(fun(X30,X23),'_dst_x0x1x1_3'(s(fun(X24,fun(X22,X23)),'_dst_x00x00x11_3'(s(fun(X25,fun(X26,X23)),X27),s(X25,X28),s(fun(X24,fun(X22,X26)),X29))),s(fun(X30,X24),X31),s(fun(X30,X22),X33))),s(X30,X32))) = s(X23,happ(s(fun(X22,X23),'_dst_x0x0x1_3'(s(fun(X25,fun(X26,X23)),X27),s(X25,X28),s(fun(X22,X26),happ(s(fun(X24,fun(X22,X26)),X29),s(X24,happ(s(fun(X30,X24),X31),s(X30,X32))))))),s(X22,happ(s(fun(X30,X22),X33),s(X30,X32)))))) )),
  inference(superposition,[],[f196,f199])).
fof(f7077,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s(X13,happ(s(fun(X18,X13),'_dst_x0x1x1_3'(s(fun(X14,fun(X12,X13)),'_dst_x00x11_2'(s(fun(X15,X13),X16),s(fun(X14,fun(X12,X15)),X17))),s(fun(X18,X14),X19),s(fun(X18,X12),X21))),s(X18,X20))) = s(X13,happ(s(fun(X12,X13),'_dst_x0x1_2'(s(fun(X15,X13),X16),s(fun(X12,X15),happ(s(fun(X14,fun(X12,X15)),X17),s(X14,happ(s(fun(X18,X14),X19),s(X18,X20))))))),s(X12,happ(s(fun(X18,X12),X21),s(X18,X20)))))) )),
  inference(superposition,[],[f196,f193])).
fof(f7076,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X1,happ(s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(X0,happ(s(fun(X8,X0),X11),s(X8,X10))))) = s(X1,happ(s(fun(X8,X1),'_dst_x0x1x1_3'(s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(fun(X8,X2),X9),s(fun(X8,X0),X11))),s(X8,X10)))) )),
  inference(superposition,[],[f196,f198])).
fof(f7075,plain,(
  ( ! [X280,X277,X279,X275,X276,X278] : (s(X277,happ(s(fun('type.closSem.state'(X275),X277),'_dst_x0x1x1_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun(X278,X277)),X279),s(fun('type.closSem.state'(X275),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X275),X278),X280))),s('type.closSem.state'(X275),X276))) = s(X277,happ(s(fun(X278,X277),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun(X278,X277)),X279),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X275),X276))))),s(X278,happ(s(fun('type.closSem.state'(X275),X278),X280),s('type.closSem.state'(X275),X276)))))) )),
  inference(superposition,[],[f196,f160])).
fof(f7074,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (s(X271,happ(s(fun(X272,X271),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),fun(X272,X271)),X273),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270))))),s(X272,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X272),X274),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270))))) = s(X271,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X271),'_dst_x0x1x1_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269)),fun(X272,X271)),X273),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X269))),'const.closSem.evaluate_0'),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X272),X274))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X269))),X270)))) )),
  inference(superposition,[],[f196,f161])).
fof(f7073,plain,(
  ( ! [X261,X263,X257,X259,X265,X267,X260,X262,X258,X268,X264,X266] : (s(X265,happ(s(fun(X260,X265),'_dst_x0x1x1_3'(s(fun(fun(X257,fun(X258,X259)),fun(X266,X265)),X267),s(fun(X260,fun(X257,fun(X258,X259))),'_dst_x000x111_2'(s(fun(X261,X259),X262),s(fun(X260,fun(X257,fun(X258,X261))),X263))),s(fun(X260,X266),X268))),s(X260,X264))) = s(X265,happ(s(fun(X266,X265),happ(s(fun(fun(X257,fun(X258,X259)),fun(X266,X265)),X267),s(fun(X257,fun(X258,X259)),'_dst_x00x11_2'(s(fun(X261,X259),X262),s(fun(X257,fun(X258,X261)),happ(s(fun(X260,fun(X257,fun(X258,X261))),X263),s(X260,X264))))))),s(X266,happ(s(fun(X260,X266),X268),s(X260,X264)))))) )),
  inference(superposition,[],[f196,f197])).
fof(f7072,plain,(
  ( ! [X255,X253,X251,X249,X246,X244,X254,X256,X250,X248,X252,X247,X245] : (s(X253,happ(s(fun(X246,X253),'_dst_x0x1x1_3'(s(fun(fun(X244,X245),fun(X254,X253)),X255),s(fun(X246,fun(X244,X245)),'_dst_x00x00x11_3'(s(fun(X247,fun(X248,X245)),X249),s(X247,X250),s(fun(X246,fun(X244,X248)),X251))),s(fun(X246,X254),X256))),s(X246,X252))) = s(X253,happ(s(fun(X254,X253),happ(s(fun(fun(X244,X245),fun(X254,X253)),X255),s(fun(X244,X245),'_dst_x0x0x1_3'(s(fun(X247,fun(X248,X245)),X249),s(X247,X250),s(fun(X244,X248),happ(s(fun(X246,fun(X244,X248)),X251),s(X246,X252))))))),s(X254,happ(s(fun(X246,X254),X256),s(X246,X252)))))) )),
  inference(superposition,[],[f196,f199])).
fof(f7071,plain,(
  ( ! [X243,X241,X239,X237,X235,X233,X242,X240,X238,X236,X234] : (s(X240,happ(s(fun(X235,X240),'_dst_x0x1x1_3'(s(fun(fun(X233,X234),fun(X241,X240)),X242),s(fun(X235,fun(X233,X234)),'_dst_x00x11_2'(s(fun(X236,X234),X237),s(fun(X235,fun(X233,X236)),X238))),s(fun(X235,X241),X243))),s(X235,X239))) = s(X240,happ(s(fun(X241,X240),happ(s(fun(fun(X233,X234),fun(X241,X240)),X242),s(fun(X233,X234),'_dst_x0x1_2'(s(fun(X236,X234),X237),s(fun(X233,X236),happ(s(fun(X235,fun(X233,X236)),X238),s(X235,X239))))))),s(X241,happ(s(fun(X235,X241),X243),s(X235,X239)))))) )),
  inference(superposition,[],[f196,f193])).
fof(f7070,plain,(
  ( ! [X222,X220,X231,X229,X227,X225,X223,X221,X230,X228,X226,X224,X232] : (s(X229,happ(s(fun(X222,X229),'_dst_x0x1x1_3'(s(fun(fun(X220,X221),fun(X230,X229)),X231),s(fun(X222,fun(X220,X221)),'_dst_x00x01x00_3'(s(fun(X223,fun(X224,X221)),X225),s(fun(X220,X223),X226),s(X224,X227))),s(fun(X222,X230),X232))),s(X222,X228))) = s(X229,happ(s(fun(X230,X229),happ(s(fun(fun(X220,X221),fun(X230,X229)),X231),s(fun(X220,X221),'_dst_x0x1x0_3'(s(fun(X223,fun(X224,X221)),X225),s(fun(X220,X223),X226),s(X224,X227))))),s(X230,happ(s(fun(X222,X230),X232),s(X222,X228)))))) )),
  inference(superposition,[],[f196,f198])).
fof(f7067,plain,(
  ( ! [X198,X204,X202,X200,X199,X197,X203,X201] : (s(X201,happ(s(fun(X197,X201),'_dst_x0x1x1_3'(s(fun(bool,fun(X202,X201)),X203),s(fun(X197,bool),X198),s(fun(X197,X202),X204))),s(X197,X199))) = s(X201,happ(s(fun(X202,X201),happ(s(fun(bool,fun(X202,X201)),X203),s(bool,X200))),s(X202,happ(s(fun(X197,X202),X204),s(X197,X199))))) | s(bool,'T_0') = s(bool,happ(s(fun(X197,bool),X198),s(X197,X199))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f196,f223])).
fof(f7065,plain,(
  ( ! [X187,X185,X182,X188,X186,X184,X183,X181] : (s(X184,happ(s(fun(X181,X184),'_dst_x0x1x1_3'(s(fun(bool,fun(X185,X184)),X186),s(fun(X181,bool),X182),s(fun(X181,X185),X187))),s(X181,X183))) = s(X184,happ(s(fun(X185,X184),happ(s(fun(bool,fun(X185,X184)),X186),s(bool,'T_0'))),s(X185,happ(s(fun(X181,X185),X187),s(X181,X183))))) | s(bool,X188) = s(bool,happ(s(fun(X181,bool),X182),s(X181,X183))) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f196,f223])).
fof(f7309,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(X160,happ(s(fun(bool,X160),'_dst_x0x0x1_3'(s(fun(bool,fun(X161,X160)),X162),s(bool,'T_0'),s(fun(bool,X161),X163))),s(bool,X159))) = s(X160,happ(s(fun(bool,X160),'_dst_x0x1x1_3'(s(fun(bool,fun(X161,X160)),X162),s(fun(bool,bool),'_i_0_0'),s(fun(bool,X161),X163))),s(bool,X159))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(forward_demodulation,[],[f7061,f194])).
fof(f7061,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(X160,happ(s(fun(bool,X160),'_dst_x0x1x1_3'(s(fun(bool,fun(X161,X160)),X162),s(fun(bool,bool),'_i_0_0'),s(fun(bool,X161),X163))),s(bool,X159))) = s(X160,happ(s(fun(X161,X160),happ(s(fun(bool,fun(X161,X160)),X162),s(bool,'T_0'))),s(X161,happ(s(fun(bool,X161),X163),s(bool,X159))))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f196,f235])).
fof(f7059,plain,(
  ( ! [X146,X151,X149,X147,X145,X150,X148] : (s(X148,happ(s(fun(X149,X148),happ(s(fun(bool,fun(X149,X148)),X150),s(bool,'T_0'))),s(X149,happ(s(fun(X145,X149),X151),s(X145,X147))))) = s(X148,happ(s(fun(X145,X148),'_dst_x0x1x1_3'(s(fun(bool,fun(X149,X148)),X150),s(fun(X145,bool),'_dst_x0_1'(s(bool,X146))),s(fun(X145,X149),X151))),s(X145,X147))) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f196,f267])).
fof(f7058,plain,(
  ( ! [X144,X142,X140,X138,X143,X141,X139] : (s(X141,happ(s(fun(X142,X141),happ(s(fun(bool,fun(X142,X141)),X143),s(bool,'F_0'))),s(X142,happ(s(fun(X138,X142),X144),s(X138,X140))))) = s(X141,happ(s(fun(X138,X141),'_dst_x0x1x1_3'(s(fun(bool,fun(X142,X141)),X143),s(fun(X138,bool),'_dst_x0_1'(s(bool,X139))),s(fun(X138,X142),X144))),s(X138,X140))) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f196,f614])).
fof(f7057,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(X134,happ(s(fun(X135,X134),happ(s(fun(bool,fun(X135,X134)),X136),s(bool,'T_0'))),s(X135,happ(s(fun(X132,X135),X137),s(X132,X133))))) = s(X134,happ(s(fun(X132,X134),'_dst_x0x1x1_3'(s(fun(bool,fun(X135,X134)),X136),s(fun(X132,bool),happ(s(fun(X132,fun(X132,bool)),'=_0'),s(X132,X133))),s(fun(X132,X135),X137))),s(X132,X133)))) )),
  inference(superposition,[],[f196,f744])).
fof(f7056,plain,(
  ( ! [X127,X125,X123,X130,X128,X126,X124,X131,X129] : (s(X128,happ(s(fun(X129,X128),happ(s(fun(X123,fun(X129,X128)),X130),s(X123,'const.pair.UNCURRY_2'(s(fun(X124,fun(X125,X123)),X126),s('type.pair.prod'(X124,X125),X127))))),s(X129,happ(s(fun('type.pair.prod'(X124,X125),X129),X131),s('type.pair.prod'(X124,X125),X127))))) = s(X128,happ(s(fun('type.pair.prod'(X124,X125),X128),'_dst_x0x1x1_3'(s(fun(X123,fun(X129,X128)),X130),s(fun('type.pair.prod'(X124,X125),X123),happ(s(fun(fun(X124,fun(X125,X123)),fun('type.pair.prod'(X124,X125),X123)),'const.pair.UNCURRY_0'),s(fun(X124,fun(X125,X123)),X126))),s(fun('type.pair.prod'(X124,X125),X129),X131))),s('type.pair.prod'(X124,X125),X127)))) )),
  inference(superposition,[],[f196,f180])).
fof(f7055,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (s(X119,happ(s(fun(X120,X119),happ(s(fun(X114,fun(X120,X119)),X121),s(X114,'const.pair.pair_CASE_2'(s('type.pair.prod'(X115,X116),X117),s(fun(X115,fun(X116,X114)),X118))))),s(X120,happ(s(fun(fun(X115,fun(X116,X114)),X120),X122),s(fun(X115,fun(X116,X114)),X118))))) = s(X119,happ(s(fun(fun(X115,fun(X116,X114)),X119),'_dst_x0x1x1_3'(s(fun(X114,fun(X120,X119)),X121),s(fun(fun(X115,fun(X116,X114)),X114),happ(s(fun('type.pair.prod'(X115,X116),fun(fun(X115,fun(X116,X114)),X114)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X115,X116),X117))),s(fun(fun(X115,fun(X116,X114)),X120),X122))),s(fun(X115,fun(X116,X114)),X118)))) )),
  inference(superposition,[],[f196,f181])).
fof(f7054,plain,(
  ( ! [X103,X111,X109,X107,X105,X112,X102,X110,X108,X106,X104,X113] : (s(X110,happ(s(fun(X103,X110),'_dst_x0x1x1_3'(s(fun(X102,fun(X111,X110)),X112),s(fun(X103,X102),'_dst_x0x0x1_3'(s(fun(X104,fun(X105,X102)),X106),s(X104,X107),s(fun(X103,X105),X108))),s(fun(X103,X111),X113))),s(X103,X109))) = s(X110,happ(s(fun(X111,X110),happ(s(fun(X102,fun(X111,X110)),X112),s(X102,happ(s(fun(X105,X102),happ(s(fun(X104,fun(X105,X102)),X106),s(X104,X107))),s(X105,happ(s(fun(X103,X105),X108),s(X103,X109))))))),s(X111,happ(s(fun(X103,X111),X113),s(X103,X109)))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7053,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (s(X98,happ(s(fun(X99,X98),happ(s(fun(X94,fun(X99,X98)),X100),s(X94,X96))),s(X99,happ(s(fun(X95,X99),X101),s(X95,X97))))) = s(X98,happ(s(fun(X95,X98),'_dst_x0x1x1_3'(s(fun(X94,fun(X99,X98)),X100),s(fun(X95,X94),'_dst_x0_1'(s(X94,X96))),s(fun(X95,X99),X101))),s(X95,X97)))) )),
  inference(superposition,[],[f196,f170])).
fof(f7052,plain,(
  ( ! [X92,X90,X88,X87,X85,X83,X93,X91,X89,X86,X84,X82] : (s(X90,happ(s(fun(X83,X90),'_dst_x0x1x1_3'(s(fun(X82,fun(X91,X90)),X92),s(fun(X83,X82),'_dst_x0x1x0_3'(s(fun(X84,fun(X85,X82)),X86),s(fun(X83,X84),X87),s(X85,X88))),s(fun(X83,X91),X93))),s(X83,X89))) = s(X90,happ(s(fun(X91,X90),happ(s(fun(X82,fun(X91,X90)),X92),s(X82,happ(s(fun(X85,X82),happ(s(fun(X84,fun(X85,X82)),X86),s(X84,happ(s(fun(X83,X84),X87),s(X83,X89))))),s(X85,X88))))),s(X91,happ(s(fun(X83,X91),X93),s(X83,X89)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7051,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73] : (s(X78,happ(s(fun(X71,X78),'_dst_x0x1x1_3'(s(fun(X70,fun(X79,X78)),X80),s(fun(X71,X70),happ(s(fun(X72,fun(X71,X70)),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76))))),s(fun(X71,X79),X81))),s(X71,happ(s(fun(X74,X71),X77),s(X74,X76))))) = s(X78,happ(s(fun(X79,X78),happ(s(fun(X70,fun(X79,X78)),X80),s(X70,happ(s(fun(X74,X70),'_dst_x0x1x1_3'(s(fun(X72,fun(X71,X70)),X73),s(fun(X74,X72),X75),s(fun(X74,X71),X77))),s(X74,X76))))),s(X79,happ(s(fun(X71,X79),X81),s(X71,happ(s(fun(X74,X71),X77),s(X74,X76)))))))) )),
  inference(superposition,[],[f196,f196])).
fof(f7050,plain,(
  ( ! [X61,X59,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s(X66,happ(s(fun(X59,X66),'_dst_x0x1x1_3'(s(fun(X58,fun(X67,X66)),X68),s(fun(X59,X58),happ(s(fun(X60,fun(X59,X58)),X61),s(X60,happ(s(fun(X62,X60),X63),s(X62,X64))))),s(fun(X59,X67),X69))),s(X59,X65))) = s(X66,happ(s(fun(X67,X66),happ(s(fun(X58,fun(X67,X66)),X68),s(X58,happ(s(fun(X62,X58),'_dst_x0x1x0_3'(s(fun(X60,fun(X59,X58)),X61),s(fun(X62,X60),X63),s(X59,X65))),s(X62,X64))))),s(X67,happ(s(fun(X59,X67),X69),s(X59,X65)))))) )),
  inference(superposition,[],[f196,f195])).
fof(f7049,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (s(X54,happ(s(fun(X47,X54),'_dst_x0x1x1_3'(s(fun(X46,fun(X55,X54)),X56),s(fun(X47,X46),happ(s(fun(X48,fun(X47,X46)),X49),s(X48,X50))),s(fun(X47,X55),X57))),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53))))) = s(X54,happ(s(fun(X55,X54),happ(s(fun(X46,fun(X55,X54)),X56),s(X46,happ(s(fun(X51,X46),'_dst_x0x0x1_3'(s(fun(X48,fun(X47,X46)),X49),s(X48,X50),s(fun(X51,X47),X52))),s(X51,X53))))),s(X55,happ(s(fun(X47,X55),X57),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53)))))))) )),
  inference(superposition,[],[f196,f194])).
fof(f7307,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X2,bool),'_dst_x0x1x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X3))),s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X3))))),s(X2,X4))))) )),
  inference(backward_demodulation,[],[f7047,f1526])).
fof(f7047,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (s(X32,happ(s(fun(X27,X32),'_dst_x0x1x1_3'(s(fun(X26,fun(X33,X32)),X34),s(fun(X27,X26),happ(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))),s(fun(X27,X33),X35))),s(X27,X31))) = s(X32,happ(s(fun(X33,X32),happ(s(fun(X26,fun(X33,X32)),X34),s(X26,'const.pair.pair_CASE_2'(s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))),s(fun(X28,fun(X27,X26)),X29))))),s(X33,happ(s(fun(X27,X33),X35),s(X27,X31)))))) )),
  inference(superposition,[],[f196,f184])).
fof(f7046,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X22,happ(s(fun(X17,X22),'_dst_x0x1x1_3'(s(fun(X16,fun(X23,X22)),X24),s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X19),s(fun(X17,X18),X20))),s(fun(X17,X23),X25))),s(X17,X21))) = s(X22,happ(s(fun(X23,X22),happ(s(fun(X16,fun(X23,X22)),X24),s(X16,happ(s(fun(X18,X16),X19),s(X18,happ(s(fun(X17,X18),X20),s(X17,X21))))))),s(X23,happ(s(fun(X17,X23),X25),s(X17,X21)))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7306,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X2,bool),'_dst_x0x1x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X2,X0),X3),s(fun(X2,X0),X3))),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))))) )),
  inference(backward_demodulation,[],[f7045,f1044])).
fof(f7045,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(X12,happ(s(fun(X7,X12),'_dst_x0x1x1_3'(s(fun(X6,fun(X13,X12)),X14),s(fun(X7,X6),X8),s(fun(X7,X13),X15))),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))) = s(X12,happ(s(fun(X13,X12),happ(s(fun(X6,fun(X13,X12)),X14),s(X6,happ(s(fun(X9,X6),'_dst_x0x1_2'(s(fun(X7,X6),X8),s(fun(X9,X7),X10))),s(X9,X11))))),s(X13,happ(s(fun(X7,X13),X15),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11)))))))) )),
  inference(superposition,[],[f196,f182])).
fof(f7044,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,happ(s(fun(X0,X2),'_dst_x0x1x1_3'(s(fun(X0,fun(X3,X2)),X4),s(fun(X0,X0),'_i_0_0'),s(fun(X0,X3),X5))),s(X0,X1))) = s(X2,happ(s(fun(X3,X2),happ(s(fun(X0,fun(X3,X2)),X4),s(X0,X1))),s(X3,happ(s(fun(X0,X3),X5),s(X0,X1)))))) )),
  inference(superposition,[],[f196,f159])).
fof(f7303,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X4,bool),'_dst_x0x1x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X4,X0),'_dst_x0x1x1_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7))),s(fun(X4,X0),'_dst_x0x1x1_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7))))),s(X4,X6))))) )),
  inference(forward_demodulation,[],[f7302,f196])).
fof(f7302,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7))),s(X4,X6))))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7))),s(X4,X6))))))) )),
  inference(forward_demodulation,[],[f6999,f196])).
fof(f6999,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,happ(s(fun(X4,X1),X7),s(X4,X6))))))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x1_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(fun(X4,X1),X7))),s(X4,X6))))))) )),
  inference(unit_resulting_resolution,[],[f196,f173])).
fof(f6700,plain,(
  ( ! [X350,X352,X354,X349,X351,X347,X353,X355,X348] : (p(s(bool,happ(s(fun(fun(X347,fun(X348,X349)),bool),happ(s(fun(fun(X347,fun(X348,X349)),fun(fun(X347,fun(X348,X349)),bool)),'=_0'),s(fun(X347,fun(X348,X349)),X355))),s(fun(X347,fun(X348,X349)),'_dst_x00x11_2'(s(fun(X351,X349),X352),s(fun(X347,fun(X348,X351)),happ(s(fun(X350,fun(X347,fun(X348,X351))),X353),s(X350,X354)))))))) | s(fun(X347,fun(X348,X349)),X355) != s(fun(X347,fun(X348,X349)),happ(s(fun(X350,fun(X347,fun(X348,X349))),'_dst_x000x111_2'(s(fun(X351,X349),X352),s(fun(X350,fun(X347,fun(X348,X351))),X353))),s(X350,X354)))) )),
  inference(superposition,[],[f173,f197])).
fof(f6699,plain,(
  ( ! [X344,X346,X341,X343,X339,X345,X340,X342,X338] : (p(s(bool,happ(s(fun(fun(X338,fun(X339,X340)),bool),happ(s(fun(fun(X338,fun(X339,X340)),fun(fun(X338,fun(X339,X340)),bool)),'=_0'),s(fun(X338,fun(X339,X340)),'_dst_x00x11_2'(s(fun(X342,X340),X343),s(fun(X338,fun(X339,X342)),happ(s(fun(X341,fun(X338,fun(X339,X342))),X344),s(X341,X345))))))),s(fun(X338,fun(X339,X340)),X346)))) | s(fun(X338,fun(X339,X340)),X346) != s(fun(X338,fun(X339,X340)),happ(s(fun(X341,fun(X338,fun(X339,X340))),'_dst_x000x111_2'(s(fun(X342,X340),X343),s(fun(X341,fun(X338,fun(X339,X342))),X344))),s(X341,X345)))) )),
  inference(superposition,[],[f173,f197])).
fof(f6698,plain,(
  ( ! [X333,X335,X329,X331,X337,X332,X334,X330,X336] : (~p(s(bool,happ(s(fun(fun(X329,fun(X330,X331)),bool),happ(s(fun(fun(X329,fun(X330,X331)),fun(fun(X329,fun(X330,X331)),bool)),'=_0'),s(fun(X329,fun(X330,X331)),X337))),s(fun(X329,fun(X330,X331)),'_dst_x00x11_2'(s(fun(X333,X331),X334),s(fun(X329,fun(X330,X333)),happ(s(fun(X332,fun(X329,fun(X330,X333))),X335),s(X332,X336)))))))) | s(fun(X329,fun(X330,X331)),X337) = s(fun(X329,fun(X330,X331)),happ(s(fun(X332,fun(X329,fun(X330,X331))),'_dst_x000x111_2'(s(fun(X333,X331),X334),s(fun(X332,fun(X329,fun(X330,X333))),X335))),s(X332,X336)))) )),
  inference(superposition,[],[f172,f197])).
fof(f6697,plain,(
  ( ! [X325,X327,X321,X323,X324,X326,X320,X322,X328] : (~p(s(bool,happ(s(fun(fun(X320,fun(X321,X322)),bool),happ(s(fun(fun(X320,fun(X321,X322)),fun(fun(X320,fun(X321,X322)),bool)),'=_0'),s(fun(X320,fun(X321,X322)),'_dst_x00x11_2'(s(fun(X324,X322),X325),s(fun(X320,fun(X321,X324)),happ(s(fun(X323,fun(X320,fun(X321,X324))),X326),s(X323,X327))))))),s(fun(X320,fun(X321,X322)),X328)))) | s(fun(X320,fun(X321,X322)),X328) = s(fun(X320,fun(X321,X322)),happ(s(fun(X323,fun(X320,fun(X321,X322))),'_dst_x000x111_2'(s(fun(X324,X322),X325),s(fun(X323,fun(X320,fun(X321,X324))),X326))),s(X323,X327)))) )),
  inference(superposition,[],[f172,f197])).
fof(f6692,plain,(
  ( ! [X280,X277,X279,X273,X275,X281,X276,X278,X272,X274] : (s(X280,happ(s(fun(X275,X280),'_dst_x0x1_2'(s(fun(fun(X272,fun(X273,X274)),X280),X281),s(fun(X275,fun(X272,fun(X273,X274))),'_dst_x000x111_2'(s(fun(X276,X274),X277),s(fun(X275,fun(X272,fun(X273,X276))),X278))))),s(X275,X279))) = s(X280,happ(s(fun(fun(X272,fun(X273,X274)),X280),X281),s(fun(X272,fun(X273,X274)),'_dst_x00x11_2'(s(fun(X276,X274),X277),s(fun(X272,fun(X273,X276)),happ(s(fun(X275,fun(X272,fun(X273,X276))),X278),s(X275,X279)))))))) )),
  inference(superposition,[],[f182,f197])).
fof(f6691,plain,(
  ( ! [X261,X263,X269,X271,X265,X267,X260,X262,X268,X270,X264,X266] : (s(X268,happ(s(fun(X263,X268),'_dst_x0x0x1_3'(s(fun(X269,fun(fun(X260,fun(X261,X262)),X268)),X270),s(X269,X271),s(fun(X263,fun(X260,fun(X261,X262))),'_dst_x000x111_2'(s(fun(X264,X262),X265),s(fun(X263,fun(X260,fun(X261,X264))),X266))))),s(X263,X267))) = s(X268,happ(s(fun(fun(X260,fun(X261,X262)),X268),happ(s(fun(X269,fun(fun(X260,fun(X261,X262)),X268)),X270),s(X269,X271))),s(fun(X260,fun(X261,X262)),'_dst_x00x11_2'(s(fun(X264,X262),X265),s(fun(X260,fun(X261,X264)),happ(s(fun(X263,fun(X260,fun(X261,X264))),X266),s(X263,X267)))))))) )),
  inference(superposition,[],[f194,f197])).
fof(f6690,plain,(
  ( ! [X255,X253,X251,X249,X257,X259,X254,X256,X250,X248,X258,X252] : (s(X256,happ(s(fun(X251,X256),'_dst_x0x1x0_3'(s(fun(fun(X248,fun(X249,X250)),fun(X257,X256)),X258),s(fun(X251,fun(X248,fun(X249,X250))),'_dst_x000x111_2'(s(fun(X252,X250),X253),s(fun(X251,fun(X248,fun(X249,X252))),X254))),s(X257,X259))),s(X251,X255))) = s(X256,happ(s(fun(X257,X256),happ(s(fun(fun(X248,fun(X249,X250)),fun(X257,X256)),X258),s(fun(X248,fun(X249,X250)),'_dst_x00x11_2'(s(fun(X252,X250),X253),s(fun(X248,fun(X249,X252)),happ(s(fun(X251,fun(X248,fun(X249,X252))),X254),s(X251,X255))))))),s(X257,X259)))) )),
  inference(superposition,[],[f195,f197])).
fof(f6673,plain,(
  ( ! [X61,X59,X57,X54,X52,X60,X58,X56,X55,X53,X51] : (s(fun(X52,X53),happ(s(fun(X58,fun(X52,X53)),'_dst_x0x1x0_3'(s(fun(X54,fun(X51,fun(X52,X53))),'_dst_x000x111_2'(s(fun(X55,X53),X56),s(fun(X54,fun(X51,fun(X52,X55))),X57))),s(fun(X58,X54),X59),s(X51,X61))),s(X58,X60))) = s(fun(X52,X53),happ(s(fun(X51,fun(X52,X53)),'_dst_x00x11_2'(s(fun(X55,X53),X56),s(fun(X51,fun(X52,X55)),happ(s(fun(X54,fun(X51,fun(X52,X55))),X57),s(X54,happ(s(fun(X58,X54),X59),s(X58,X60))))))),s(X51,X61)))) )),
  inference(superposition,[],[f195,f197])).
fof(f6672,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (s(fun(X43,X44),happ(s(fun(X42,fun(X43,X44)),'_dst_x00x11_2'(s(fun(X46,X44),X47),s(fun(X42,fun(X43,X46)),happ(s(fun(X45,fun(X42,fun(X43,X46))),X48),s(X45,X49))))),s(X42,X50))) = s(fun(X43,X44),'const.pair.UNCURRY_2'(s(fun(X45,fun(X42,fun(X43,X44))),'_dst_x000x111_2'(s(fun(X46,X44),X47),s(fun(X45,fun(X42,fun(X43,X46))),X48))),s('type.pair.prod'(X45,X42),'const.pair.,_2'(s(X45,X49),s(X42,X50)))))) )),
  inference(superposition,[],[f183,f197])).
fof(f6671,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X40] : (s(fun(X34,X35),happ(s(fun(X33,fun(X34,X35)),'_dst_x00x11_2'(s(fun(X37,X35),X38),s(fun(X33,fun(X34,X37)),happ(s(fun(X36,fun(X33,fun(X34,X37))),X39),s(X36,X40))))),s(X33,X41))) = s(fun(X34,X35),'const.pair.pair_CASE_2'(s('type.pair.prod'(X36,X33),'const.pair.,_2'(s(X36,X40),s(X33,X41))),s(fun(X36,fun(X33,fun(X34,X35))),'_dst_x000x111_2'(s(fun(X37,X35),X38),s(fun(X36,fun(X33,fun(X34,X37))),X39)))))) )),
  inference(superposition,[],[f184,f197])).
fof(f6670,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32] : (s(fun(X23,X31),'_dst_x0x1_2'(s(fun(fun(X24,X25),X31),X32),s(fun(X23,fun(X24,X25)),'_dst_x00x11_2'(s(fun(X27,X25),X28),s(fun(X23,fun(X24,X27)),happ(s(fun(X26,fun(X23,fun(X24,X27))),X29),s(X26,X30))))))) = s(fun(X23,X31),happ(s(fun(X26,fun(X23,X31)),'_dst_x00x11_2'(s(fun(fun(X24,X25),X31),X32),s(fun(X26,fun(X23,fun(X24,X25))),'_dst_x000x111_2'(s(fun(X27,X25),X28),s(fun(X26,fun(X23,fun(X24,X27))),X29))))),s(X26,X30)))) )),
  inference(superposition,[],[f193,f197])).
fof(f6669,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(fun(X13,X14),happ(s(fun(X20,fun(X13,X14)),'_dst_x0x0x1_3'(s(fun(X15,fun(X12,fun(X13,X14))),'_dst_x000x111_2'(s(fun(X16,X14),X17),s(fun(X15,fun(X12,fun(X13,X16))),X18))),s(X15,X19),s(fun(X20,X12),X21))),s(X20,X22))) = s(fun(X13,X14),happ(s(fun(X12,fun(X13,X14)),'_dst_x00x11_2'(s(fun(X16,X14),X17),s(fun(X12,fun(X13,X16)),happ(s(fun(X15,fun(X12,fun(X13,X16))),X18),s(X15,X19))))),s(X12,happ(s(fun(X20,X12),X21),s(X20,X22)))))) )),
  inference(superposition,[],[f194,f197])).
fof(f6668,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,X8),'_dst_x0x0x1_3'(s(fun(X9,fun(fun(X1,X2),X8)),X10),s(X9,X11),s(fun(X0,fun(X1,X2)),'_dst_x00x11_2'(s(fun(X4,X2),X5),s(fun(X0,fun(X1,X4)),happ(s(fun(X3,fun(X0,fun(X1,X4))),X6),s(X3,X7))))))) = s(fun(X0,X8),happ(s(fun(X3,fun(X0,X8)),'_dst_x00x00x11_3'(s(fun(X9,fun(fun(X1,X2),X8)),X10),s(X9,X11),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x111_2'(s(fun(X4,X2),X5),s(fun(X3,fun(X0,fun(X1,X4))),X6))))),s(X3,X7)))) )),
  inference(superposition,[],[f199,f197])).
fof(f6647,plain,(
  ( ! [X255,X253,X251,X257,X254,X252,X250,X258,X256] : (p(s(bool,happ(s(fun(fun(X250,fun(X251,X252)),bool),happ(s(fun(fun(X250,fun(X251,X252)),fun(fun(X250,fun(X251,X252)),bool)),'=_0'),s(fun(X250,fun(X251,X252)),X258))),s(fun(X250,fun(X251,X252)),happ(s(fun(X255,fun(X250,fun(X251,X252))),'_dst_x000x111_2'(s(fun(X253,X252),X254),s(fun(X255,fun(X250,fun(X251,X253))),X256))),s(X255,X257)))))) | s(fun(X250,fun(X251,X252)),X258) != s(fun(X250,fun(X251,X252)),'_dst_x00x11_2'(s(fun(X253,X252),X254),s(fun(X250,fun(X251,X253)),happ(s(fun(X255,fun(X250,fun(X251,X253))),X256),s(X255,X257)))))) )),
  inference(superposition,[],[f173,f197])).
fof(f6646,plain,(
  ( ! [X243,X241,X249,X246,X244,X242,X248,X247,X245] : (p(s(bool,happ(s(fun(fun(X241,fun(X242,X243)),bool),happ(s(fun(fun(X241,fun(X242,X243)),fun(fun(X241,fun(X242,X243)),bool)),'=_0'),s(fun(X241,fun(X242,X243)),happ(s(fun(X246,fun(X241,fun(X242,X243))),'_dst_x000x111_2'(s(fun(X244,X243),X245),s(fun(X246,fun(X241,fun(X242,X244))),X247))),s(X246,X248))))),s(fun(X241,fun(X242,X243)),X249)))) | s(fun(X241,fun(X242,X243)),X249) != s(fun(X241,fun(X242,X243)),'_dst_x00x11_2'(s(fun(X244,X243),X245),s(fun(X241,fun(X242,X244)),happ(s(fun(X246,fun(X241,fun(X242,X244))),X247),s(X246,X248)))))) )),
  inference(superposition,[],[f173,f197])).
fof(f6645,plain,(
  ( ! [X239,X237,X235,X233,X240,X238,X236,X234,X232] : (~p(s(bool,happ(s(fun(fun(X232,fun(X233,X234)),bool),happ(s(fun(fun(X232,fun(X233,X234)),fun(fun(X232,fun(X233,X234)),bool)),'=_0'),s(fun(X232,fun(X233,X234)),X240))),s(fun(X232,fun(X233,X234)),happ(s(fun(X237,fun(X232,fun(X233,X234))),'_dst_x000x111_2'(s(fun(X235,X234),X236),s(fun(X237,fun(X232,fun(X233,X235))),X238))),s(X237,X239)))))) | s(fun(X232,fun(X233,X234)),X240) = s(fun(X232,fun(X233,X234)),'_dst_x00x11_2'(s(fun(X235,X234),X236),s(fun(X232,fun(X233,X235)),happ(s(fun(X237,fun(X232,fun(X233,X235))),X238),s(X237,X239)))))) )),
  inference(superposition,[],[f172,f197])).
fof(f6644,plain,(
  ( ! [X231,X229,X227,X225,X223,X230,X228,X226,X224] : (~p(s(bool,happ(s(fun(fun(X223,fun(X224,X225)),bool),happ(s(fun(fun(X223,fun(X224,X225)),fun(fun(X223,fun(X224,X225)),bool)),'=_0'),s(fun(X223,fun(X224,X225)),happ(s(fun(X228,fun(X223,fun(X224,X225))),'_dst_x000x111_2'(s(fun(X226,X225),X227),s(fun(X228,fun(X223,fun(X224,X226))),X229))),s(X228,X230))))),s(fun(X223,fun(X224,X225)),X231)))) | s(fun(X223,fun(X224,X225)),X231) = s(fun(X223,fun(X224,X225)),'_dst_x00x11_2'(s(fun(X226,X225),X227),s(fun(X223,fun(X224,X226)),happ(s(fun(X228,fun(X223,fun(X224,X226))),X229),s(X228,X230)))))) )),
  inference(superposition,[],[f172,f197])).
fof(f6624,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X2),'_dst_x0x1_2'(s(fun(X3,X2),X4),s(fun(X1,X3),happ(s(fun(X0,fun(X1,X3)),happ(s(fun(X5,fun(X0,fun(X1,X3))),X6),s(X5,X7))),s(X0,X8))))) = s(fun(X1,X2),happ(s(fun(X0,fun(X1,X2)),happ(s(fun(X5,fun(X0,fun(X1,X2))),'_dst_x000x111_2'(s(fun(X3,X2),X4),s(fun(X5,fun(X0,fun(X1,X3))),X6))),s(X5,X7))),s(X0,X8)))) )),
  inference(superposition,[],[f193,f197])).
fof(f6623,plain,(
  ( ! [X14,X12,X10,X8,X17,X15,X13,X11,X9,X16] : (s(fun(X8,fun(X9,X10)),'_dst_x00x11_2'(s(fun(X12,X10),X13),s(fun(X8,fun(X9,X12)),happ(s(fun(X11,fun(X8,fun(X9,X12))),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(fun(X8,fun(X9,X10)),happ(s(fun(X15,fun(X8,fun(X9,X10))),'_dst_x0x1_2'(s(fun(X11,fun(X8,fun(X9,X10))),'_dst_x000x111_2'(s(fun(X12,X10),X13),s(fun(X11,fun(X8,fun(X9,X12))),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f197,f182])).
fof(f6536,plain,(
  ( ! [X158,X156,X163,X161,X159,X157,X155,X162,X160] : (s(fun(X155,fun(X156,X162)),'_dst_x00x11_2'(s(fun(X157,X162),X163),s(fun(X155,fun(X156,X157)),'const.pair.UNCURRY_2'(s(fun(X158,fun(X159,fun(X155,fun(X156,X157)))),X160),s('type.pair.prod'(X158,X159),X161))))) = s(fun(X155,fun(X156,X162)),happ(s(fun('type.pair.prod'(X158,X159),fun(X155,fun(X156,X162))),'_dst_x000x111_2'(s(fun(X157,X162),X163),s(fun('type.pair.prod'(X158,X159),fun(X155,fun(X156,X157))),happ(s(fun(fun(X158,fun(X159,fun(X155,fun(X156,X157)))),fun('type.pair.prod'(X158,X159),fun(X155,fun(X156,X157)))),'const.pair.UNCURRY_0'),s(fun(X158,fun(X159,fun(X155,fun(X156,X157)))),X160))))),s('type.pair.prod'(X158,X159),X161)))) )),
  inference(superposition,[],[f197,f180])).
fof(f6535,plain,(
  ( ! [X144,X134,X142,X140,X138,X136,X145,X135,X143,X141,X139,X137] : (s(fun(X134,fun(X135,X144)),happ(s(fun(X137,fun(X134,fun(X135,X144))),'_dst_x000x111_2'(s(fun(X136,X144),X145),s(fun(X137,fun(X134,fun(X135,X136))),'_dst_x0x0x1_3'(s(fun(X138,fun(X139,fun(X134,fun(X135,X136)))),X140),s(X138,X141),s(fun(X137,X139),X142))))),s(X137,X143))) = s(fun(X134,fun(X135,X144)),'_dst_x00x11_2'(s(fun(X136,X144),X145),s(fun(X134,fun(X135,X136)),happ(s(fun(X139,fun(X134,fun(X135,X136))),happ(s(fun(X138,fun(X139,fun(X134,fun(X135,X136)))),X140),s(X138,X141))),s(X139,happ(s(fun(X137,X139),X142),s(X137,X143)))))))) )),
  inference(superposition,[],[f197,f194])).
fof(f6534,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (s(fun(X126,fun(X127,X132)),'_dst_x00x11_2'(s(fun(X128,X132),X133),s(fun(X126,fun(X127,X128)),X130))) = s(fun(X126,fun(X127,X132)),happ(s(fun(X129,fun(X126,fun(X127,X132))),'_dst_x000x111_2'(s(fun(X128,X132),X133),s(fun(X129,fun(X126,fun(X127,X128))),'_dst_x0_1'(s(fun(X126,fun(X127,X128)),X130))))),s(X129,X131)))) )),
  inference(superposition,[],[f197,f170])).
fof(f6533,plain,(
  ( ! [X125,X123,X121,X118,X116,X114,X124,X122,X120,X119,X117,X115] : (s(fun(X114,fun(X115,X124)),happ(s(fun(X117,fun(X114,fun(X115,X124))),'_dst_x000x111_2'(s(fun(X116,X124),X125),s(fun(X117,fun(X114,fun(X115,X116))),'_dst_x0x1x0_3'(s(fun(X118,fun(X119,fun(X114,fun(X115,X116)))),X120),s(fun(X117,X118),X121),s(X119,X122))))),s(X117,X123))) = s(fun(X114,fun(X115,X124)),'_dst_x00x11_2'(s(fun(X116,X124),X125),s(fun(X114,fun(X115,X116)),happ(s(fun(X119,fun(X114,fun(X115,X116))),happ(s(fun(X118,fun(X119,fun(X114,fun(X115,X116)))),X120),s(X118,happ(s(fun(X117,X118),X121),s(X117,X123))))),s(X119,X122)))))) )),
  inference(superposition,[],[f197,f195])).
fof(f6532,plain,(
  ( ! [X103,X111,X109,X107,X105,X112,X102,X110,X108,X106,X104,X113] : (s(fun(X102,fun(X103,X112)),'_dst_x00x11_2'(s(fun(X104,X112),X113),s(fun(X102,fun(X103,X104)),happ(s(fun(X108,fun(X102,fun(X103,X104))),'_dst_x0x1x0_3'(s(fun(X106,fun(X105,fun(X102,fun(X103,X104)))),X107),s(fun(X108,X106),X109),s(X105,X111))),s(X108,X110))))) = s(fun(X102,fun(X103,X112)),happ(s(fun(X105,fun(X102,fun(X103,X112))),'_dst_x000x111_2'(s(fun(X104,X112),X113),s(fun(X105,fun(X102,fun(X103,X104))),happ(s(fun(X106,fun(X105,fun(X102,fun(X103,X104)))),X107),s(X106,happ(s(fun(X108,X106),X109),s(X108,X110))))))),s(X105,X111)))) )),
  inference(superposition,[],[f197,f195])).
fof(f6531,plain,(
  ( ! [X94,X92,X90,X101,X99,X97,X95,X93,X91,X100,X98,X96] : (s(fun(X90,fun(X91,X100)),'_dst_x00x11_2'(s(fun(X92,X100),X101),s(fun(X90,fun(X91,X92)),happ(s(fun(X97,fun(X90,fun(X91,X92))),'_dst_x0x0x1_3'(s(fun(X94,fun(X93,fun(X90,fun(X91,X92)))),X95),s(X94,X96),s(fun(X97,X93),X98))),s(X97,X99))))) = s(fun(X90,fun(X91,X100)),happ(s(fun(X93,fun(X90,fun(X91,X100))),'_dst_x000x111_2'(s(fun(X92,X100),X101),s(fun(X93,fun(X90,fun(X91,X92))),happ(s(fun(X94,fun(X93,fun(X90,fun(X91,X92)))),X95),s(X94,X96))))),s(X93,happ(s(fun(X97,X93),X98),s(X97,X99)))))) )),
  inference(superposition,[],[f197,f194])).
fof(f6529,plain,(
  ( ! [X70,X78,X76,X74,X72,X71,X79,X77,X75,X73] : (s(fun(X70,fun(X71,X78)),'_dst_x00x11_2'(s(fun(X72,X78),X79),s(fun(X70,fun(X71,X72)),'const.pair.pair_CASE_2'(s('type.pair.prod'(X74,X73),'const.pair.,_2'(s(X74,X76),s(X73,X77))),s(fun(X74,fun(X73,fun(X70,fun(X71,X72)))),X75))))) = s(fun(X70,fun(X71,X78)),happ(s(fun(X73,fun(X70,fun(X71,X78))),'_dst_x000x111_2'(s(fun(X72,X78),X79),s(fun(X73,fun(X70,fun(X71,X72))),happ(s(fun(X74,fun(X73,fun(X70,fun(X71,X72)))),X75),s(X74,X76))))),s(X73,X77)))) )),
  inference(superposition,[],[f197,f184])).
fof(f6528,plain,(
  ( ! [X61,X68,X66,X64,X62,X60,X69,X67,X65,X63] : (s(fun(X60,fun(X61,X68)),'_dst_x00x11_2'(s(fun(X62,X68),X69),s(fun(X60,fun(X61,X62)),happ(s(fun(X64,fun(X60,fun(X61,X62))),X65),s(X64,happ(s(fun(X63,X64),X66),s(X63,X67))))))) = s(fun(X60,fun(X61,X68)),happ(s(fun(X63,fun(X60,fun(X61,X68))),'_dst_x000x111_2'(s(fun(X62,X68),X69),s(fun(X63,fun(X60,fun(X61,X62))),'_dst_x0x1_2'(s(fun(X64,fun(X60,fun(X61,X62))),X65),s(fun(X63,X64),X66))))),s(X63,X67)))) )),
  inference(superposition,[],[f197,f182])).
fof(f6527,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51] : (s(fun(X50,fun(X51,X58)),happ(s(fun(X53,fun(X50,fun(X51,X58))),'_dst_x000x111_2'(s(fun(X52,X58),X59),s(fun(X53,fun(X50,fun(X51,X52))),X54))),s(X53,happ(s(fun(X55,X53),X56),s(X55,X57))))) = s(fun(X50,fun(X51,X58)),'_dst_x00x11_2'(s(fun(X52,X58),X59),s(fun(X50,fun(X51,X52)),happ(s(fun(X55,fun(X50,fun(X51,X52))),'_dst_x0x1_2'(s(fun(X53,fun(X50,fun(X51,X52))),X54),s(fun(X55,X53),X56))),s(X55,X57)))))) )),
  inference(superposition,[],[f197,f182])).
fof(f6526,plain,(
  ( ! [X39,X37,X35,X33,X43,X41,X38,X36,X34,X32,X42,X40] : (s(fun(X32,fun(X33,X42)),'_dst_x00x11_2'(s(fun(X34,X42),X43),s(fun(X32,fun(X33,X34)),'_dst_x0x0x1_3'(s(fun(X36,fun(X37,fun(X33,X34))),X38),s(X36,X39),s(fun(X32,X37),happ(s(fun(X35,fun(X32,X37)),X40),s(X35,X41))))))) = s(fun(X32,fun(X33,X42)),happ(s(fun(X35,fun(X32,fun(X33,X42))),'_dst_x000x111_2'(s(fun(X34,X42),X43),s(fun(X35,fun(X32,fun(X33,X34))),'_dst_x00x00x11_3'(s(fun(X36,fun(X37,fun(X33,X34))),X38),s(X36,X39),s(fun(X35,fun(X32,X37)),X40))))),s(X35,X41)))) )),
  inference(superposition,[],[f197,f199])).
fof(f6525,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X22] : (s(fun(X22,fun(X23,X30)),'_dst_x00x11_2'(s(fun(X24,X30),X31),s(fun(X22,fun(X23,X24)),'_dst_x0x1_2'(s(fun(X26,fun(X23,X24)),X27),s(fun(X22,X26),happ(s(fun(X25,fun(X22,X26)),X28),s(X25,X29))))))) = s(fun(X22,fun(X23,X30)),happ(s(fun(X25,fun(X22,fun(X23,X30))),'_dst_x000x111_2'(s(fun(X24,X30),X31),s(fun(X25,fun(X22,fun(X23,X24))),'_dst_x00x11_2'(s(fun(X26,fun(X23,X24)),X27),s(fun(X25,fun(X22,X26)),X28))))),s(X25,X29)))) )),
  inference(superposition,[],[f197,f193])).
fof(f6524,plain,(
  ( ! [X14,X12,X10,X21,X19,X17,X15,X13,X11,X20,X18,X16] : (s(fun(X10,fun(X11,X20)),'_dst_x00x11_2'(s(fun(X12,X20),X21),s(fun(X10,fun(X11,X12)),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,fun(X11,X12))),X16),s(fun(X10,X14),X17),s(X15,X18))))) = s(fun(X10,fun(X11,X20)),happ(s(fun(X13,fun(X10,fun(X11,X20))),'_dst_x000x111_2'(s(fun(X12,X20),X21),s(fun(X13,fun(X10,fun(X11,X12))),'_dst_x00x01x00_3'(s(fun(X14,fun(X15,fun(X11,X12))),X16),s(fun(X10,X14),X17),s(X15,X18))))),s(X13,X19)))) )),
  inference(superposition,[],[f197,f198])).
fof(f6523,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X0,fun(X1,X8)),'_dst_x00x11_2'(s(fun(X2,X8),X9),s(fun(X0,fun(X1,X2)),'_dst_x00x11_2'(s(fun(X4,X2),X5),s(fun(X0,fun(X1,X4)),happ(s(fun(X3,fun(X0,fun(X1,X4))),X6),s(X3,X7))))))) = s(fun(X0,fun(X1,X8)),happ(s(fun(X3,fun(X0,fun(X1,X8))),'_dst_x000x111_2'(s(fun(X2,X8),X9),s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x111_2'(s(fun(X4,X2),X5),s(fun(X3,fun(X0,fun(X1,X4))),X6))))),s(X3,X7)))) )),
  inference(superposition,[],[f197,f197])).
fof(f6521,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),'_dst_x00x11_2'(s(fun(X3,X2),X4),s(fun(X0,fun(X1,X3)),happ(s(fun(X5,fun(X0,fun(X1,X3))),X6),s(X5,X7))))))),s(fun(X0,fun(X1,X2)),happ(s(fun(X5,fun(X0,fun(X1,X2))),'_dst_x000x111_2'(s(fun(X3,X2),X4),s(fun(X5,fun(X0,fun(X1,X3))),X6))),s(X5,X7))))))) )),
  inference(unit_resulting_resolution,[],[f197,f173])).
fof(f6522,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,fun(X1,X2)),bool),happ(s(fun(fun(X0,fun(X1,X2)),fun(fun(X0,fun(X1,X2)),bool)),'=_0'),s(fun(X0,fun(X1,X2)),happ(s(fun(X3,fun(X0,fun(X1,X2))),'_dst_x000x111_2'(s(fun(X4,X2),X5),s(fun(X3,fun(X0,fun(X1,X4))),X6))),s(X3,X7))))),s(fun(X0,fun(X1,X2)),'_dst_x00x11_2'(s(fun(X4,X2),X5),s(fun(X0,fun(X1,X4)),happ(s(fun(X3,fun(X0,fun(X1,X4))),X6),s(X3,X7))))))))) )),
  inference(unit_resulting_resolution,[],[f197,f173])).
fof(f6044,plain,(
  ( ! [X372,X374,X368,X370,X376,X367,X373,X375,X369,X371] : (p(s(bool,happ(s(fun(fun(X367,X368),bool),happ(s(fun(fun(X367,X368),fun(fun(X367,X368),bool)),'=_0'),s(fun(X367,X368),X376))),s(fun(X367,X368),'_dst_x0x0x1_3'(s(fun(X370,fun(X371,X368)),X372),s(X370,X373),s(fun(X367,X371),happ(s(fun(X369,fun(X367,X371)),X374),s(X369,X375)))))))) | s(fun(X367,X368),X376) != s(fun(X367,X368),happ(s(fun(X369,fun(X367,X368)),'_dst_x00x00x11_3'(s(fun(X370,fun(X371,X368)),X372),s(X370,X373),s(fun(X369,fun(X367,X371)),X374))),s(X369,X375)))) )),
  inference(superposition,[],[f173,f199])).
fof(f6043,plain,(
  ( ! [X358,X364,X366,X360,X362,X357,X359,X365,X361,X363] : (p(s(bool,happ(s(fun(fun(X357,X358),bool),happ(s(fun(fun(X357,X358),fun(fun(X357,X358),bool)),'=_0'),s(fun(X357,X358),'_dst_x0x0x1_3'(s(fun(X360,fun(X361,X358)),X362),s(X360,X363),s(fun(X357,X361),happ(s(fun(X359,fun(X357,X361)),X364),s(X359,X365))))))),s(fun(X357,X358),X366)))) | s(fun(X357,X358),X366) != s(fun(X357,X358),happ(s(fun(X359,fun(X357,X358)),'_dst_x00x00x11_3'(s(fun(X360,fun(X361,X358)),X362),s(X360,X363),s(fun(X359,fun(X357,X361)),X364))),s(X359,X365)))) )),
  inference(superposition,[],[f173,f199])).
fof(f6042,plain,(
  ( ! [X350,X356,X352,X354,X349,X351,X347,X353,X355,X348] : (~p(s(bool,happ(s(fun(fun(X347,X348),bool),happ(s(fun(fun(X347,X348),fun(fun(X347,X348),bool)),'=_0'),s(fun(X347,X348),X356))),s(fun(X347,X348),'_dst_x0x0x1_3'(s(fun(X350,fun(X351,X348)),X352),s(X350,X353),s(fun(X347,X351),happ(s(fun(X349,fun(X347,X351)),X354),s(X349,X355)))))))) | s(fun(X347,X348),X356) = s(fun(X347,X348),happ(s(fun(X349,fun(X347,X348)),'_dst_x00x00x11_3'(s(fun(X350,fun(X351,X348)),X352),s(X350,X353),s(fun(X349,fun(X347,X351)),X354))),s(X349,X355)))) )),
  inference(superposition,[],[f172,f199])).
fof(f6041,plain,(
  ( ! [X344,X346,X341,X343,X337,X339,X345,X340,X342,X338] : (~p(s(bool,happ(s(fun(fun(X337,X338),bool),happ(s(fun(fun(X337,X338),fun(fun(X337,X338),bool)),'=_0'),s(fun(X337,X338),'_dst_x0x0x1_3'(s(fun(X340,fun(X341,X338)),X342),s(X340,X343),s(fun(X337,X341),happ(s(fun(X339,fun(X337,X341)),X344),s(X339,X345))))))),s(fun(X337,X338),X346)))) | s(fun(X337,X338),X346) = s(fun(X337,X338),happ(s(fun(X339,fun(X337,X338)),'_dst_x00x00x11_3'(s(fun(X340,fun(X341,X338)),X342),s(X340,X343),s(fun(X339,fun(X337,X341)),X344))),s(X339,X345)))) )),
  inference(superposition,[],[f172,f199])).
fof(f6036,plain,(
  ( ! [X292,X294,X288,X290,X285,X287,X293,X289,X291,X284,X286] : (s(X293,happ(s(fun(fun(X284,X285),X293),X294),s(fun(X284,X285),'_dst_x0x0x1_3'(s(fun(X287,fun(X288,X285)),X289),s(X287,X290),s(fun(X284,X288),happ(s(fun(X286,fun(X284,X288)),X291),s(X286,X292))))))) = s(X293,happ(s(fun(X286,X293),'_dst_x0x1_2'(s(fun(fun(X284,X285),X293),X294),s(fun(X286,fun(X284,X285)),'_dst_x00x00x11_3'(s(fun(X287,fun(X288,X285)),X289),s(X287,X290),s(fun(X286,fun(X284,X288)),X291))))),s(X286,X292)))) )),
  inference(superposition,[],[f182,f199])).
fof(f6035,plain,(
  ( ! [X280,X282,X271,X277,X279,X273,X275,X281,X283,X276,X278,X272,X274] : (s(X280,happ(s(fun(X273,X280),'_dst_x0x0x1_3'(s(fun(X281,fun(fun(X271,X272),X280)),X282),s(X281,X283),s(fun(X273,fun(X271,X272)),'_dst_x00x00x11_3'(s(fun(X274,fun(X275,X272)),X276),s(X274,X277),s(fun(X273,fun(X271,X275)),X278))))),s(X273,X279))) = s(X280,happ(s(fun(fun(X271,X272),X280),happ(s(fun(X281,fun(fun(X271,X272),X280)),X282),s(X281,X283))),s(fun(X271,X272),'_dst_x0x0x1_3'(s(fun(X274,fun(X275,X272)),X276),s(X274,X277),s(fun(X271,X275),happ(s(fun(X273,fun(X271,X275)),X278),s(X273,X279)))))))) )),
  inference(superposition,[],[f194,f199])).
fof(f6034,plain,(
  ( ! [X261,X263,X259,X269,X265,X267,X260,X262,X258,X268,X270,X264,X266] : (s(X267,happ(s(fun(X260,X267),'_dst_x0x1x0_3'(s(fun(fun(X258,X259),fun(X268,X267)),X269),s(fun(X260,fun(X258,X259)),'_dst_x00x00x11_3'(s(fun(X261,fun(X262,X259)),X263),s(X261,X264),s(fun(X260,fun(X258,X262)),X265))),s(X268,X270))),s(X260,X266))) = s(X267,happ(s(fun(X268,X267),happ(s(fun(fun(X258,X259),fun(X268,X267)),X269),s(fun(X258,X259),'_dst_x0x0x1_3'(s(fun(X261,fun(X262,X259)),X263),s(X261,X264),s(fun(X258,X262),happ(s(fun(X260,fun(X258,X262)),X265),s(X260,X266))))))),s(X268,X270)))) )),
  inference(superposition,[],[f195,f199])).
fof(f6033,plain,(
  ( ! [X255,X253,X251,X249,X257,X254,X252,X250,X248,X256,X247] : (~p(s(bool,happ(s(fun('type.pair.prod'(X247,X248),bool),'_dst_x0x0x1_3'(s(fun(X250,fun(X251,bool)),X252),s(X250,X253),s(fun('type.pair.prod'(X247,X248),X251),happ(s(fun(X249,fun('type.pair.prod'(X247,X248),X251)),X254),s(X249,X255))))),s('type.pair.prod'(X247,X248),sK7(happ(s(fun(X249,fun('type.pair.prod'(X247,X248),bool)),'_dst_x00x00x11_3'(s(fun(X250,fun(X251,bool)),X252),s(X250,X253),s(fun(X249,fun('type.pair.prod'(X247,X248),X251)),X254))),s(X249,X255)),X248,X247))))) | p(s(bool,happ(s(fun('type.pair.prod'(X247,X248),bool),happ(s(fun(X249,fun('type.pair.prod'(X247,X248),bool)),'_dst_x00x00x11_3'(s(fun(X250,fun(X251,bool)),X252),s(X250,X253),s(fun(X249,fun('type.pair.prod'(X247,X248),X251)),X254))),s(X249,X255))),s('type.pair.prod'(X247,X248),'const.pair.,_2'(s(X247,X256),s(X248,X257))))))) )),
  inference(superposition,[],[f168,f199])).
fof(f6032,plain,(
  ( ! [X243,X241,X239,X237,X246,X244,X242,X240,X238,X245] : (~p(s(bool,happ(s(fun('type.pair.prod'(X237,X238),bool),'_dst_x0x0x1_3'(s(fun(X240,fun(X241,bool)),X242),s(X240,X243),s(fun('type.pair.prod'(X237,X238),X241),happ(s(fun(X239,fun('type.pair.prod'(X237,X238),X241)),X244),s(X239,X245))))),s('type.pair.prod'(X237,X238),'const.pair.,_2'(s(X237,sK8(happ(s(fun(X239,fun('type.pair.prod'(X237,X238),bool)),'_dst_x00x00x11_3'(s(fun(X240,fun(X241,bool)),X242),s(X240,X243),s(fun(X239,fun('type.pair.prod'(X237,X238),X241)),X244))),s(X239,X245)),X238,X237)),s(X238,sK9(happ(s(fun(X239,fun('type.pair.prod'(X237,X238),bool)),'_dst_x00x00x11_3'(s(fun(X240,fun(X241,bool)),X242),s(X240,X243),s(fun(X239,fun('type.pair.prod'(X237,X238),X241)),X244))),s(X239,X245)),X238,X237))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X237,X238),bool),happ(s(fun(X239,fun('type.pair.prod'(X237,X238),bool)),'_dst_x00x00x11_3'(s(fun(X240,fun(X241,bool)),X242),s(X240,X243),s(fun(X239,fun('type.pair.prod'(X237,X238),X241)),X244))),s(X239,X245))),s('type.pair.prod'(X237,X238),X246))))) )),
  inference(superposition,[],[f169,f199])).
fof(f6017,plain,(
  ( ! [X47,X45,X43,X54,X52,X50,X48,X46,X44,X53,X51,X49] : (s(X44,happ(s(fun(X51,X44),'_dst_x0x1x0_3'(s(fun(X45,fun(X43,X44)),'_dst_x00x00x11_3'(s(fun(X46,fun(X47,X44)),X48),s(X46,X49),s(fun(X45,fun(X43,X47)),X50))),s(fun(X51,X45),X52),s(X43,X54))),s(X51,X53))) = s(X44,happ(s(fun(X43,X44),'_dst_x0x0x1_3'(s(fun(X46,fun(X47,X44)),X48),s(X46,X49),s(fun(X43,X47),happ(s(fun(X45,fun(X43,X47)),X50),s(X45,happ(s(fun(X51,X45),X52),s(X51,X53))))))),s(X43,X54)))) )),
  inference(superposition,[],[f195,f199])).
fof(f6016,plain,(
  ( ! [X39,X37,X35,X33,X41,X38,X36,X34,X42,X40] : (s(X34,happ(s(fun(X33,X34),'_dst_x0x0x1_3'(s(fun(X36,fun(X37,X34)),X38),s(X36,X39),s(fun(X33,X37),happ(s(fun(X35,fun(X33,X37)),X40),s(X35,X41))))),s(X33,X42))) = s(X34,'const.pair.UNCURRY_2'(s(fun(X35,fun(X33,X34)),'_dst_x00x00x11_3'(s(fun(X36,fun(X37,X34)),X38),s(X36,X39),s(fun(X35,fun(X33,X37)),X40))),s('type.pair.prod'(X35,X33),'const.pair.,_2'(s(X35,X41),s(X33,X42)))))) )),
  inference(superposition,[],[f183,f199])).
fof(f6015,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25,X32] : (s(X24,happ(s(fun(X23,X24),'_dst_x0x0x1_3'(s(fun(X26,fun(X27,X24)),X28),s(X26,X29),s(fun(X23,X27),happ(s(fun(X25,fun(X23,X27)),X30),s(X25,X31))))),s(X23,X32))) = s(X24,'const.pair.pair_CASE_2'(s('type.pair.prod'(X25,X23),'const.pair.,_2'(s(X25,X31),s(X23,X32))),s(fun(X25,fun(X23,X24)),'_dst_x00x00x11_3'(s(fun(X26,fun(X27,X24)),X28),s(X26,X29),s(fun(X25,fun(X23,X27)),X30)))))) )),
  inference(superposition,[],[f184,f199])).
fof(f6014,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(fun(X12,X21),'_dst_x0x1_2'(s(fun(X13,X21),X22),s(fun(X12,X13),'_dst_x0x0x1_3'(s(fun(X15,fun(X16,X13)),X17),s(X15,X18),s(fun(X12,X16),happ(s(fun(X14,fun(X12,X16)),X19),s(X14,X20))))))) = s(fun(X12,X21),happ(s(fun(X14,fun(X12,X21)),'_dst_x00x11_2'(s(fun(X13,X21),X22),s(fun(X14,fun(X12,X13)),'_dst_x00x00x11_3'(s(fun(X15,fun(X16,X13)),X17),s(X15,X18),s(fun(X14,fun(X12,X16)),X19))))),s(X14,X20)))) )),
  inference(superposition,[],[f193,f199])).
fof(f6013,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X1,happ(s(fun(X9,X1),'_dst_x0x0x1_3'(s(fun(X2,fun(X0,X1)),'_dst_x00x00x11_3'(s(fun(X3,fun(X4,X1)),X5),s(X3,X6),s(fun(X2,fun(X0,X4)),X7))),s(X2,X8),s(fun(X9,X0),X10))),s(X9,X11))) = s(X1,happ(s(fun(X0,X1),'_dst_x0x0x1_3'(s(fun(X3,fun(X4,X1)),X5),s(X3,X6),s(fun(X0,X4),happ(s(fun(X2,fun(X0,X4)),X7),s(X2,X8))))),s(X0,happ(s(fun(X9,X0),X10),s(X9,X11)))))) )),
  inference(superposition,[],[f194,f199])).
fof(f5995,plain,(
  ( ! [X280,X282,X277,X279,X275,X281,X283,X276,X278,X274] : (p(s(bool,happ(s(fun(fun(X274,X275),bool),happ(s(fun(fun(X274,X275),fun(fun(X274,X275),bool)),'=_0'),s(fun(X274,X275),X283))),s(fun(X274,X275),happ(s(fun(X280,fun(X274,X275)),'_dst_x00x00x11_3'(s(fun(X276,fun(X277,X275)),X278),s(X276,X279),s(fun(X280,fun(X274,X277)),X281))),s(X280,X282)))))) | s(fun(X274,X275),X283) != s(fun(X274,X275),'_dst_x0x0x1_3'(s(fun(X276,fun(X277,X275)),X278),s(X276,X279),s(fun(X274,X277),happ(s(fun(X280,fun(X274,X277)),X281),s(X280,X282)))))) )),
  inference(superposition,[],[f173,f199])).
fof(f5994,plain,(
  ( ! [X269,X271,X265,X267,X273,X268,X270,X264,X266,X272] : (p(s(bool,happ(s(fun(fun(X264,X265),bool),happ(s(fun(fun(X264,X265),fun(fun(X264,X265),bool)),'=_0'),s(fun(X264,X265),happ(s(fun(X270,fun(X264,X265)),'_dst_x00x00x11_3'(s(fun(X266,fun(X267,X265)),X268),s(X266,X269),s(fun(X270,fun(X264,X267)),X271))),s(X270,X272))))),s(fun(X264,X265),X273)))) | s(fun(X264,X265),X273) != s(fun(X264,X265),'_dst_x0x0x1_3'(s(fun(X266,fun(X267,X265)),X268),s(X266,X269),s(fun(X264,X267),happ(s(fun(X270,fun(X264,X267)),X271),s(X270,X272)))))) )),
  inference(superposition,[],[f173,f199])).
fof(f5993,plain,(
  ( ! [X255,X263,X257,X259,X261,X260,X254,X256,X258,X262] : (~p(s(bool,happ(s(fun(fun(X254,X255),bool),happ(s(fun(fun(X254,X255),fun(fun(X254,X255),bool)),'=_0'),s(fun(X254,X255),X263))),s(fun(X254,X255),happ(s(fun(X260,fun(X254,X255)),'_dst_x00x00x11_3'(s(fun(X256,fun(X257,X255)),X258),s(X256,X259),s(fun(X260,fun(X254,X257)),X261))),s(X260,X262)))))) | s(fun(X254,X255),X263) = s(fun(X254,X255),'_dst_x0x0x1_3'(s(fun(X256,fun(X257,X255)),X258),s(X256,X259),s(fun(X254,X257),happ(s(fun(X260,fun(X254,X257)),X261),s(X260,X262)))))) )),
  inference(superposition,[],[f172,f199])).
fof(f5992,plain,(
  ( ! [X253,X251,X249,X246,X244,X252,X250,X248,X247,X245] : (~p(s(bool,happ(s(fun(fun(X244,X245),bool),happ(s(fun(fun(X244,X245),fun(fun(X244,X245),bool)),'=_0'),s(fun(X244,X245),happ(s(fun(X250,fun(X244,X245)),'_dst_x00x00x11_3'(s(fun(X246,fun(X247,X245)),X248),s(X246,X249),s(fun(X250,fun(X244,X247)),X251))),s(X250,X252))))),s(fun(X244,X245),X253)))) | s(fun(X244,X245),X253) = s(fun(X244,X245),'_dst_x0x0x1_3'(s(fun(X246,fun(X247,X245)),X248),s(X246,X249),s(fun(X244,X247),happ(s(fun(X250,fun(X244,X247)),X251),s(X250,X252)))))) )),
  inference(superposition,[],[f172,f199])).
fof(f5988,plain,(
  ( ! [X212,X210,X208,X206,X204,X202,X211,X209,X207,X205,X203] : (~p(s(bool,happ(s(fun('type.pair.prod'(X202,X203),bool),happ(s(fun(X208,fun('type.pair.prod'(X202,X203),bool)),'_dst_x00x00x11_3'(s(fun(X204,fun(X205,bool)),X206),s(X204,X207),s(fun(X208,fun('type.pair.prod'(X202,X203),X205)),X209))),s(X208,X210))),s('type.pair.prod'(X202,X203),sK7('_dst_x0x0x1_3'(s(fun(X204,fun(X205,bool)),X206),s(X204,X207),s(fun('type.pair.prod'(X202,X203),X205),happ(s(fun(X208,fun('type.pair.prod'(X202,X203),X205)),X209),s(X208,X210)))),X203,X202))))) | p(s(bool,happ(s(fun('type.pair.prod'(X202,X203),bool),'_dst_x0x0x1_3'(s(fun(X204,fun(X205,bool)),X206),s(X204,X207),s(fun('type.pair.prod'(X202,X203),X205),happ(s(fun(X208,fun('type.pair.prod'(X202,X203),X205)),X209),s(X208,X210))))),s('type.pair.prod'(X202,X203),'const.pair.,_2'(s(X202,X211),s(X203,X212))))))) )),
  inference(superposition,[],[f168,f199])).
fof(f5987,plain,(
  ( ! [X198,X196,X194,X192,X200,X199,X197,X195,X193,X201] : (~p(s(bool,happ(s(fun('type.pair.prod'(X192,X193),bool),happ(s(fun(X198,fun('type.pair.prod'(X192,X193),bool)),'_dst_x00x00x11_3'(s(fun(X194,fun(X195,bool)),X196),s(X194,X197),s(fun(X198,fun('type.pair.prod'(X192,X193),X195)),X199))),s(X198,X200))),s('type.pair.prod'(X192,X193),'const.pair.,_2'(s(X192,sK8('_dst_x0x0x1_3'(s(fun(X194,fun(X195,bool)),X196),s(X194,X197),s(fun('type.pair.prod'(X192,X193),X195),happ(s(fun(X198,fun('type.pair.prod'(X192,X193),X195)),X199),s(X198,X200)))),X193,X192)),s(X193,sK9('_dst_x0x0x1_3'(s(fun(X194,fun(X195,bool)),X196),s(X194,X197),s(fun('type.pair.prod'(X192,X193),X195),happ(s(fun(X198,fun('type.pair.prod'(X192,X193),X195)),X199),s(X198,X200)))),X193,X192))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X192,X193),bool),'_dst_x0x0x1_3'(s(fun(X194,fun(X195,bool)),X196),s(X194,X197),s(fun('type.pair.prod'(X192,X193),X195),happ(s(fun(X198,fun('type.pair.prod'(X192,X193),X195)),X199),s(X198,X200))))),s('type.pair.prod'(X192,X193),X201))))) )),
  inference(superposition,[],[f169,f199])).
fof(f5972,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),happ(s(fun(X6,fun(X0,X1)),'_dst_x00x00x11_3'(s(fun(X2,fun(X3,X1)),X4),s(X2,X5),s(fun(X6,fun(X0,X3)),X7))),s(X6,X8))),s(X0,X9))) = s(X1,happ(s(fun(X3,X1),happ(s(fun(X2,fun(X3,X1)),X4),s(X2,X5))),s(X3,happ(s(fun(X0,X3),happ(s(fun(X6,fun(X0,X3)),X7),s(X6,X8))),s(X0,X9)))))) )),
  inference(superposition,[],[f194,f199])).
fof(f5971,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X9,X18,X16] : (s(fun(X9,X10),'_dst_x0x0x1_3'(s(fun(X12,fun(X13,X10)),X14),s(X12,X15),s(fun(X9,X13),happ(s(fun(X11,fun(X9,X13)),X16),s(X11,happ(s(fun(X17,X11),X18),s(X17,X19))))))) = s(fun(X9,X10),happ(s(fun(X17,fun(X9,X10)),'_dst_x0x1_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x00x11_3'(s(fun(X12,fun(X13,X10)),X14),s(X12,X15),s(fun(X11,fun(X9,X13)),X16))),s(fun(X17,X11),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f199,f182])).
fof(f5852,plain,(
  ( ! [X167,X165,X163,X161,X159,X166,X164,X162,X160,X168] : (s(fun(X159,X165),'_dst_x0x0x1_3'(s(fun(X166,fun(X160,X165)),X167),s(X166,X168),s(fun(X159,X160),'const.pair.UNCURRY_2'(s(fun(X161,fun(X162,fun(X159,X160))),X163),s('type.pair.prod'(X161,X162),X164))))) = s(fun(X159,X165),happ(s(fun('type.pair.prod'(X161,X162),fun(X159,X165)),'_dst_x00x00x11_3'(s(fun(X166,fun(X160,X165)),X167),s(X166,X168),s(fun('type.pair.prod'(X161,X162),fun(X159,X160)),happ(s(fun(fun(X161,fun(X162,fun(X159,X160))),fun('type.pair.prod'(X161,X162),fun(X159,X160))),'const.pair.UNCURRY_0'),s(fun(X161,fun(X162,fun(X159,X160))),X163))))),s('type.pair.prod'(X161,X162),X164)))) )),
  inference(superposition,[],[f199,f180])).
fof(f5851,plain,(
  ( ! [X146,X144,X142,X140,X138,X136,X147,X145,X143,X141,X139,X137,X148] : (s(fun(X136,X145),happ(s(fun(X138,fun(X136,X145)),'_dst_x00x00x11_3'(s(fun(X146,fun(X137,X145)),X147),s(X146,X148),s(fun(X138,fun(X136,X137)),'_dst_x0x0x1_3'(s(fun(X139,fun(X140,fun(X136,X137))),X141),s(X139,X142),s(fun(X138,X140),X143))))),s(X138,X144))) = s(fun(X136,X145),'_dst_x0x0x1_3'(s(fun(X146,fun(X137,X145)),X147),s(X146,X148),s(fun(X136,X137),happ(s(fun(X140,fun(X136,X137)),happ(s(fun(X139,fun(X140,fun(X136,X137))),X141),s(X139,X142))),s(X140,happ(s(fun(X138,X140),X143),s(X138,X144)))))))) )),
  inference(superposition,[],[f199,f194])).
fof(f5850,plain,(
  ( ! [X127,X134,X132,X130,X128,X135,X133,X131,X129] : (s(fun(X127,X132),'_dst_x0x0x1_3'(s(fun(X133,fun(X128,X132)),X134),s(X133,X135),s(fun(X127,X128),X130))) = s(fun(X127,X132),happ(s(fun(X129,fun(X127,X132)),'_dst_x00x00x11_3'(s(fun(X133,fun(X128,X132)),X134),s(X133,X135),s(fun(X129,fun(X127,X128)),'_dst_x0_1'(s(fun(X127,X128),X130))))),s(X129,X131)))) )),
  inference(superposition,[],[f199,f170])).
fof(f5849,plain,(
  ( ! [X125,X123,X121,X118,X116,X114,X126,X124,X122,X120,X119,X117,X115] : (s(fun(X114,X123),happ(s(fun(X116,fun(X114,X123)),'_dst_x00x00x11_3'(s(fun(X124,fun(X115,X123)),X125),s(X124,X126),s(fun(X116,fun(X114,X115)),'_dst_x0x1x0_3'(s(fun(X117,fun(X118,fun(X114,X115))),X119),s(fun(X116,X117),X120),s(X118,X121))))),s(X116,X122))) = s(fun(X114,X123),'_dst_x0x0x1_3'(s(fun(X124,fun(X115,X123)),X125),s(X124,X126),s(fun(X114,X115),happ(s(fun(X118,fun(X114,X115)),happ(s(fun(X117,fun(X118,fun(X114,X115))),X119),s(X117,happ(s(fun(X116,X117),X120),s(X116,X122))))),s(X118,X121)))))) )),
  inference(superposition,[],[f199,f195])).
fof(f5848,plain,(
  ( ! [X103,X101,X111,X109,X107,X105,X112,X102,X110,X108,X106,X104,X113] : (s(fun(X101,X110),'_dst_x0x0x1_3'(s(fun(X111,fun(X102,X110)),X112),s(X111,X113),s(fun(X101,X102),happ(s(fun(X106,fun(X101,X102)),'_dst_x0x1x0_3'(s(fun(X104,fun(X103,fun(X101,X102))),X105),s(fun(X106,X104),X107),s(X103,X109))),s(X106,X108))))) = s(fun(X101,X110),happ(s(fun(X103,fun(X101,X110)),'_dst_x00x00x11_3'(s(fun(X111,fun(X102,X110)),X112),s(X111,X113),s(fun(X103,fun(X101,X102)),happ(s(fun(X104,fun(X103,fun(X101,X102))),X105),s(X104,happ(s(fun(X106,X104),X107),s(X106,X108))))))),s(X103,X109)))) )),
  inference(superposition,[],[f199,f195])).
fof(f5847,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X95,X93,X91,X89,X100,X98,X96] : (s(fun(X88,X97),'_dst_x0x0x1_3'(s(fun(X98,fun(X89,X97)),X99),s(X98,X100),s(fun(X88,X89),happ(s(fun(X94,fun(X88,X89)),'_dst_x0x0x1_3'(s(fun(X91,fun(X90,fun(X88,X89))),X92),s(X91,X93),s(fun(X94,X90),X95))),s(X94,X96))))) = s(fun(X88,X97),happ(s(fun(X90,fun(X88,X97)),'_dst_x00x00x11_3'(s(fun(X98,fun(X89,X97)),X99),s(X98,X100),s(fun(X90,fun(X88,X89)),happ(s(fun(X91,fun(X90,fun(X88,X89))),X92),s(X91,X93))))),s(X90,happ(s(fun(X94,X90),X95),s(X94,X96)))))) )),
  inference(superposition,[],[f199,f194])).
fof(f5845,plain,(
  ( ! [X70,X68,X66,X76,X74,X72,X71,X69,X67,X75,X73] : (s(fun(X66,X73),'_dst_x0x0x1_3'(s(fun(X74,fun(X67,X73)),X75),s(X74,X76),s(fun(X66,X67),'const.pair.pair_CASE_2'(s('type.pair.prod'(X69,X68),'const.pair.,_2'(s(X69,X71),s(X68,X72))),s(fun(X69,fun(X68,fun(X66,X67))),X70))))) = s(fun(X66,X73),happ(s(fun(X68,fun(X66,X73)),'_dst_x00x00x11_3'(s(fun(X74,fun(X67,X73)),X75),s(X74,X76),s(fun(X68,fun(X66,X67)),happ(s(fun(X69,fun(X68,fun(X66,X67))),X70),s(X69,X71))))),s(X68,X72)))) )),
  inference(superposition,[],[f199,f184])).
fof(f5844,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X65,X55,X63] : (s(fun(X55,X62),'_dst_x0x0x1_3'(s(fun(X63,fun(X56,X62)),X64),s(X63,X65),s(fun(X55,X56),happ(s(fun(X58,fun(X55,X56)),X59),s(X58,happ(s(fun(X57,X58),X60),s(X57,X61))))))) = s(fun(X55,X62),happ(s(fun(X57,fun(X55,X62)),'_dst_x00x00x11_3'(s(fun(X63,fun(X56,X62)),X64),s(X63,X65),s(fun(X57,fun(X55,X56)),'_dst_x0x1_2'(s(fun(X58,fun(X55,X56)),X59),s(fun(X57,X58),X60))))),s(X57,X61)))) )),
  inference(superposition,[],[f199,f182])).
fof(f5843,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X44,X53,X51,X49] : (s(fun(X44,X51),happ(s(fun(X46,fun(X44,X51)),'_dst_x00x00x11_3'(s(fun(X52,fun(X45,X51)),X53),s(X52,X54),s(fun(X46,fun(X44,X45)),X47))),s(X46,happ(s(fun(X48,X46),X49),s(X48,X50))))) = s(fun(X44,X51),'_dst_x0x0x1_3'(s(fun(X52,fun(X45,X51)),X53),s(X52,X54),s(fun(X44,X45),happ(s(fun(X48,fun(X44,X45)),'_dst_x0x1_2'(s(fun(X46,fun(X44,X45)),X47),s(fun(X48,X46),X49))),s(X48,X50)))))) )),
  inference(superposition,[],[f199,f182])).
fof(f5842,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X36,X34,X32] : (s(fun(X24,X33),'_dst_x0x0x1_3'(s(fun(X34,fun(X25,X33)),X35),s(X34,X36),s(fun(X24,X25),'_dst_x0x0x1_3'(s(fun(X27,fun(X28,X25)),X29),s(X27,X30),s(fun(X24,X28),happ(s(fun(X26,fun(X24,X28)),X31),s(X26,X32))))))) = s(fun(X24,X33),happ(s(fun(X26,fun(X24,X33)),'_dst_x00x00x11_3'(s(fun(X34,fun(X25,X33)),X35),s(X34,X36),s(fun(X26,fun(X24,X25)),'_dst_x00x00x11_3'(s(fun(X27,fun(X28,X25)),X29),s(X27,X30),s(fun(X26,fun(X24,X28)),X31))))),s(X26,X32)))) )),
  inference(superposition,[],[f199,f199])).
fof(f5841,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(fun(X13,X20),'_dst_x0x0x1_3'(s(fun(X21,fun(X14,X20)),X22),s(X21,X23),s(fun(X13,X14),'_dst_x0x1_2'(s(fun(X16,X14),X17),s(fun(X13,X16),happ(s(fun(X15,fun(X13,X16)),X18),s(X15,X19))))))) = s(fun(X13,X20),happ(s(fun(X15,fun(X13,X20)),'_dst_x00x00x11_3'(s(fun(X21,fun(X14,X20)),X22),s(X21,X23),s(fun(X15,fun(X13,X14)),'_dst_x00x11_2'(s(fun(X16,X14),X17),s(fun(X15,fun(X13,X16)),X18))))),s(X15,X19)))) )),
  inference(superposition,[],[f199,f193])).
fof(f5840,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(fun(X0,X9),'_dst_x0x0x1_3'(s(fun(X10,fun(X1,X9)),X11),s(X10,X12),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))))) = s(fun(X0,X9),happ(s(fun(X2,fun(X0,X9)),'_dst_x00x00x11_3'(s(fun(X10,fun(X1,X9)),X11),s(X10,X12),s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))))),s(X2,X8)))) )),
  inference(superposition,[],[f199,f198])).
fof(f5838,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x0x1_3'(s(fun(X2,fun(X3,X1)),X4),s(X2,X5),s(fun(X0,X3),happ(s(fun(X6,fun(X0,X3)),X7),s(X6,X8))))))),s(fun(X0,X1),happ(s(fun(X6,fun(X0,X1)),'_dst_x00x00x11_3'(s(fun(X2,fun(X3,X1)),X4),s(X2,X5),s(fun(X6,fun(X0,X3)),X7))),s(X6,X8))))))) )),
  inference(unit_resulting_resolution,[],[f199,f173])).
fof(f5839,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x00x00x11_3'(s(fun(X3,fun(X4,X1)),X5),s(X3,X6),s(fun(X2,fun(X0,X4)),X7))),s(X2,X8))))),s(fun(X0,X1),'_dst_x0x0x1_3'(s(fun(X3,fun(X4,X1)),X5),s(X3,X6),s(fun(X0,X4),happ(s(fun(X2,fun(X0,X4)),X7),s(X2,X8))))))))) )),
  inference(unit_resulting_resolution,[],[f199,f173])).
fof(f5609,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(X16,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X14)),X16),happ(s(fun(X17,fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X14)),X16)),X18),s(X17,X19))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X14)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X14))),X15))))) = s(X16,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X14))),X16),'_dst_x0x0x1_3'(s(fun(X17,fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X14)),X16)),X18),s(X17,X19),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X14))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X14))),'const.closSem.evaluate_0'))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X14))),X15)))) )),
  inference(superposition,[],[f194,f161])).
fof(f5608,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(X10,happ(s(fun(X11,X10),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8)),fun(X11,X10)),X12),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X8))),X9))))),s(X11,X13))) = s(X10,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X8))),X10),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8)),fun(X11,X10)),X12),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X8))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X8))),'const.closSem.evaluate_0'),s(X11,X13))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X8))),X9)))) )),
  inference(superposition,[],[f195,f161])).
fof(f5635,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool),'_dst_x0x0x1_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool)),'=_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))) )),
  inference(forward_demodulation,[],[f5590,f194])).
fof(f5590,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool)),'=_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))))) )),
  inference(unit_resulting_resolution,[],[f161,f173])).
fof(f5634,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool)),'=_0'),s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))) )),
  inference(forward_demodulation,[],[f5591,f195])).
fof(f5591,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),bool)),'=_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))))) )),
  inference(unit_resulting_resolution,[],[f161,f173])).
fof(f5364,plain,(
  ( ! [X705,X699,X707,X708,X700,X704,X702,X706,X709,X701,X703] : (~p(s(bool,happ(s(fun('type.pair.prod'(X699,X700),bool),happ(s(fun(X704,fun('type.pair.prod'(X699,X700),bool)),'_dst_x0x1x0_3'(s(fun(X702,fun(X701,fun('type.pair.prod'(X699,X700),bool))),X703),s(fun(X704,X702),X705),s(X701,X707))),s(X704,X706))),s('type.pair.prod'(X699,X700),sK7(happ(s(fun(X701,fun('type.pair.prod'(X699,X700),bool)),happ(s(fun(X702,fun(X701,fun('type.pair.prod'(X699,X700),bool))),X703),s(X702,happ(s(fun(X704,X702),X705),s(X704,X706))))),s(X701,X707)),X700,X699))))) | p(s(bool,happ(s(fun('type.pair.prod'(X699,X700),bool),happ(s(fun(X701,fun('type.pair.prod'(X699,X700),bool)),happ(s(fun(X702,fun(X701,fun('type.pair.prod'(X699,X700),bool))),X703),s(X702,happ(s(fun(X704,X702),X705),s(X704,X706))))),s(X701,X707))),s('type.pair.prod'(X699,X700),'const.pair.,_2'(s(X699,X708),s(X700,X709))))))) )),
  inference(superposition,[],[f168,f195])).
fof(f5363,plain,(
  ( ! [X697,X693,X695,X689,X691,X696,X698,X692,X694,X690] : (~p(s(bool,happ(s(fun('type.pair.prod'(X689,X690),bool),happ(s(fun(X694,fun('type.pair.prod'(X689,X690),bool)),'_dst_x0x1x0_3'(s(fun(X692,fun(X691,fun('type.pair.prod'(X689,X690),bool))),X693),s(fun(X694,X692),X695),s(X691,X697))),s(X694,X696))),s('type.pair.prod'(X689,X690),'const.pair.,_2'(s(X689,sK8(happ(s(fun(X691,fun('type.pair.prod'(X689,X690),bool)),happ(s(fun(X692,fun(X691,fun('type.pair.prod'(X689,X690),bool))),X693),s(X692,happ(s(fun(X694,X692),X695),s(X694,X696))))),s(X691,X697)),X690,X689)),s(X690,sK9(happ(s(fun(X691,fun('type.pair.prod'(X689,X690),bool)),happ(s(fun(X692,fun(X691,fun('type.pair.prod'(X689,X690),bool))),X693),s(X692,happ(s(fun(X694,X692),X695),s(X694,X696))))),s(X691,X697)),X690,X689))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X689,X690),bool),happ(s(fun(X691,fun('type.pair.prod'(X689,X690),bool)),happ(s(fun(X692,fun(X691,fun('type.pair.prod'(X689,X690),bool))),X693),s(X692,happ(s(fun(X694,X692),X695),s(X694,X696))))),s(X691,X697))),s('type.pair.prod'(X689,X690),X698))))) )),
  inference(superposition,[],[f169,f195])).
fof(f5350,plain,(
  ( ! [X532,X534,X528,X530,X526,X533,X529,X531,X525,X527] : (s(X526,happ(s(fun(X525,X526),happ(s(fun(X530,fun(X525,X526)),'_dst_x0x1x0_3'(s(fun(X528,fun(X527,fun(X525,X526))),X529),s(fun(X530,X528),X531),s(X527,X533))),s(X530,X532))),s(X525,X534))) = s(X526,'const.pair.UNCURRY_2'(s(fun(X527,fun(X525,X526)),happ(s(fun(X528,fun(X527,fun(X525,X526))),X529),s(X528,happ(s(fun(X530,X528),X531),s(X530,X532))))),s('type.pair.prod'(X527,X525),'const.pair.,_2'(s(X527,X533),s(X525,X534)))))) )),
  inference(superposition,[],[f183,f195])).
fof(f5349,plain,(
  ( ! [X524,X520,X522,X516,X518,X521,X523,X517,X519,X515] : (s(X516,happ(s(fun(X515,X516),happ(s(fun(X520,fun(X515,X516)),'_dst_x0x1x0_3'(s(fun(X518,fun(X517,fun(X515,X516))),X519),s(fun(X520,X518),X521),s(X517,X523))),s(X520,X522))),s(X515,X524))) = s(X516,'const.pair.pair_CASE_2'(s('type.pair.prod'(X517,X515),'const.pair.,_2'(s(X517,X523),s(X515,X524))),s(fun(X517,fun(X515,X516)),happ(s(fun(X518,fun(X517,fun(X515,X516))),X519),s(X518,happ(s(fun(X520,X518),X521),s(X520,X522)))))))) )),
  inference(superposition,[],[f184,f195])).
fof(f5348,plain,(
  ( ! [X509,X511,X505,X507,X512,X514,X508,X510,X504,X506,X513] : (s(fun(X504,X513),'_dst_x0x1_2'(s(fun(X505,X513),X514),s(fun(X504,X505),happ(s(fun(X509,fun(X504,X505)),'_dst_x0x1x0_3'(s(fun(X507,fun(X506,fun(X504,X505))),X508),s(fun(X509,X507),X510),s(X506,X512))),s(X509,X511))))) = s(fun(X504,X513),happ(s(fun(X506,fun(X504,X513)),'_dst_x00x11_2'(s(fun(X505,X513),X514),s(fun(X506,fun(X504,X505)),happ(s(fun(X507,fun(X506,fun(X504,X505))),X508),s(X507,happ(s(fun(X509,X507),X510),s(X509,X511))))))),s(X506,X512)))) )),
  inference(superposition,[],[f193,f195])).
fof(f5347,plain,(
  ( ! [X499,X492,X494,X500,X502,X496,X498,X493,X495,X501,X503,X497] : (s(X493,happ(s(fun(X492,X493),happ(s(fun(X497,fun(X492,X493)),'_dst_x0x1x0_3'(s(fun(X495,fun(X494,fun(X492,X493))),X496),s(fun(X497,X495),X498),s(X494,X500))),s(X497,X499))),s(X492,happ(s(fun(X501,X492),X502),s(X501,X503))))) = s(X493,happ(s(fun(X501,X493),'_dst_x0x0x1_3'(s(fun(X494,fun(X492,X493)),happ(s(fun(X495,fun(X494,fun(X492,X493))),X496),s(X495,happ(s(fun(X497,X495),X498),s(X497,X499))))),s(X494,X500),s(fun(X501,X492),X502))),s(X501,X503)))) )),
  inference(superposition,[],[f194,f195])).
fof(f5345,plain,(
  ( ! [X476,X478,X480,X482,X475,X477,X479,X481] : (s(bool,'F_0') != s(bool,happ(s(fun(X478,bool),'_dst_x0x1x0_3'(s(fun(X476,fun(X475,bool)),X477),s(fun(X478,X476),X479),s(X475,X481))),s(X478,X480))) | s(bool,'T_0') = s(bool,X482) | s(bool,X482) = s(bool,happ(s(fun(X475,bool),happ(s(fun(X476,fun(X475,bool)),X477),s(X476,happ(s(fun(X478,X476),X479),s(X478,X480))))),s(X475,X481)))) )),
  inference(superposition,[],[f1176,f195])).
fof(f5344,plain,(
  ( ! [X472,X474,X469,X471,X473,X468,X470] : (p(s(bool,happ(s(fun(X471,bool),'_dst_x0x1x0_3'(s(fun(X469,fun(X468,bool)),X470),s(fun(X471,X469),X472),s(X468,X474))),s(X471,X473)))) | s(bool,'F_0') = s(bool,happ(s(fun(X468,bool),happ(s(fun(X469,fun(X468,bool)),X470),s(X469,happ(s(fun(X471,X469),X472),s(X471,X473))))),s(X468,X474)))) )),
  inference(superposition,[],[f1018,f195])).
fof(f5343,plain,(
  ( ! [X466,X461,X463,X465,X467,X460,X462,X464] : (s(bool,X467) != s(bool,happ(s(fun(X463,bool),'_dst_x0x1x0_3'(s(fun(X461,fun(X460,bool)),X462),s(fun(X463,X461),X464),s(X460,X466))),s(X463,X465))) | s(bool,'T_0') = s(bool,X467) | s(bool,'F_0') = s(bool,happ(s(fun(X460,bool),happ(s(fun(X461,fun(X460,bool)),X462),s(X461,happ(s(fun(X463,X461),X464),s(X463,X465))))),s(X460,X466)))) )),
  inference(superposition,[],[f638,f195])).
fof(f5342,plain,(
  ( ! [X453,X455,X457,X459,X452,X454,X456,X458] : (s(bool,X459) != s(bool,happ(s(fun(X455,bool),'_dst_x0x1x0_3'(s(fun(X453,fun(X452,bool)),X454),s(fun(X455,X453),X456),s(X452,X458))),s(X455,X457))) | s(bool,'T_0') = s(bool,happ(s(fun(X452,bool),happ(s(fun(X453,fun(X452,bool)),X454),s(X453,happ(s(fun(X455,X453),X456),s(X455,X457))))),s(X452,X458))) | s(bool,'F_0') = s(bool,X459)) )),
  inference(superposition,[],[f638,f195])).
fof(f5341,plain,(
  ( ! [X445,X447,X449,X451,X444,X446,X448,X450] : (p(s(bool,happ(s(fun(X447,bool),'_dst_x0x1x0_3'(s(fun(X445,fun(X444,bool)),X446),s(fun(X447,X445),X448),s(X444,X450))),s(X447,X449)))) | s(bool,X451) != s(bool,happ(s(fun(X444,bool),happ(s(fun(X445,fun(X444,bool)),X446),s(X445,happ(s(fun(X447,X445),X448),s(X447,X449))))),s(X444,X450))) | s(bool,'F_0') = s(bool,X451)) )),
  inference(superposition,[],[f616,f195])).
fof(f5340,plain,(
  ( ! [X441,X443,X436,X438,X440,X442,X437,X439,X435] : (s(bool,'F_0') = s(bool,happ(s(fun(X442,bool),'_dst_x0_1'(s(bool,happ(s(fun(X438,bool),'_dst_x0x1x0_3'(s(fun(X436,fun(X435,bool)),X437),s(fun(X438,X436),X439),s(X435,X441))),s(X438,X440))))),s(X442,X443))) | s(bool,'T_0') = s(bool,happ(s(fun(X435,bool),happ(s(fun(X436,fun(X435,bool)),X437),s(X436,happ(s(fun(X438,X436),X439),s(X438,X440))))),s(X435,X441)))) )),
  inference(superposition,[],[f614,f195])).
fof(f5339,plain,(
  ( ! [X428,X430,X426,X432,X434,X429,X431,X427,X433] : (p(s(bool,happ(s(fun(X429,bool),'_dst_x0x1x0_3'(s(fun(X427,fun(X426,bool)),X428),s(fun(X429,X427),X430),s(X426,X432))),s(X429,X431)))) | s(bool,X433) != s(bool,happ(s(fun(X426,bool),happ(s(fun(X427,fun(X426,bool)),X428),s(X427,happ(s(fun(X429,X427),X430),s(X429,X431))))),s(X426,X432))) | s(bool,X433) = s(bool,X434) | s(bool,'T_0') = s(bool,X434)) )),
  inference(superposition,[],[f489,f195])).
fof(f5337,plain,(
  ( ! [X412,X414,X416,X418,X413,X415,X411,X417] : (p(s(bool,happ(s(fun(X414,bool),'_dst_x0x1x0_3'(s(fun(X412,fun(X411,bool)),X413),s(fun(X414,X412),X415),s(X411,X417))),s(X414,X416)))) | s(bool,X418) = s(bool,happ(s(fun(X411,bool),happ(s(fun(X412,fun(X411,bool)),X413),s(X412,happ(s(fun(X414,X412),X415),s(X414,X416))))),s(X411,X417))) | s(bool,'T_0') = s(bool,X418)) )),
  inference(superposition,[],[f347,f195])).
fof(f5336,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (s(bool,'T_0') = s(bool,happ(s(fun(X409,bool),'_dst_x0_1'(s(bool,happ(s(fun(X405,bool),'_dst_x0x1x0_3'(s(fun(X403,fun(X402,bool)),X404),s(fun(X405,X403),X406),s(X402,X408))),s(X405,X407))))),s(X409,X410))) | s(bool,'F_0') = s(bool,happ(s(fun(X402,bool),happ(s(fun(X403,fun(X402,bool)),X404),s(X403,happ(s(fun(X405,X403),X406),s(X405,X407))))),s(X402,X408)))) )),
  inference(superposition,[],[f267,f195])).
fof(f5335,plain,(
  ( ! [X397,X399,X395,X401,X396,X398,X400] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X398,bool),'_dst_x0x1x0_3'(s(fun(X396,fun(X395,bool)),X397),s(fun(X398,X396),X399),s(X395,X401))),s(X398,X400))))) | s(bool,'F_0') = s(bool,happ(s(fun(X395,bool),happ(s(fun(X396,fun(X395,bool)),X397),s(X396,happ(s(fun(X398,X396),X399),s(X398,X400))))),s(X395,X401)))) )),
  inference(superposition,[],[f235,f195])).
fof(f5334,plain,(
  ( ! [X389,X391,X393,X388,X390,X392,X394] : (~p(s(bool,happ(s(fun(X391,bool),'_dst_x0x1x0_3'(s(fun(X389,fun(X388,bool)),X390),s(fun(X391,X389),X392),s(X388,X394))),s(X391,X393)))) | s(bool,'T_0') = s(bool,happ(s(fun(X388,bool),happ(s(fun(X389,fun(X388,bool)),X390),s(X389,happ(s(fun(X391,X389),X392),s(X391,X393))))),s(X388,X394)))) )),
  inference(superposition,[],[f225,f195])).
fof(f5329,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X355,X348] : (~p(s(bool,happ(s(fun(X351,bool),'_dst_x0x1x0_3'(s(fun(X349,fun(X348,bool)),X350),s(fun(X351,X349),X352),s(X348,X354))),s(X351,X353)))) | s(bool,X355) != s(bool,happ(s(fun(X348,bool),happ(s(fun(X349,fun(X348,bool)),X350),s(X349,happ(s(fun(X351,X349),X352),s(X351,X353))))),s(X348,X354))) | p(s(bool,X355))) )),
  inference(superposition,[],[f215,f195])).
fof(f5324,plain,(
  ( ! [X315,X316,X318,X317,X319] : (~p(s(bool,happ(s(fun(X316,bool),'_dst_x0x1x0_3'(s(fun(X315,fun(X315,bool)),'=_0'),s(fun(X316,X315),X317),s(X315,X319))),s(X316,X318)))) | s(X315,X319) = s(X315,happ(s(fun(X316,X315),X317),s(X316,X318)))) )),
  inference(superposition,[],[f172,f195])).
fof(f5323,plain,(
  ( ! [X313,X310,X312,X314,X311] : (p(s(bool,happ(s(fun(X311,bool),'_dst_x0x1x0_3'(s(fun(X310,fun(X310,bool)),'=_0'),s(fun(X311,X310),X312),s(X310,X314))),s(X311,X313)))) | s(X310,X314) != s(X310,happ(s(fun(X311,X310),X312),s(X311,X313)))) )),
  inference(superposition,[],[f173,f195])).
fof(f5307,plain,(
  ( ! [X127,X134,X132,X130,X128,X135,X133,X131,X129] : (p(s(bool,happ(s(fun(X127,bool),happ(s(fun(X127,fun(X127,bool)),'=_0'),s(X127,X135))),s(X127,happ(s(fun(X131,X127),'_dst_x0x1x0_3'(s(fun(X129,fun(X128,X127)),X130),s(fun(X131,X129),X132),s(X128,X134))),s(X131,X133)))))) | s(X127,X135) != s(X127,happ(s(fun(X128,X127),happ(s(fun(X129,fun(X128,X127)),X130),s(X129,happ(s(fun(X131,X129),X132),s(X131,X133))))),s(X128,X134)))) )),
  inference(superposition,[],[f173,f195])).
fof(f5306,plain,(
  ( ! [X125,X123,X121,X118,X126,X124,X122,X120,X119] : (p(s(bool,happ(s(fun(X118,bool),happ(s(fun(X118,fun(X118,bool)),'=_0'),s(X118,happ(s(fun(X122,X118),'_dst_x0x1x0_3'(s(fun(X120,fun(X119,X118)),X121),s(fun(X122,X120),X123),s(X119,X125))),s(X122,X124))))),s(X118,X126)))) | s(X118,X126) != s(X118,happ(s(fun(X119,X118),happ(s(fun(X120,fun(X119,X118)),X121),s(X120,happ(s(fun(X122,X120),X123),s(X122,X124))))),s(X119,X125)))) )),
  inference(superposition,[],[f173,f195])).
fof(f5305,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),'=_0'),s(X109,X117))),s(X109,happ(s(fun(X113,X109),'_dst_x0x1x0_3'(s(fun(X111,fun(X110,X109)),X112),s(fun(X113,X111),X114),s(X110,X116))),s(X113,X115)))))) | s(X109,X117) = s(X109,happ(s(fun(X110,X109),happ(s(fun(X111,fun(X110,X109)),X112),s(X111,happ(s(fun(X113,X111),X114),s(X113,X115))))),s(X110,X116)))) )),
  inference(superposition,[],[f172,f195])).
fof(f5304,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),'=_0'),s(X100,happ(s(fun(X104,X100),'_dst_x0x1x0_3'(s(fun(X102,fun(X101,X100)),X103),s(fun(X104,X102),X105),s(X101,X107))),s(X104,X106))))),s(X100,X108)))) | s(X100,X108) = s(X100,happ(s(fun(X101,X100),happ(s(fun(X102,fun(X101,X100)),X103),s(X102,happ(s(fun(X104,X102),X105),s(X104,X106))))),s(X101,X107)))) )),
  inference(superposition,[],[f172,f195])).
fof(f5294,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s(X20,happ(s(fun(X12,X20),X21),s(X12,happ(s(fun(X16,X12),'_dst_x0x1x0_3'(s(fun(X14,fun(X13,X12)),X15),s(fun(X16,X14),X17),s(X13,X19))),s(X16,X18))))) = s(X20,happ(s(fun(X13,X20),'_dst_x0x1_2'(s(fun(X12,X20),X21),s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))),s(X13,X19)))) )),
  inference(superposition,[],[f182,f195])).
fof(f5293,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X8,happ(s(fun(X0,X8),happ(s(fun(X9,fun(X0,X8)),X10),s(X9,X11))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6))))) = s(X8,happ(s(fun(X1,X8),'_dst_x0x0x1_3'(s(fun(X9,fun(X0,X8)),X10),s(X9,X11),s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))))),s(X1,X7)))) )),
  inference(superposition,[],[f194,f195])).
fof(f5289,plain,(
  ( ! [X652,X654,X648,X650,X646,X653,X655,X649,X651,X645,X647] : (~p(s(bool,happ(s(fun('type.pair.prod'(X645,X646),bool),happ(s(fun(X649,fun('type.pair.prod'(X645,X646),bool)),happ(s(fun(X648,fun(X649,fun('type.pair.prod'(X645,X646),bool))),X650),s(X648,happ(s(fun(X647,X648),X651),s(X647,X653))))),s(X649,X652))),s('type.pair.prod'(X645,X646),sK7(happ(s(fun(X647,fun('type.pair.prod'(X645,X646),bool)),'_dst_x0x1x0_3'(s(fun(X648,fun(X649,fun('type.pair.prod'(X645,X646),bool))),X650),s(fun(X647,X648),X651),s(X649,X652))),s(X647,X653)),X646,X645))))) | p(s(bool,happ(s(fun('type.pair.prod'(X645,X646),bool),happ(s(fun(X647,fun('type.pair.prod'(X645,X646),bool)),'_dst_x0x1x0_3'(s(fun(X648,fun(X649,fun('type.pair.prod'(X645,X646),bool))),X650),s(fun(X647,X648),X651),s(X649,X652))),s(X647,X653))),s('type.pair.prod'(X645,X646),'const.pair.,_2'(s(X645,X654),s(X646,X655))))))) )),
  inference(superposition,[],[f168,f195])).
fof(f5288,plain,(
  ( ! [X639,X635,X644,X640,X642,X636,X638,X641,X643,X637] : (~p(s(bool,happ(s(fun('type.pair.prod'(X635,X636),bool),happ(s(fun(X639,fun('type.pair.prod'(X635,X636),bool)),happ(s(fun(X638,fun(X639,fun('type.pair.prod'(X635,X636),bool))),X640),s(X638,happ(s(fun(X637,X638),X641),s(X637,X643))))),s(X639,X642))),s('type.pair.prod'(X635,X636),'const.pair.,_2'(s(X635,sK8(happ(s(fun(X637,fun('type.pair.prod'(X635,X636),bool)),'_dst_x0x1x0_3'(s(fun(X638,fun(X639,fun('type.pair.prod'(X635,X636),bool))),X640),s(fun(X637,X638),X641),s(X639,X642))),s(X637,X643)),X636,X635)),s(X636,sK9(happ(s(fun(X637,fun('type.pair.prod'(X635,X636),bool)),'_dst_x0x1x0_3'(s(fun(X638,fun(X639,fun('type.pair.prod'(X635,X636),bool))),X640),s(fun(X637,X638),X641),s(X639,X642))),s(X637,X643)),X636,X635))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X635,X636),bool),happ(s(fun(X637,fun('type.pair.prod'(X635,X636),bool)),'_dst_x0x1x0_3'(s(fun(X638,fun(X639,fun('type.pair.prod'(X635,X636),bool))),X640),s(fun(X637,X638),X641),s(X639,X642))),s(X637,X643))),s('type.pair.prod'(X635,X636),X644))))) )),
  inference(superposition,[],[f169,f195])).
fof(f5275,plain,(
  ( ! [X476,X478,X472,X474,X480,X471,X475,X477,X479,X473] : (s(X472,'const.pair.UNCURRY_2'(s(fun(X473,fun(X471,X472)),'_dst_x0x1x0_3'(s(fun(X474,fun(X475,fun(X471,X472))),X476),s(fun(X473,X474),X477),s(X475,X478))),s('type.pair.prod'(X473,X471),'const.pair.,_2'(s(X473,X479),s(X471,X480))))) = s(X472,happ(s(fun(X471,X472),happ(s(fun(X475,fun(X471,X472)),happ(s(fun(X474,fun(X475,fun(X471,X472))),X476),s(X474,happ(s(fun(X473,X474),X477),s(X473,X479))))),s(X475,X478))),s(X471,X480)))) )),
  inference(superposition,[],[f183,f195])).
fof(f5274,plain,(
  ( ! [X466,X461,X463,X469,X465,X467,X462,X468,X470,X464] : (s(X462,'const.pair.pair_CASE_2'(s('type.pair.prod'(X463,X461),'const.pair.,_2'(s(X463,X469),s(X461,X470))),s(fun(X463,fun(X461,X462)),'_dst_x0x1x0_3'(s(fun(X464,fun(X465,fun(X461,X462))),X466),s(fun(X463,X464),X467),s(X465,X468))))) = s(X462,happ(s(fun(X461,X462),happ(s(fun(X465,fun(X461,X462)),happ(s(fun(X464,fun(X465,fun(X461,X462))),X466),s(X464,happ(s(fun(X463,X464),X467),s(X463,X469))))),s(X465,X468))),s(X461,X470)))) )),
  inference(superposition,[],[f184,f195])).
fof(f5273,plain,(
  ( ! [X453,X455,X451,X457,X459,X452,X454,X450,X460,X456,X458] : (s(fun(X450,X459),happ(s(fun(X452,fun(X450,X459)),'_dst_x00x11_2'(s(fun(X451,X459),X460),s(fun(X452,fun(X450,X451)),'_dst_x0x1x0_3'(s(fun(X453,fun(X454,fun(X450,X451))),X455),s(fun(X452,X453),X456),s(X454,X457))))),s(X452,X458))) = s(fun(X450,X459),'_dst_x0x1_2'(s(fun(X451,X459),X460),s(fun(X450,X451),happ(s(fun(X454,fun(X450,X451)),happ(s(fun(X453,fun(X454,fun(X450,X451))),X455),s(X453,happ(s(fun(X452,X453),X456),s(X452,X458))))),s(X454,X457)))))) )),
  inference(superposition,[],[f193,f195])).
fof(f5272,plain,(
  ( ! [X445,X447,X441,X443,X449,X438,X444,X446,X440,X442,X448,X439] : (s(X439,happ(s(fun(X447,X439),'_dst_x0x0x1_3'(s(fun(X440,fun(X438,X439)),'_dst_x0x1x0_3'(s(fun(X441,fun(X442,fun(X438,X439))),X443),s(fun(X440,X441),X444),s(X442,X445))),s(X440,X446),s(fun(X447,X438),X448))),s(X447,X449))) = s(X439,happ(s(fun(X438,X439),happ(s(fun(X442,fun(X438,X439)),happ(s(fun(X441,fun(X442,fun(X438,X439))),X443),s(X441,happ(s(fun(X440,X441),X444),s(X440,X446))))),s(X442,X445))),s(X438,happ(s(fun(X447,X438),X448),s(X447,X449)))))) )),
  inference(superposition,[],[f194,f195])).
fof(f5270,plain,(
  ( ! [X422,X428,X424,X426,X421,X423,X425,X427] : (s(bool,'F_0') != s(bool,happ(s(fun(X423,bool),happ(s(fun(X422,fun(X423,bool)),X424),s(X422,happ(s(fun(X421,X422),X425),s(X421,X427))))),s(X423,X426))) | s(bool,'T_0') = s(bool,X428) | s(bool,X428) = s(bool,happ(s(fun(X421,bool),'_dst_x0x1x0_3'(s(fun(X422,fun(X423,bool)),X424),s(fun(X421,X422),X425),s(X423,X426))),s(X421,X427)))) )),
  inference(superposition,[],[f1176,f195])).
fof(f5269,plain,(
  ( ! [X414,X420,X416,X418,X415,X417,X419] : (p(s(bool,happ(s(fun(X416,bool),happ(s(fun(X415,fun(X416,bool)),X417),s(X415,happ(s(fun(X414,X415),X418),s(X414,X420))))),s(X416,X419)))) | s(bool,'F_0') = s(bool,happ(s(fun(X414,bool),'_dst_x0x1x0_3'(s(fun(X415,fun(X416,bool)),X417),s(fun(X414,X415),X418),s(X416,X419))),s(X414,X420)))) )),
  inference(superposition,[],[f1018,f195])).
fof(f5268,plain,(
  ( ! [X412,X408,X410,X407,X413,X409,X411,X406] : (s(bool,X413) != s(bool,happ(s(fun(X408,bool),happ(s(fun(X407,fun(X408,bool)),X409),s(X407,happ(s(fun(X406,X407),X410),s(X406,X412))))),s(X408,X411))) | s(bool,'T_0') = s(bool,X413) | s(bool,'F_0') = s(bool,happ(s(fun(X406,bool),'_dst_x0x1x0_3'(s(fun(X407,fun(X408,bool)),X409),s(fun(X406,X407),X410),s(X408,X411))),s(X406,X412)))) )),
  inference(superposition,[],[f638,f195])).
fof(f5267,plain,(
  ( ! [X399,X405,X401,X403,X398,X404,X400,X402] : (s(bool,X405) != s(bool,happ(s(fun(X400,bool),happ(s(fun(X399,fun(X400,bool)),X401),s(X399,happ(s(fun(X398,X399),X402),s(X398,X404))))),s(X400,X403))) | s(bool,'T_0') = s(bool,happ(s(fun(X398,bool),'_dst_x0x1x0_3'(s(fun(X399,fun(X400,bool)),X401),s(fun(X398,X399),X402),s(X400,X403))),s(X398,X404))) | s(bool,'F_0') = s(bool,X405)) )),
  inference(superposition,[],[f638,f195])).
fof(f5266,plain,(
  ( ! [X391,X397,X393,X395,X390,X396,X392,X394] : (p(s(bool,happ(s(fun(X392,bool),happ(s(fun(X391,fun(X392,bool)),X393),s(X391,happ(s(fun(X390,X391),X394),s(X390,X396))))),s(X392,X395)))) | s(bool,X397) != s(bool,happ(s(fun(X390,bool),'_dst_x0x1x0_3'(s(fun(X391,fun(X392,bool)),X393),s(fun(X390,X391),X394),s(X392,X395))),s(X390,X396))) | s(bool,'F_0') = s(bool,X397)) )),
  inference(superposition,[],[f616,f195])).
fof(f5265,plain,(
  ( ! [X383,X389,X385,X387,X382,X388,X384,X386,X381] : (s(bool,'F_0') = s(bool,happ(s(fun(X388,bool),'_dst_x0_1'(s(bool,happ(s(fun(X383,bool),happ(s(fun(X382,fun(X383,bool)),X384),s(X382,happ(s(fun(X381,X382),X385),s(X381,X387))))),s(X383,X386))))),s(X388,X389))) | s(bool,'T_0') = s(bool,happ(s(fun(X381,bool),'_dst_x0x1x0_3'(s(fun(X382,fun(X383,bool)),X384),s(fun(X381,X382),X385),s(X383,X386))),s(X381,X387)))) )),
  inference(superposition,[],[f614,f195])).
fof(f5264,plain,(
  ( ! [X377,X379,X372,X374,X380,X376,X378,X373,X375] : (p(s(bool,happ(s(fun(X374,bool),happ(s(fun(X373,fun(X374,bool)),X375),s(X373,happ(s(fun(X372,X373),X376),s(X372,X378))))),s(X374,X377)))) | s(bool,X379) != s(bool,happ(s(fun(X372,bool),'_dst_x0x1x0_3'(s(fun(X373,fun(X374,bool)),X375),s(fun(X372,X373),X376),s(X374,X377))),s(X372,X378))) | s(bool,X379) = s(bool,X380) | s(bool,'T_0') = s(bool,X380)) )),
  inference(superposition,[],[f489,f195])).
fof(f5262,plain,(
  ( ! [X358,X364,X360,X362,X357,X359,X361,X363] : (p(s(bool,happ(s(fun(X359,bool),happ(s(fun(X358,fun(X359,bool)),X360),s(X358,happ(s(fun(X357,X358),X361),s(X357,X363))))),s(X359,X362)))) | s(bool,X364) = s(bool,happ(s(fun(X357,bool),'_dst_x0x1x0_3'(s(fun(X358,fun(X359,bool)),X360),s(fun(X357,X358),X361),s(X359,X362))),s(X357,X363))) | s(bool,'T_0') = s(bool,X364)) )),
  inference(superposition,[],[f347,f195])).
fof(f5261,plain,(
  ( ! [X350,X356,X352,X354,X349,X351,X353,X355,X348] : (s(bool,'T_0') = s(bool,happ(s(fun(X355,bool),'_dst_x0_1'(s(bool,happ(s(fun(X350,bool),happ(s(fun(X349,fun(X350,bool)),X351),s(X349,happ(s(fun(X348,X349),X352),s(X348,X354))))),s(X350,X353))))),s(X355,X356))) | s(bool,'F_0') = s(bool,happ(s(fun(X348,bool),'_dst_x0x1x0_3'(s(fun(X349,fun(X350,bool)),X351),s(fun(X348,X349),X352),s(X350,X353))),s(X348,X354)))) )),
  inference(superposition,[],[f267,f195])).
fof(f5260,plain,(
  ( ! [X344,X346,X341,X343,X345,X347,X342] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X343,bool),happ(s(fun(X342,fun(X343,bool)),X344),s(X342,happ(s(fun(X341,X342),X345),s(X341,X347))))),s(X343,X346))))) | s(bool,'F_0') = s(bool,happ(s(fun(X341,bool),'_dst_x0x1x0_3'(s(fun(X342,fun(X343,bool)),X344),s(fun(X341,X342),X345),s(X343,X346))),s(X341,X347)))) )),
  inference(superposition,[],[f235,f195])).
fof(f5259,plain,(
  ( ! [X335,X337,X339,X334,X340,X336,X338] : (~p(s(bool,happ(s(fun(X336,bool),happ(s(fun(X335,fun(X336,bool)),X337),s(X335,happ(s(fun(X334,X335),X338),s(X334,X340))))),s(X336,X339)))) | s(bool,'T_0') = s(bool,happ(s(fun(X334,bool),'_dst_x0x1x0_3'(s(fun(X335,fun(X336,bool)),X337),s(fun(X334,X335),X338),s(X336,X339))),s(X334,X340)))) )),
  inference(superposition,[],[f225,f195])).
fof(f5254,plain,(
  ( ! [X294,X300,X296,X298,X295,X301,X297,X299] : (~p(s(bool,happ(s(fun(X296,bool),happ(s(fun(X295,fun(X296,bool)),X297),s(X295,happ(s(fun(X294,X295),X298),s(X294,X300))))),s(X296,X299)))) | s(bool,X301) != s(bool,happ(s(fun(X294,bool),'_dst_x0x1x0_3'(s(fun(X295,fun(X296,bool)),X297),s(fun(X294,X295),X298),s(X296,X299))),s(X294,X300))) | p(s(bool,X301))) )),
  inference(superposition,[],[f215,f195])).
fof(f5252,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X284,X286] : (~p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X281,fun(X282,bool)),X283),s(X281,happ(s(fun('type.pair.prod'(X279,X280),X281),X284),s('type.pair.prod'(X279,X280),'const.pair.,_2'(s(X279,sK8('_dst_x0x1x0_3'(s(fun(X281,fun(X282,bool)),X283),s(fun('type.pair.prod'(X279,X280),X281),X284),s(X282,X285)),X280,X279)),s(X280,sK9('_dst_x0x1x0_3'(s(fun(X281,fun(X282,bool)),X283),s(fun('type.pair.prod'(X279,X280),X281),X284),s(X282,X285)),X280,X279)))))))),s(X282,X285)))) | p(s(bool,happ(s(fun('type.pair.prod'(X279,X280),bool),'_dst_x0x1x0_3'(s(fun(X281,fun(X282,bool)),X283),s(fun('type.pair.prod'(X279,X280),X281),X284),s(X282,X285))),s('type.pair.prod'(X279,X280),X286))))) )),
  inference(superposition,[],[f169,f195])).
fof(f5251,plain,(
  ( ! [X271,X277,X273,X275,X270,X276,X278,X272,X274] : (~p(s(bool,happ(s(fun(X273,bool),happ(s(fun(X272,fun(X273,bool)),X274),s(X272,happ(s(fun('type.pair.prod'(X270,X271),X272),X275),s('type.pair.prod'(X270,X271),sK7('_dst_x0x1x0_3'(s(fun(X272,fun(X273,bool)),X274),s(fun('type.pair.prod'(X270,X271),X272),X275),s(X273,X276)),X271,X270)))))),s(X273,X276)))) | p(s(bool,happ(s(fun('type.pair.prod'(X270,X271),bool),'_dst_x0x1x0_3'(s(fun(X272,fun(X273,bool)),X274),s(fun('type.pair.prod'(X270,X271),X272),X275),s(X273,X276))),s('type.pair.prod'(X270,X271),'const.pair.,_2'(s(X270,X277),s(X271,X278))))))) )),
  inference(superposition,[],[f168,f195])).
fof(f5235,plain,(
  ( ! [X94,X92,X90,X88,X87,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X87,bool),happ(s(fun(X87,fun(X87,bool)),'=_0'),s(X87,X95))),s(X87,happ(s(fun(X90,X87),happ(s(fun(X89,fun(X90,X87)),X91),s(X89,happ(s(fun(X88,X89),X92),s(X88,X94))))),s(X90,X93)))))) | s(X87,X95) != s(X87,happ(s(fun(X88,X87),'_dst_x0x1x0_3'(s(fun(X89,fun(X90,X87)),X91),s(fun(X88,X89),X92),s(X90,X93))),s(X88,X94)))) )),
  inference(superposition,[],[f173,f195])).
fof(f5234,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),'=_0'),s(X78,happ(s(fun(X81,X78),happ(s(fun(X80,fun(X81,X78)),X82),s(X80,happ(s(fun(X79,X80),X83),s(X79,X85))))),s(X81,X84))))),s(X78,X86)))) | s(X78,X86) != s(X78,happ(s(fun(X79,X78),'_dst_x0x1x0_3'(s(fun(X80,fun(X81,X78)),X82),s(fun(X79,X80),X83),s(X81,X84))),s(X79,X85)))) )),
  inference(superposition,[],[f173,f195])).
fof(f5233,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),'=_0'),s(X69,X77))),s(X69,happ(s(fun(X72,X69),happ(s(fun(X71,fun(X72,X69)),X73),s(X71,happ(s(fun(X70,X71),X74),s(X70,X76))))),s(X72,X75)))))) | s(X69,X77) = s(X69,happ(s(fun(X70,X69),'_dst_x0x1x0_3'(s(fun(X71,fun(X72,X69)),X73),s(fun(X70,X71),X74),s(X72,X75))),s(X70,X76)))) )),
  inference(superposition,[],[f172,f195])).
fof(f5232,plain,(
  ( ! [X61,X68,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),happ(s(fun(X60,fun(X60,bool)),'=_0'),s(X60,happ(s(fun(X63,X60),happ(s(fun(X62,fun(X63,X60)),X64),s(X62,happ(s(fun(X61,X62),X65),s(X61,X67))))),s(X63,X66))))),s(X60,X68)))) | s(X60,X68) = s(X60,happ(s(fun(X61,X60),'_dst_x0x1x0_3'(s(fun(X62,fun(X63,X60)),X64),s(fun(X61,X62),X65),s(X63,X66))),s(X61,X67)))) )),
  inference(superposition,[],[f172,f195])).
fof(f5227,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s(X20,happ(s(fun(X13,X20),'_dst_x0x1_2'(s(fun(X12,X20),X21),s(fun(X13,X12),'_dst_x0x1x0_3'(s(fun(X14,fun(X15,X12)),X16),s(fun(X13,X14),X17),s(X15,X18))))),s(X13,X19))) = s(X20,happ(s(fun(X12,X20),X21),s(X12,happ(s(fun(X15,X12),happ(s(fun(X14,fun(X15,X12)),X16),s(X14,happ(s(fun(X13,X14),X17),s(X13,X19))))),s(X15,X18)))))) )),
  inference(superposition,[],[f182,f195])).
fof(f5226,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X8,happ(s(fun(X1,X8),'_dst_x0x0x1_3'(s(fun(X9,fun(X0,X8)),X10),s(X9,X11),s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))))),s(X1,X7))) = s(X8,happ(s(fun(X0,X8),happ(s(fun(X9,fun(X0,X8)),X10),s(X9,X11))),s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),X4),s(X2,happ(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6)))))) )),
  inference(superposition,[],[f194,f195])).
fof(f5225,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'F_0') = s(bool,happ(s(fun(X97,bool),'_dst_x0x1x0_3'(s(fun(X95,fun(X94,bool)),X96),s(fun(X97,X95),X98),s(X94,X100))),s(X97,X99))) | s(bool,'T_0') = s(bool,happ(s(fun(X94,bool),happ(s(fun(X95,fun(X94,bool)),X96),s(X95,happ(s(fun(X97,X95),X98),s(X97,X99))))),s(X94,X100)))) )),
  inference(superposition,[],[f195,f143])).
fof(f5220,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') = s(bool,happ(s(fun(X59,bool),'_dst_x0x1x0_3'(s(fun(X58,fun(X58,bool)),'=_0'),s(fun(X59,X58),X60),s(X58,happ(s(fun(X59,X58),X60),s(X59,X61))))),s(X59,X61)))) )),
  inference(superposition,[],[f195,f744])).
fof(f5219,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(X51,'const.pair.UNCURRY_2'(s(fun(X52,fun(X53,X51)),happ(s(fun(X54,fun(X52,fun(X53,X51))),X55),s(X54,X56))),s('type.pair.prod'(X52,X53),X57))) = s(X51,happ(s(fun(X54,X51),'_dst_x0x1x0_3'(s(fun(fun(X52,fun(X53,X51)),fun('type.pair.prod'(X52,X53),X51)),'const.pair.UNCURRY_0'),s(fun(X54,fun(X52,fun(X53,X51))),X55),s('type.pair.prod'(X52,X53),X57))),s(X54,X56)))) )),
  inference(superposition,[],[f195,f180])).
fof(f5218,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (s(X44,'const.pair.pair_CASE_2'(s('type.pair.prod'(X45,X46),happ(s(fun(X47,'type.pair.prod'(X45,X46)),X48),s(X47,X49))),s(fun(X45,fun(X46,X44)),X50))) = s(X44,happ(s(fun(X47,X44),'_dst_x0x1x0_3'(s(fun('type.pair.prod'(X45,X46),fun(fun(X45,fun(X46,X44)),X44)),'const.pair.pair_CASE_0'),s(fun(X47,'type.pair.prod'(X45,X46)),X48),s(fun(X45,fun(X46,X44)),X50))),s(X47,X49)))) )),
  inference(superposition,[],[f195,f181])).
fof(f5217,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X34,happ(s(fun(X38,X34),'_dst_x0x1x0_3'(s(fun(X36,fun(X35,X34)),X37),s(fun(X38,X36),X39),s(X35,happ(s(fun(X41,X35),X42),s(X41,X43))))),s(X38,X40))) = s(X34,happ(s(fun(X41,X34),'_dst_x0x1_2'(s(fun(X35,X34),happ(s(fun(X36,fun(X35,X34)),X37),s(X36,happ(s(fun(X38,X36),X39),s(X38,X40))))),s(fun(X41,X35),X42))),s(X41,X43)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5216,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,happ(s(fun(X28,X24),'_dst_x0x1x0_3'(s(fun(X26,fun(X25,X24)),X27),s(fun(X28,X26),X29),s(X25,happ(s(fun(X31,X25),X32),s(X31,X33))))),s(X28,X30))) = s(X24,happ(s(fun(X31,X24),'_dst_x0x0x1_3'(s(fun(X26,fun(X25,X24)),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30))),s(fun(X31,X25),X32))),s(X31,X33)))) )),
  inference(superposition,[],[f195,f194])).
fof(f5215,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X20,X16),'_dst_x0x1x0_3'(s(fun(X18,fun(X17,X16)),X19),s(fun(X20,X18),X21),s(X17,X23))),s(X20,X22))) = s(X16,'const.pair.UNCURRY_2'(s(fun(X18,fun(X17,X16)),X19),s('type.pair.prod'(X18,X17),'const.pair.,_2'(s(X18,happ(s(fun(X20,X18),X21),s(X20,X22))),s(X17,X23)))))) )),
  inference(superposition,[],[f195,f183])).
fof(f5214,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X12,X8),'_dst_x0x1x0_3'(s(fun(X10,fun(X9,X8)),X11),s(fun(X12,X10),X13),s(X9,X15))),s(X12,X14))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X10,X9),'const.pair.,_2'(s(X10,happ(s(fun(X12,X10),X13),s(X12,X14))),s(X9,X15))),s(fun(X10,fun(X9,X8)),X11)))) )),
  inference(superposition,[],[f195,f184])).
fof(f5180,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129,X137] : (s(X130,happ(s(fun(X129,X130),'const.pair.UNCURRY_2'(s(fun(X131,fun(X132,fun(X129,X130))),X133),s('type.pair.prod'(X131,X132),happ(s(fun(X134,'type.pair.prod'(X131,X132)),X135),s(X134,X136))))),s(X129,X137))) = s(X130,happ(s(fun(X134,X130),'_dst_x0x1x0_3'(s(fun('type.pair.prod'(X131,X132),fun(X129,X130)),happ(s(fun(fun(X131,fun(X132,fun(X129,X130))),fun('type.pair.prod'(X131,X132),fun(X129,X130))),'const.pair.UNCURRY_0'),s(fun(X131,fun(X132,fun(X129,X130))),X133))),s(fun(X134,'type.pair.prod'(X131,X132)),X135),s(X129,X137))),s(X134,X136)))) )),
  inference(superposition,[],[f195,f180])).
fof(f5179,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120] : (s(X121,happ(s(fun(X120,X121),'const.pair.pair_CASE_2'(s('type.pair.prod'(X122,X123),X124),s(fun(X122,fun(X123,fun(X120,X121))),happ(s(fun(X125,fun(X122,fun(X123,fun(X120,X121)))),X126),s(X125,X127))))),s(X120,X128))) = s(X121,happ(s(fun(X125,X121),'_dst_x0x1x0_3'(s(fun(fun(X122,fun(X123,fun(X120,X121))),fun(X120,X121)),happ(s(fun('type.pair.prod'(X122,X123),fun(fun(X122,fun(X123,fun(X120,X121))),fun(X120,X121))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X122,X123),X124))),s(fun(X125,fun(X122,fun(X123,fun(X120,X121)))),X126),s(X120,X128))),s(X125,X127)))) )),
  inference(superposition,[],[f195,f181])).
fof(f5178,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X109,happ(s(fun(X116,X109),'_dst_x0x1x0_3'(s(fun(X110,fun(X108,X109)),'_dst_x0x0x1_3'(s(fun(X111,fun(X112,fun(X108,X109))),X113),s(X111,X114),s(fun(X110,X112),X115))),s(fun(X116,X110),X117),s(X108,X119))),s(X116,X118))) = s(X109,happ(s(fun(X108,X109),happ(s(fun(X112,fun(X108,X109)),happ(s(fun(X111,fun(X112,fun(X108,X109))),X113),s(X111,X114))),s(X112,happ(s(fun(X110,X112),X115),s(X110,happ(s(fun(X116,X110),X117),s(X116,X118))))))),s(X108,X119)))) )),
  inference(superposition,[],[f195,f194])).
fof(f5177,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (s(X101,happ(s(fun(X100,X101),X103),s(X100,X107))) = s(X101,happ(s(fun(X104,X101),'_dst_x0x1x0_3'(s(fun(X102,fun(X100,X101)),'_dst_x0_1'(s(fun(X100,X101),X103))),s(fun(X104,X102),X105),s(X100,X107))),s(X104,X106)))) )),
  inference(superposition,[],[f195,f170])).
fof(f5176,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X95,X93,X91,X89,X98,X96] : (s(X89,happ(s(fun(X96,X89),'_dst_x0x1x0_3'(s(fun(X90,fun(X88,X89)),'_dst_x0x1x0_3'(s(fun(X91,fun(X92,fun(X88,X89))),X93),s(fun(X90,X91),X94),s(X92,X95))),s(fun(X96,X90),X97),s(X88,X99))),s(X96,X98))) = s(X89,happ(s(fun(X88,X89),happ(s(fun(X92,fun(X88,X89)),happ(s(fun(X91,fun(X92,fun(X88,X89))),X93),s(X91,happ(s(fun(X90,X91),X94),s(X90,happ(s(fun(X96,X90),X97),s(X96,X98))))))),s(X92,X95))),s(X88,X99)))) )),
  inference(superposition,[],[f195,f195])).
fof(f5175,plain,(
  ( ! [X80,X78,X76,X87,X85,X83,X81,X79,X77,X86,X84,X82] : (s(X77,happ(s(fun(X76,X77),happ(s(fun(X81,fun(X76,X77)),'_dst_x0x1x0_3'(s(fun(X79,fun(X78,fun(X76,X77))),X80),s(fun(X81,X79),X82),s(X78,happ(s(fun(X84,X78),X85),s(X84,X86))))),s(X81,X83))),s(X76,X87))) = s(X77,happ(s(fun(X84,X77),'_dst_x0x1x0_3'(s(fun(X78,fun(X76,X77)),happ(s(fun(X79,fun(X78,fun(X76,X77))),X80),s(X79,happ(s(fun(X81,X79),X82),s(X81,X83))))),s(fun(X84,X78),X85),s(X76,X87))),s(X84,X86)))) )),
  inference(superposition,[],[f195,f195])).
fof(f5173,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X65,X63] : (s(X57,happ(s(fun(X62,X57),'_dst_x0x1x0_3'(s(fun(X58,fun(X56,X57)),happ(s(fun(X59,fun(X58,fun(X56,X57))),X60),s(X59,X61))),s(fun(X62,X58),X63),s(X56,X65))),s(X62,X64))) = s(X57,happ(s(fun(X56,X57),'const.pair.pair_CASE_2'(s('type.pair.prod'(X59,X58),'const.pair.,_2'(s(X59,X61),s(X58,happ(s(fun(X62,X58),X63),s(X62,X64))))),s(fun(X59,fun(X58,fun(X56,X57))),X60))),s(X56,X65)))) )),
  inference(superposition,[],[f195,f184])).
fof(f5172,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),happ(s(fun(X52,fun(X46,X47)),'_dst_x0x0x1_3'(s(fun(X49,fun(X48,fun(X46,X47))),X50),s(X49,X51),s(fun(X52,X48),X53))),s(X52,X54))),s(X46,X55))) = s(X47,happ(s(fun(X52,X47),'_dst_x0x1x0_3'(s(fun(X48,fun(X46,X47)),happ(s(fun(X49,fun(X48,fun(X46,X47))),X50),s(X49,X51))),s(fun(X52,X48),X53),s(X46,X55))),s(X52,X54)))) )),
  inference(superposition,[],[f195,f194])).
fof(f5171,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (s(X37,happ(s(fun(X36,X37),happ(s(fun(X39,fun(X36,X37)),X40),s(X39,happ(s(fun(X38,X39),X41),s(X38,happ(s(fun(X42,X38),X43),s(X42,X44))))))),s(X36,X45))) = s(X37,happ(s(fun(X42,X37),'_dst_x0x1x0_3'(s(fun(X38,fun(X36,X37)),'_dst_x0x1_2'(s(fun(X39,fun(X36,X37)),X40),s(fun(X38,X39),X41))),s(fun(X42,X38),X43),s(X36,X45))),s(X42,X44)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5390,plain,(
  ( ! [X187,X185,X182,X180,X186,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun(X183,bool),'_dst_x0x1x0_3'(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(fun(X183,X181),X184),s('type.pair.prod'(X179,X180),sK7(happ(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185)))),X180,X179)))),s(X183,X185)))) | p(s(bool,happ(s(fun('type.pair.prod'(X179,X180),bool),happ(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))),s('type.pair.prod'(X179,X180),'const.pair.,_2'(s(X179,X186),s(X180,X187))))))) )),
  inference(backward_demodulation,[],[f5170,f1115])).
fof(f5389,plain,(
  ( ! [X158,X156,X154,X152,X151,X157,X155,X153,X150] : (p(s(bool,happ(s(fun(X152,bool),'_dst_x0x1x0_3'(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(fun(X152,X153),X155),s('type.pair.prod'(X150,X151),'const.pair.,_2'(s(X150,X157),s(X151,X158))))),s(X152,X156)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X150,X151),bool),happ(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(X153,happ(s(fun(X152,X153),X155),s(X152,X156))))),s('type.pair.prod'(X150,X151),sK7(happ(s(fun(X152,fun('type.pair.prod'(X150,X151),bool)),'_dst_x0x1_2'(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(fun(X152,X153),X155))),s(X152,X156)),X151,X150)))))) )),
  inference(backward_demodulation,[],[f5170,f1142])).
fof(f5388,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),'_dst_x0x1x0_3'(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(fun(X8,X6),X9),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK8(happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))),X5,X4)),s(X5,sK9(happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))),X5,X4)))))),s(X8,X10)))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10))))),s('type.pair.prod'(X4,X5),X11))))) )),
  inference(backward_demodulation,[],[f5170,f1827])).
fof(f5387,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (p(s(bool,happ(s(fun(X14,bool),'_dst_x0x1x0_3'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17),s('type.pair.prod'(X12,X13),X19))),s(X14,X18)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X12,X13),bool),happ(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))),s('type.pair.prod'(X12,X13),'const.pair.,_2'(s(X12,sK8(happ(s(fun(X14,fun('type.pair.prod'(X12,X13),bool)),'_dst_x0x1_2'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17))),s(X14,X18)),X13,X12)),s(X13,sK9(happ(s(fun(X14,fun('type.pair.prod'(X12,X13),bool)),'_dst_x0x1_2'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17))),s(X14,X18)),X13,X12)))))))) )),
  inference(backward_demodulation,[],[f5170,f1828])).
fof(f5170,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X29,happ(s(fun(X32,X29),'_dst_x0x1x0_3'(s(fun(X30,fun(X28,X29)),X31),s(fun(X32,X30),X33),s(X28,X35))),s(X32,X34))) = s(X29,happ(s(fun(X28,X29),happ(s(fun(X32,fun(X28,X29)),'_dst_x0x1_2'(s(fun(X30,fun(X28,X29)),X31),s(fun(X32,X30),X33))),s(X32,X34))),s(X28,X35)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5169,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(X23,happ(s(fun(X22,X23),happ(s(fun(X24,fun(X22,X23)),X25),s(X24,X26))),s(X22,X27))) = s(X23,happ(s(fun(X24,X23),'_dst_x0x1x0_3'(s(fun(fun(X22,X23),fun(X22,X23)),'_i_0_0'),s(fun(X24,fun(X22,X23)),X25),s(X22,X27))),s(X24,X26)))) )),
  inference(superposition,[],[f195,f159])).
fof(f5168,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s(X13,happ(s(fun(X18,X13),'_dst_x0x1x0_3'(s(fun(X14,fun(X12,X13)),'_dst_x00x11_2'(s(fun(X15,X13),X16),s(fun(X14,fun(X12,X15)),X17))),s(fun(X18,X14),X19),s(X12,X21))),s(X18,X20))) = s(X13,happ(s(fun(X12,X13),'_dst_x0x1_2'(s(fun(X15,X13),X16),s(fun(X12,X15),happ(s(fun(X14,fun(X12,X15)),X17),s(X14,happ(s(fun(X18,X14),X19),s(X18,X20))))))),s(X12,X21)))) )),
  inference(superposition,[],[f195,f193])).
fof(f5167,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X1,happ(s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(X0,X11))) = s(X1,happ(s(fun(X8,X1),'_dst_x0x1x0_3'(s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(fun(X8,X2),X9),s(X0,X11))),s(X8,X10)))) )),
  inference(superposition,[],[f195,f198])).
fof(f5166,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (s(X218,happ(s(fun(X219,X218),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun(X219,X218)),X220),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X216),X217))))),s(X219,X221))) = s(X218,happ(s(fun('type.closSem.state'(X216),X218),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun(X219,X218)),X220),s(fun('type.closSem.state'(X216),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(X219,X221))),s('type.closSem.state'(X216),X217)))) )),
  inference(superposition,[],[f195,f160])).
fof(f5165,plain,(
  ( ! [X212,X210,X208,X206,X215,X213,X211,X209,X207,X205,X214] : (s(X212,happ(s(fun(X207,X212),'_dst_x0x1x0_3'(s(fun(fun(X205,X206),fun(X213,X212)),X214),s(fun(X207,fun(X205,X206)),'_dst_x00x11_2'(s(fun(X208,X206),X209),s(fun(X207,fun(X205,X208)),X210))),s(X213,X215))),s(X207,X211))) = s(X212,happ(s(fun(X213,X212),happ(s(fun(fun(X205,X206),fun(X213,X212)),X214),s(fun(X205,X206),'_dst_x0x1_2'(s(fun(X208,X206),X209),s(fun(X205,X208),happ(s(fun(X207,fun(X205,X208)),X210),s(X207,X211))))))),s(X213,X215)))) )),
  inference(superposition,[],[f195,f193])).
fof(f5164,plain,(
  ( ! [X198,X196,X194,X192,X204,X202,X200,X199,X197,X195,X193,X203,X201] : (s(X201,happ(s(fun(X202,X201),happ(s(fun(fun(X192,X193),fun(X202,X201)),X203),s(fun(X192,X193),'_dst_x0x1x0_3'(s(fun(X195,fun(X196,X193)),X197),s(fun(X192,X195),X198),s(X196,X199))))),s(X202,X204))) = s(X201,happ(s(fun(X194,X201),'_dst_x0x1x0_3'(s(fun(fun(X192,X193),fun(X202,X201)),X203),s(fun(X194,fun(X192,X193)),'_dst_x00x01x00_3'(s(fun(X195,fun(X196,X193)),X197),s(fun(X192,X195),X198),s(X196,X199))),s(X202,X204))),s(X194,X200)))) )),
  inference(superposition,[],[f195,f198])).
fof(f5163,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186] : (s(X188,happ(s(fun(X189,X188),happ(s(fun(bool,fun(X189,X188)),X190),s(bool,'F_0'))),s(X189,X191))) = s(X188,happ(s(fun(X185,X188),'_dst_x0x1x0_3'(s(fun(bool,fun(X189,X188)),X190),s(fun(X185,bool),X186),s(X189,X191))),s(X185,X187))) | s(bool,'T_0') = s(bool,happ(s(fun(X185,bool),X186),s(X185,X187)))) )),
  inference(superposition,[],[f195,f143])).
fof(f5161,plain,(
  ( ! [X175,X173,X171,X169,X176,X174,X172,X170] : (s(X173,happ(s(fun(X174,X173),happ(s(fun(bool,fun(X174,X173)),X175),s(bool,X172))),s(X174,X176))) = s(X173,happ(s(fun(X169,X173),'_dst_x0x1x0_3'(s(fun(bool,fun(X174,X173)),X175),s(fun(X169,bool),X170),s(X174,X176))),s(X169,X171))) | s(bool,'T_0') = s(bool,happ(s(fun(X169,bool),X170),s(X169,X171))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f195,f223])).
fof(f5159,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X153,X160] : (s(X156,happ(s(fun(X157,X156),happ(s(fun(bool,fun(X157,X156)),X158),s(bool,'T_0'))),s(X157,X159))) = s(X156,happ(s(fun(X153,X156),'_dst_x0x1x0_3'(s(fun(bool,fun(X157,X156)),X158),s(fun(X153,bool),X154),s(X157,X159))),s(X153,X155))) | s(bool,X160) = s(bool,happ(s(fun(X153,bool),X154),s(X153,X155))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f195,f223])).
fof(f5156,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(X141,happ(s(fun(X142,X141),happ(s(fun(bool,fun(X142,X141)),X143),s(bool,'T_0'))),s(X142,X144))) = s(X141,happ(s(fun(bool,X141),'_dst_x0x1x0_3'(s(fun(bool,fun(X142,X141)),X143),s(fun(bool,bool),'_i_0_0'),s(X142,X144))),s(bool,X140))) | s(bool,'F_0') = s(bool,X140)) )),
  inference(superposition,[],[f195,f235])).
fof(f5155,plain,(
  ( ! [X134,X138,X136,X135,X133,X139,X137] : (s(X136,happ(s(fun(X137,X136),happ(s(fun(bool,fun(X137,X136)),X138),s(bool,'T_0'))),s(X137,X139))) = s(X136,happ(s(fun(X133,X136),'_dst_x0x1x0_3'(s(fun(bool,fun(X137,X136)),X138),s(fun(X133,bool),'_dst_x0_1'(s(bool,X134))),s(X137,X139))),s(X133,X135))) | s(bool,'F_0') = s(bool,X134)) )),
  inference(superposition,[],[f195,f267])).
fof(f5154,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (s(X129,happ(s(fun(X130,X129),happ(s(fun(bool,fun(X130,X129)),X131),s(bool,'F_0'))),s(X130,X132))) = s(X129,happ(s(fun(X126,X129),'_dst_x0x1x0_3'(s(fun(bool,fun(X130,X129)),X131),s(fun(X126,bool),'_dst_x0_1'(s(bool,X127))),s(X130,X132))),s(X126,X128))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f195,f614])).
fof(f5153,plain,(
  ( ! [X125,X123,X121,X124,X122,X120] : (s(X122,happ(s(fun(X123,X122),happ(s(fun(bool,fun(X123,X122)),X124),s(bool,'T_0'))),s(X123,X125))) = s(X122,happ(s(fun(X120,X122),'_dst_x0x1x0_3'(s(fun(bool,fun(X123,X122)),X124),s(fun(X120,bool),happ(s(fun(X120,fun(X120,bool)),'=_0'),s(X120,X121))),s(X123,X125))),s(X120,X121)))) )),
  inference(superposition,[],[f195,f744])).
fof(f5152,plain,(
  ( ! [X111,X118,X116,X114,X112,X119,X117,X115,X113] : (s(X116,happ(s(fun(X117,X116),happ(s(fun(X111,fun(X117,X116)),X118),s(X111,'const.pair.UNCURRY_2'(s(fun(X112,fun(X113,X111)),X114),s('type.pair.prod'(X112,X113),X115))))),s(X117,X119))) = s(X116,happ(s(fun('type.pair.prod'(X112,X113),X116),'_dst_x0x1x0_3'(s(fun(X111,fun(X117,X116)),X118),s(fun('type.pair.prod'(X112,X113),X111),happ(s(fun(fun(X112,fun(X113,X111)),fun('type.pair.prod'(X112,X113),X111)),'const.pair.UNCURRY_0'),s(fun(X112,fun(X113,X111)),X114))),s(X117,X119))),s('type.pair.prod'(X112,X113),X115)))) )),
  inference(superposition,[],[f195,f180])).
fof(f5151,plain,(
  ( ! [X103,X109,X107,X105,X102,X110,X108,X106,X104] : (s(X107,happ(s(fun(X108,X107),happ(s(fun(X102,fun(X108,X107)),X109),s(X102,'const.pair.pair_CASE_2'(s('type.pair.prod'(X103,X104),X105),s(fun(X103,fun(X104,X102)),X106))))),s(X108,X110))) = s(X107,happ(s(fun(fun(X103,fun(X104,X102)),X107),'_dst_x0x1x0_3'(s(fun(X102,fun(X108,X107)),X109),s(fun(fun(X103,fun(X104,X102)),X102),happ(s(fun('type.pair.prod'(X103,X104),fun(fun(X103,fun(X104,X102)),X102)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X103,X104),X105))),s(X108,X110))),s(fun(X103,fun(X104,X102)),X106)))) )),
  inference(superposition,[],[f195,f181])).
fof(f5150,plain,(
  ( ! [X94,X92,X90,X101,X99,X97,X95,X93,X91,X100,X98,X96] : (s(X98,happ(s(fun(X91,X98),'_dst_x0x1x0_3'(s(fun(X90,fun(X99,X98)),X100),s(fun(X91,X90),'_dst_x0x0x1_3'(s(fun(X92,fun(X93,X90)),X94),s(X92,X95),s(fun(X91,X93),X96))),s(X99,X101))),s(X91,X97))) = s(X98,happ(s(fun(X99,X98),happ(s(fun(X90,fun(X99,X98)),X100),s(X90,happ(s(fun(X93,X90),happ(s(fun(X92,fun(X93,X90)),X94),s(X92,X95))),s(X93,happ(s(fun(X91,X93),X96),s(X91,X97))))))),s(X99,X101)))) )),
  inference(superposition,[],[f195,f194])).
fof(f5149,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (s(X86,happ(s(fun(X87,X86),happ(s(fun(X82,fun(X87,X86)),X88),s(X82,X84))),s(X87,X89))) = s(X86,happ(s(fun(X83,X86),'_dst_x0x1x0_3'(s(fun(X82,fun(X87,X86)),X88),s(fun(X83,X82),'_dst_x0_1'(s(X82,X84))),s(X87,X89))),s(X83,X85)))) )),
  inference(superposition,[],[f195,f170])).
fof(f5148,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73] : (s(X78,happ(s(fun(X71,X78),'_dst_x0x1x0_3'(s(fun(X70,fun(X79,X78)),X80),s(fun(X71,X70),'_dst_x0x1x0_3'(s(fun(X72,fun(X73,X70)),X74),s(fun(X71,X72),X75),s(X73,X76))),s(X79,X81))),s(X71,X77))) = s(X78,happ(s(fun(X79,X78),happ(s(fun(X70,fun(X79,X78)),X80),s(X70,happ(s(fun(X73,X70),happ(s(fun(X72,fun(X73,X70)),X74),s(X72,happ(s(fun(X71,X72),X75),s(X71,X77))))),s(X73,X76))))),s(X79,X81)))) )),
  inference(superposition,[],[f195,f195])).
fof(f5147,plain,(
  ( ! [X61,X59,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s(X66,happ(s(fun(X59,X66),'_dst_x0x1x0_3'(s(fun(X58,fun(X67,X66)),X68),s(fun(X59,X58),happ(s(fun(X60,fun(X59,X58)),X61),s(X60,happ(s(fun(X62,X60),X63),s(X62,X64))))),s(X67,X69))),s(X59,X65))) = s(X66,happ(s(fun(X67,X66),happ(s(fun(X58,fun(X67,X66)),X68),s(X58,happ(s(fun(X62,X58),'_dst_x0x1x0_3'(s(fun(X60,fun(X59,X58)),X61),s(fun(X62,X60),X63),s(X59,X65))),s(X62,X64))))),s(X67,X69)))) )),
  inference(superposition,[],[f195,f195])).
fof(f5146,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (s(X54,happ(s(fun(X47,X54),'_dst_x0x1x0_3'(s(fun(X46,fun(X55,X54)),X56),s(fun(X47,X46),happ(s(fun(X48,fun(X47,X46)),X49),s(X48,X50))),s(X55,X57))),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53))))) = s(X54,happ(s(fun(X55,X54),happ(s(fun(X46,fun(X55,X54)),X56),s(X46,happ(s(fun(X51,X46),'_dst_x0x0x1_3'(s(fun(X48,fun(X47,X46)),X49),s(X48,X50),s(fun(X51,X47),X52))),s(X51,X53))))),s(X55,X57)))) )),
  inference(superposition,[],[f195,f194])).
fof(f5144,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (s(X32,happ(s(fun(X27,X32),'_dst_x0x1x0_3'(s(fun(X26,fun(X33,X32)),X34),s(fun(X27,X26),happ(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))),s(X33,X35))),s(X27,X31))) = s(X32,happ(s(fun(X33,X32),happ(s(fun(X26,fun(X33,X32)),X34),s(X26,'const.pair.pair_CASE_2'(s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))),s(fun(X28,fun(X27,X26)),X29))))),s(X33,X35)))) )),
  inference(superposition,[],[f195,f184])).
fof(f5143,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X22,happ(s(fun(X17,X22),'_dst_x0x1x0_3'(s(fun(X16,fun(X23,X22)),X24),s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X19),s(fun(X17,X18),X20))),s(X23,X25))),s(X17,X21))) = s(X22,happ(s(fun(X23,X22),happ(s(fun(X16,fun(X23,X22)),X24),s(X16,happ(s(fun(X18,X16),X19),s(X18,happ(s(fun(X17,X18),X20),s(X17,X21))))))),s(X23,X25)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5142,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(X12,happ(s(fun(X7,X12),'_dst_x0x1x0_3'(s(fun(X6,fun(X13,X12)),X14),s(fun(X7,X6),X8),s(X13,X15))),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))) = s(X12,happ(s(fun(X13,X12),happ(s(fun(X6,fun(X13,X12)),X14),s(X6,happ(s(fun(X9,X6),'_dst_x0x1_2'(s(fun(X7,X6),X8),s(fun(X9,X7),X10))),s(X9,X11))))),s(X13,X15)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5141,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,happ(s(fun(X3,X2),happ(s(fun(X0,fun(X3,X2)),X4),s(X0,X1))),s(X3,X5))) = s(X2,happ(s(fun(X0,X2),'_dst_x0x1x0_3'(s(fun(X0,fun(X3,X2)),X4),s(fun(X0,X0),'_i_0_0'),s(X3,X5))),s(X0,X1)))) )),
  inference(superposition,[],[f195,f159])).
fof(f5073,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,happ(s(fun(X50,X51),X54),s(X50,X56))))),s(X52,X55))) | s(bool,'T_0') = s(bool,happ(s(fun(X50,bool),'_dst_x0x1x0_3'(s(fun(X51,fun(X52,bool)),X53),s(fun(X50,X51),X54),s(X52,X55))),s(X50,X56)))) )),
  inference(superposition,[],[f195,f143])).
fof(f5071,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,X41) = s(bool,happ(s(fun(X36,bool),happ(s(fun(X35,fun(X36,bool)),X37),s(X35,happ(s(fun(X34,X35),X38),s(X34,X40))))),s(X36,X39))) | s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),'_dst_x0x1x0_3'(s(fun(X35,fun(X36,bool)),X37),s(fun(X34,X35),X38),s(X36,X39))),s(X34,X40))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f195,f223])).
fof(f5069,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),happ(s(fun(X19,fun(X20,bool)),X21),s(X19,happ(s(fun(X18,X19),X22),s(X18,X24))))),s(X20,X23))) | s(bool,X25) = s(bool,happ(s(fun(X18,bool),'_dst_x0x1x0_3'(s(fun(X19,fun(X20,bool)),X21),s(fun(X18,X19),X22),s(X20,X23))),s(X18,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f195,f223])).
fof(f5068,plain,(
  ( ! [X14,X12,X10,X8,X17,X15,X13,X11,X9,X16] : (s(X8,happ(s(fun(X15,X8),'_dst_x0x1_2'(s(fun(X9,X8),'_dst_x0x1x0_3'(s(fun(X10,fun(X11,X8)),X12),s(fun(X9,X10),X13),s(X11,X14))),s(fun(X15,X9),X16))),s(X15,X17))) = s(X8,happ(s(fun(X11,X8),happ(s(fun(X10,fun(X11,X8)),X12),s(X10,happ(s(fun(X9,X10),X13),s(X9,happ(s(fun(X15,X9),X16),s(X15,X17))))))),s(X11,X14)))) )),
  inference(superposition,[],[f195,f182])).
fof(f5066,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X3,X0)),X4),s(fun(X1,X2),X5),s(X3,X6))),s(X1,X7))))),s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),X4),s(X2,happ(s(fun(X1,X2),X5),s(X1,X7))))),s(X3,X6))))))) )),
  inference(unit_resulting_resolution,[],[f195,f173])).
fof(f5067,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))))),s(X0,happ(s(fun(X4,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5),s(X1,X7))),s(X4,X6))))))) )),
  inference(unit_resulting_resolution,[],[f195,f173])).
fof(f4862,plain,(
  ( ! [X600,X596,X598,X592,X594,X590,X597,X599,X593,X595,X591] : (~p(s(bool,happ(s(fun('type.pair.prod'(X590,X591),bool),happ(s(fun(X596,fun('type.pair.prod'(X590,X591),bool)),'_dst_x0x0x1_3'(s(fun(X593,fun(X592,fun('type.pair.prod'(X590,X591),bool))),X594),s(X593,X595),s(fun(X596,X592),X597))),s(X596,X598))),s('type.pair.prod'(X590,X591),sK7(happ(s(fun(X592,fun('type.pair.prod'(X590,X591),bool)),happ(s(fun(X593,fun(X592,fun('type.pair.prod'(X590,X591),bool))),X594),s(X593,X595))),s(X592,happ(s(fun(X596,X592),X597),s(X596,X598)))),X591,X590))))) | p(s(bool,happ(s(fun('type.pair.prod'(X590,X591),bool),happ(s(fun(X592,fun('type.pair.prod'(X590,X591),bool)),happ(s(fun(X593,fun(X592,fun('type.pair.prod'(X590,X591),bool))),X594),s(X593,X595))),s(X592,happ(s(fun(X596,X592),X597),s(X596,X598))))),s('type.pair.prod'(X590,X591),'const.pair.,_2'(s(X590,X599),s(X591,X600))))))) )),
  inference(superposition,[],[f168,f194])).
fof(f4861,plain,(
  ( ! [X588,X584,X586,X580,X582,X589,X585,X587,X581,X583] : (~p(s(bool,happ(s(fun('type.pair.prod'(X580,X581),bool),happ(s(fun(X586,fun('type.pair.prod'(X580,X581),bool)),'_dst_x0x0x1_3'(s(fun(X583,fun(X582,fun('type.pair.prod'(X580,X581),bool))),X584),s(X583,X585),s(fun(X586,X582),X587))),s(X586,X588))),s('type.pair.prod'(X580,X581),'const.pair.,_2'(s(X580,sK8(happ(s(fun(X582,fun('type.pair.prod'(X580,X581),bool)),happ(s(fun(X583,fun(X582,fun('type.pair.prod'(X580,X581),bool))),X584),s(X583,X585))),s(X582,happ(s(fun(X586,X582),X587),s(X586,X588)))),X581,X580)),s(X581,sK9(happ(s(fun(X582,fun('type.pair.prod'(X580,X581),bool)),happ(s(fun(X583,fun(X582,fun('type.pair.prod'(X580,X581),bool))),X584),s(X583,X585))),s(X582,happ(s(fun(X586,X582),X587),s(X586,X588)))),X581,X580))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X580,X581),bool),happ(s(fun(X582,fun('type.pair.prod'(X580,X581),bool)),happ(s(fun(X583,fun(X582,fun('type.pair.prod'(X580,X581),bool))),X584),s(X583,X585))),s(X582,happ(s(fun(X586,X582),X587),s(X586,X588))))),s('type.pair.prod'(X580,X581),X589))))) )),
  inference(superposition,[],[f169,f194])).
fof(f4850,plain,(
  ( ! [X445,X447,X453,X449,X451,X444,X446,X452,X448,X450] : (s(X445,happ(s(fun(X444,X445),happ(s(fun(X450,fun(X444,X445)),'_dst_x0x0x1_3'(s(fun(X447,fun(X446,fun(X444,X445))),X448),s(X447,X449),s(fun(X450,X446),X451))),s(X450,X452))),s(X444,X453))) = s(X445,'const.pair.UNCURRY_2'(s(fun(X446,fun(X444,X445)),happ(s(fun(X447,fun(X446,fun(X444,X445))),X448),s(X447,X449))),s('type.pair.prod'(X446,X444),'const.pair.,_2'(s(X446,happ(s(fun(X450,X446),X451),s(X450,X452))),s(X444,X453)))))) )),
  inference(superposition,[],[f183,f194])).
fof(f4849,plain,(
  ( ! [X441,X443,X436,X438,X434,X440,X442,X437,X439,X435] : (s(X435,happ(s(fun(X434,X435),happ(s(fun(X440,fun(X434,X435)),'_dst_x0x0x1_3'(s(fun(X437,fun(X436,fun(X434,X435))),X438),s(X437,X439),s(fun(X440,X436),X441))),s(X440,X442))),s(X434,X443))) = s(X435,'const.pair.pair_CASE_2'(s('type.pair.prod'(X436,X434),'const.pair.,_2'(s(X436,happ(s(fun(X440,X436),X441),s(X440,X442))),s(X434,X443))),s(fun(X436,fun(X434,X435)),happ(s(fun(X437,fun(X436,fun(X434,X435))),X438),s(X437,X439)))))) )),
  inference(superposition,[],[f184,f194])).
fof(f4848,plain,(
  ( ! [X428,X430,X424,X426,X432,X423,X429,X431,X425,X427,X433] : (s(fun(X423,X432),'_dst_x0x1_2'(s(fun(X424,X432),X433),s(fun(X423,X424),happ(s(fun(X429,fun(X423,X424)),'_dst_x0x0x1_3'(s(fun(X426,fun(X425,fun(X423,X424))),X427),s(X426,X428),s(fun(X429,X425),X430))),s(X429,X431))))) = s(fun(X423,X432),happ(s(fun(X425,fun(X423,X432)),'_dst_x00x11_2'(s(fun(X424,X432),X433),s(fun(X425,fun(X423,X424)),happ(s(fun(X426,fun(X425,fun(X423,X424))),X427),s(X426,X428))))),s(X425,happ(s(fun(X429,X425),X430),s(X429,X431)))))) )),
  inference(superposition,[],[f193,f194])).
fof(f4846,plain,(
  ( ! [X412,X408,X410,X407,X413,X409,X411,X406] : (s(bool,'F_0') != s(bool,happ(s(fun(X410,bool),'_dst_x0x0x1_3'(s(fun(X407,fun(X406,bool)),X408),s(X407,X409),s(fun(X410,X406),X411))),s(X410,X412))) | s(bool,'T_0') = s(bool,X413) | s(bool,X413) = s(bool,happ(s(fun(X406,bool),happ(s(fun(X407,fun(X406,bool)),X408),s(X407,X409))),s(X406,happ(s(fun(X410,X406),X411),s(X410,X412)))))) )),
  inference(superposition,[],[f1176,f194])).
fof(f4845,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (p(s(bool,happ(s(fun(X403,bool),'_dst_x0x0x1_3'(s(fun(X400,fun(X399,bool)),X401),s(X400,X402),s(fun(X403,X399),X404))),s(X403,X405)))) | s(bool,'F_0') = s(bool,happ(s(fun(X399,bool),happ(s(fun(X400,fun(X399,bool)),X401),s(X400,X402))),s(X399,happ(s(fun(X403,X399),X404),s(X403,X405)))))) )),
  inference(superposition,[],[f1018,f194])).
fof(f4844,plain,(
  ( ! [X391,X397,X393,X395,X396,X398,X392,X394] : (s(bool,X398) != s(bool,happ(s(fun(X395,bool),'_dst_x0x0x1_3'(s(fun(X392,fun(X391,bool)),X393),s(X392,X394),s(fun(X395,X391),X396))),s(X395,X397))) | s(bool,'T_0') = s(bool,X398) | s(bool,'F_0') = s(bool,happ(s(fun(X391,bool),happ(s(fun(X392,fun(X391,bool)),X393),s(X392,X394))),s(X391,happ(s(fun(X395,X391),X396),s(X395,X397)))))) )),
  inference(superposition,[],[f638,f194])).
fof(f4843,plain,(
  ( ! [X383,X389,X385,X387,X388,X390,X384,X386] : (s(bool,X390) != s(bool,happ(s(fun(X387,bool),'_dst_x0x0x1_3'(s(fun(X384,fun(X383,bool)),X385),s(X384,X386),s(fun(X387,X383),X388))),s(X387,X389))) | s(bool,'T_0') = s(bool,happ(s(fun(X383,bool),happ(s(fun(X384,fun(X383,bool)),X385),s(X384,X386))),s(X383,happ(s(fun(X387,X383),X388),s(X387,X389))))) | s(bool,'F_0') = s(bool,X390)) )),
  inference(superposition,[],[f638,f194])).
fof(f4842,plain,(
  ( ! [X377,X379,X380,X382,X376,X378,X375,X381] : (p(s(bool,happ(s(fun(X379,bool),'_dst_x0x0x1_3'(s(fun(X376,fun(X375,bool)),X377),s(X376,X378),s(fun(X379,X375),X380))),s(X379,X381)))) | s(bool,X382) != s(bool,happ(s(fun(X375,bool),happ(s(fun(X376,fun(X375,bool)),X377),s(X376,X378))),s(X375,happ(s(fun(X379,X375),X380),s(X379,X381))))) | s(bool,'F_0') = s(bool,X382)) )),
  inference(superposition,[],[f616,f194])).
fof(f4841,plain,(
  ( ! [X366,X372,X374,X368,X370,X367,X373,X369,X371] : (s(bool,'F_0') = s(bool,happ(s(fun(X373,bool),'_dst_x0_1'(s(bool,happ(s(fun(X370,bool),'_dst_x0x0x1_3'(s(fun(X367,fun(X366,bool)),X368),s(X367,X369),s(fun(X370,X366),X371))),s(X370,X372))))),s(X373,X374))) | s(bool,'T_0') = s(bool,happ(s(fun(X366,bool),happ(s(fun(X367,fun(X366,bool)),X368),s(X367,X369))),s(X366,happ(s(fun(X370,X366),X371),s(X370,X372)))))) )),
  inference(superposition,[],[f614,f194])).
fof(f4840,plain,(
  ( ! [X358,X364,X360,X362,X357,X359,X365,X361,X363] : (p(s(bool,happ(s(fun(X361,bool),'_dst_x0x0x1_3'(s(fun(X358,fun(X357,bool)),X359),s(X358,X360),s(fun(X361,X357),X362))),s(X361,X363)))) | s(bool,X364) != s(bool,happ(s(fun(X357,bool),happ(s(fun(X358,fun(X357,bool)),X359),s(X358,X360))),s(X357,happ(s(fun(X361,X357),X362),s(X361,X363))))) | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f489,f194])).
fof(f4838,plain,(
  ( ! [X344,X346,X343,X349,X345,X347,X342,X348] : (p(s(bool,happ(s(fun(X346,bool),'_dst_x0x0x1_3'(s(fun(X343,fun(X342,bool)),X344),s(X343,X345),s(fun(X346,X342),X347))),s(X346,X348)))) | s(bool,X349) = s(bool,happ(s(fun(X342,bool),happ(s(fun(X343,fun(X342,bool)),X344),s(X343,X345))),s(X342,happ(s(fun(X346,X342),X347),s(X346,X348))))) | s(bool,'T_0') = s(bool,X349)) )),
  inference(superposition,[],[f347,f194])).
fof(f4837,plain,(
  ( ! [X333,X335,X341,X337,X339,X334,X340,X336,X338] : (s(bool,'T_0') = s(bool,happ(s(fun(X340,bool),'_dst_x0_1'(s(bool,happ(s(fun(X337,bool),'_dst_x0x0x1_3'(s(fun(X334,fun(X333,bool)),X335),s(X334,X336),s(fun(X337,X333),X338))),s(X337,X339))))),s(X340,X341))) | s(bool,'F_0') = s(bool,happ(s(fun(X333,bool),happ(s(fun(X334,fun(X333,bool)),X335),s(X334,X336))),s(X333,happ(s(fun(X337,X333),X338),s(X337,X339)))))) )),
  inference(superposition,[],[f267,f194])).
fof(f4836,plain,(
  ( ! [X327,X329,X331,X326,X332,X328,X330] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X330,bool),'_dst_x0x0x1_3'(s(fun(X327,fun(X326,bool)),X328),s(X327,X329),s(fun(X330,X326),X331))),s(X330,X332))))) | s(bool,'F_0') = s(bool,happ(s(fun(X326,bool),happ(s(fun(X327,fun(X326,bool)),X328),s(X327,X329))),s(X326,happ(s(fun(X330,X326),X331),s(X330,X332)))))) )),
  inference(superposition,[],[f235,f194])).
fof(f4835,plain,(
  ( ! [X325,X321,X323,X324,X320,X322,X319] : (~p(s(bool,happ(s(fun(X323,bool),'_dst_x0x0x1_3'(s(fun(X320,fun(X319,bool)),X321),s(X320,X322),s(fun(X323,X319),X324))),s(X323,X325)))) | s(bool,'T_0') = s(bool,happ(s(fun(X319,bool),happ(s(fun(X320,fun(X319,bool)),X321),s(X320,X322))),s(X319,happ(s(fun(X323,X319),X324),s(X323,X325)))))) )),
  inference(superposition,[],[f225,f194])).
fof(f4830,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X284,X286] : (~p(s(bool,happ(s(fun(X283,bool),'_dst_x0x0x1_3'(s(fun(X280,fun(X279,bool)),X281),s(X280,X282),s(fun(X283,X279),X284))),s(X283,X285)))) | s(bool,X286) != s(bool,happ(s(fun(X279,bool),happ(s(fun(X280,fun(X279,bool)),X281),s(X280,X282))),s(X279,happ(s(fun(X283,X279),X284),s(X283,X285))))) | p(s(bool,X286))) )),
  inference(superposition,[],[f215,f194])).
fof(f4812,plain,(
  ( ! [X103,X111,X109,X107,X105,X110,X108,X106,X104] : (p(s(bool,happ(s(fun(X103,bool),happ(s(fun(X103,fun(X103,bool)),'=_0'),s(X103,X111))),s(X103,happ(s(fun(X108,X103),'_dst_x0x0x1_3'(s(fun(X105,fun(X104,X103)),X106),s(X105,X107),s(fun(X108,X104),X109))),s(X108,X110)))))) | s(X103,X111) != s(X103,happ(s(fun(X104,X103),happ(s(fun(X105,fun(X104,X103)),X106),s(X105,X107))),s(X104,happ(s(fun(X108,X104),X109),s(X108,X110)))))) )),
  inference(superposition,[],[f173,f194])).
fof(f4811,plain,(
  ( ! [X94,X101,X99,X97,X95,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),'=_0'),s(X94,happ(s(fun(X99,X94),'_dst_x0x0x1_3'(s(fun(X96,fun(X95,X94)),X97),s(X96,X98),s(fun(X99,X95),X100))),s(X99,X101))))),s(X94,X102)))) | s(X94,X102) != s(X94,happ(s(fun(X95,X94),happ(s(fun(X96,fun(X95,X94)),X97),s(X96,X98))),s(X95,happ(s(fun(X99,X95),X100),s(X99,X101)))))) )),
  inference(superposition,[],[f173,f194])).
fof(f4810,plain,(
  ( ! [X92,X90,X88,X87,X85,X93,X91,X89,X86] : (~p(s(bool,happ(s(fun(X85,bool),happ(s(fun(X85,fun(X85,bool)),'=_0'),s(X85,X93))),s(X85,happ(s(fun(X90,X85),'_dst_x0x0x1_3'(s(fun(X87,fun(X86,X85)),X88),s(X87,X89),s(fun(X90,X86),X91))),s(X90,X92)))))) | s(X85,X93) = s(X85,happ(s(fun(X86,X85),happ(s(fun(X87,fun(X86,X85)),X88),s(X87,X89))),s(X86,happ(s(fun(X90,X86),X91),s(X90,X92)))))) )),
  inference(superposition,[],[f172,f194])).
fof(f4809,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X84,X82] : (~p(s(bool,happ(s(fun(X76,bool),happ(s(fun(X76,fun(X76,bool)),'=_0'),s(X76,happ(s(fun(X81,X76),'_dst_x0x0x1_3'(s(fun(X78,fun(X77,X76)),X79),s(X78,X80),s(fun(X81,X77),X82))),s(X81,X83))))),s(X76,X84)))) | s(X76,X84) = s(X76,happ(s(fun(X77,X76),happ(s(fun(X78,fun(X77,X76)),X79),s(X78,X80))),s(X77,happ(s(fun(X81,X77),X82),s(X81,X83)))))) )),
  inference(superposition,[],[f172,f194])).
fof(f4801,plain,(
  ( ! [X14,X12,X10,X8,X17,X15,X13,X11,X9,X16] : (s(X16,happ(s(fun(X8,X16),X17),s(X8,happ(s(fun(X13,X8),'_dst_x0x0x1_3'(s(fun(X10,fun(X9,X8)),X11),s(X10,X12),s(fun(X13,X9),X14))),s(X13,X15))))) = s(X16,happ(s(fun(X9,X16),'_dst_x0x1_2'(s(fun(X8,X16),X17),s(fun(X9,X8),happ(s(fun(X10,fun(X9,X8)),X11),s(X10,X12))))),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f182,f194])).
fof(f4796,plain,(
  ( ! [X573,X575,X569,X571,X565,X567,X572,X574,X568,X570,X566] : (~p(s(bool,happ(s(fun('type.pair.prod'(X565,X566),bool),happ(s(fun(X569,fun('type.pair.prod'(X565,X566),bool)),happ(s(fun(X568,fun(X569,fun('type.pair.prod'(X565,X566),bool))),X570),s(X568,X571))),s(X569,happ(s(fun(X567,X569),X572),s(X567,X573))))),s('type.pair.prod'(X565,X566),sK7(happ(s(fun(X567,fun('type.pair.prod'(X565,X566),bool)),'_dst_x0x0x1_3'(s(fun(X568,fun(X569,fun('type.pair.prod'(X565,X566),bool))),X570),s(X568,X571),s(fun(X567,X569),X572))),s(X567,X573)),X566,X565))))) | p(s(bool,happ(s(fun('type.pair.prod'(X565,X566),bool),happ(s(fun(X567,fun('type.pair.prod'(X565,X566),bool)),'_dst_x0x0x1_3'(s(fun(X568,fun(X569,fun('type.pair.prod'(X565,X566),bool))),X570),s(X568,X571),s(fun(X567,X569),X572))),s(X567,X573))),s('type.pair.prod'(X565,X566),'const.pair.,_2'(s(X565,X574),s(X566,X575))))))) )),
  inference(superposition,[],[f168,f194])).
fof(f4795,plain,(
  ( ! [X561,X563,X557,X559,X555,X564,X560,X562,X556,X558] : (~p(s(bool,happ(s(fun('type.pair.prod'(X555,X556),bool),happ(s(fun(X559,fun('type.pair.prod'(X555,X556),bool)),happ(s(fun(X558,fun(X559,fun('type.pair.prod'(X555,X556),bool))),X560),s(X558,X561))),s(X559,happ(s(fun(X557,X559),X562),s(X557,X563))))),s('type.pair.prod'(X555,X556),'const.pair.,_2'(s(X555,sK8(happ(s(fun(X557,fun('type.pair.prod'(X555,X556),bool)),'_dst_x0x0x1_3'(s(fun(X558,fun(X559,fun('type.pair.prod'(X555,X556),bool))),X560),s(X558,X561),s(fun(X557,X559),X562))),s(X557,X563)),X556,X555)),s(X556,sK9(happ(s(fun(X557,fun('type.pair.prod'(X555,X556),bool)),'_dst_x0x0x1_3'(s(fun(X558,fun(X559,fun('type.pair.prod'(X555,X556),bool))),X560),s(X558,X561),s(fun(X557,X559),X562))),s(X557,X563)),X556,X555))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X555,X556),bool),happ(s(fun(X557,fun('type.pair.prod'(X555,X556),bool)),'_dst_x0x0x1_3'(s(fun(X558,fun(X559,fun('type.pair.prod'(X555,X556),bool))),X560),s(X558,X561),s(fun(X557,X559),X562))),s(X557,X563))),s('type.pair.prod'(X555,X556),X564))))) )),
  inference(superposition,[],[f169,f194])).
fof(f4784,plain,(
  ( ! [X420,X422,X428,X424,X426,X421,X423,X419,X425,X427] : (s(X420,'const.pair.UNCURRY_2'(s(fun(X421,fun(X419,X420)),'_dst_x0x0x1_3'(s(fun(X422,fun(X423,fun(X419,X420))),X424),s(X422,X425),s(fun(X421,X423),X426))),s('type.pair.prod'(X421,X419),'const.pair.,_2'(s(X421,X427),s(X419,X428))))) = s(X420,happ(s(fun(X419,X420),happ(s(fun(X423,fun(X419,X420)),happ(s(fun(X422,fun(X423,fun(X419,X420))),X424),s(X422,X425))),s(X423,happ(s(fun(X421,X423),X426),s(X421,X427))))),s(X419,X428)))) )),
  inference(superposition,[],[f183,f194])).
fof(f4783,plain,(
  ( ! [X412,X414,X410,X416,X418,X413,X415,X409,X411,X417] : (s(X410,'const.pair.pair_CASE_2'(s('type.pair.prod'(X411,X409),'const.pair.,_2'(s(X411,X417),s(X409,X418))),s(fun(X411,fun(X409,X410)),'_dst_x0x0x1_3'(s(fun(X412,fun(X413,fun(X409,X410))),X414),s(X412,X415),s(fun(X411,X413),X416))))) = s(X410,happ(s(fun(X409,X410),happ(s(fun(X413,fun(X409,X410)),happ(s(fun(X412,fun(X413,fun(X409,X410))),X414),s(X412,X415))),s(X413,happ(s(fun(X411,X413),X416),s(X411,X417))))),s(X409,X418)))) )),
  inference(superposition,[],[f184,f194])).
fof(f4782,plain,(
  ( ! [X408,X399,X405,X407,X401,X403,X398,X404,X406,X400,X402] : (s(fun(X398,X407),happ(s(fun(X400,fun(X398,X407)),'_dst_x00x11_2'(s(fun(X399,X407),X408),s(fun(X400,fun(X398,X399)),'_dst_x0x0x1_3'(s(fun(X401,fun(X402,fun(X398,X399))),X403),s(X401,X404),s(fun(X400,X402),X405))))),s(X400,X406))) = s(fun(X398,X407),'_dst_x0x1_2'(s(fun(X399,X407),X408),s(fun(X398,X399),happ(s(fun(X402,fun(X398,X399)),happ(s(fun(X401,fun(X402,fun(X398,X399))),X403),s(X401,X404))),s(X402,happ(s(fun(X400,X402),X405),s(X400,X406)))))))) )),
  inference(superposition,[],[f193,f194])).
fof(f4780,plain,(
  ( ! [X383,X385,X387,X382,X388,X384,X386,X381] : (s(bool,'F_0') != s(bool,happ(s(fun(X383,bool),happ(s(fun(X382,fun(X383,bool)),X384),s(X382,X385))),s(X383,happ(s(fun(X381,X383),X386),s(X381,X387))))) | s(bool,'T_0') = s(bool,X388) | s(bool,X388) = s(bool,happ(s(fun(X381,bool),'_dst_x0x0x1_3'(s(fun(X382,fun(X383,bool)),X384),s(X382,X385),s(fun(X381,X383),X386))),s(X381,X387)))) )),
  inference(superposition,[],[f1176,f194])).
fof(f4779,plain,(
  ( ! [X377,X379,X374,X380,X376,X378,X375] : (p(s(bool,happ(s(fun(X376,bool),happ(s(fun(X375,fun(X376,bool)),X377),s(X375,X378))),s(X376,happ(s(fun(X374,X376),X379),s(X374,X380)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X374,bool),'_dst_x0x0x1_3'(s(fun(X375,fun(X376,bool)),X377),s(X375,X378),s(fun(X374,X376),X379))),s(X374,X380)))) )),
  inference(superposition,[],[f1018,f194])).
fof(f4778,plain,(
  ( ! [X366,X372,X368,X370,X367,X373,X369,X371] : (s(bool,X373) != s(bool,happ(s(fun(X368,bool),happ(s(fun(X367,fun(X368,bool)),X369),s(X367,X370))),s(X368,happ(s(fun(X366,X368),X371),s(X366,X372))))) | s(bool,'T_0') = s(bool,X373) | s(bool,'F_0') = s(bool,happ(s(fun(X366,bool),'_dst_x0x0x1_3'(s(fun(X367,fun(X368,bool)),X369),s(X367,X370),s(fun(X366,X368),X371))),s(X366,X372)))) )),
  inference(superposition,[],[f638,f194])).
fof(f4777,plain,(
  ( ! [X358,X364,X360,X362,X359,X365,X361,X363] : (s(bool,X365) != s(bool,happ(s(fun(X360,bool),happ(s(fun(X359,fun(X360,bool)),X361),s(X359,X362))),s(X360,happ(s(fun(X358,X360),X363),s(X358,X364))))) | s(bool,'T_0') = s(bool,happ(s(fun(X358,bool),'_dst_x0x0x1_3'(s(fun(X359,fun(X360,bool)),X361),s(X359,X362),s(fun(X358,X360),X363))),s(X358,X364))) | s(bool,'F_0') = s(bool,X365)) )),
  inference(superposition,[],[f638,f194])).
fof(f4776,plain,(
  ( ! [X350,X356,X352,X354,X351,X357,X353,X355] : (p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X351,fun(X352,bool)),X353),s(X351,X354))),s(X352,happ(s(fun(X350,X352),X355),s(X350,X356)))))) | s(bool,X357) != s(bool,happ(s(fun(X350,bool),'_dst_x0x0x1_3'(s(fun(X351,fun(X352,bool)),X353),s(X351,X354),s(fun(X350,X352),X355))),s(X350,X356))) | s(bool,'F_0') = s(bool,X357)) )),
  inference(superposition,[],[f616,f194])).
fof(f4775,plain,(
  ( ! [X344,X346,X341,X343,X349,X345,X347,X342,X348] : (s(bool,'F_0') = s(bool,happ(s(fun(X348,bool),'_dst_x0_1'(s(bool,happ(s(fun(X343,bool),happ(s(fun(X342,fun(X343,bool)),X344),s(X342,X345))),s(X343,happ(s(fun(X341,X343),X346),s(X341,X347))))))),s(X348,X349))) | s(bool,'T_0') = s(bool,happ(s(fun(X341,bool),'_dst_x0x0x1_3'(s(fun(X342,fun(X343,bool)),X344),s(X342,X345),s(fun(X341,X343),X346))),s(X341,X347)))) )),
  inference(superposition,[],[f614,f194])).
fof(f4774,plain,(
  ( ! [X333,X335,X337,X339,X332,X334,X340,X336,X338] : (p(s(bool,happ(s(fun(X334,bool),happ(s(fun(X333,fun(X334,bool)),X335),s(X333,X336))),s(X334,happ(s(fun(X332,X334),X337),s(X332,X338)))))) | s(bool,X339) != s(bool,happ(s(fun(X332,bool),'_dst_x0x0x1_3'(s(fun(X333,fun(X334,bool)),X335),s(X333,X336),s(fun(X332,X334),X337))),s(X332,X338))) | s(bool,X339) = s(bool,X340) | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f489,f194])).
fof(f4772,plain,(
  ( ! [X321,X323,X318,X324,X320,X322,X317,X319] : (p(s(bool,happ(s(fun(X319,bool),happ(s(fun(X318,fun(X319,bool)),X320),s(X318,X321))),s(X319,happ(s(fun(X317,X319),X322),s(X317,X323)))))) | s(bool,X324) = s(bool,happ(s(fun(X317,bool),'_dst_x0x0x1_3'(s(fun(X318,fun(X319,bool)),X320),s(X318,X321),s(fun(X317,X319),X322))),s(X317,X323))) | s(bool,'T_0') = s(bool,X324)) )),
  inference(superposition,[],[f347,f194])).
fof(f4771,plain,(
  ( ! [X313,X315,X308,X310,X316,X312,X314,X309,X311] : (s(bool,'T_0') = s(bool,happ(s(fun(X315,bool),'_dst_x0_1'(s(bool,happ(s(fun(X310,bool),happ(s(fun(X309,fun(X310,bool)),X311),s(X309,X312))),s(X310,happ(s(fun(X308,X310),X313),s(X308,X314))))))),s(X315,X316))) | s(bool,'F_0') = s(bool,happ(s(fun(X308,bool),'_dst_x0x0x1_3'(s(fun(X309,fun(X310,bool)),X311),s(X309,X312),s(fun(X308,X310),X313))),s(X308,X314)))) )),
  inference(superposition,[],[f267,f194])).
fof(f4770,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X303,bool),happ(s(fun(X302,fun(X303,bool)),X304),s(X302,X305))),s(X303,happ(s(fun(X301,X303),X306),s(X301,X307))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X301,bool),'_dst_x0x0x1_3'(s(fun(X302,fun(X303,bool)),X304),s(X302,X305),s(fun(X301,X303),X306))),s(X301,X307)))) )),
  inference(superposition,[],[f235,f194])).
fof(f4769,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (~p(s(bool,happ(s(fun(X296,bool),happ(s(fun(X295,fun(X296,bool)),X297),s(X295,X298))),s(X296,happ(s(fun(X294,X296),X299),s(X294,X300)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X294,bool),'_dst_x0x0x1_3'(s(fun(X295,fun(X296,bool)),X297),s(X295,X298),s(fun(X294,X296),X299))),s(X294,X300)))) )),
  inference(superposition,[],[f225,f194])).
fof(f4764,plain,(
  ( ! [X255,X261,X257,X259,X260,X254,X256,X258] : (~p(s(bool,happ(s(fun(X256,bool),happ(s(fun(X255,fun(X256,bool)),X257),s(X255,X258))),s(X256,happ(s(fun(X254,X256),X259),s(X254,X260)))))) | s(bool,X261) != s(bool,happ(s(fun(X254,bool),'_dst_x0x0x1_3'(s(fun(X255,fun(X256,bool)),X257),s(X255,X258),s(fun(X254,X256),X259))),s(X254,X260))) | p(s(bool,X261))) )),
  inference(superposition,[],[f215,f194])).
fof(f4762,plain,(
  ( ! [X243,X241,X239,X246,X244,X242,X240,X245] : (~p(s(bool,happ(s(fun(X242,bool),happ(s(fun(X241,fun(X242,bool)),X243),s(X241,X244))),s(X242,happ(s(fun('type.pair.prod'(X239,X240),X242),X245),s('type.pair.prod'(X239,X240),'const.pair.,_2'(s(X239,sK8('_dst_x0x0x1_3'(s(fun(X241,fun(X242,bool)),X243),s(X241,X244),s(fun('type.pair.prod'(X239,X240),X242),X245)),X240,X239)),s(X240,sK9('_dst_x0x0x1_3'(s(fun(X241,fun(X242,bool)),X243),s(X241,X244),s(fun('type.pair.prod'(X239,X240),X242),X245)),X240,X239))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X239,X240),bool),'_dst_x0x0x1_3'(s(fun(X241,fun(X242,bool)),X243),s(X241,X244),s(fun('type.pair.prod'(X239,X240),X242),X245))),s('type.pair.prod'(X239,X240),X246))))) )),
  inference(superposition,[],[f169,f194])).
fof(f4761,plain,(
  ( ! [X231,X237,X235,X233,X230,X238,X236,X234,X232] : (~p(s(bool,happ(s(fun(X233,bool),happ(s(fun(X232,fun(X233,bool)),X234),s(X232,X235))),s(X233,happ(s(fun('type.pair.prod'(X230,X231),X233),X236),s('type.pair.prod'(X230,X231),sK7('_dst_x0x0x1_3'(s(fun(X232,fun(X233,bool)),X234),s(X232,X235),s(fun('type.pair.prod'(X230,X231),X233),X236)),X231,X230))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X230,X231),bool),'_dst_x0x0x1_3'(s(fun(X232,fun(X233,bool)),X234),s(X232,X235),s(fun('type.pair.prod'(X230,X231),X233),X236))),s('type.pair.prod'(X230,X231),'const.pair.,_2'(s(X230,X237),s(X231,X238))))))) )),
  inference(superposition,[],[f168,f194])).
fof(f4747,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X75,X82] : (p(s(bool,happ(s(fun(X75,bool),happ(s(fun(X75,fun(X75,bool)),'=_0'),s(X75,X83))),s(X75,happ(s(fun(X78,X75),happ(s(fun(X77,fun(X78,X75)),X79),s(X77,X80))),s(X78,happ(s(fun(X76,X78),X81),s(X76,X82)))))))) | s(X75,X83) != s(X75,happ(s(fun(X76,X75),'_dst_x0x0x1_3'(s(fun(X77,fun(X78,X75)),X79),s(X77,X80),s(fun(X76,X78),X81))),s(X76,X82)))) )),
  inference(superposition,[],[f173,f194])).
fof(f4746,plain,(
  ( ! [X70,X68,X66,X74,X72,X71,X69,X67,X73] : (p(s(bool,happ(s(fun(X66,bool),happ(s(fun(X66,fun(X66,bool)),'=_0'),s(X66,happ(s(fun(X69,X66),happ(s(fun(X68,fun(X69,X66)),X70),s(X68,X71))),s(X69,happ(s(fun(X67,X69),X72),s(X67,X73))))))),s(X66,X74)))) | s(X66,X74) != s(X66,happ(s(fun(X67,X66),'_dst_x0x0x1_3'(s(fun(X68,fun(X69,X66)),X70),s(X68,X71),s(fun(X67,X69),X72))),s(X67,X73)))) )),
  inference(superposition,[],[f173,f194])).
fof(f4745,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X65,X63] : (~p(s(bool,happ(s(fun(X57,bool),happ(s(fun(X57,fun(X57,bool)),'=_0'),s(X57,X65))),s(X57,happ(s(fun(X60,X57),happ(s(fun(X59,fun(X60,X57)),X61),s(X59,X62))),s(X60,happ(s(fun(X58,X60),X63),s(X58,X64)))))))) | s(X57,X65) = s(X57,happ(s(fun(X58,X57),'_dst_x0x0x1_3'(s(fun(X59,fun(X60,X57)),X61),s(X59,X62),s(fun(X58,X60),X63))),s(X58,X64)))) )),
  inference(superposition,[],[f172,f194])).
fof(f4744,plain,(
  ( ! [X54,X52,X50,X48,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),happ(s(fun(X48,fun(X48,bool)),'=_0'),s(X48,happ(s(fun(X51,X48),happ(s(fun(X50,fun(X51,X48)),X52),s(X50,X53))),s(X51,happ(s(fun(X49,X51),X54),s(X49,X55))))))),s(X48,X56)))) | s(X48,X56) = s(X48,happ(s(fun(X49,X48),'_dst_x0x0x1_3'(s(fun(X50,fun(X51,X48)),X52),s(X50,X53),s(fun(X49,X51),X54))),s(X49,X55)))) )),
  inference(superposition,[],[f172,f194])).
fof(f4739,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X8,happ(s(fun(X1,X8),'_dst_x0x1_2'(s(fun(X0,X8),X9),s(fun(X1,X0),'_dst_x0x0x1_3'(s(fun(X2,fun(X3,X0)),X4),s(X2,X5),s(fun(X1,X3),X6))))),s(X1,X7))) = s(X8,happ(s(fun(X0,X8),X9),s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),X4),s(X2,X5))),s(X3,happ(s(fun(X1,X3),X6),s(X1,X7)))))))) )),
  inference(superposition,[],[f182,f194])).
fof(f4738,plain,(
  ( ! [X88,X87,X85,X83,X86,X84,X82] : (s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),'_dst_x0x0x1_3'(s(fun(X83,fun(X82,bool)),X84),s(X83,X85),s(fun(X86,X82),X87))),s(X86,X88))) | s(bool,'T_0') = s(bool,happ(s(fun(X82,bool),happ(s(fun(X83,fun(X82,bool)),X84),s(X83,X85))),s(X82,happ(s(fun(X86,X82),X87),s(X86,X88)))))) )),
  inference(superposition,[],[f194,f143])).
fof(f4733,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),'_dst_x0x0x1_3'(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,happ(s(fun(X47,X46),X48),s(X47,X49))),s(fun(X47,X46),X48))),s(X47,X49)))) )),
  inference(superposition,[],[f194,f744])).
fof(f4732,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (s(X39,'const.pair.UNCURRY_2'(s(fun(X40,fun(X41,X39)),X42),s('type.pair.prod'(X40,X41),happ(s(fun(X43,'type.pair.prod'(X40,X41)),X44),s(X43,X45))))) = s(X39,happ(s(fun(X43,X39),'_dst_x0x0x1_3'(s(fun(fun(X40,fun(X41,X39)),fun('type.pair.prod'(X40,X41),X39)),'const.pair.UNCURRY_0'),s(fun(X40,fun(X41,X39)),X42),s(fun(X43,'type.pair.prod'(X40,X41)),X44))),s(X43,X45)))) )),
  inference(superposition,[],[f194,f180])).
fof(f4731,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (s(X32,'const.pair.pair_CASE_2'(s('type.pair.prod'(X33,X34),X35),s(fun(X33,fun(X34,X32)),happ(s(fun(X36,fun(X33,fun(X34,X32))),X37),s(X36,X38))))) = s(X32,happ(s(fun(X36,X32),'_dst_x0x0x1_3'(s(fun('type.pair.prod'(X33,X34),fun(fun(X33,fun(X34,X32)),X32)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X33,X34),X35),s(fun(X36,fun(X33,fun(X34,X32))),X37))),s(X36,X38)))) )),
  inference(superposition,[],[f194,f181])).
fof(f4883,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,happ(s(fun(X69,bool),'_dst_x0x0x1_3'(s(fun(X67,fun(X67,bool)),'=_0'),s(X67,X68),s(fun(X69,X67),X70))),s(X69,X71)))) | s(X67,X68) != s(X67,happ(s(fun(X69,X67),X70),s(X69,X71)))) )),
  inference(backward_demodulation,[],[f4730,f1097])).
fof(f4882,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,happ(s(fun(X74,bool),'_dst_x0x0x1_3'(s(fun(X72,fun(X72,bool)),'=_0'),s(X72,X73),s(fun(X74,X72),X75))),s(X74,X76)))) | s(X72,X73) = s(X72,happ(s(fun(X74,X72),X75),s(X74,X76)))) )),
  inference(backward_demodulation,[],[f4730,f1098])).
fof(f4730,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(X24,happ(s(fun(X29,X24),'_dst_x0x0x1_3'(s(fun(X26,fun(X25,X24)),X27),s(X26,X28),s(fun(X29,X25),X30))),s(X29,X31))) = s(X24,happ(s(fun(X29,X24),'_dst_x0x1_2'(s(fun(X25,X24),happ(s(fun(X26,fun(X25,X24)),X27),s(X26,X28))),s(fun(X29,X25),X30))),s(X29,X31)))) )),
  inference(superposition,[],[f194,f182])).
fof(f4729,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X21,X16),'_dst_x0x0x1_3'(s(fun(X18,fun(X17,X16)),X19),s(X18,X20),s(fun(X21,X17),X22))),s(X21,X23))) = s(X16,'const.pair.UNCURRY_2'(s(fun(X18,fun(X17,X16)),X19),s('type.pair.prod'(X18,X17),'const.pair.,_2'(s(X18,X20),s(X17,happ(s(fun(X21,X17),X22),s(X21,X23)))))))) )),
  inference(superposition,[],[f194,f183])).
fof(f4728,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X13,X8),'_dst_x0x0x1_3'(s(fun(X10,fun(X9,X8)),X11),s(X10,X12),s(fun(X13,X9),X14))),s(X13,X15))) = s(X8,'const.pair.pair_CASE_2'(s('type.pair.prod'(X10,X9),'const.pair.,_2'(s(X10,X12),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15))))),s(fun(X10,fun(X9,X8)),X11)))) )),
  inference(superposition,[],[f194,f184])).
fof(f4727,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (s(X190,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X190),happ(s(fun(X191,fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X190)),X192),s(X191,X193))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X188),X189))))) = s(X190,happ(s(fun('type.closSem.state'(X188),X190),'_dst_x0x0x1_3'(s(fun(X191,fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X190)),X192),s(X191,X193),s(fun('type.closSem.state'(X188),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))),s('type.closSem.state'(X188),X189)))) )),
  inference(superposition,[],[f194,f160])).
fof(f4726,plain,(
  ( ! [X177,X187,X185,X182,X180,X178,X186,X184,X183,X181,X179] : (s(X184,happ(s(fun(X179,X184),'_dst_x0x0x1_3'(s(fun(X185,fun(fun(X177,X178),X184)),X186),s(X185,X187),s(fun(X179,fun(X177,X178)),'_dst_x00x11_2'(s(fun(X180,X178),X181),s(fun(X179,fun(X177,X180)),X182))))),s(X179,X183))) = s(X184,happ(s(fun(fun(X177,X178),X184),happ(s(fun(X185,fun(fun(X177,X178),X184)),X186),s(X185,X187))),s(fun(X177,X178),'_dst_x0x1_2'(s(fun(X180,X178),X181),s(fun(X177,X180),happ(s(fun(X179,fun(X177,X180)),X182),s(X179,X183)))))))) )),
  inference(superposition,[],[f194,f193])).
fof(f4725,plain,(
  ( ! [X167,X165,X175,X173,X171,X169,X176,X166,X164,X174,X172,X170,X168] : (s(X173,happ(s(fun(fun(X164,X165),X173),happ(s(fun(X174,fun(fun(X164,X165),X173)),X175),s(X174,X176))),s(fun(X164,X165),'_dst_x0x1x0_3'(s(fun(X167,fun(X168,X165)),X169),s(fun(X164,X167),X170),s(X168,X171))))) = s(X173,happ(s(fun(X166,X173),'_dst_x0x0x1_3'(s(fun(X174,fun(fun(X164,X165),X173)),X175),s(X174,X176),s(fun(X166,fun(X164,X165)),'_dst_x00x01x00_3'(s(fun(X167,fun(X168,X165)),X169),s(fun(X164,X167),X170),s(X168,X171))))),s(X166,X172)))) )),
  inference(superposition,[],[f194,f198])).
fof(f4724,plain,(
  ( ! [X158,X163,X161,X159,X157,X162,X160] : (s(X160,happ(s(fun(bool,X160),happ(s(fun(X161,fun(bool,X160)),X162),s(X161,X163))),s(bool,'F_0'))) = s(X160,happ(s(fun(X157,X160),'_dst_x0x0x1_3'(s(fun(X161,fun(bool,X160)),X162),s(X161,X163),s(fun(X157,bool),X158))),s(X157,X159))) | s(bool,'T_0') = s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) )),
  inference(superposition,[],[f194,f143])).
fof(f4722,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (s(X145,happ(s(fun(bool,X145),happ(s(fun(X146,fun(bool,X145)),X147),s(X146,X148))),s(bool,X144))) = s(X145,happ(s(fun(X141,X145),'_dst_x0x0x1_3'(s(fun(X146,fun(bool,X145)),X147),s(X146,X148),s(fun(X141,bool),X142))),s(X141,X143))) | s(bool,'T_0') = s(bool,happ(s(fun(X141,bool),X142),s(X141,X143))) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f194,f223])).
fof(f4720,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (s(X128,happ(s(fun(bool,X128),happ(s(fun(X129,fun(bool,X128)),X130),s(X129,X131))),s(bool,'T_0'))) = s(X128,happ(s(fun(X125,X128),'_dst_x0x0x1_3'(s(fun(X129,fun(bool,X128)),X130),s(X129,X131),s(fun(X125,bool),X126))),s(X125,X127))) | s(bool,X132) = s(bool,happ(s(fun(X125,bool),X126),s(X125,X127))) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f194,f223])).
fof(f4718,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(X117,happ(s(fun(bool,X117),happ(s(fun(X118,fun(bool,X117)),X119),s(X118,X120))),s(bool,'T_0'))) = s(X117,happ(s(fun(bool,X117),'_dst_x0x0x1_3'(s(fun(X118,fun(bool,X117)),X119),s(X118,X120),s(fun(bool,bool),'_i_0_0'))),s(bool,X116))) | s(bool,'F_0') = s(bool,X116)) )),
  inference(superposition,[],[f194,f235])).
fof(f4717,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (s(X112,happ(s(fun(bool,X112),happ(s(fun(X113,fun(bool,X112)),X114),s(X113,X115))),s(bool,'T_0'))) = s(X112,happ(s(fun(X109,X112),'_dst_x0x0x1_3'(s(fun(X113,fun(bool,X112)),X114),s(X113,X115),s(fun(X109,bool),'_dst_x0_1'(s(bool,X110))))),s(X109,X111))) | s(bool,'F_0') = s(bool,X110)) )),
  inference(superposition,[],[f194,f267])).
fof(f4716,plain,(
  ( ! [X103,X107,X105,X102,X108,X106,X104] : (s(X105,happ(s(fun(bool,X105),happ(s(fun(X106,fun(bool,X105)),X107),s(X106,X108))),s(bool,'F_0'))) = s(X105,happ(s(fun(X102,X105),'_dst_x0x0x1_3'(s(fun(X106,fun(bool,X105)),X107),s(X106,X108),s(fun(X102,bool),'_dst_x0_1'(s(bool,X103))))),s(X102,X104))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f194,f614])).
fof(f4715,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s(X98,happ(s(fun(bool,X98),happ(s(fun(X99,fun(bool,X98)),X100),s(X99,X101))),s(bool,'T_0'))) = s(X98,happ(s(fun(X96,X98),'_dst_x0x0x1_3'(s(fun(X99,fun(bool,X98)),X100),s(X99,X101),s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),'=_0'),s(X96,X97))))),s(X96,X97)))) )),
  inference(superposition,[],[f194,f744])).
fof(f4714,plain,(
  ( ! [X94,X92,X90,X88,X87,X95,X93,X91,X89] : (s(X92,happ(s(fun(X87,X92),happ(s(fun(X93,fun(X87,X92)),X94),s(X93,X95))),s(X87,'const.pair.UNCURRY_2'(s(fun(X88,fun(X89,X87)),X90),s('type.pair.prod'(X88,X89),X91))))) = s(X92,happ(s(fun('type.pair.prod'(X88,X89),X92),'_dst_x0x0x1_3'(s(fun(X93,fun(X87,X92)),X94),s(X93,X95),s(fun('type.pair.prod'(X88,X89),X87),happ(s(fun(fun(X88,fun(X89,X87)),fun('type.pair.prod'(X88,X89),X87)),'const.pair.UNCURRY_0'),s(fun(X88,fun(X89,X87)),X90))))),s('type.pair.prod'(X88,X89),X91)))) )),
  inference(superposition,[],[f194,f180])).
fof(f4713,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X86,X84,X82] : (s(X83,happ(s(fun(X78,X83),happ(s(fun(X84,fun(X78,X83)),X85),s(X84,X86))),s(X78,'const.pair.pair_CASE_2'(s('type.pair.prod'(X79,X80),X81),s(fun(X79,fun(X80,X78)),X82))))) = s(X83,happ(s(fun(fun(X79,fun(X80,X78)),X83),'_dst_x0x0x1_3'(s(fun(X84,fun(X78,X83)),X85),s(X84,X86),s(fun(fun(X79,fun(X80,X78)),X78),happ(s(fun('type.pair.prod'(X79,X80),fun(fun(X79,fun(X80,X78)),X78)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X79,X80),X81))))),s(fun(X79,fun(X80,X78)),X82)))) )),
  inference(superposition,[],[f194,f181])).
fof(f4712,plain,(
  ( ! [X70,X68,X66,X76,X74,X72,X71,X69,X67,X77,X75,X73] : (s(X74,happ(s(fun(X67,X74),'_dst_x0x0x1_3'(s(fun(X75,fun(X66,X74)),X76),s(X75,X77),s(fun(X67,X66),'_dst_x0x0x1_3'(s(fun(X68,fun(X69,X66)),X70),s(X68,X71),s(fun(X67,X69),X72))))),s(X67,X73))) = s(X74,happ(s(fun(X66,X74),happ(s(fun(X75,fun(X66,X74)),X76),s(X75,X77))),s(X66,happ(s(fun(X69,X66),happ(s(fun(X68,fun(X69,X66)),X70),s(X68,X71))),s(X69,happ(s(fun(X67,X69),X72),s(X67,X73)))))))) )),
  inference(superposition,[],[f194,f194])).
fof(f4711,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (s(X62,happ(s(fun(X58,X62),happ(s(fun(X63,fun(X58,X62)),X64),s(X63,X65))),s(X58,X60))) = s(X62,happ(s(fun(X59,X62),'_dst_x0x0x1_3'(s(fun(X63,fun(X58,X62)),X64),s(X63,X65),s(fun(X59,X58),'_dst_x0_1'(s(X58,X60))))),s(X59,X61)))) )),
  inference(superposition,[],[f194,f170])).
fof(f4710,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (s(X54,happ(s(fun(X46,X54),happ(s(fun(X55,fun(X46,X54)),X56),s(X55,X57))),s(X46,happ(s(fun(X51,X46),'_dst_x0x0x1_3'(s(fun(X48,fun(X47,X46)),X49),s(X48,X50),s(fun(X51,X47),X52))),s(X51,X53))))) = s(X54,happ(s(fun(X47,X54),'_dst_x0x0x1_3'(s(fun(X55,fun(X46,X54)),X56),s(X55,X57),s(fun(X47,X46),happ(s(fun(X48,fun(X47,X46)),X49),s(X48,X50))))),s(X47,happ(s(fun(X51,X47),X52),s(X51,X53)))))) )),
  inference(superposition,[],[f194,f194])).
fof(f4708,plain,(
  ( ! [X30,X28,X26,X35,X33,X31,X29,X27,X34,X32] : (s(X32,happ(s(fun(X27,X32),'_dst_x0x0x1_3'(s(fun(X33,fun(X26,X32)),X34),s(X33,X35),s(fun(X27,X26),happ(s(fun(X28,fun(X27,X26)),X29),s(X28,X30))))),s(X27,X31))) = s(X32,happ(s(fun(X26,X32),happ(s(fun(X33,fun(X26,X32)),X34),s(X33,X35))),s(X26,'const.pair.pair_CASE_2'(s('type.pair.prod'(X28,X27),'const.pair.,_2'(s(X28,X30),s(X27,X31))),s(fun(X28,fun(X27,X26)),X29)))))) )),
  inference(superposition,[],[f194,f184])).
fof(f4707,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X22,happ(s(fun(X17,X22),'_dst_x0x0x1_3'(s(fun(X23,fun(X16,X22)),X24),s(X23,X25),s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X19),s(fun(X17,X18),X20))))),s(X17,X21))) = s(X22,happ(s(fun(X16,X22),happ(s(fun(X23,fun(X16,X22)),X24),s(X23,X25))),s(X16,happ(s(fun(X18,X16),X19),s(X18,happ(s(fun(X17,X18),X20),s(X17,X21)))))))) )),
  inference(superposition,[],[f194,f182])).
fof(f4706,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(X12,happ(s(fun(X7,X12),'_dst_x0x0x1_3'(s(fun(X13,fun(X6,X12)),X14),s(X13,X15),s(fun(X7,X6),X8))),s(X7,happ(s(fun(X9,X7),X10),s(X9,X11))))) = s(X12,happ(s(fun(X6,X12),happ(s(fun(X13,fun(X6,X12)),X14),s(X13,X15))),s(X6,happ(s(fun(X9,X6),'_dst_x0x1_2'(s(fun(X7,X6),X8),s(fun(X9,X7),X10))),s(X9,X11)))))) )),
  inference(superposition,[],[f194,f182])).
fof(f4705,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,happ(s(fun(X0,X2),happ(s(fun(X3,fun(X0,X2)),X4),s(X3,X5))),s(X0,X1))) = s(X2,happ(s(fun(X0,X2),'_dst_x0x0x1_3'(s(fun(X3,fun(X0,X2)),X4),s(X3,X5),s(fun(X0,X0),'_i_0_0'))),s(X0,X1)))) )),
  inference(superposition,[],[f194,f159])).
fof(f4659,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X117,X115,X113] : (s(X110,happ(s(fun(X109,X110),'const.pair.UNCURRY_2'(s(fun(X111,fun(X112,fun(X109,X110))),X113),s('type.pair.prod'(X111,X112),X114))),s(X109,happ(s(fun(X115,X109),X116),s(X115,X117))))) = s(X110,happ(s(fun(X115,X110),'_dst_x0x0x1_3'(s(fun('type.pair.prod'(X111,X112),fun(X109,X110)),happ(s(fun(fun(X111,fun(X112,fun(X109,X110))),fun('type.pair.prod'(X111,X112),fun(X109,X110))),'const.pair.UNCURRY_0'),s(fun(X111,fun(X112,fun(X109,X110))),X113))),s('type.pair.prod'(X111,X112),X114),s(fun(X115,X109),X116))),s(X115,X117)))) )),
  inference(superposition,[],[f194,f180])).
fof(f4658,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X108,X106,X104] : (s(X101,happ(s(fun(X100,X101),'const.pair.pair_CASE_2'(s('type.pair.prod'(X102,X103),X104),s(fun(X102,fun(X103,fun(X100,X101))),X105))),s(X100,happ(s(fun(X106,X100),X107),s(X106,X108))))) = s(X101,happ(s(fun(X106,X101),'_dst_x0x0x1_3'(s(fun(fun(X102,fun(X103,fun(X100,X101))),fun(X100,X101)),happ(s(fun('type.pair.prod'(X102,X103),fun(fun(X102,fun(X103,fun(X100,X101))),fun(X100,X101))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X102,X103),X104))),s(fun(X102,fun(X103,fun(X100,X101))),X105),s(fun(X106,X100),X107))),s(X106,X108)))) )),
  inference(superposition,[],[f194,f181])).
fof(f4657,plain,(
  ( ! [X94,X92,X90,X88,X99,X97,X95,X93,X91,X89,X98,X96] : (s(X89,happ(s(fun(X97,X89),'_dst_x0x0x1_3'(s(fun(X90,fun(X88,X89)),'_dst_x0x0x1_3'(s(fun(X91,fun(X92,fun(X88,X89))),X93),s(X91,X94),s(fun(X90,X92),X95))),s(X90,X96),s(fun(X97,X88),X98))),s(X97,X99))) = s(X89,happ(s(fun(X88,X89),happ(s(fun(X92,fun(X88,X89)),happ(s(fun(X91,fun(X92,fun(X88,X89))),X93),s(X91,X94))),s(X92,happ(s(fun(X90,X92),X95),s(X90,X96))))),s(X88,happ(s(fun(X97,X88),X98),s(X97,X99)))))) )),
  inference(superposition,[],[f194,f194])).
fof(f4656,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (s(X81,happ(s(fun(X80,X81),X83),s(X80,happ(s(fun(X85,X80),X86),s(X85,X87))))) = s(X81,happ(s(fun(X85,X81),'_dst_x0x0x1_3'(s(fun(X82,fun(X80,X81)),'_dst_x0_1'(s(fun(X80,X81),X83))),s(X82,X84),s(fun(X85,X80),X86))),s(X85,X87)))) )),
  inference(superposition,[],[f194,f170])).
fof(f4655,plain,(
  ( ! [X70,X68,X78,X76,X74,X72,X71,X69,X79,X77,X75,X73] : (s(X69,happ(s(fun(X68,X69),happ(s(fun(X74,fun(X68,X69)),'_dst_x0x0x1_3'(s(fun(X71,fun(X70,fun(X68,X69))),X72),s(X71,X73),s(fun(X74,X70),X75))),s(X74,X76))),s(X68,happ(s(fun(X77,X68),X78),s(X77,X79))))) = s(X69,happ(s(fun(X77,X69),'_dst_x0x0x1_3'(s(fun(X70,fun(X68,X69)),happ(s(fun(X71,fun(X70,fun(X68,X69))),X72),s(X71,X73))),s(X70,happ(s(fun(X74,X70),X75),s(X74,X76))),s(fun(X77,X68),X78))),s(X77,X79)))) )),
  inference(superposition,[],[f194,f194])).
fof(f4653,plain,(
  ( ! [X57,X54,X52,X50,X48,X56,X55,X53,X51,X49] : (s(X49,happ(s(fun(X55,X49),'_dst_x0x0x1_3'(s(fun(X50,fun(X48,X49)),happ(s(fun(X51,fun(X50,fun(X48,X49))),X52),s(X51,X53))),s(X50,X54),s(fun(X55,X48),X56))),s(X55,X57))) = s(X49,happ(s(fun(X48,X49),'const.pair.pair_CASE_2'(s('type.pair.prod'(X51,X50),'const.pair.,_2'(s(X51,X53),s(X50,X54))),s(fun(X51,fun(X50,fun(X48,X49))),X52))),s(X48,happ(s(fun(X55,X48),X56),s(X55,X57)))))) )),
  inference(superposition,[],[f194,f184])).
fof(f4652,plain,(
  ( ! [X39,X47,X45,X43,X41,X38,X46,X44,X42,X40] : (s(X39,happ(s(fun(X38,X39),happ(s(fun(X41,fun(X38,X39)),X42),s(X41,happ(s(fun(X40,X41),X43),s(X40,X44))))),s(X38,happ(s(fun(X45,X38),X46),s(X45,X47))))) = s(X39,happ(s(fun(X45,X39),'_dst_x0x0x1_3'(s(fun(X40,fun(X38,X39)),'_dst_x0x1_2'(s(fun(X41,fun(X38,X39)),X42),s(fun(X40,X41),X43))),s(X40,X44),s(fun(X45,X38),X46))),s(X45,X47)))) )),
  inference(superposition,[],[f194,f182])).
fof(f4651,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X36,X34,X32] : (s(X29,happ(s(fun(X35,X29),'_dst_x0x0x1_3'(s(fun(X30,fun(X28,X29)),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34))),s(fun(X35,X28),X36))),s(X35,X37))) = s(X29,happ(s(fun(X28,X29),happ(s(fun(X32,fun(X28,X29)),'_dst_x0x1_2'(s(fun(X30,fun(X28,X29)),X31),s(fun(X32,X30),X33))),s(X32,X34))),s(X28,happ(s(fun(X35,X28),X36),s(X35,X37)))))) )),
  inference(superposition,[],[f194,f182])).
fof(f4650,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(X23,happ(s(fun(X22,X23),X24),s(X22,happ(s(fun(X25,X22),X26),s(X25,X27))))) = s(X23,happ(s(fun(X25,X23),'_dst_x0x0x1_3'(s(fun(fun(X22,X23),fun(X22,X23)),'_i_0_0'),s(fun(X22,X23),X24),s(fun(X25,X22),X26))),s(X25,X27)))) )),
  inference(superposition,[],[f194,f159])).
fof(f4649,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s(X13,happ(s(fun(X19,X13),'_dst_x0x0x1_3'(s(fun(X14,fun(X12,X13)),'_dst_x00x11_2'(s(fun(X15,X13),X16),s(fun(X14,fun(X12,X15)),X17))),s(X14,X18),s(fun(X19,X12),X20))),s(X19,X21))) = s(X13,happ(s(fun(X12,X13),'_dst_x0x1_2'(s(fun(X15,X13),X16),s(fun(X12,X15),happ(s(fun(X14,fun(X12,X15)),X17),s(X14,X18))))),s(X12,happ(s(fun(X19,X12),X20),s(X19,X21)))))) )),
  inference(superposition,[],[f194,f193])).
fof(f4648,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X1,happ(s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(X0,happ(s(fun(X9,X0),X10),s(X9,X11))))) = s(X1,happ(s(fun(X9,X1),'_dst_x0x0x1_3'(s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(X2,X8),s(fun(X9,X0),X10))),s(X9,X11)))) )),
  inference(superposition,[],[f194,f198])).
fof(f4602,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X52,bool),happ(s(fun(X51,fun(X52,bool)),X53),s(X51,X54))),s(X52,happ(s(fun(X50,X52),X55),s(X50,X56))))) | s(bool,'T_0') = s(bool,happ(s(fun(X50,bool),'_dst_x0x0x1_3'(s(fun(X51,fun(X52,bool)),X53),s(X51,X54),s(fun(X50,X52),X55))),s(X50,X56)))) )),
  inference(superposition,[],[f194,f143])).
fof(f4600,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (s(bool,X41) = s(bool,happ(s(fun(X36,bool),happ(s(fun(X35,fun(X36,bool)),X37),s(X35,X38))),s(X36,happ(s(fun(X34,X36),X39),s(X34,X40))))) | s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),'_dst_x0x0x1_3'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38),s(fun(X34,X36),X39))),s(X34,X40))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f194,f223])).
fof(f4598,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),happ(s(fun(X19,fun(X20,bool)),X21),s(X19,X22))),s(X20,happ(s(fun(X18,X20),X23),s(X18,X24))))) | s(bool,X25) = s(bool,happ(s(fun(X18,bool),'_dst_x0x0x1_3'(s(fun(X19,fun(X20,bool)),X21),s(X19,X22),s(fun(X18,X20),X23))),s(X18,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f194,f223])).
fof(f4597,plain,(
  ( ! [X14,X12,X10,X8,X17,X15,X13,X11,X9,X16] : (s(X8,happ(s(fun(X15,X8),'_dst_x0x1_2'(s(fun(X9,X8),'_dst_x0x0x1_3'(s(fun(X10,fun(X11,X8)),X12),s(X10,X13),s(fun(X9,X11),X14))),s(fun(X15,X9),X16))),s(X15,X17))) = s(X8,happ(s(fun(X11,X8),happ(s(fun(X10,fun(X11,X8)),X12),s(X10,X13))),s(X11,happ(s(fun(X9,X11),X14),s(X9,happ(s(fun(X15,X9),X16),s(X15,X17)))))))) )),
  inference(superposition,[],[f194,f182])).
fof(f4595,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x0x1_3'(s(fun(X2,fun(X3,X0)),X4),s(X2,X5),s(fun(X1,X3),X6))),s(X1,X7))))),s(X0,happ(s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),X4),s(X2,X5))),s(X3,happ(s(fun(X1,X3),X6),s(X1,X7))))))))) )),
  inference(unit_resulting_resolution,[],[f194,f173])).
fof(f4596,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,happ(s(fun(X5,X1),X6),s(X5,X7))))))),s(X0,happ(s(fun(X5,X0),'_dst_x0x0x1_3'(s(fun(X2,fun(X1,X0)),X3),s(X2,X4),s(fun(X5,X1),X6))),s(X5,X7))))))) )),
  inference(unit_resulting_resolution,[],[f194,f173])).
fof(f142,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.closProps.evaluate_app_const',aHL_FALSITY)).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f159,f143])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f170,f143])).
fof(f543,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f172,f143])).
fof(f548,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X15)))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f172,f143])).
fof(f554,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f172,f143])).
fof(f559,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,'F_0')))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f172,f143])).
fof(f614,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))),s(X16,X18))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(forward_demodulation,[],[f588,f159])).
fof(f588,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'F_0') = s(bool,happ(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f235,f170])).
fof(f616,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(forward_demodulation,[],[f615,f159])).
fof(f615,plain,(
  ( ! [X2,X1] : (s(bool,X2) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f594,f140])).
fof(f594,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X1))) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f215,f235])).
fof(f617,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,X15) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f601,f159])).
fof(f601,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,'T_0')))) | s(bool,X15) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X14))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f172,f235])).
fof(f618,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,X17) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(forward_demodulation,[],[f602,f159])).
fof(f602,plain,(
  ( ! [X17,X16] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X17)))) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X16))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f172,f235])).
fof(f653,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X13)))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f173,f143])).
fof(f656,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X20))),s(bool,X21)))) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f173,f143])).
fof(f665,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X13))),s(bool,'F_0')))) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f173,f143])).
fof(f668,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X21))),s(bool,X20)))) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f173,f143])).
fof(f684,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X23)))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f657,f159])).
fof(f657,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X23)))) | s(bool,X23) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f173,f235])).
fof(f688,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f669,f159])).
fof(f669,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X23))),s(bool,'T_0')))) | s(bool,X23) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f173,f235])).
fof(f710,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f638,f159])).
fof(f711,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),'_dst_x0_1'(s(bool,X25))),s(X24,X26))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f638,f170])).
fof(f721,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) )),
  inference(superposition,[],[f638,f159])).
fof(f722,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X27) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,X27))),s(X26,X28)))) )),
  inference(superposition,[],[f638,f170])).
fof(f795,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK10(X16,X17,X18,bool))) )),
  inference(superposition,[],[f171,f143])).
fof(f800,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK10(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK10(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f171,f143])).
fof(f956,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,X42))),s(X44,X45))) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),'_dst_x0_1'(s(bool,X42))),s(X41,X43)))) )),
  inference(superposition,[],[f267,f170])).
fof(f964,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X23))),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f267,f223])).
fof(f1018,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2)) )),
  inference(forward_demodulation,[],[f1017,f170])).
fof(f1017,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0_1'(s(bool,X2))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f975,f140])).
fof(f975,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),'_dst_x0_1'(s(bool,X2))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f168,f267])).
fof(f1051,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),'_dst_x0x1_2'(s(fun(X27,bool),'_dst_x0_1'(s(bool,X28))),s(fun(X29,X27),X30))),s(X29,X31))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f182,f267])).
fof(f1056,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'F_0') = s(bool,happ(s(fun(X58,bool),'_dst_x0x1_2'(s(fun(X56,bool),X57),s(fun(X58,X56),X59))),s(X58,X60))) | s(bool,'T_0') = s(bool,happ(s(fun(X56,bool),X57),s(X56,happ(s(fun(X58,X56),X59),s(X58,X60)))))) )),
  inference(superposition,[],[f182,f143])).
fof(f1087,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(X39,bool),X40),s(X39,happ(s(fun(X38,X39),X41),s(X38,X42))))) | s(bool,'T_0') = s(bool,happ(s(fun(X38,bool),'_dst_x0x1_2'(s(fun(X39,bool),X40),s(fun(X38,X39),X41))),s(X38,X42)))) )),
  inference(superposition,[],[f182,f143])).
fof(f1107,plain,(
  ( ! [X127,X125,X128,X126,X129] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X127,bool),'_dst_x0x1_2'(s(fun(X125,bool),X126),s(fun(X127,X125),X128))),s(X127,X129))))) | s(bool,'F_0') = s(bool,happ(s(fun(X125,bool),X126),s(X125,happ(s(fun(X127,X125),X128),s(X127,X129)))))) )),
  inference(superposition,[],[f235,f182])).
fof(f1108,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131] : (s(bool,'T_0') = s(bool,happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,happ(s(fun(X132,bool),'_dst_x0x1_2'(s(fun(X130,bool),X131),s(fun(X132,X130),X133))),s(X132,X134))))),s(X135,X136))) | s(bool,'F_0') = s(bool,happ(s(fun(X130,bool),X131),s(X130,happ(s(fun(X132,X130),X133),s(X132,X134)))))) )),
  inference(superposition,[],[f267,f182])).
fof(f1111,plain,(
  ( ! [X152,X151,X149,X153,X150,X148] : (p(s(bool,happ(s(fun(X150,bool),'_dst_x0x1_2'(s(fun(X148,bool),X149),s(fun(X150,X148),X151))),s(X150,X152)))) | s(bool,X153) != s(bool,happ(s(fun(X148,bool),X149),s(X148,happ(s(fun(X150,X148),X151),s(X150,X152))))) | s(bool,'F_0') = s(bool,X153)) )),
  inference(superposition,[],[f616,f182])).
fof(f1112,plain,(
  ( ! [X158,X156,X154,X159,X157,X155] : (s(bool,X159) != s(bool,happ(s(fun(X156,bool),'_dst_x0x1_2'(s(fun(X154,bool),X155),s(fun(X156,X154),X157))),s(X156,X158))) | s(bool,'T_0') = s(bool,happ(s(fun(X154,bool),X155),s(X154,happ(s(fun(X156,X154),X157),s(X156,X158))))) | s(bool,'F_0') = s(bool,X159)) )),
  inference(superposition,[],[f638,f182])).
fof(f1113,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (s(bool,X165) != s(bool,happ(s(fun(X162,bool),'_dst_x0x1_2'(s(fun(X160,bool),X161),s(fun(X162,X160),X163))),s(X162,X164))) | s(bool,'T_0') = s(bool,X165) | s(bool,'F_0') = s(bool,happ(s(fun(X160,bool),X161),s(X160,happ(s(fun(X162,X160),X163),s(X162,X164)))))) )),
  inference(superposition,[],[f638,f182])).
fof(f1114,plain,(
  ( ! [X167,X169,X166,X170,X168] : (p(s(bool,happ(s(fun(X168,bool),'_dst_x0x1_2'(s(fun(X166,bool),X167),s(fun(X168,X166),X169))),s(X168,X170)))) | s(bool,'F_0') = s(bool,happ(s(fun(X166,bool),X167),s(X166,happ(s(fun(X168,X166),X169),s(X168,X170)))))) )),
  inference(superposition,[],[f1018,f182])).
fof(f1134,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X105,bool),X106),s(X105,happ(s(fun(X104,X105),X107),s(X104,X108))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X104,bool),'_dst_x0x1_2'(s(fun(X105,bool),X106),s(fun(X104,X105),X107))),s(X104,X108)))) )),
  inference(superposition,[],[f235,f182])).
fof(f1135,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X114,bool),'_dst_x0_1'(s(bool,happ(s(fun(X110,bool),X111),s(X110,happ(s(fun(X109,X110),X112),s(X109,X113))))))),s(X114,X115))) | s(bool,'F_0') = s(bool,happ(s(fun(X109,bool),'_dst_x0x1_2'(s(fun(X110,bool),X111),s(fun(X109,X110),X112))),s(X109,X113)))) )),
  inference(superposition,[],[f267,f182])).
fof(f1138,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (p(s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X127,X128),X130),s(X127,X131)))))) | s(bool,X132) != s(bool,happ(s(fun(X127,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X127,X128),X130))),s(X127,X131))) | s(bool,'F_0') = s(bool,X132)) )),
  inference(superposition,[],[f616,f182])).
fof(f1139,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (s(bool,X138) != s(bool,happ(s(fun(X134,bool),X135),s(X134,happ(s(fun(X133,X134),X136),s(X133,X137))))) | s(bool,'T_0') = s(bool,happ(s(fun(X133,bool),'_dst_x0x1_2'(s(fun(X134,bool),X135),s(fun(X133,X134),X136))),s(X133,X137))) | s(bool,'F_0') = s(bool,X138)) )),
  inference(superposition,[],[f638,f182])).
fof(f1140,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (s(bool,X144) != s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X139,X140),X142),s(X139,X143))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X139,X140),X142))),s(X139,X143)))) )),
  inference(superposition,[],[f638,f182])).
fof(f1141,plain,(
  ( ! [X146,X149,X147,X145,X148] : (p(s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X145,X146),X148),s(X145,X149)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X145,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X145,X146),X148))),s(X145,X149)))) )),
  inference(superposition,[],[f1018,f182])).
fof(f1196,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X25) | s(bool,X25) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) )),
  inference(superposition,[],[f1176,f159])).
fof(f1197,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') = s(bool,X29) | s(bool,X29) = s(bool,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,X27))),s(X26,X28)))) )),
  inference(superposition,[],[f1176,f170])).
fof(f1198,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,happ(s(fun(X32,bool),'_dst_x0x1_2'(s(fun(X30,bool),X31),s(fun(X32,X30),X33))),s(X32,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X35) = s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34)))))) )),
  inference(superposition,[],[f1176,f182])).
fof(f1199,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'F_0') != s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X36,X37),X39),s(X36,X40))))) | s(bool,'T_0') = s(bool,X41) | s(bool,X41) = s(bool,happ(s(fun(X36,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X36,X37),X39))),s(X36,X40)))) )),
  inference(superposition,[],[f1176,f182])).
fof(f1269,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X32),s(X30,X33))),s(X31,X34))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X30,fun(X31,bool)),X32),s('type.pair.prod'(X30,X31),'const.pair.,_2'(s(X30,X33),s(X31,X34)))))) )),
  inference(superposition,[],[f183,f143])).
fof(f1301,plain,(
  ( ! [X111,X109,X107,X110,X108] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X108,bool),happ(s(fun(X107,fun(X108,bool)),X109),s(X107,X110))),s(X108,X111))))) | s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X107,fun(X108,bool)),X109),s('type.pair.prod'(X107,X108),'const.pair.,_2'(s(X107,X110),s(X108,X111)))))) )),
  inference(superposition,[],[f235,f183])).
fof(f1302,plain,(
  ( ! [X118,X116,X114,X112,X117,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X117,bool),'_dst_x0_1'(s(bool,happ(s(fun(X113,bool),happ(s(fun(X112,fun(X113,bool)),X114),s(X112,X115))),s(X113,X116))))),s(X117,X118))) | s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X112,fun(X113,bool)),X114),s('type.pair.prod'(X112,X113),'const.pair.,_2'(s(X112,X115),s(X113,X116)))))) )),
  inference(superposition,[],[f267,f183])).
fof(f1307,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (s(bool,X148) != s(bool,happ(s(fun(X144,bool),happ(s(fun(X143,fun(X144,bool)),X145),s(X143,X146))),s(X144,X147))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X143,fun(X144,bool)),X145),s('type.pair.prod'(X143,X144),'const.pair.,_2'(s(X143,X146),s(X144,X147))))) | s(bool,'F_0') = s(bool,X148)) )),
  inference(superposition,[],[f638,f183])).
fof(f1308,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,X154) != s(bool,happ(s(fun(X150,bool),happ(s(fun(X149,fun(X150,bool)),X151),s(X149,X152))),s(X150,X153))) | s(bool,'T_0') = s(bool,X154) | s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X149,fun(X150,bool)),X151),s('type.pair.prod'(X149,X150),'const.pair.,_2'(s(X149,X152),s(X150,X153)))))) )),
  inference(superposition,[],[f638,f183])).
fof(f1310,plain,(
  ( ! [X165,X163,X161,X164,X162,X160] : (s(bool,'F_0') != s(bool,happ(s(fun(X161,bool),happ(s(fun(X160,fun(X161,bool)),X162),s(X160,X163))),s(X161,X164))) | s(bool,'T_0') = s(bool,X165) | s(bool,X165) = s(bool,'const.pair.UNCURRY_2'(s(fun(X160,fun(X161,bool)),X162),s('type.pair.prod'(X160,X161),'const.pair.,_2'(s(X160,X163),s(X161,X164)))))) )),
  inference(superposition,[],[f1176,f183])).
fof(f1343,plain,(
  ( ! [X47,X45,X50,X48,X46,X44,X49] : (s(bool,'F_0') = s(bool,happ(s(fun(X49,bool),'_dst_x0_1'(s(bool,happ(s(fun(X46,bool),'_dst_x0x1_2'(s(fun(X44,bool),X45),s(fun(X46,X44),X47))),s(X46,X48))))),s(X49,X50))) | s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48)))))) )),
  inference(superposition,[],[f614,f182])).
fof(f1344,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,happ(s(fun(X56,bool),'_dst_x0_1'(s(bool,happ(s(fun(X52,bool),X53),s(X52,happ(s(fun(X51,X52),X54),s(X51,X55))))))),s(X56,X57))) | s(bool,'T_0') = s(bool,happ(s(fun(X51,bool),'_dst_x0x1_2'(s(fun(X52,bool),X53),s(fun(X51,X52),X54))),s(X51,X55)))) )),
  inference(superposition,[],[f614,f182])).
fof(f1346,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,happ(s(fun(X73,bool),'_dst_x0_1'(s(bool,happ(s(fun(X69,bool),happ(s(fun(X68,fun(X69,bool)),X70),s(X68,X71))),s(X69,X72))))),s(X73,X74))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X68,fun(X69,bool)),X70),s('type.pair.prod'(X68,X69),'const.pair.,_2'(s(X68,X71),s(X69,X72)))))) )),
  inference(superposition,[],[f614,f183])).
fof(f1372,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),'_dst_x0x1_2'(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29))),s(fun(X30,X28),X31))),s(X30,X32))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f614,f182])).
fof(f1424,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) != s(bool,happ(s(fun(X54,bool),'_dst_x0_1'(s(bool,X55))),s(X54,X56))) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1386,f142])).
fof(f1386,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,'F_0')) | s(bool,X57) != s(bool,happ(s(fun(X54,bool),'_dst_x0_1'(s(bool,X55))),s(X54,X56))) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f616,f614])).
fof(f1590,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X31,bool),happ(s(fun(X30,fun(X31,bool)),X34),s(X30,X32))),s(X31,X33))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X30,X31),'const.pair.,_2'(s(X30,X32),s(X31,X33))),s(fun(X30,fun(X31,bool)),X34)))) )),
  inference(superposition,[],[f184,f143])).
fof(f1633,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X141,bool),happ(s(fun(X140,fun(X141,bool)),X144),s(X140,X142))),s(X141,X143))))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X140,X141),'const.pair.,_2'(s(X140,X142),s(X141,X143))),s(fun(X140,fun(X141,bool)),X144)))) )),
  inference(superposition,[],[f235,f184])).
fof(f1634,plain,(
  ( ! [X146,X151,X149,X147,X145,X150,X148] : (s(bool,'T_0') = s(bool,happ(s(fun(X150,bool),'_dst_x0_1'(s(bool,happ(s(fun(X146,bool),happ(s(fun(X145,fun(X146,bool)),X149),s(X145,X147))),s(X146,X148))))),s(X150,X151))) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X145,X146),'const.pair.,_2'(s(X145,X147),s(X146,X148))),s(fun(X145,fun(X146,bool)),X149)))) )),
  inference(superposition,[],[f267,f184])).
fof(f1638,plain,(
  ( ! [X175,X173,X171,X176,X174,X172,X170] : (s(bool,'F_0') = s(bool,happ(s(fun(X175,bool),'_dst_x0_1'(s(bool,happ(s(fun(X171,bool),happ(s(fun(X170,fun(X171,bool)),X174),s(X170,X172))),s(X171,X173))))),s(X175,X176))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X170,X171),'const.pair.,_2'(s(X170,X172),s(X171,X173))),s(fun(X170,fun(X171,bool)),X174)))) )),
  inference(superposition,[],[f614,f184])).
fof(f1640,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,X188) != s(bool,happ(s(fun(X184,bool),happ(s(fun(X183,fun(X184,bool)),X187),s(X183,X185))),s(X184,X186))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X183,X184),'const.pair.,_2'(s(X183,X185),s(X184,X186))),s(fun(X183,fun(X184,bool)),X187))) | s(bool,'F_0') = s(bool,X188)) )),
  inference(superposition,[],[f638,f184])).
fof(f1641,plain,(
  ( ! [X191,X189,X194,X192,X190,X193] : (s(bool,X194) != s(bool,happ(s(fun(X190,bool),happ(s(fun(X189,fun(X190,bool)),X193),s(X189,X191))),s(X190,X192))) | s(bool,'T_0') = s(bool,X194) | s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X189,X190),'const.pair.,_2'(s(X189,X191),s(X190,X192))),s(fun(X189,fun(X190,bool)),X193)))) )),
  inference(superposition,[],[f638,f184])).
fof(f1643,plain,(
  ( ! [X204,X202,X200,X205,X203,X201] : (s(bool,'F_0') != s(bool,happ(s(fun(X201,bool),happ(s(fun(X200,fun(X201,bool)),X204),s(X200,X202))),s(X201,X203))) | s(bool,'T_0') = s(bool,X205) | s(bool,X205) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X200,X201),'const.pair.,_2'(s(X200,X202),s(X201,X203))),s(fun(X200,fun(X201,bool)),X204)))) )),
  inference(superposition,[],[f1176,f184])).
fof(f1838,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X17),bool),X16),s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,'F_0'),s(X17,sK9(X16,X17,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X17),bool),X16),s('type.pair.prod'(bool,X17),X18)))) | s(bool,'T_0') = s(bool,sK8(X16,X17,bool))) )),
  inference(superposition,[],[f169,f143])).
fof(f1843,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X17,bool),bool),X16),s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,sK8(X16,bool,X17)),s(bool,'F_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X17,bool),bool),X16),s('type.pair.prod'(X17,bool),X18)))) | s(bool,'T_0') = s(bool,sK9(X16,bool,X17))) )),
  inference(superposition,[],[f169,f143])).
fof(f2046,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X42,fun(X43,bool)),X44),s('type.pair.prod'(X42,X43),X45))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun(fun(X42,fun(X43,bool)),fun('type.pair.prod'(X42,X43),bool)),'const.pair.UNCURRY_0'),s(fun(X42,fun(X43,bool)),X44))),s('type.pair.prod'(X42,X43),X45)))) )),
  inference(superposition,[],[f180,f143])).
fof(f2074,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0_1'(s(bool,'const.pair.UNCURRY_2'(s(fun(X175,fun(X176,bool)),X177),s('type.pair.prod'(X175,X176),X178))))),s(X179,X180))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X175,X176),bool),happ(s(fun(fun(X175,fun(X176,bool)),fun('type.pair.prod'(X175,X176),bool)),'const.pair.UNCURRY_0'),s(fun(X175,fun(X176,bool)),X177))),s('type.pair.prod'(X175,X176),X178)))) )),
  inference(superposition,[],[f267,f180])).
fof(f2078,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'F_0') = s(bool,happ(s(fun(X200,bool),'_dst_x0_1'(s(bool,'const.pair.UNCURRY_2'(s(fun(X196,fun(X197,bool)),X198),s('type.pair.prod'(X196,X197),X199))))),s(X200,X201))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X196,X197),bool),happ(s(fun(fun(X196,fun(X197,bool)),fun('type.pair.prod'(X196,X197),bool)),'const.pair.UNCURRY_0'),s(fun(X196,fun(X197,bool)),X198))),s('type.pair.prod'(X196,X197),X199)))) )),
  inference(superposition,[],[f614,f180])).
fof(f2080,plain,(
  ( ! [X210,X208,X211,X209,X207] : (s(bool,X211) != s(bool,'const.pair.UNCURRY_2'(s(fun(X207,fun(X208,bool)),X209),s('type.pair.prod'(X207,X208),X210))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X207,X208),bool),happ(s(fun(fun(X207,fun(X208,bool)),fun('type.pair.prod'(X207,X208),bool)),'const.pair.UNCURRY_0'),s(fun(X207,fun(X208,bool)),X209))),s('type.pair.prod'(X207,X208),X210))) | s(bool,'F_0') = s(bool,X211)) )),
  inference(superposition,[],[f638,f180])).
fof(f2081,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,X216) != s(bool,'const.pair.UNCURRY_2'(s(fun(X212,fun(X213,bool)),X214),s('type.pair.prod'(X212,X213),X215))) | s(bool,'T_0') = s(bool,X216) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X212,X213),bool),happ(s(fun(fun(X212,fun(X213,bool)),fun('type.pair.prod'(X212,X213),bool)),'const.pair.UNCURRY_0'),s(fun(X212,fun(X213,bool)),X214))),s('type.pair.prod'(X212,X213),X215)))) )),
  inference(superposition,[],[f638,f180])).
fof(f2083,plain,(
  ( ! [X222,X225,X223,X221,X224] : (s(bool,'F_0') != s(bool,'const.pair.UNCURRY_2'(s(fun(X221,fun(X222,bool)),X223),s('type.pair.prod'(X221,X222),X224))) | s(bool,'T_0') = s(bool,X225) | s(bool,X225) = s(bool,happ(s(fun('type.pair.prod'(X221,X222),bool),happ(s(fun(fun(X221,fun(X222,bool)),fun('type.pair.prod'(X221,X222),bool)),'const.pair.UNCURRY_0'),s(fun(X221,fun(X222,bool)),X223))),s('type.pair.prod'(X221,X222),X224)))) )),
  inference(superposition,[],[f1176,f180])).
fof(f2107,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X171,fun(X172,bool)),X173),s('type.pair.prod'(X171,X172),X174))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X171,X172),bool),happ(s(fun(fun(X171,fun(X172,bool)),fun('type.pair.prod'(X171,X172),bool)),'const.pair.UNCURRY_0'),s(fun(X171,fun(X172,bool)),X173))),s('type.pair.prod'(X171,X172),X174)))) )),
  inference(forward_demodulation,[],[f2073,f159])).
fof(f2073,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.UNCURRY_2'(s(fun(X171,fun(X172,bool)),X173),s('type.pair.prod'(X171,X172),X174))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(X171,X172),bool),happ(s(fun(fun(X171,fun(X172,bool)),fun('type.pair.prod'(X171,X172),bool)),'const.pair.UNCURRY_0'),s(fun(X171,fun(X172,bool)),X173))),s('type.pair.prod'(X171,X172),X174)))) )),
  inference(superposition,[],[f235,f180])).
fof(f2658,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X42,X43),X44),s(fun(X42,fun(X43,bool)),X45))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X42,fun(X43,bool)),bool),happ(s(fun('type.pair.prod'(X42,X43),fun(fun(X42,fun(X43,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X42,X43),X44))),s(fun(X42,fun(X43,bool)),X45)))) )),
  inference(superposition,[],[f181,f143])).
fof(f2689,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (s(bool,'T_0') = s(bool,happ(s(fun(X205,bool),'_dst_x0_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X201,X202),X203),s(fun(X201,fun(X202,bool)),X204))))),s(X205,X206))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X201,fun(X202,bool)),bool),happ(s(fun('type.pair.prod'(X201,X202),fun(fun(X201,fun(X202,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X201,X202),X203))),s(fun(X201,fun(X202,bool)),X204)))) )),
  inference(superposition,[],[f267,f181])).
fof(f2693,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (s(bool,'F_0') = s(bool,happ(s(fun(X226,bool),'_dst_x0_1'(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X222,X223),X224),s(fun(X222,fun(X223,bool)),X225))))),s(X226,X227))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X222,fun(X223,bool)),bool),happ(s(fun('type.pair.prod'(X222,X223),fun(fun(X222,fun(X223,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X222,X223),X224))),s(fun(X222,fun(X223,bool)),X225)))) )),
  inference(superposition,[],[f614,f181])).
fof(f2695,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,X237) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X233,X234),X235),s(fun(X233,fun(X234,bool)),X236))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X233,fun(X234,bool)),bool),happ(s(fun('type.pair.prod'(X233,X234),fun(fun(X233,fun(X234,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X233,X234),X235))),s(fun(X233,fun(X234,bool)),X236))) | s(bool,'F_0') = s(bool,X237)) )),
  inference(superposition,[],[f638,f181])).
fof(f2696,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,X242) != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X238,X239),X240),s(fun(X238,fun(X239,bool)),X241))) | s(bool,'T_0') = s(bool,X242) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X238,fun(X239,bool)),bool),happ(s(fun('type.pair.prod'(X238,X239),fun(fun(X238,fun(X239,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X238,X239),X240))),s(fun(X238,fun(X239,bool)),X241)))) )),
  inference(superposition,[],[f638,f181])).
fof(f2698,plain,(
  ( ! [X251,X249,X250,X248,X247] : (s(bool,'F_0') != s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X247,X248),X249),s(fun(X247,fun(X248,bool)),X250))) | s(bool,'T_0') = s(bool,X251) | s(bool,X251) = s(bool,happ(s(fun(fun(X247,fun(X248,bool)),bool),happ(s(fun('type.pair.prod'(X247,X248),fun(fun(X247,fun(X248,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X247,X248),X249))),s(fun(X247,fun(X248,bool)),X250)))) )),
  inference(superposition,[],[f1176,f181])).
fof(f2721,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X197,X198),X199),s(fun(X197,fun(X198,bool)),X200))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X197,fun(X198,bool)),bool),happ(s(fun('type.pair.prod'(X197,X198),fun(fun(X197,fun(X198,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X197,X198),X199))),s(fun(X197,fun(X198,bool)),X200)))) )),
  inference(forward_demodulation,[],[f2688,f159])).
fof(f2688,plain,(
  ( ! [X198,X200,X199,X197] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X197,X198),X199),s(fun(X197,fun(X198,bool)),X200))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X197,fun(X198,bool)),bool),happ(s(fun('type.pair.prod'(X197,X198),fun(fun(X197,fun(X198,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X197,X198),X199))),s(fun(X197,fun(X198,bool)),X200)))) )),
  inference(superposition,[],[f235,f181])).
fof(f798,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK10(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK10(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f171,f223])).
fof(f796,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK10(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK10(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f171,f223])).
fof(f680,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f646,f195])).
fof(f646,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f159,f173])).
fof(f683,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f643,f194])).
fof(f643,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f159,f173])).
fof(f172,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X44,X45,X46] : (p(s(bool,happ(s(fun(X44,bool),happ(s(fun(X44,fun(X44,bool)),'=_0'),s(X44,X45))),s(X44,X46)))) <=> s(X44,X45) = s(X44,X46))),
  file('thm.closProps.evaluate_app_const',arity_0_HL_EQ)).
fof(f173,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f133])).
fof(f1101,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | s(bool,X95) != s(bool,happ(s(fun(X90,bool),X91),s(X90,happ(s(fun(X92,X90),X93),s(X92,X94))))) | p(s(bool,X95))) )),
  inference(superposition,[],[f215,f182])).
fof(f1106,plain,(
  ( ! [X123,X121,X124,X122,X120] : (~p(s(bool,happ(s(fun(X122,bool),'_dst_x0x1_2'(s(fun(X120,bool),X121),s(fun(X122,X120),X123))),s(X122,X124)))) | s(bool,'T_0') = s(bool,happ(s(fun(X120,bool),X121),s(X120,happ(s(fun(X122,X120),X123),s(X122,X124)))))) )),
  inference(superposition,[],[f225,f182])).
fof(f1109,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (p(s(bool,happ(s(fun(X139,bool),'_dst_x0x1_2'(s(fun(X137,bool),X138),s(fun(X139,X137),X140))),s(X139,X141)))) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),X138),s(X137,happ(s(fun(X139,X137),X140),s(X139,X141))))) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f347,f182])).
fof(f1128,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,happ(s(fun(X69,X70),X72),s(X69,X73)))))) | s(bool,X74) != s(bool,happ(s(fun(X69,bool),'_dst_x0x1_2'(s(fun(X70,bool),X71),s(fun(X69,X70),X72))),s(X69,X73))) | p(s(bool,X74))) )),
  inference(superposition,[],[f215,f182])).
fof(f1133,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,happ(s(fun(X100,bool),X101),s(X100,happ(s(fun(X99,X100),X102),s(X99,X103)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X99,bool),'_dst_x0x1_2'(s(fun(X100,bool),X101),s(fun(X99,X100),X102))),s(X99,X103)))) )),
  inference(superposition,[],[f225,f182])).
fof(f1136,plain,(
  ( ! [X121,X118,X116,X120,X119,X117] : (p(s(bool,happ(s(fun(X117,bool),X118),s(X117,happ(s(fun(X116,X117),X119),s(X116,X120)))))) | s(bool,X121) = s(bool,happ(s(fun(X116,bool),'_dst_x0x1_2'(s(fun(X117,bool),X118),s(fun(X116,X117),X119))),s(X116,X120))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f347,f182])).
fof(f1165,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),'_dst_x0x1_2'(s(fun(X46,bool),X47),s(fun(X48,X46),X49))),s(X48,X50)))) | s(bool,X51) != s(bool,happ(s(fun(X46,bool),X47),s(X46,happ(s(fun(X48,X46),X49),s(X48,X50))))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f489,f182])).
fof(f1166,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X54,bool),X55),s(X54,happ(s(fun(X53,X54),X56),s(X53,X57)))))) | s(bool,X58) != s(bool,happ(s(fun(X53,bool),'_dst_x0x1_2'(s(fun(X54,bool),X55),s(fun(X53,X54),X56))),s(X53,X57))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f489,f182])).
fof(f1098,plain,(
  ( ! [X76,X74,X72,X75,X73] : (~p(s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X72,bool),happ(s(fun(X72,fun(X72,bool)),'=_0'),s(X72,X73))),s(fun(X74,X72),X75))),s(X74,X76)))) | s(X72,X73) = s(X72,happ(s(fun(X74,X72),X75),s(X74,X76)))) )),
  inference(superposition,[],[f172,f182])).
fof(f1097,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,happ(s(fun(X69,bool),'_dst_x0x1_2'(s(fun(X67,bool),happ(s(fun(X67,fun(X67,bool)),'=_0'),s(X67,X68))),s(fun(X69,X67),X70))),s(X69,X71)))) | s(X67,X68) != s(X67,happ(s(fun(X69,X67),X70),s(X69,X71)))) )),
  inference(superposition,[],[f173,f182])).
fof(f1265,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X9))),s(X7,X10))) | s(bool,X11) = s(bool,'const.pair.UNCURRY_2'(s(fun(X6,fun(X7,bool)),X8),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f183,f223])).
fof(f1267,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X23) = s(bool,happ(s(fun(X19,bool),happ(s(fun(X18,fun(X19,bool)),X20),s(X18,X21))),s(X19,X22))) | s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X18,fun(X19,bool)),X20),s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,X21),s(X19,X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f183,f223])).
fof(f1586,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X8))),s(X7,X9))) | s(bool,X11) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X8),s(X7,X9))),s(fun(X6,fun(X7,bool)),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f184,f223])).
fof(f1588,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X23) = s(bool,happ(s(fun(X19,bool),happ(s(fun(X18,fun(X19,bool)),X22),s(X18,X20))),s(X19,X21))) | s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X18,X19),'const.pair.,_2'(s(X18,X20),s(X19,X21))),s(fun(X18,fun(X19,bool)),X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f184,f223])).
fof(f2102,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),happ(s(fun(X129,fun(X130,bool)),X131),s(X129,X132))),s(X130,X133)))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X129,fun(X130,bool)),X131),s('type.pair.prod'(X129,X130),sK7(happ(s(fun(fun(X129,fun(X130,bool)),fun('type.pair.prod'(X129,X130),bool)),'const.pair.UNCURRY_0'),s(fun(X129,fun(X130,bool)),X131)),X130,X129)))))) )),
  inference(forward_demodulation,[],[f2101,f183])).
fof(f2101,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X129,fun(X130,bool)),X131),s('type.pair.prod'(X129,X130),'const.pair.,_2'(s(X129,X132),s(X130,X133)))))) | ~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X129,fun(X130,bool)),X131),s('type.pair.prod'(X129,X130),sK7(happ(s(fun(fun(X129,fun(X130,bool)),fun('type.pair.prod'(X129,X130),bool)),'const.pair.UNCURRY_0'),s(fun(X129,fun(X130,bool)),X131)),X130,X129)))))) )),
  inference(forward_demodulation,[],[f2064,f180])).
fof(f2064,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X129,fun(X130,bool)),X131),s('type.pair.prod'(X129,X130),sK7(happ(s(fun(fun(X129,fun(X130,bool)),fun('type.pair.prod'(X129,X130),bool)),'const.pair.UNCURRY_0'),s(fun(X129,fun(X130,bool)),X131)),X130,X129))))) | p(s(bool,happ(s(fun('type.pair.prod'(X129,X130),bool),happ(s(fun(fun(X129,fun(X130,bool)),fun('type.pair.prod'(X129,X130),bool)),'const.pair.UNCURRY_0'),s(fun(X129,fun(X130,bool)),X131))),s('type.pair.prod'(X129,X130),'const.pair.,_2'(s(X129,X132),s(X130,X133))))))) )),
  inference(superposition,[],[f168,f180])).
fof(f1094,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (p(s(bool,happ(s(fun(X46,bool),happ(s(fun(X46,fun(X46,bool)),'=_0'),s(X46,happ(s(fun(X49,X46),'_dst_x0x1_2'(s(fun(X47,X46),X48),s(fun(X49,X47),X50))),s(X49,X51))))),s(X46,X52)))) | s(X46,X52) != s(X46,happ(s(fun(X47,X46),X48),s(X47,happ(s(fun(X49,X47),X50),s(X49,X51)))))) )),
  inference(superposition,[],[f173,f182])).
fof(f1092,plain,(
  ( ! [X37,X35,X33,X38,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),'=_0'),s(X32,happ(s(fun(X35,X32),'_dst_x0x1_2'(s(fun(X33,X32),X34),s(fun(X35,X33),X36))),s(X35,X37))))),s(X32,X38)))) | s(X32,X38) = s(X32,happ(s(fun(X33,X32),X34),s(X33,happ(s(fun(X35,X33),X36),s(X35,X37)))))) )),
  inference(superposition,[],[f172,f182])).
fof(f1123,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),'=_0'),s(X36,happ(s(fun(X38,X36),X39),s(X38,happ(s(fun(X37,X38),X40),s(X37,X41))))))),s(X36,X42)))) | s(X36,X42) != s(X36,happ(s(fun(X37,X36),'_dst_x0x1_2'(s(fun(X38,X36),X39),s(fun(X37,X38),X40))),s(X37,X41)))) )),
  inference(superposition,[],[f173,f182])).
fof(f1121,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,happ(s(fun(X24,X22),X25),s(X24,happ(s(fun(X23,X24),X26),s(X23,X27))))))),s(X22,X28)))) | s(X22,X28) = s(X22,happ(s(fun(X23,X22),'_dst_x0x1_2'(s(fun(X24,X22),X25),s(fun(X23,X24),X26))),s(X23,X27)))) )),
  inference(superposition,[],[f172,f182])).
fof(f1666,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(X36,X42) != s(X36,'const.pair.pair_CASE_2'(s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,X40),s(X38,X41))),s(fun(X37,fun(X38,X36)),X39))) | p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),'=_0'),s(X36,happ(s(fun(X38,X36),happ(s(fun(X37,fun(X38,X36)),X39),s(X37,X40))),s(X38,X41))))),s(X36,X42))))) )),
  inference(backward_demodulation,[],[f1532,f1290])).
fof(f1664,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s(X22,X28) = s(X22,'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X24),'const.pair.,_2'(s(X23,X26),s(X24,X27))),s(fun(X23,fun(X24,X22)),X25))) | ~p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,happ(s(fun(X24,X22),happ(s(fun(X23,fun(X24,X22)),X25),s(X23,X26))),s(X24,X27))))),s(X22,X28))))) )),
  inference(backward_demodulation,[],[f1532,f1288])).
fof(f1662,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X52,fun(X52,bool)),'=_0'),s(X52,'const.pair.pair_CASE_2'(s('type.pair.prod'(X54,X53),'const.pair.,_2'(s(X54,X56),s(X53,X57))),s(fun(X54,fun(X53,X52)),X55))))),s(X52,X58)))) | s(X52,X58) != s(X52,happ(s(fun(X53,X52),happ(s(fun(X54,fun(X53,X52)),X55),s(X54,X56))),s(X53,X57)))) )),
  inference(backward_demodulation,[],[f1532,f1277])).
fof(f1660,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X38,fun(X38,bool)),'=_0'),s(X38,'const.pair.pair_CASE_2'(s('type.pair.prod'(X40,X39),'const.pair.,_2'(s(X40,X42),s(X39,X43))),s(fun(X40,fun(X39,X38)),X41))))),s(X38,X44)))) | s(X38,X44) = s(X38,happ(s(fun(X39,X38),happ(s(fun(X40,fun(X39,X38)),X41),s(X40,X42))),s(X39,X43)))) )),
  inference(backward_demodulation,[],[f1532,f1275])).
fof(f1095,plain,(
  ( ! [X59,X57,X54,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X53,bool),happ(s(fun(X53,fun(X53,bool)),'=_0'),s(X53,X59))),s(X53,happ(s(fun(X56,X53),'_dst_x0x1_2'(s(fun(X54,X53),X55),s(fun(X56,X54),X57))),s(X56,X58)))))) | s(X53,X59) != s(X53,happ(s(fun(X54,X53),X55),s(X54,happ(s(fun(X56,X54),X57),s(X56,X58)))))) )),
  inference(superposition,[],[f173,f182])).
fof(f1093,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),happ(s(fun(X39,fun(X39,bool)),'=_0'),s(X39,X45))),s(X39,happ(s(fun(X42,X39),'_dst_x0x1_2'(s(fun(X40,X39),X41),s(fun(X42,X40),X43))),s(X42,X44)))))) | s(X39,X45) = s(X39,happ(s(fun(X40,X39),X41),s(X40,happ(s(fun(X42,X40),X43),s(X42,X44)))))) )),
  inference(superposition,[],[f172,f182])).
fof(f1124,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X49))),s(X43,happ(s(fun(X45,X43),X46),s(X45,happ(s(fun(X44,X45),X47),s(X44,X48)))))))) | s(X43,X49) != s(X43,happ(s(fun(X44,X43),'_dst_x0x1_2'(s(fun(X45,X43),X46),s(fun(X44,X45),X47))),s(X44,X48)))) )),
  inference(superposition,[],[f173,f182])).
fof(f1122,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (~p(s(bool,happ(s(fun(X29,bool),happ(s(fun(X29,fun(X29,bool)),'=_0'),s(X29,X35))),s(X29,happ(s(fun(X31,X29),X32),s(X31,happ(s(fun(X30,X31),X33),s(X30,X34)))))))) | s(X29,X35) = s(X29,happ(s(fun(X30,X29),'_dst_x0x1_2'(s(fun(X31,X29),X32),s(fun(X30,X31),X33))),s(X30,X34)))) )),
  inference(superposition,[],[f172,f182])).
fof(f1667,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(X43,X49) != s(X43,'const.pair.pair_CASE_2'(s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48))),s(fun(X44,fun(X45,X43)),X46))) | p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X49))),s(X43,happ(s(fun(X45,X43),happ(s(fun(X44,fun(X45,X43)),X46),s(X44,X47))),s(X45,X48))))))) )),
  inference(backward_demodulation,[],[f1532,f1291])).
fof(f1665,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (s(X29,X35) = s(X29,'const.pair.pair_CASE_2'(s('type.pair.prod'(X30,X31),'const.pair.,_2'(s(X30,X33),s(X31,X34))),s(fun(X30,fun(X31,X29)),X32))) | ~p(s(bool,happ(s(fun(X29,bool),happ(s(fun(X29,fun(X29,bool)),'=_0'),s(X29,X35))),s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X32),s(X30,X33))),s(X31,X34))))))) )),
  inference(backward_demodulation,[],[f1532,f1289])).
fof(f1044,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s(X0,happ(s(fun(X2,X0),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f182,f173])).
fof(f1043,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))))),s(X0,happ(s(fun(X3,X0),'_dst_x0x1_2'(s(fun(X1,X0),X2),s(fun(X3,X1),X4))),s(X3,X5))))))) )),
  inference(unit_resulting_resolution,[],[f182,f173])).
fof(f1526,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X3),s(X2,X4))),s(fun(X1,fun(X2,X0)),X5))))),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X3))),s(X2,X4))))))) )),
  inference(unit_resulting_resolution,[],[f184,f173])).
fof(f1126,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,happ(s(fun(X59,bool),X60),s(X59,happ(s(fun('type.pair.prod'(X57,X58),X59),X61),s('type.pair.prod'(X57,X58),sK7('_dst_x0x1_2'(s(fun(X59,bool),X60),s(fun('type.pair.prod'(X57,X58),X59),X61)),X58,X57))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X57,X58),bool),'_dst_x0x1_2'(s(fun(X59,bool),X60),s(fun('type.pair.prod'(X57,X58),X59),X61))),s('type.pair.prod'(X57,X58),'const.pair.,_2'(s(X57,X62),s(X58,X63))))))) )),
  inference(superposition,[],[f168,f182])).
fof(f1851,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X34,bool),X35),s(X34,happ(s(fun('type.pair.prod'(X32,X33),X34),X36),s('type.pair.prod'(X32,X33),'const.pair.,_2'(s(X32,sK8('_dst_x0x1_2'(s(fun(X34,bool),X35),s(fun('type.pair.prod'(X32,X33),X34),X36)),X33,X32)),s(X33,sK9('_dst_x0x1_2'(s(fun(X34,bool),X35),s(fun('type.pair.prod'(X32,X33),X34),X36)),X33,X32))))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),'_dst_x0x1_2'(s(fun(X34,bool),X35),s(fun('type.pair.prod'(X32,X33),X34),X36))),s('type.pair.prod'(X32,X33),X37))))) )),
  inference(superposition,[],[f169,f182])).
fof(f1525,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))),s(fun(X2,fun(X1,X0)),X3))))))) )),
  inference(unit_resulting_resolution,[],[f184,f173])).
fof(f1663,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X59,fun(X59,bool)),'=_0'),s(X59,X65))),s(X59,'const.pair.pair_CASE_2'(s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64))),s(fun(X61,fun(X60,X59)),X62)))))) | s(X59,X65) != s(X59,happ(s(fun(X60,X59),happ(s(fun(X61,fun(X60,X59)),X62),s(X61,X63))),s(X60,X64)))) )),
  inference(backward_demodulation,[],[f1532,f1278])).
fof(f1661,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (~p(s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),'=_0'),s(X45,X51))),s(X45,'const.pair.pair_CASE_2'(s('type.pair.prod'(X47,X46),'const.pair.,_2'(s(X47,X49),s(X46,X50))),s(fun(X47,fun(X46,X45)),X48)))))) | s(X45,X51) = s(X45,happ(s(fun(X46,X45),happ(s(fun(X47,fun(X46,X45)),X48),s(X47,X49))),s(X46,X50)))) )),
  inference(backward_demodulation,[],[f1532,f1276])).
fof(f2104,plain,(
  ( ! [X134,X136,X135,X137] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun(X134,fun(X135,bool)),X136),s('type.pair.prod'(X134,X135),X137)))) | ~p(s(bool,happ(s(fun(X135,bool),happ(s(fun(X134,fun(X135,bool)),X136),s(X134,sK8(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134)))),s(X135,sK9(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134)))))) )),
  inference(forward_demodulation,[],[f2103,f180])).
fof(f2103,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,happ(s(fun(X135,bool),happ(s(fun(X134,fun(X135,bool)),X136),s(X134,sK8(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134)))),s(X135,sK9(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134))))) | p(s(bool,happ(s(fun('type.pair.prod'(X134,X135),bool),happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136))),s('type.pair.prod'(X134,X135),X137))))) )),
  inference(forward_demodulation,[],[f2065,f183])).
fof(f2065,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun(X134,fun(X135,bool)),X136),s('type.pair.prod'(X134,X135),'const.pair.,_2'(s(X134,sK8(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134)),s(X135,sK9(happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136)),X135,X134))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X134,X135),bool),happ(s(fun(fun(X134,fun(X135,bool)),fun('type.pair.prod'(X134,X135),bool)),'const.pair.UNCURRY_0'),s(fun(X134,fun(X135,bool)),X136))),s('type.pair.prod'(X134,X135),X137))))) )),
  inference(superposition,[],[f169,f180])).
fof(f649,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X1)))) | s(bool,X0) != s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f173,f223])).
fof(f544,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X3)))) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f172,f223])).
fof(f546,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X9))),s(bool,X10)))) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f172,f223])).
fof(f549,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X16))),s(bool,X17)))) | s(bool,'T_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f172,f223])).
fof(f651,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X7))),s(bool,X8)))) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f173,f223])).
fof(f654,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X14))),s(bool,X15)))) | s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f173,f223])).
fof(f663,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X8))),s(bool,X7)))) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f173,f223])).
fof(f666,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X15))),s(bool,X14)))) | s(bool,'T_0') != s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f173,f223])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X1))),s(bool,'T_0')))) | s(bool,X0) != s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f173,f223])).
fof(f555,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X3))),s(bool,'T_0')))) | s(bool,X2) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f172,f223])).
fof(f3001,plain,(
  ( ! [X156,X154,X152,X151,X157,X155,X153,X150] : (p(s(bool,happ(s(fun(fun(X150,X151),bool),happ(s(fun(fun(X150,X151),fun(fun(X150,X151),bool)),'=_0'),s(fun(X150,X151),happ(s(fun(X154,fun(X150,X151)),'_dst_x00x11_2'(s(fun(X152,X151),X153),s(fun(X154,fun(X150,X152)),X155))),s(X154,X156))))),s(fun(X150,X151),X157)))) | s(fun(X150,X151),X157) != s(fun(X150,X151),'_dst_x0x1_2'(s(fun(X152,X151),X153),s(fun(X150,X152),happ(s(fun(X154,fun(X150,X152)),X155),s(X154,X156)))))) )),
  inference(superposition,[],[f173,f193])).
fof(f2999,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (~p(s(bool,happ(s(fun(fun(X134,X135),bool),happ(s(fun(fun(X134,X135),fun(fun(X134,X135),bool)),'=_0'),s(fun(X134,X135),happ(s(fun(X138,fun(X134,X135)),'_dst_x00x11_2'(s(fun(X136,X135),X137),s(fun(X138,fun(X134,X136)),X139))),s(X138,X140))))),s(fun(X134,X135),X141)))) | s(fun(X134,X135),X141) = s(fun(X134,X135),'_dst_x0x1_2'(s(fun(X136,X135),X137),s(fun(X134,X136),happ(s(fun(X138,fun(X134,X136)),X139),s(X138,X140)))))) )),
  inference(superposition,[],[f172,f193])).
fof(f3034,plain,(
  ( ! [X177,X182,X180,X178,X176,X183,X181,X179] : (p(s(bool,happ(s(fun(fun(X176,X177),bool),happ(s(fun(fun(X176,X177),fun(fun(X176,X177),bool)),'=_0'),s(fun(X176,X177),'_dst_x0x1_2'(s(fun(X179,X177),X180),s(fun(X176,X179),happ(s(fun(X178,fun(X176,X179)),X181),s(X178,X182))))))),s(fun(X176,X177),X183)))) | s(fun(X176,X177),X183) != s(fun(X176,X177),happ(s(fun(X178,fun(X176,X177)),'_dst_x00x11_2'(s(fun(X179,X177),X180),s(fun(X178,fun(X176,X179)),X181))),s(X178,X182)))) )),
  inference(superposition,[],[f173,f193])).
fof(f3032,plain,(
  ( ! [X167,X165,X163,X161,X166,X164,X162,X160] : (~p(s(bool,happ(s(fun(fun(X160,X161),bool),happ(s(fun(fun(X160,X161),fun(fun(X160,X161),bool)),'=_0'),s(fun(X160,X161),'_dst_x0x1_2'(s(fun(X163,X161),X164),s(fun(X160,X163),happ(s(fun(X162,fun(X160,X163)),X165),s(X162,X166))))))),s(fun(X160,X161),X167)))) | s(fun(X160,X161),X167) = s(fun(X160,X161),happ(s(fun(X162,fun(X160,X161)),'_dst_x00x11_2'(s(fun(X163,X161),X164),s(fun(X162,fun(X160,X163)),X165))),s(X162,X166)))) )),
  inference(superposition,[],[f172,f193])).
fof(f2916,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X3,X1),X4),s(fun(X2,fun(X0,X3)),X5))),s(X2,X6))))),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X0,X3),happ(s(fun(X2,fun(X0,X3)),X5),s(X2,X6))))))))) )),
  inference(unit_resulting_resolution,[],[f193,f173])).
fof(f3035,plain,(
  ( ! [X191,X189,X187,X185,X190,X188,X186,X184] : (p(s(bool,happ(s(fun(fun(X184,X185),bool),happ(s(fun(fun(X184,X185),fun(fun(X184,X185),bool)),'=_0'),s(fun(X184,X185),X191))),s(fun(X184,X185),'_dst_x0x1_2'(s(fun(X187,X185),X188),s(fun(X184,X187),happ(s(fun(X186,fun(X184,X187)),X189),s(X186,X190)))))))) | s(fun(X184,X185),X191) != s(fun(X184,X185),happ(s(fun(X186,fun(X184,X185)),'_dst_x00x11_2'(s(fun(X187,X185),X188),s(fun(X186,fun(X184,X187)),X189))),s(X186,X190)))) )),
  inference(superposition,[],[f173,f193])).
fof(f3033,plain,(
  ( ! [X175,X173,X171,X169,X174,X172,X170,X168] : (~p(s(bool,happ(s(fun(fun(X168,X169),bool),happ(s(fun(fun(X168,X169),fun(fun(X168,X169),bool)),'=_0'),s(fun(X168,X169),X175))),s(fun(X168,X169),'_dst_x0x1_2'(s(fun(X171,X169),X172),s(fun(X168,X171),happ(s(fun(X170,fun(X168,X171)),X173),s(X170,X174)))))))) | s(fun(X168,X169),X175) = s(fun(X168,X169),happ(s(fun(X170,fun(X168,X169)),'_dst_x00x11_2'(s(fun(X171,X169),X172),s(fun(X170,fun(X168,X171)),X173))),s(X170,X174)))) )),
  inference(superposition,[],[f172,f193])).
fof(f2915,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(fun(X0,X1),bool),happ(s(fun(fun(X0,X1),fun(fun(X0,X1),bool)),'=_0'),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X3),s(fun(X0,X2),happ(s(fun(X4,fun(X0,X2)),X5),s(X4,X6))))))),s(fun(X0,X1),happ(s(fun(X4,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X2,X1),X3),s(fun(X4,fun(X0,X2)),X5))),s(X4,X6))))))) )),
  inference(unit_resulting_resolution,[],[f193,f173])).
fof(f3002,plain,(
  ( ! [X158,X165,X163,X161,X159,X164,X162,X160] : (p(s(bool,happ(s(fun(fun(X158,X159),bool),happ(s(fun(fun(X158,X159),fun(fun(X158,X159),bool)),'=_0'),s(fun(X158,X159),X165))),s(fun(X158,X159),happ(s(fun(X162,fun(X158,X159)),'_dst_x00x11_2'(s(fun(X160,X159),X161),s(fun(X162,fun(X158,X160)),X163))),s(X162,X164)))))) | s(fun(X158,X159),X165) != s(fun(X158,X159),'_dst_x0x1_2'(s(fun(X160,X159),X161),s(fun(X158,X160),happ(s(fun(X162,fun(X158,X160)),X163),s(X162,X164)))))) )),
  inference(superposition,[],[f173,f193])).
fof(f3000,plain,(
  ( ! [X146,X144,X142,X149,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(fun(X142,X143),bool),happ(s(fun(fun(X142,X143),fun(fun(X142,X143),bool)),'=_0'),s(fun(X142,X143),X149))),s(fun(X142,X143),happ(s(fun(X146,fun(X142,X143)),'_dst_x00x11_2'(s(fun(X144,X143),X145),s(fun(X146,fun(X142,X144)),X147))),s(X146,X148)))))) | s(fun(X142,X143),X149) = s(fun(X142,X143),'_dst_x0x1_2'(s(fun(X144,X143),X145),s(fun(X142,X144),happ(s(fun(X146,fun(X142,X144)),X147),s(X146,X148)))))) )),
  inference(superposition,[],[f172,f193])).
fof(f1859,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),X10),s('type.pair.prod'(X8,X9),X11)))) | s(bool,X12) = s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),X10),s('type.pair.prod'(X8,X9),'const.pair.,_2'(s(X8,sK8(X10,X9,X8)),s(X9,sK9(X10,X9,X8)))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1846,f140])).
fof(f1846,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),X10),s('type.pair.prod'(X8,X9),X11)))) | s(bool,X12) = s(bool,happ(s(fun('type.pair.prod'(X8,X9),bool),X10),s('type.pair.prod'(X8,X9),'const.pair.,_2'(s(X8,sK8(X10,X9,X8)),s(X9,sK9(X10,X9,X8)))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f169,f223])).
fof(f1855,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,happ(s(fun('type.pair.prod'(X36,X37),bool),X39),s('type.pair.prod'(X36,X37),X41)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X36,X37),bool),X39),s('type.pair.prod'(X36,X37),'const.pair.,_2'(s(X36,sK8(happ(s(fun(X38,fun('type.pair.prod'(X36,X37),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X36,X37),bool),X39))),s(X38,X40)),X37,X36)),s(X37,sK9(happ(s(fun(X38,fun('type.pair.prod'(X36,X37),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X36,X37),bool),X39))),s(X38,X40)),X37,X36)))))))) )),
  inference(forward_demodulation,[],[f1831,f170])).
fof(f1831,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,happ(s(fun('type.pair.prod'(X36,X37),bool),X39),s('type.pair.prod'(X36,X37),'const.pair.,_2'(s(X36,sK8(happ(s(fun(X38,fun('type.pair.prod'(X36,X37),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X36,X37),bool),X39))),s(X38,X40)),X37,X36)),s(X37,sK9(happ(s(fun(X38,fun('type.pair.prod'(X36,X37),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X36,X37),bool),X39))),s(X38,X40)),X37,X36))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X36,X37),bool),happ(s(fun(X38,fun('type.pair.prod'(X36,X37),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X36,X37),bool),X39))),s(X38,X40))),s('type.pair.prod'(X36,X37),X41))))) )),
  inference(superposition,[],[f169,f170])).
fof(f1853,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0)),s(X1,sK9(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f1826,f159])).
fof(f1826,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0)),s(X1,sK9(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(superposition,[],[f169,f159])).
fof(f169,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(X2,X1,X0)),s(X1,sK9(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(X2,X1,X0))))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK8(X2,X1,X0)),s(X1,sK9(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X27] : ! [X29] : (! [X30] : p(s(bool,happ(s(fun('type.pair.prod'(X6,X27),bool),X29),s('type.pair.prod'(X6,X27),X30)))) <=> ! [X31,X32] : p(s(bool,happ(s(fun('type.pair.prod'(X6,X27),bool),X29),s('type.pair.prod'(X6,X27),'const.pair.,_2'(s(X6,X31),s(X27,X32)))))))),
  file('thm.closProps.evaluate_app_const','thm.pair.FORALL_PROD')).
fof(f942,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f934,f140])).
fof(f934,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f168,f223])).
fof(f941,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X8),s('type.pair.prod'(X5,X6),'const.pair.,_2'(s(X5,X10),s(X6,X11)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X8),s('type.pair.prod'(X5,X6),sK7(happ(s(fun(X7,fun('type.pair.prod'(X5,X6),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X5,X6),bool),X8))),s(X7,X9)),X6,X5)))))) )),
  inference(forward_demodulation,[],[f933,f170])).
fof(f933,plain,(
  ( ! [X6,X10,X8,X7,X5,X11,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),X8),s('type.pair.prod'(X5,X6),sK7(happ(s(fun(X7,fun('type.pair.prod'(X5,X6),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X5,X6),bool),X8))),s(X7,X9)),X6,X5))))) | p(s(bool,happ(s(fun('type.pair.prod'(X5,X6),bool),happ(s(fun(X7,fun('type.pair.prod'(X5,X6),bool)),'_dst_x0_1'(s(fun('type.pair.prod'(X5,X6),bool),X8))),s(X7,X9))),s('type.pair.prod'(X5,X6),'const.pair.,_2'(s(X5,X10),s(X6,X11))))))) )),
  inference(superposition,[],[f168,f170])).
fof(f940,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0)))))) )),
  inference(forward_demodulation,[],[f932,f159])).
fof(f932,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2)),X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),happ(s(fun(fun('type.pair.prod'(X0,X1),bool),fun('type.pair.prod'(X0,X1),bool)),'_i_0_0'),s(fun('type.pair.prod'(X0,X1),bool),X2))),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X4))))))) )),
  inference(superposition,[],[f168,f159])).
fof(f168,plain,(
  ( ! [X6,X2,X0,X7,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK7(X2,X1,X0))))) | p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f2994,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (~p(s(bool,happ(s(fun('type.pair.prod'(X92,X93),bool),happ(s(fun(X96,fun('type.pair.prod'(X92,X93),bool)),'_dst_x00x11_2'(s(fun(X94,bool),X95),s(fun(X96,fun('type.pair.prod'(X92,X93),X94)),X97))),s(X96,X98))),s('type.pair.prod'(X92,X93),'const.pair.,_2'(s(X92,sK8('_dst_x0x1_2'(s(fun(X94,bool),X95),s(fun('type.pair.prod'(X92,X93),X94),happ(s(fun(X96,fun('type.pair.prod'(X92,X93),X94)),X97),s(X96,X98)))),X93,X92)),s(X93,sK9('_dst_x0x1_2'(s(fun(X94,bool),X95),s(fun('type.pair.prod'(X92,X93),X94),happ(s(fun(X96,fun('type.pair.prod'(X92,X93),X94)),X97),s(X96,X98)))),X93,X92))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X92,X93),bool),'_dst_x0x1_2'(s(fun(X94,bool),X95),s(fun('type.pair.prod'(X92,X93),X94),happ(s(fun(X96,fun('type.pair.prod'(X92,X93),X94)),X97),s(X96,X98))))),s('type.pair.prod'(X92,X93),X99))))) )),
  inference(superposition,[],[f169,f193])).
fof(f2995,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,happ(s(fun('type.pair.prod'(X100,X101),bool),happ(s(fun(X104,fun('type.pair.prod'(X100,X101),bool)),'_dst_x00x11_2'(s(fun(X102,bool),X103),s(fun(X104,fun('type.pair.prod'(X100,X101),X102)),X105))),s(X104,X106))),s('type.pair.prod'(X100,X101),sK7('_dst_x0x1_2'(s(fun(X102,bool),X103),s(fun('type.pair.prod'(X100,X101),X102),happ(s(fun(X104,fun('type.pair.prod'(X100,X101),X102)),X105),s(X104,X106)))),X101,X100))))) | p(s(bool,happ(s(fun('type.pair.prod'(X100,X101),bool),'_dst_x0x1_2'(s(fun(X102,bool),X103),s(fun('type.pair.prod'(X100,X101),X102),happ(s(fun(X104,fun('type.pair.prod'(X100,X101),X102)),X105),s(X104,X106))))),s('type.pair.prod'(X100,X101),'const.pair.,_2'(s(X100,X107),s(X101,X108))))))) )),
  inference(superposition,[],[f168,f193])).
fof(f3025,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (~p(s(bool,happ(s(fun('type.pair.prod'(X100,X101),bool),'_dst_x0x1_2'(s(fun(X103,bool),X104),s(fun('type.pair.prod'(X100,X101),X103),happ(s(fun(X102,fun('type.pair.prod'(X100,X101),X103)),X105),s(X102,X106))))),s('type.pair.prod'(X100,X101),'const.pair.,_2'(s(X100,sK8(happ(s(fun(X102,fun('type.pair.prod'(X100,X101),bool)),'_dst_x00x11_2'(s(fun(X103,bool),X104),s(fun(X102,fun('type.pair.prod'(X100,X101),X103)),X105))),s(X102,X106)),X101,X100)),s(X101,sK9(happ(s(fun(X102,fun('type.pair.prod'(X100,X101),bool)),'_dst_x00x11_2'(s(fun(X103,bool),X104),s(fun(X102,fun('type.pair.prod'(X100,X101),X103)),X105))),s(X102,X106)),X101,X100))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X100,X101),bool),happ(s(fun(X102,fun('type.pair.prod'(X100,X101),bool)),'_dst_x00x11_2'(s(fun(X103,bool),X104),s(fun(X102,fun('type.pair.prod'(X100,X101),X103)),X105))),s(X102,X106))),s('type.pair.prod'(X100,X101),X107))))) )),
  inference(superposition,[],[f169,f193])).
fof(f3026,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X115,X113] : (~p(s(bool,happ(s(fun('type.pair.prod'(X108,X109),bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun('type.pair.prod'(X108,X109),X111),happ(s(fun(X110,fun('type.pair.prod'(X108,X109),X111)),X113),s(X110,X114))))),s('type.pair.prod'(X108,X109),sK7(happ(s(fun(X110,fun('type.pair.prod'(X108,X109),bool)),'_dst_x00x11_2'(s(fun(X111,bool),X112),s(fun(X110,fun('type.pair.prod'(X108,X109),X111)),X113))),s(X110,X114)),X109,X108))))) | p(s(bool,happ(s(fun('type.pair.prod'(X108,X109),bool),happ(s(fun(X110,fun('type.pair.prod'(X108,X109),bool)),'_dst_x00x11_2'(s(fun(X111,bool),X112),s(fun(X110,fun('type.pair.prod'(X108,X109),X111)),X113))),s(X110,X114))),s('type.pair.prod'(X108,X109),'const.pair.,_2'(s(X108,X115),s(X109,X116))))))) )),
  inference(superposition,[],[f168,f193])).
fof(f1115,plain,(
  ( ! [X187,X185,X182,X180,X186,X184,X183,X181,X179] : (~p(s(bool,happ(s(fun('type.pair.prod'(X179,X180),bool),happ(s(fun(X183,fun('type.pair.prod'(X179,X180),bool)),'_dst_x0x1_2'(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(fun(X183,X181),X184))),s(X183,X185))),s('type.pair.prod'(X179,X180),sK7(happ(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185)))),X180,X179))))) | p(s(bool,happ(s(fun('type.pair.prod'(X179,X180),bool),happ(s(fun(X181,fun('type.pair.prod'(X179,X180),bool)),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))),s('type.pair.prod'(X179,X180),'const.pair.,_2'(s(X179,X186),s(X180,X187))))))) )),
  inference(superposition,[],[f168,f182])).
fof(f1142,plain,(
  ( ! [X158,X156,X154,X152,X151,X157,X155,X153,X150] : (~p(s(bool,happ(s(fun('type.pair.prod'(X150,X151),bool),happ(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(X153,happ(s(fun(X152,X153),X155),s(X152,X156))))),s('type.pair.prod'(X150,X151),sK7(happ(s(fun(X152,fun('type.pair.prod'(X150,X151),bool)),'_dst_x0x1_2'(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(fun(X152,X153),X155))),s(X152,X156)),X151,X150))))) | p(s(bool,happ(s(fun('type.pair.prod'(X150,X151),bool),happ(s(fun(X152,fun('type.pair.prod'(X150,X151),bool)),'_dst_x0x1_2'(s(fun(X153,fun('type.pair.prod'(X150,X151),bool)),X154),s(fun(X152,X153),X155))),s(X152,X156))),s('type.pair.prod'(X150,X151),'const.pair.,_2'(s(X150,X157),s(X151,X158))))))) )),
  inference(superposition,[],[f168,f182])).
fof(f1647,plain,(
  ( ! [X243,X241,X239,X237,X244,X242,X240,X238,X236] : (~p(s(bool,happ(s(fun('type.pair.prod'(X236,X237),bool),happ(s(fun(X239,fun('type.pair.prod'(X236,X237),bool)),happ(s(fun(X238,fun(X239,fun('type.pair.prod'(X236,X237),bool))),X242),s(X238,X240))),s(X239,X241))),s('type.pair.prod'(X236,X237),sK7('const.pair.pair_CASE_2'(s('type.pair.prod'(X238,X239),'const.pair.,_2'(s(X238,X240),s(X239,X241))),s(fun(X238,fun(X239,fun('type.pair.prod'(X236,X237),bool))),X242)),X237,X236))))) | p(s(bool,happ(s(fun('type.pair.prod'(X236,X237),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X238,X239),'const.pair.,_2'(s(X238,X240),s(X239,X241))),s(fun(X238,fun(X239,fun('type.pair.prod'(X236,X237),bool))),X242))),s('type.pair.prod'(X236,X237),'const.pair.,_2'(s(X236,X243),s(X237,X244))))))) )),
  inference(superposition,[],[f168,f184])).
fof(f1670,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X186,X193] : (p(s(bool,happ(s(fun('type.pair.prod'(X186,X187),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X188,X189),'const.pair.,_2'(s(X188,X191),s(X189,X192))),s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190))),s('type.pair.prod'(X186,X187),'const.pair.,_2'(s(X186,X193),s(X187,X194)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X186,X187),bool),happ(s(fun(X189,fun('type.pair.prod'(X186,X187),bool)),happ(s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190),s(X188,X191))),s(X189,X192))),s('type.pair.prod'(X186,X187),sK7('const.pair.UNCURRY_2'(s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190),s('type.pair.prod'(X188,X189),'const.pair.,_2'(s(X188,X191),s(X189,X192)))),X187,X186)))))) )),
  inference(backward_demodulation,[],[f1532,f1313])).
fof(f1671,plain,(
  ( ! [X206,X204,X202,X200,X199,X207,X205,X203,X201] : (~p(s(bool,happ(s(fun('type.pair.prod'(X199,X200),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X202,X201),'const.pair.,_2'(s(X202,X204),s(X201,X205))),s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203))),s('type.pair.prod'(X199,X200),sK7(happ(s(fun(X201,fun('type.pair.prod'(X199,X200),bool)),happ(s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203),s(X202,X204))),s(X201,X205)),X200,X199))))) | p(s(bool,happ(s(fun('type.pair.prod'(X199,X200),bool),happ(s(fun(X201,fun('type.pair.prod'(X199,X200),bool)),happ(s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203),s(X202,X204))),s(X201,X205))),s('type.pair.prod'(X199,X200),'const.pair.,_2'(s(X199,X206),s(X200,X207))))))) )),
  inference(backward_demodulation,[],[f1532,f1283])).
fof(f1827,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(X8,fun('type.pair.prod'(X4,X5),bool)),'_dst_x0x1_2'(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(fun(X8,X6),X9))),s(X8,X10))),s('type.pair.prod'(X4,X5),'const.pair.,_2'(s(X4,sK8(happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))),X5,X4)),s(X5,sK9(happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10)))),X5,X4))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X4,X5),bool),happ(s(fun(X6,fun('type.pair.prod'(X4,X5),bool)),X7),s(X6,happ(s(fun(X8,X6),X9),s(X8,X10))))),s('type.pair.prod'(X4,X5),X11))))) )),
  inference(superposition,[],[f169,f182])).
fof(f1828,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (~p(s(bool,happ(s(fun('type.pair.prod'(X12,X13),bool),happ(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))),s('type.pair.prod'(X12,X13),'const.pair.,_2'(s(X12,sK8(happ(s(fun(X14,fun('type.pair.prod'(X12,X13),bool)),'_dst_x0x1_2'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17))),s(X14,X18)),X13,X12)),s(X13,sK9(happ(s(fun(X14,fun('type.pair.prod'(X12,X13),bool)),'_dst_x0x1_2'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17))),s(X14,X18)),X13,X12))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X12,X13),bool),happ(s(fun(X14,fun('type.pair.prod'(X12,X13),bool)),'_dst_x0x1_2'(s(fun(X15,fun('type.pair.prod'(X12,X13),bool)),X16),s(fun(X14,X15),X17))),s(X14,X18))),s('type.pair.prod'(X12,X13),X19))))) )),
  inference(superposition,[],[f169,f182])).
fof(f1833,plain,(
  ( ! [X57,X54,X52,X50,X56,X55,X53,X51] : (~p(s(bool,happ(s(fun('type.pair.prod'(X50,X51),bool),happ(s(fun(X53,fun('type.pair.prod'(X50,X51),bool)),happ(s(fun(X52,fun(X53,fun('type.pair.prod'(X50,X51),bool))),X56),s(X52,X54))),s(X53,X55))),s('type.pair.prod'(X50,X51),'const.pair.,_2'(s(X50,sK8('const.pair.pair_CASE_2'(s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X54),s(X53,X55))),s(fun(X52,fun(X53,fun('type.pair.prod'(X50,X51),bool))),X56)),X51,X50)),s(X51,sK9('const.pair.pair_CASE_2'(s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X54),s(X53,X55))),s(fun(X52,fun(X53,fun('type.pair.prod'(X50,X51),bool))),X56)),X51,X50))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X50,X51),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X52,X53),'const.pair.,_2'(s(X52,X54),s(X53,X55))),s(fun(X52,fun(X53,fun('type.pair.prod'(X50,X51),bool))),X56))),s('type.pair.prod'(X50,X51),X57))))) )),
  inference(superposition,[],[f169,f184])).
fof(f1856,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48))),s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46))),s('type.pair.prod'(X42,X43),X49)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun(X45,fun('type.pair.prod'(X42,X43),bool)),happ(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s(X44,X47))),s(X45,X48))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK8('const.pair.UNCURRY_2'(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48)))),X43,X42)),s(X43,sK9('const.pair.UNCURRY_2'(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48)))),X43,X42)))))))) )),
  inference(forward_demodulation,[],[f1832,f1532])).
fof(f1832,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42,X49] : (~p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),happ(s(fun(X45,fun('type.pair.prod'(X42,X43),bool)),happ(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s(X44,X47))),s(X45,X48))),s('type.pair.prod'(X42,X43),'const.pair.,_2'(s(X42,sK8('const.pair.UNCURRY_2'(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48)))),X43,X42)),s(X43,sK9('const.pair.UNCURRY_2'(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48)))),X43,X42))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X42,X43),bool),'const.pair.UNCURRY_2'(s(fun(X44,fun(X45,fun('type.pair.prod'(X42,X43),bool))),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48))))),s('type.pair.prod'(X42,X43),X49))))) )),
  inference(superposition,[],[f169,f183])).
fof(f1829,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (~p(s(bool,happ(s(fun('type.pair.prod'(X20,X21),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X22),'const.pair.,_2'(s(X23,X25),s(X22,X26))),s(fun(X23,fun(X22,fun('type.pair.prod'(X20,X21),bool))),X24))),s('type.pair.prod'(X20,X21),'const.pair.,_2'(s(X20,sK8(happ(s(fun(X22,fun('type.pair.prod'(X20,X21),bool)),happ(s(fun(X23,fun(X22,fun('type.pair.prod'(X20,X21),bool))),X24),s(X23,X25))),s(X22,X26)),X21,X20)),s(X21,sK9(happ(s(fun(X22,fun('type.pair.prod'(X20,X21),bool)),happ(s(fun(X23,fun(X22,fun('type.pair.prod'(X20,X21),bool))),X24),s(X23,X25))),s(X22,X26)),X21,X20))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X20,X21),bool),happ(s(fun(X22,fun('type.pair.prod'(X20,X21),bool)),happ(s(fun(X23,fun(X22,fun('type.pair.prod'(X20,X21),bool))),X24),s(X23,X25))),s(X22,X26))),s('type.pair.prod'(X20,X21),X27))))) )),
  inference(superposition,[],[f169,f184])).
fof(f2431,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,happ(s(fun('type.pair.prod'(X79,X80),bool),'_dst_x0x1x0_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s('type.pair.prod'(X79,X80),X88)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X79,X80),bool),'_dst_x0x1x0_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s('type.pair.prod'(X79,X80),'const.pair.,_2'(s(X79,sK8(happ(s(fun(X81,fun('type.pair.prod'(X79,X80),bool)),'_dst_x00x01x00_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s(X81,X87)),X80,X79)),s(X80,sK9(happ(s(fun(X81,fun('type.pair.prod'(X79,X80),bool)),'_dst_x00x01x00_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s(X81,X87)),X80,X79)))))))) )),
  inference(forward_demodulation,[],[f2398,f198])).
fof(f2398,plain,(
  ( ! [X80,X88,X87,X85,X83,X81,X79,X86,X84,X82] : (~p(s(bool,happ(s(fun('type.pair.prod'(X79,X80),bool),'_dst_x0x1x0_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s('type.pair.prod'(X79,X80),'const.pair.,_2'(s(X79,sK8(happ(s(fun(X81,fun('type.pair.prod'(X79,X80),bool)),'_dst_x00x01x00_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s(X81,X87)),X80,X79)),s(X80,sK9(happ(s(fun(X81,fun('type.pair.prod'(X79,X80),bool)),'_dst_x00x01x00_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s(X81,X87)),X80,X79))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X79,X80),bool),happ(s(fun(X81,fun('type.pair.prod'(X79,X80),bool)),'_dst_x00x01x00_3'(s(fun(X82,fun(X83,bool)),X84),s(fun('type.pair.prod'(X79,X80),X82),X85),s(X83,X86))),s(X81,X87))),s('type.pair.prod'(X79,X80),X88))))) )),
  inference(superposition,[],[f169,f198])).
fof(f2432,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X98,X96] : (p(s(bool,happ(s(fun('type.pair.prod'(X89,X90),bool),'_dst_x0x1x0_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s('type.pair.prod'(X89,X90),'const.pair.,_2'(s(X89,X98),s(X90,X99)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X89,X90),bool),'_dst_x0x1x0_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s('type.pair.prod'(X89,X90),sK7(happ(s(fun(X91,fun('type.pair.prod'(X89,X90),bool)),'_dst_x00x01x00_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s(X91,X97)),X90,X89)))))) )),
  inference(forward_demodulation,[],[f2399,f198])).
fof(f2399,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X98,X96] : (~p(s(bool,happ(s(fun('type.pair.prod'(X89,X90),bool),'_dst_x0x1x0_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s('type.pair.prod'(X89,X90),sK7(happ(s(fun(X91,fun('type.pair.prod'(X89,X90),bool)),'_dst_x00x01x00_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s(X91,X97)),X90,X89))))) | p(s(bool,happ(s(fun('type.pair.prod'(X89,X90),bool),happ(s(fun(X91,fun('type.pair.prod'(X89,X90),bool)),'_dst_x00x01x00_3'(s(fun(X92,fun(X93,bool)),X94),s(fun('type.pair.prod'(X89,X90),X92),X95),s(X93,X96))),s(X91,X97))),s('type.pair.prod'(X89,X90),'const.pair.,_2'(s(X89,X98),s(X90,X99))))))) )),
  inference(superposition,[],[f168,f198])).
fof(f2727,plain,(
  ( ! [X344,X346,X341,X343,X345,X340,X342] : (p(s(bool,happ(s(fun('type.pair.prod'(X340,X341),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X342,X343),X344),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345))),s('type.pair.prod'(X340,X341),X346)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X340,X341),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X342,X343),X344),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345))),s('type.pair.prod'(X340,X341),'const.pair.,_2'(s(X340,sK8(happ(s(fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool)),happ(s(fun('type.pair.prod'(X342,X343),fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X342,X343),X344))),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345)),X341,X340)),s(X341,sK9(happ(s(fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool)),happ(s(fun('type.pair.prod'(X342,X343),fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X342,X343),X344))),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345)),X341,X340)))))))) )),
  inference(forward_demodulation,[],[f2707,f181])).
fof(f2707,plain,(
  ( ! [X344,X346,X341,X343,X345,X340,X342] : (~p(s(bool,happ(s(fun('type.pair.prod'(X340,X341),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X342,X343),X344),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345))),s('type.pair.prod'(X340,X341),'const.pair.,_2'(s(X340,sK8(happ(s(fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool)),happ(s(fun('type.pair.prod'(X342,X343),fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X342,X343),X344))),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345)),X341,X340)),s(X341,sK9(happ(s(fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool)),happ(s(fun('type.pair.prod'(X342,X343),fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X342,X343),X344))),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345)),X341,X340))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X340,X341),bool),happ(s(fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool)),happ(s(fun('type.pair.prod'(X342,X343),fun(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),fun('type.pair.prod'(X340,X341),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X342,X343),X344))),s(fun(X342,fun(X343,fun('type.pair.prod'(X340,X341),bool))),X345))),s('type.pair.prod'(X340,X341),X346))))) )),
  inference(superposition,[],[f169,f181])).
fof(f2728,plain,(
  ( ! [X350,X352,X354,X349,X351,X347,X353,X348] : (p(s(bool,happ(s(fun('type.pair.prod'(X347,X348),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,X350),X351),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352))),s('type.pair.prod'(X347,X348),'const.pair.,_2'(s(X347,X353),s(X348,X354)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X347,X348),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,X350),X351),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352))),s('type.pair.prod'(X347,X348),sK7(happ(s(fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool)),happ(s(fun('type.pair.prod'(X349,X350),fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X349,X350),X351))),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352)),X348,X347)))))) )),
  inference(forward_demodulation,[],[f2708,f181])).
fof(f2708,plain,(
  ( ! [X350,X352,X354,X349,X351,X347,X353,X348] : (~p(s(bool,happ(s(fun('type.pair.prod'(X347,X348),bool),'const.pair.pair_CASE_2'(s('type.pair.prod'(X349,X350),X351),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352))),s('type.pair.prod'(X347,X348),sK7(happ(s(fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool)),happ(s(fun('type.pair.prod'(X349,X350),fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X349,X350),X351))),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352)),X348,X347))))) | p(s(bool,happ(s(fun('type.pair.prod'(X347,X348),bool),happ(s(fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool)),happ(s(fun('type.pair.prod'(X349,X350),fun(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),fun('type.pair.prod'(X347,X348),bool))),'const.pair.pair_CASE_0'),s('type.pair.prod'(X349,X350),X351))),s(fun(X349,fun(X350,fun('type.pair.prod'(X347,X348),bool))),X352))),s('type.pair.prod'(X347,X348),'const.pair.,_2'(s(X347,X353),s(X348,X354))))))) )),
  inference(superposition,[],[f168,f181])).
fof(f2113,plain,(
  ( ! [X280,X282,X277,X279,X281,X276,X278] : (p(s(bool,happ(s(fun('type.pair.prod'(X276,X277),bool),'const.pair.UNCURRY_2'(s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280),s('type.pair.prod'(X278,X279),X281))),s('type.pair.prod'(X276,X277),X282)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X276,X277),bool),'const.pair.UNCURRY_2'(s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280),s('type.pair.prod'(X278,X279),X281))),s('type.pair.prod'(X276,X277),'const.pair.,_2'(s(X276,sK8(happ(s(fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool)),happ(s(fun(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool))),'const.pair.UNCURRY_0'),s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280))),s('type.pair.prod'(X278,X279),X281)),X277,X276)),s(X277,sK9(happ(s(fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool)),happ(s(fun(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool))),'const.pair.UNCURRY_0'),s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280))),s('type.pair.prod'(X278,X279),X281)),X277,X276)))))))) )),
  inference(forward_demodulation,[],[f2090,f180])).
fof(f2090,plain,(
  ( ! [X280,X282,X277,X279,X281,X276,X278] : (~p(s(bool,happ(s(fun('type.pair.prod'(X276,X277),bool),'const.pair.UNCURRY_2'(s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280),s('type.pair.prod'(X278,X279),X281))),s('type.pair.prod'(X276,X277),'const.pair.,_2'(s(X276,sK8(happ(s(fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool)),happ(s(fun(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool))),'const.pair.UNCURRY_0'),s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280))),s('type.pair.prod'(X278,X279),X281)),X277,X276)),s(X277,sK9(happ(s(fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool)),happ(s(fun(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool))),'const.pair.UNCURRY_0'),s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280))),s('type.pair.prod'(X278,X279),X281)),X277,X276))))))) | p(s(bool,happ(s(fun('type.pair.prod'(X276,X277),bool),happ(s(fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool)),happ(s(fun(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),fun('type.pair.prod'(X278,X279),fun('type.pair.prod'(X276,X277),bool))),'const.pair.UNCURRY_0'),s(fun(X278,fun(X279,fun('type.pair.prod'(X276,X277),bool))),X280))),s('type.pair.prod'(X278,X279),X281))),s('type.pair.prod'(X276,X277),X282))))) )),
  inference(superposition,[],[f169,f180])).
fof(f2114,plain,(
  ( ! [X288,X290,X285,X287,X283,X289,X284,X286] : (p(s(bool,happ(s(fun('type.pair.prod'(X283,X284),bool),'const.pair.UNCURRY_2'(s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287),s('type.pair.prod'(X285,X286),X288))),s('type.pair.prod'(X283,X284),'const.pair.,_2'(s(X283,X289),s(X284,X290)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X283,X284),bool),'const.pair.UNCURRY_2'(s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287),s('type.pair.prod'(X285,X286),X288))),s('type.pair.prod'(X283,X284),sK7(happ(s(fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool)),happ(s(fun(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool))),'const.pair.UNCURRY_0'),s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287))),s('type.pair.prod'(X285,X286),X288)),X284,X283)))))) )),
  inference(forward_demodulation,[],[f2091,f180])).
fof(f2091,plain,(
  ( ! [X288,X290,X285,X287,X283,X289,X284,X286] : (~p(s(bool,happ(s(fun('type.pair.prod'(X283,X284),bool),'const.pair.UNCURRY_2'(s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287),s('type.pair.prod'(X285,X286),X288))),s('type.pair.prod'(X283,X284),sK7(happ(s(fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool)),happ(s(fun(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool))),'const.pair.UNCURRY_0'),s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287))),s('type.pair.prod'(X285,X286),X288)),X284,X283))))) | p(s(bool,happ(s(fun('type.pair.prod'(X283,X284),bool),happ(s(fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool)),happ(s(fun(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),fun('type.pair.prod'(X285,X286),fun('type.pair.prod'(X283,X284),bool))),'const.pair.UNCURRY_0'),s(fun(X285,fun(X286,fun('type.pair.prod'(X283,X284),bool))),X287))),s('type.pair.prod'(X285,X286),X288))),s('type.pair.prod'(X283,X284),'const.pair.,_2'(s(X283,X289),s(X284,X290))))))) )),
  inference(superposition,[],[f168,f180])).
fof(f1839,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,sK8(X0,bool,X1)),s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.pair.prod'(X1,bool),bool),X0),s('type.pair.prod'(X1,bool),X2)))) | s(bool,X3) = s(bool,sK9(X0,bool,X1)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f169,f223])).
fof(f1841,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(X9,bool),bool),X8),s('type.pair.prod'(X9,bool),'const.pair.,_2'(s(X9,sK8(X8,bool,X9)),s(bool,X10)))))) | p(s(bool,happ(s(fun('type.pair.prod'(X9,bool),bool),X8),s('type.pair.prod'(X9,bool),X11)))) | s(bool,'T_0') = s(bool,sK9(X8,bool,X9)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f169,f223])).
fof(f1834,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,sK9(X0,X1,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X1),bool),X0),s('type.pair.prod'(bool,X1),X2)))) | s(bool,X3) = s(bool,sK8(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f169,f223])).
fof(f1836,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),X8),s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,sK9(X8,X9,bool))))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),X8),s('type.pair.prod'(bool,X9),X11)))) | s(bool,'T_0') = s(bool,sK8(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f169,f223])).
fof(f221,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X1),s('type.closSem.v',X2),s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),X4))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X4))))))))))) )),
  inference(forward_demodulation,[],[f167,f160])).
fof(f167,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X1),s('type.closSem.v',X2),s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),X4))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X4))))))))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4] : p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X1),s('type.closSem.v',X2),s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),X4))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X4))))))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X1),s('type.closSem.v',X2),s('type.list.list'('type.closSem.v'),X3),s('type.closSem.state'(X0),X4))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X0),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X4))))))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X21] : ! [X22,X23,X24,X25] : p(s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X22),s('type.closSem.v',X23),s('type.list.list'('type.closSem.v'),X24),s('type.closSem.state'(X21),X25))),s(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X21),bool)),'_dst_x0_1'(s(fun('type.closSem.state'(X21),bool),'_dst_x0x1x0_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X21),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X21),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X21),X25))))))))))),
  file('thm.closProps.evaluate_app_const','thm.closProps.closProps_59_prove_ltb|split|1')).
fof(f2654,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X23),X24),s(fun(X22,fun(X23,bool)),X25))) | s(bool,X26) = s(bool,happ(s(fun(fun(X22,fun(X23,bool)),bool),happ(s(fun('type.pair.prod'(X22,X23),fun(fun(X22,fun(X23,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X22,X23),X24))),s(fun(X22,fun(X23,bool)),X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f181,f223])).
fof(f2656,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,X36) = s(bool,'const.pair.pair_CASE_2'(s('type.pair.prod'(X32,X33),X34),s(fun(X32,fun(X33,bool)),X35))) | s(bool,'T_0') = s(bool,happ(s(fun(fun(X32,fun(X33,bool)),bool),happ(s(fun('type.pair.prod'(X32,X33),fun(fun(X32,fun(X33,bool)),bool)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X32,X33),X34))),s(fun(X32,fun(X33,bool)),X35))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f181,f223])).
fof(f222,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'_dst_x0x1_2'(s(fun(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool)),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x111x001_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),bool)),'const.pair.pair_CASE_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))))),'_dst_x000x111_2'(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))))),'_dst_x000x100x011_3'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))))),'const.pair.,_0'),s(fun('type.list.list'('type.closLang.exp'),'type.list.list'('type.closLang.exp')),'_i_0_0'),s(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))),'const.pair.,_0'))))),s(fun('type.closSem.state'(X0),fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x001x100_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))) )),
  inference(backward_demodulation,[],[f180,f162])).
fof(f2042,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun(X22,fun(X23,bool)),X24),s('type.pair.prod'(X22,X23),X25))) | s(bool,X26) = s(bool,happ(s(fun('type.pair.prod'(X22,X23),bool),happ(s(fun(fun(X22,fun(X23,bool)),fun('type.pair.prod'(X22,X23),bool)),'const.pair.UNCURRY_0'),s(fun(X22,fun(X23,bool)),X24))),s('type.pair.prod'(X22,X23),X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f180,f223])).
fof(f2044,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,X36) = s(bool,'const.pair.UNCURRY_2'(s(fun(X32,fun(X33,bool)),X34),s('type.pair.prod'(X32,X33),X35))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(X32,X33),bool),happ(s(fun(fun(X32,fun(X33,bool)),fun('type.pair.prod'(X32,X33),bool)),'const.pair.UNCURRY_0'),s(fun(X32,fun(X33,bool)),X34))),s('type.pair.prod'(X32,X33),X35))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f180,f223])).
fof(f3027,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119,X117] : (s(X124,happ(s(fun(fun(X117,X118),X124),X125),s(fun(X117,X118),'_dst_x0x1_2'(s(fun(X120,X118),X121),s(fun(X117,X120),happ(s(fun(X119,fun(X117,X120)),X122),s(X119,X123))))))) = s(X124,happ(s(fun(X119,X124),'_dst_x0x1_2'(s(fun(fun(X117,X118),X124),X125),s(fun(X119,fun(X117,X118)),'_dst_x00x11_2'(s(fun(X120,X118),X121),s(fun(X119,fun(X117,X120)),X122))))),s(X119,X123)))) )),
  inference(superposition,[],[f182,f193])).
fof(f3016,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X9,happ(s(fun(X8,X9),'_dst_x0x1_2'(s(fun(X11,X9),X12),s(fun(X8,X11),happ(s(fun(X10,fun(X8,X11)),X13),s(X10,X14))))),s(X8,X15))) = s(X9,'const.pair.UNCURRY_2'(s(fun(X10,fun(X8,X9)),'_dst_x00x11_2'(s(fun(X11,X9),X12),s(fun(X10,fun(X8,X11)),X13))),s('type.pair.prod'(X10,X8),'const.pair.,_2'(s(X10,X14),s(X8,X15)))))) )),
  inference(superposition,[],[f183,f193])).
fof(f3015,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X3,X1),X4),s(fun(X0,X3),happ(s(fun(X2,fun(X0,X3)),X5),s(X2,X6))))),s(X0,X7))) = s(X1,'const.pair.pair_CASE_2'(s('type.pair.prod'(X2,X0),'const.pair.,_2'(s(X2,X6),s(X0,X7))),s(fun(X2,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X3,X1),X4),s(fun(X2,fun(X0,X3)),X5)))))) )),
  inference(superposition,[],[f184,f193])).
fof(f2985,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X3),s(X2,happ(s(fun(X0,X2),happ(s(fun(X4,fun(X0,X2)),X5),s(X4,X6))),s(X0,X7))))) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X4,fun(X0,X1)),'_dst_x00x11_2'(s(fun(X2,X1),X3),s(fun(X4,fun(X0,X2)),X5))),s(X4,X6))),s(X0,X7)))) )),
  inference(superposition,[],[f182,f193])).
fof(f2984,plain,(
  ( ! [X14,X12,X10,X8,X7,X15,X13,X11,X9] : (s(fun(X7,X8),'_dst_x0x1_2'(s(fun(X10,X8),X11),s(fun(X7,X10),happ(s(fun(X9,fun(X7,X10)),X12),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15))))))) = s(fun(X7,X8),happ(s(fun(X13,fun(X7,X8)),'_dst_x0x1_2'(s(fun(X9,fun(X7,X8)),'_dst_x00x11_2'(s(fun(X10,X8),X11),s(fun(X9,fun(X7,X10)),X12))),s(fun(X13,X9),X14))),s(X13,X15)))) )),
  inference(superposition,[],[f193,f182])).
fof(f2924,plain,(
  ( ! [X80,X78,X76,X83,X81,X79,X77,X82] : (s(fun(X76,X82),'_dst_x0x1_2'(s(fun(X77,X82),X83),s(fun(X76,X77),'const.pair.UNCURRY_2'(s(fun(X78,fun(X79,fun(X76,X77))),X80),s('type.pair.prod'(X78,X79),X81))))) = s(fun(X76,X82),happ(s(fun('type.pair.prod'(X78,X79),fun(X76,X82)),'_dst_x00x11_2'(s(fun(X77,X82),X83),s(fun('type.pair.prod'(X78,X79),fun(X76,X77)),happ(s(fun(fun(X78,fun(X79,fun(X76,X77))),fun('type.pair.prod'(X78,X79),fun(X76,X77))),'const.pair.UNCURRY_0'),s(fun(X78,fun(X79,fun(X76,X77))),X80))))),s('type.pair.prod'(X78,X79),X81)))) )),
  inference(superposition,[],[f193,f180])).
fof(f2923,plain,(
  ( ! [X61,X66,X64,X62,X67,X65,X63] : (s(fun(X61,X66),'_dst_x0x1_2'(s(fun(X62,X66),X67),s(fun(X61,X62),X64))) = s(fun(X61,X66),happ(s(fun(X63,fun(X61,X66)),'_dst_x00x11_2'(s(fun(X62,X66),X67),s(fun(X63,fun(X61,X62)),'_dst_x0_1'(s(fun(X61,X62),X64))))),s(X63,X65)))) )),
  inference(superposition,[],[f193,f170])).
fof(f2921,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X51,X49] : (s(fun(X43,X50),'_dst_x0x1_2'(s(fun(X44,X50),X51),s(fun(X43,X44),'const.pair.pair_CASE_2'(s('type.pair.prod'(X46,X45),'const.pair.,_2'(s(X46,X48),s(X45,X49))),s(fun(X46,fun(X45,fun(X43,X44))),X47))))) = s(fun(X43,X50),happ(s(fun(X45,fun(X43,X50)),'_dst_x00x11_2'(s(fun(X44,X50),X51),s(fun(X45,fun(X43,X44)),happ(s(fun(X46,fun(X45,fun(X43,X44))),X47),s(X46,X48))))),s(X45,X49)))) )),
  inference(superposition,[],[f193,f184])).
fof(f2920,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (s(fun(X34,X41),'_dst_x0x1_2'(s(fun(X35,X41),X42),s(fun(X34,X35),happ(s(fun(X37,fun(X34,X35)),X38),s(X37,happ(s(fun(X36,X37),X39),s(X36,X40))))))) = s(fun(X34,X41),happ(s(fun(X36,fun(X34,X41)),'_dst_x00x11_2'(s(fun(X35,X41),X42),s(fun(X36,fun(X34,X35)),'_dst_x0x1_2'(s(fun(X37,fun(X34,X35)),X38),s(fun(X36,X37),X39))))),s(X36,X40)))) )),
  inference(superposition,[],[f193,f182])).
fof(f2919,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X25,X32] : (s(fun(X25,X32),happ(s(fun(X27,fun(X25,X32)),'_dst_x00x11_2'(s(fun(X26,X32),X33),s(fun(X27,fun(X25,X26)),X28))),s(X27,happ(s(fun(X29,X27),X30),s(X29,X31))))) = s(fun(X25,X32),'_dst_x0x1_2'(s(fun(X26,X32),X33),s(fun(X25,X26),happ(s(fun(X29,fun(X25,X26)),'_dst_x0x1_2'(s(fun(X27,fun(X25,X26)),X28),s(fun(X29,X27),X30))),s(X29,X31)))))) )),
  inference(superposition,[],[f193,f182])).
fof(f2918,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X11,X18,X16] : (s(fun(X11,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X11,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X11,X14),happ(s(fun(X13,fun(X11,X14)),X16),s(X13,X17))))))) = s(fun(X11,X18),happ(s(fun(X13,fun(X11,X18)),'_dst_x00x11_2'(s(fun(X12,X18),X19),s(fun(X13,fun(X11,X12)),'_dst_x00x11_2'(s(fun(X14,X12),X15),s(fun(X13,fun(X11,X14)),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f193,f193])).
fof(f2917,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X0,X9),'_dst_x0x1_2'(s(fun(X1,X9),X10),s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))))) = s(fun(X0,X9),happ(s(fun(X2,fun(X0,X9)),'_dst_x00x11_2'(s(fun(X1,X9),X10),s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))))),s(X2,X8)))) )),
  inference(superposition,[],[f193,f198])).
fof(f2653,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X15,'const.pair.pair_CASE_2'(s('type.pair.prod'(X16,X17),X18),s(fun(X16,fun(X17,X15)),happ(s(fun(X19,fun(X16,fun(X17,X15))),X20),s(X19,X21))))) = s(X15,happ(s(fun(X19,X15),'_dst_x0x1_2'(s(fun(fun(X16,fun(X17,X15)),X15),happ(s(fun('type.pair.prod'(X16,X17),fun(fun(X16,fun(X17,X15)),X15)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X16,X17),X18))),s(fun(X19,fun(X16,fun(X17,X15))),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f181,f182])).
fof(f2652,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(X10,'const.pair.pair_CASE_2'(s('type.pair.prod'(X11,X12),X13),s(fun(X11,fun(X12,X10)),X14))) = s(X10,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X11,X12),fun(fun(X11,fun(X12,X10)),X10)),'const.pair.pair_CASE_0'),s('type.pair.prod'('type.pair.prod'(X11,X12),fun(X11,fun(X12,X10))),'const.pair.,_2'(s('type.pair.prod'(X11,X12),X13),s(fun(X11,fun(X12,X10)),X14)))))) )),
  inference(superposition,[],[f181,f183])).
fof(f2651,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X7),X8),s(fun(X6,fun(X7,X5)),X9))) = s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X6,X7),fun(X6,fun(X7,X5))),'const.pair.,_2'(s('type.pair.prod'(X6,X7),X8),s(fun(X6,fun(X7,X5)),X9))),s(fun('type.pair.prod'(X6,X7),fun(fun(X6,fun(X7,X5)),X5)),'const.pair.pair_CASE_0')))) )),
  inference(superposition,[],[f181,f184])).
fof(f2400,plain,(
  ( ! [X103,X101,X109,X107,X105,X102,X100,X110,X108,X106,X104] : (s(X109,happ(s(fun(fun(X100,X101),X109),X110),s(fun(X100,X101),'_dst_x0x1x0_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X100,X103),X106),s(X104,X107))))) = s(X109,happ(s(fun(X102,X109),'_dst_x0x1_2'(s(fun(fun(X100,X101),X109),X110),s(fun(X102,fun(X100,X101)),'_dst_x00x01x00_3'(s(fun(X103,fun(X104,X101)),X105),s(fun(X100,X103),X106),s(X104,X107))))),s(X102,X108)))) )),
  inference(superposition,[],[f182,f198])).
fof(f2392,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),'_dst_x0x1x0_3'(s(fun(X13,fun(X14,X11)),X15),s(fun(X10,X13),X16),s(X14,X17))),s(X10,X19))) = s(X11,'const.pair.UNCURRY_2'(s(fun(X12,fun(X10,X11)),'_dst_x00x01x00_3'(s(fun(X13,fun(X14,X11)),X15),s(fun(X10,X13),X16),s(X14,X17))),s('type.pair.prod'(X12,X10),'const.pair.,_2'(s(X12,X18),s(X10,X19)))))) )),
  inference(superposition,[],[f183,f198])).
fof(f2391,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7))),s(X0,X9))) = s(X1,'const.pair.pair_CASE_2'(s('type.pair.prod'(X2,X0),'const.pair.,_2'(s(X2,X8),s(X0,X9))),s(fun(X2,fun(X0,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X0,X3),X6),s(X4,X7)))))) )),
  inference(superposition,[],[f184,f198])).
fof(f2390,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X9,X18,X16] : (s(fun(X9,X10),'_dst_x0x1x0_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),X15),s(X13,X16))) = s(fun(X9,X10),happ(s(fun(X17,fun(X9,X10)),'_dst_x0x1_2'(s(fun(X11,fun(X9,X10)),'_dst_x00x01x00_3'(s(fun(X12,fun(X13,X10)),X14),s(fun(X9,X12),X15),s(X13,X16))),s(fun(X17,X11),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f198,f182])).
fof(f2085,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234] : (s(X234,happ(s(fun(X233,X234),'const.pair.UNCURRY_2'(s(fun(X235,fun(X236,fun(X233,X234))),X237),s('type.pair.prod'(X235,X236),X238))),s(X233,X239))) = s(X234,'const.pair.UNCURRY_2'(s(fun('type.pair.prod'(X235,X236),fun(X233,X234)),happ(s(fun(fun(X235,fun(X236,fun(X233,X234))),fun('type.pair.prod'(X235,X236),fun(X233,X234))),'const.pair.UNCURRY_0'),s(fun(X235,fun(X236,fun(X233,X234))),X237))),s('type.pair.prod'('type.pair.prod'(X235,X236),X233),'const.pair.,_2'(s('type.pair.prod'(X235,X236),X238),s(X233,X239)))))) )),
  inference(superposition,[],[f183,f180])).
fof(f2084,plain,(
  ( ! [X231,X229,X227,X230,X228,X226,X232] : (s(X227,happ(s(fun(X226,X227),'const.pair.UNCURRY_2'(s(fun(X228,fun(X229,fun(X226,X227))),X230),s('type.pair.prod'(X228,X229),X231))),s(X226,X232))) = s(X227,'const.pair.pair_CASE_2'(s('type.pair.prod'('type.pair.prod'(X228,X229),X226),'const.pair.,_2'(s('type.pair.prod'(X228,X229),X231),s(X226,X232))),s(fun('type.pair.prod'(X228,X229),fun(X226,X227)),happ(s(fun(fun(X228,fun(X229,fun(X226,X227))),fun('type.pair.prod'(X228,X229),fun(X226,X227))),'const.pair.UNCURRY_0'),s(fun(X228,fun(X229,fun(X226,X227))),X230)))))) )),
  inference(superposition,[],[f184,f180])).
fof(f2047,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X5,happ(s(fun(X0,X5),X6),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))))) = s(X5,happ(s(fun('type.pair.prod'(X1,X2),X5),'_dst_x0x1_2'(s(fun(X0,X5),X6),s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(superposition,[],[f182,f180])).
fof(f2041,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X15,'const.pair.UNCURRY_2'(s(fun(X16,fun(X17,X15)),X18),s('type.pair.prod'(X16,X17),happ(s(fun(X19,'type.pair.prod'(X16,X17)),X20),s(X19,X21))))) = s(X15,happ(s(fun(X19,X15),'_dst_x0x1_2'(s(fun('type.pair.prod'(X16,X17),X15),happ(s(fun(fun(X16,fun(X17,X15)),fun('type.pair.prod'(X16,X17),X15)),'const.pair.UNCURRY_0'),s(fun(X16,fun(X17,X15)),X18))),s(fun(X19,'type.pair.prod'(X16,X17)),X20))),s(X19,X21)))) )),
  inference(superposition,[],[f180,f182])).
fof(f1533,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X12,'const.pair.pair_CASE_2'(s('type.pair.prod'(X14,X13),'const.pair.,_2'(s(X14,X16),s(X13,happ(s(fun(X17,X13),X18),s(X17,X19))))),s(fun(X14,fun(X13,X12)),X15))) = s(X12,happ(s(fun(X17,X12),'_dst_x0x1_2'(s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),X15),s(X14,X16))),s(fun(X17,X13),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f184,f182])).
fof(f1669,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X21,happ(s(fun(X20,X21),'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X22),'const.pair.,_2'(s(X23,X25),s(X22,X26))),s(fun(X23,fun(X22,fun(X20,X21))),X24))),s(X20,X27))) = s(X21,'const.pair.UNCURRY_2'(s(fun(X22,fun(X20,X21)),happ(s(fun(X23,fun(X22,fun(X20,X21))),X24),s(X23,X25))),s('type.pair.prod'(X22,X20),'const.pair.,_2'(s(X22,X26),s(X20,X27)))))) )),
  inference(backward_demodulation,[],[f1532,f1220])).
fof(f1659,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))),s(fun(X2,fun(X1,X0)),X3))))) = s(X6,happ(s(fun(X1,X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))))),s(X1,X5)))) )),
  inference(backward_demodulation,[],[f1532,f1270])).
fof(f1532,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X8,X7),'const.pair.,_2'(s(X8,X10),s(X7,X11))),s(fun(X8,fun(X7,X6)),X9))) = s(X6,'const.pair.UNCURRY_2'(s(fun(X8,fun(X7,X6)),X9),s('type.pair.prod'(X8,X7),'const.pair.,_2'(s(X8,X10),s(X7,X11)))))) )),
  inference(superposition,[],[f184,f183])).
fof(f1531,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(X37,happ(s(fun(X36,X37),X39),s(X36,X41))) = s(X37,'const.pair.pair_CASE_2'(s('type.pair.prod'(X38,X36),'const.pair.,_2'(s(X38,X40),s(X36,X41))),s(fun(X38,fun(X36,X37)),'_dst_x0_1'(s(fun(X36,X37),X39)))))) )),
  inference(superposition,[],[f184,f170])).
fof(f1530,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X29,happ(s(fun(X28,X29),'const.pair.UNCURRY_2'(s(fun(X31,fun(X30,fun(X28,X29))),X32),s('type.pair.prod'(X31,X30),'const.pair.,_2'(s(X31,X33),s(X30,X34))))),s(X28,X35))) = s(X29,'const.pair.pair_CASE_2'(s('type.pair.prod'(X30,X28),'const.pair.,_2'(s(X30,X34),s(X28,X35))),s(fun(X30,fun(X28,X29)),happ(s(fun(X31,fun(X30,fun(X28,X29))),X32),s(X31,X33)))))) )),
  inference(superposition,[],[f184,f183])).
fof(f1529,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X21,happ(s(fun(X20,X21),'const.pair.pair_CASE_2'(s('type.pair.prod'(X23,X22),'const.pair.,_2'(s(X23,X25),s(X22,X26))),s(fun(X23,fun(X22,fun(X20,X21))),X24))),s(X20,X27))) = s(X21,'const.pair.pair_CASE_2'(s('type.pair.prod'(X22,X20),'const.pair.,_2'(s(X22,X26),s(X20,X27))),s(fun(X22,fun(X20,X21)),happ(s(fun(X23,fun(X22,fun(X20,X21))),X24),s(X23,X25)))))) )),
  inference(superposition,[],[f184,f184])).
fof(f1528,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X13,happ(s(fun(X12,X13),happ(s(fun(X15,fun(X12,X13)),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))),s(X12,X19))) = s(X13,'const.pair.pair_CASE_2'(s('type.pair.prod'(X14,X12),'const.pair.,_2'(s(X14,X18),s(X12,X19))),s(fun(X14,fun(X12,X13)),'_dst_x0x1_2'(s(fun(X15,fun(X12,X13)),X16),s(fun(X14,X15),X17)))))) )),
  inference(superposition,[],[f184,f182])).
fof(f1527,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X5,'const.pair.pair_CASE_2'(s('type.pair.prod'(X6,X4),'const.pair.,_2'(s(X6,happ(s(fun(X8,X6),X9),s(X8,X10))),s(X4,X11))),s(fun(X6,fun(X4,X5)),X7))) = s(X5,happ(s(fun(X4,X5),happ(s(fun(X8,fun(X4,X5)),'_dst_x0x1_2'(s(fun(X6,fun(X4,X5)),X7),s(fun(X8,X6),X9))),s(X8,X10))),s(X4,X11)))) )),
  inference(superposition,[],[f184,f182])).
fof(f1461,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),'_dst_x0x1_2'(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun(X17,X16),X18),s(X17,X19))))),s(fun(X17,X16),X18))),s(X17,X19)))) )),
  inference(superposition,[],[f744,f182])).
fof(f1449,plain,(
  ( ! [X39,X40] : (s(bool,X39) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X39))),s(bool,X39))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f744,f223])).
fof(f1423,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X52) != s(bool,happ(s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))),s(X49,X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1385,f142])).
fof(f1385,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,'F_0')) | s(bool,X52) != s(bool,happ(s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))),s(X49,X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f489,f614])).
fof(f1313,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X186,X193] : (~p(s(bool,happ(s(fun('type.pair.prod'(X186,X187),bool),happ(s(fun(X189,fun('type.pair.prod'(X186,X187),bool)),happ(s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190),s(X188,X191))),s(X189,X192))),s('type.pair.prod'(X186,X187),sK7('const.pair.UNCURRY_2'(s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190),s('type.pair.prod'(X188,X189),'const.pair.,_2'(s(X188,X191),s(X189,X192)))),X187,X186))))) | p(s(bool,happ(s(fun('type.pair.prod'(X186,X187),bool),'const.pair.UNCURRY_2'(s(fun(X188,fun(X189,fun('type.pair.prod'(X186,X187),bool))),X190),s('type.pair.prod'(X188,X189),'const.pair.,_2'(s(X188,X191),s(X189,X192))))),s('type.pair.prod'(X186,X187),'const.pair.,_2'(s(X186,X193),s(X187,X194))))))) )),
  inference(superposition,[],[f168,f183])).
fof(f1291,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X49))),s(X43,happ(s(fun(X45,X43),happ(s(fun(X44,fun(X45,X43)),X46),s(X44,X47))),s(X45,X48)))))) | s(X43,X49) != s(X43,'const.pair.UNCURRY_2'(s(fun(X44,fun(X45,X43)),X46),s('type.pair.prod'(X44,X45),'const.pair.,_2'(s(X44,X47),s(X45,X48)))))) )),
  inference(superposition,[],[f173,f183])).
fof(f1290,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),'=_0'),s(X36,happ(s(fun(X38,X36),happ(s(fun(X37,fun(X38,X36)),X39),s(X37,X40))),s(X38,X41))))),s(X36,X42)))) | s(X36,X42) != s(X36,'const.pair.UNCURRY_2'(s(fun(X37,fun(X38,X36)),X39),s('type.pair.prod'(X37,X38),'const.pair.,_2'(s(X37,X40),s(X38,X41)))))) )),
  inference(superposition,[],[f173,f183])).
fof(f1289,plain,(
  ( ! [X30,X35,X33,X31,X29,X34,X32] : (~p(s(bool,happ(s(fun(X29,bool),happ(s(fun(X29,fun(X29,bool)),'=_0'),s(X29,X35))),s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X32),s(X30,X33))),s(X31,X34)))))) | s(X29,X35) = s(X29,'const.pair.UNCURRY_2'(s(fun(X30,fun(X31,X29)),X32),s('type.pair.prod'(X30,X31),'const.pair.,_2'(s(X30,X33),s(X31,X34)))))) )),
  inference(superposition,[],[f172,f183])).
fof(f1288,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),happ(s(fun(X22,fun(X22,bool)),'=_0'),s(X22,happ(s(fun(X24,X22),happ(s(fun(X23,fun(X24,X22)),X25),s(X23,X26))),s(X24,X27))))),s(X22,X28)))) | s(X22,X28) = s(X22,'const.pair.UNCURRY_2'(s(fun(X23,fun(X24,X22)),X25),s('type.pair.prod'(X23,X24),'const.pair.,_2'(s(X23,X26),s(X24,X27)))))) )),
  inference(superposition,[],[f172,f183])).
fof(f1283,plain,(
  ( ! [X206,X204,X202,X200,X199,X207,X205,X203,X201] : (~p(s(bool,happ(s(fun('type.pair.prod'(X199,X200),bool),'const.pair.UNCURRY_2'(s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203),s('type.pair.prod'(X202,X201),'const.pair.,_2'(s(X202,X204),s(X201,X205))))),s('type.pair.prod'(X199,X200),sK7(happ(s(fun(X201,fun('type.pair.prod'(X199,X200),bool)),happ(s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203),s(X202,X204))),s(X201,X205)),X200,X199))))) | p(s(bool,happ(s(fun('type.pair.prod'(X199,X200),bool),happ(s(fun(X201,fun('type.pair.prod'(X199,X200),bool)),happ(s(fun(X202,fun(X201,fun('type.pair.prod'(X199,X200),bool))),X203),s(X202,X204))),s(X201,X205))),s('type.pair.prod'(X199,X200),'const.pair.,_2'(s(X199,X206),s(X200,X207))))))) )),
  inference(superposition,[],[f168,f183])).
fof(f1278,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (p(s(bool,happ(s(fun(X59,bool),happ(s(fun(X59,fun(X59,bool)),'=_0'),s(X59,X65))),s(X59,'const.pair.UNCURRY_2'(s(fun(X61,fun(X60,X59)),X62),s('type.pair.prod'(X61,X60),'const.pair.,_2'(s(X61,X63),s(X60,X64)))))))) | s(X59,X65) != s(X59,happ(s(fun(X60,X59),happ(s(fun(X61,fun(X60,X59)),X62),s(X61,X63))),s(X60,X64)))) )),
  inference(superposition,[],[f173,f183])).
fof(f1277,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53] : (p(s(bool,happ(s(fun(X52,bool),happ(s(fun(X52,fun(X52,bool)),'=_0'),s(X52,'const.pair.UNCURRY_2'(s(fun(X54,fun(X53,X52)),X55),s('type.pair.prod'(X54,X53),'const.pair.,_2'(s(X54,X56),s(X53,X57))))))),s(X52,X58)))) | s(X52,X58) != s(X52,happ(s(fun(X53,X52),happ(s(fun(X54,fun(X53,X52)),X55),s(X54,X56))),s(X53,X57)))) )),
  inference(superposition,[],[f173,f183])).
fof(f1276,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (~p(s(bool,happ(s(fun(X45,bool),happ(s(fun(X45,fun(X45,bool)),'=_0'),s(X45,X51))),s(X45,'const.pair.UNCURRY_2'(s(fun(X47,fun(X46,X45)),X48),s('type.pair.prod'(X47,X46),'const.pair.,_2'(s(X47,X49),s(X46,X50)))))))) | s(X45,X51) = s(X45,happ(s(fun(X46,X45),happ(s(fun(X47,fun(X46,X45)),X48),s(X47,X49))),s(X46,X50)))) )),
  inference(superposition,[],[f172,f183])).
fof(f1275,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X38,fun(X38,bool)),'=_0'),s(X38,'const.pair.UNCURRY_2'(s(fun(X40,fun(X39,X38)),X41),s('type.pair.prod'(X40,X39),'const.pair.,_2'(s(X40,X42),s(X39,X43))))))),s(X38,X44)))) | s(X38,X44) = s(X38,happ(s(fun(X39,X38),happ(s(fun(X40,fun(X39,X38)),X41),s(X40,X42))),s(X39,X43)))) )),
  inference(superposition,[],[f172,f183])).
fof(f1270,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X6,happ(s(fun(X0,X6),X7),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))))))) = s(X6,happ(s(fun(X1,X6),'_dst_x0x1_2'(s(fun(X0,X6),X7),s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))))),s(X1,X5)))) )),
  inference(superposition,[],[f182,f183])).
fof(f1222,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s(X6,'const.pair.UNCURRY_2'(s(fun(X8,fun(X7,X6)),X9),s('type.pair.prod'(X8,X7),'const.pair.,_2'(s(X8,X10),s(X7,happ(s(fun(X11,X7),X12),s(X11,X13))))))) = s(X6,happ(s(fun(X11,X6),'_dst_x0x1_2'(s(fun(X7,X6),happ(s(fun(X8,fun(X7,X6)),X9),s(X8,X10))),s(fun(X11,X7),X12))),s(X11,X13)))) )),
  inference(superposition,[],[f183,f182])).
fof(f1221,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(X29,happ(s(fun(X28,X29),X31),s(X28,X33))) = s(X29,'const.pair.UNCURRY_2'(s(fun(X30,fun(X28,X29)),'_dst_x0_1'(s(fun(X28,X29),X31))),s('type.pair.prod'(X30,X28),'const.pair.,_2'(s(X30,X32),s(X28,X33)))))) )),
  inference(superposition,[],[f183,f170])).
fof(f1220,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X21,happ(s(fun(X20,X21),'const.pair.UNCURRY_2'(s(fun(X23,fun(X22,fun(X20,X21))),X24),s('type.pair.prod'(X23,X22),'const.pair.,_2'(s(X23,X25),s(X22,X26))))),s(X20,X27))) = s(X21,'const.pair.UNCURRY_2'(s(fun(X22,fun(X20,X21)),happ(s(fun(X23,fun(X22,fun(X20,X21))),X24),s(X23,X25))),s('type.pair.prod'(X22,X20),'const.pair.,_2'(s(X22,X26),s(X20,X27)))))) )),
  inference(superposition,[],[f183,f183])).
fof(f1219,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X13,happ(s(fun(X12,X13),happ(s(fun(X15,fun(X12,X13)),X16),s(X15,happ(s(fun(X14,X15),X17),s(X14,X18))))),s(X12,X19))) = s(X13,'const.pair.UNCURRY_2'(s(fun(X14,fun(X12,X13)),'_dst_x0x1_2'(s(fun(X15,fun(X12,X13)),X16),s(fun(X14,X15),X17))),s('type.pair.prod'(X14,X12),'const.pair.,_2'(s(X14,X18),s(X12,X19)))))) )),
  inference(superposition,[],[f183,f182])).
fof(f1218,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X5,'const.pair.UNCURRY_2'(s(fun(X6,fun(X4,X5)),X7),s('type.pair.prod'(X6,X4),'const.pair.,_2'(s(X6,happ(s(fun(X8,X6),X9),s(X8,X10))),s(X4,X11))))) = s(X5,happ(s(fun(X4,X5),happ(s(fun(X8,fun(X4,X5)),'_dst_x0x1_2'(s(fun(X6,fun(X4,X5)),X7),s(fun(X8,X6),X9))),s(X8,X10))),s(X4,X11)))) )),
  inference(superposition,[],[f183,f182])).
fof(f1216,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,X4))),s(X1,X5))))),s(X0,'const.pair.UNCURRY_2'(s(fun(X2,fun(X1,X0)),X3),s('type.pair.prod'(X2,X1),'const.pair.,_2'(s(X2,X4),s(X1,X5))))))))) )),
  inference(unit_resulting_resolution,[],[f183,f173])).
fof(f1217,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5))))))),s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))))))) )),
  inference(unit_resulting_resolution,[],[f183,f173])).
fof(f1176,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f1175])).
fof(f1175,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1160,f225])).
fof(f1160,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f489,f143])).
fof(f1057,plain,(
  ( ! [X61,X64,X62,X63] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X61),happ(s(fun(X62,'type.closSem.state'(X61)),X63),s(X62,X64))))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun(X62,'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'_dst_x0x1_2'(s(fun('type.closSem.state'(X61),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun(X62,'type.closSem.state'(X61)),X63))),s(X62,X64)))) )),
  inference(superposition,[],[f182,f160])).
fof(f1054,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,X49) = s(bool,happ(s(fun(X46,bool),'_dst_x0x1_2'(s(fun(X44,bool),X45),s(fun(X46,X44),X47))),s(X46,X48))) | s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),X45),s(X44,happ(s(fun(X46,X44),X47),s(X46,X48))))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f182,f223])).
fof(f1052,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X34,bool),'_dst_x0x1_2'(s(fun(X32,bool),X33),s(fun(X34,X32),X35))),s(X34,X36))) | s(bool,X37) = s(bool,happ(s(fun(X32,bool),X33),s(X32,happ(s(fun(X34,X32),X35),s(X34,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f182,f223])).
fof(f1050,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(X16,happ(s(fun(X18,X16),X19),s(X18,happ(s(fun(X17,X18),X20),s(X17,happ(s(fun(X21,X17),X22),s(X21,X23))))))) = s(X16,happ(s(fun(X21,X16),'_dst_x0x1_2'(s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X19),s(fun(X17,X18),X20))),s(fun(X21,X17),X22))),s(X21,X23)))) )),
  inference(superposition,[],[f182,f182])).
fof(f1049,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X4,X6) = s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),'_dst_x0_1'(s(X4,X6))),s(fun(X7,X5),X8))),s(X7,X9)))) )),
  inference(superposition,[],[f182,f170])).
fof(f1048,plain,(
  ( ! [X61,X64,X62,X63] : (s(X63,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X63),X64),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X61),X62))))) = s(X63,happ(s(fun('type.closSem.state'(X61),X63),'_dst_x0x1_2'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),X63),X64),s(fun('type.closSem.state'(X61),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))),s('type.closSem.state'(X61),X62)))) )),
  inference(superposition,[],[f182,f160])).
fof(f1047,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X24,happ(s(fun(X18,X24),X25),s(X18,happ(s(fun(X20,X18),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))))) = s(X24,happ(s(fun(X19,X24),'_dst_x0x1_2'(s(fun(X18,X24),X25),s(fun(X19,X18),'_dst_x0x1_2'(s(fun(X20,X18),X21),s(fun(X19,X20),X22))))),s(X19,X23)))) )),
  inference(superposition,[],[f182,f182])).
fof(f1046,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X16,happ(s(fun(X10,X16),X17),s(X10,happ(s(fun(X13,X10),'_dst_x0x1_2'(s(fun(X11,X10),X12),s(fun(X13,X11),X14))),s(X13,X15))))) = s(X16,happ(s(fun(X11,X16),'_dst_x0x1_2'(s(fun(X10,X16),X17),s(fun(X11,X10),X12))),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))))) )),
  inference(superposition,[],[f182,f182])).
fof(f1045,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(X8,happ(s(fun(X4,X8),X9),s(X4,X6))) = s(X8,happ(s(fun(X5,X8),'_dst_x0x1_2'(s(fun(X4,X8),X9),s(fun(X5,X4),'_dst_x0_1'(s(X4,X6))))),s(X5,X7)))) )),
  inference(superposition,[],[f182,f170])).
fof(f931,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f921,f140])).
fof(f921,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f541,f223])).
fof(f918,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1)))))),
  inference(unit_resulting_resolution,[],[f140,f541,f215])).
fof(f911,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f901,f140])).
fof(f901,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f539,f223])).
fof(f898,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5)))))),
  inference(unit_resulting_resolution,[],[f140,f539,f215])).
fof(f793,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK10(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f171,f223])).
fof(f791,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK10(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f171,f223])).
fof(f744,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f679,f225])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))) | s(X0,X1) != s(X0,X2)) )),
  inference(condensation,[],[f692])).
fof(f692,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f675,f225])).
fof(f675,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X9) != s(X8,X10) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),'=_0'),s(X8,X9))),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f173,f223])).
fof(f679,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f647,f170])).
fof(f647,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X1,X0),'_dst_x0_1'(s(X0,X2))),s(X1,X3))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f170,f173])).
fof(f678,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),bool),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),bool)),'=_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))))))),
  inference(forward_demodulation,[],[f648,f138])).
fof(f648,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),bool),happ(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),bool)),'=_0'),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),sK4),s('type.closSem.v',sK3),s('type.list.list'('type.closSem.v'),sK2),s('type.closSem.state'(sK0),sK1))))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))))))),
  inference(unit_resulting_resolution,[],[f138,f173])).
fof(f638,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) )),
  inference(condensation,[],[f637])).
fof(f637,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f626,f225])).
fof(f626,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,'F_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f616,f223])).
fof(f576,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,X2) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f565,f140])).
fof(f565,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f172,f223])).
fof(f539,plain,(
  ~p(s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))))),
  inference(unit_resulting_resolution,[],[f139,f172])).
fof(f541,plain,(
  ~p(s(bool,happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool),happ(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5))))),s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))))))),
  inference(unit_resulting_resolution,[],[f139,f172])).
fof(f491,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f490])).
fof(f490,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f366,f347])).
fof(f366,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f215,f223])).
fof(f489,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f358,f140])).
fof(f358,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,X8) | p(s(bool,X8)) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f215,f223])).
fof(f486,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f485])).
fof(f485,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f348,f347])).
fof(f348,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f215,f223])).
fof(f347,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f140,f223])).
fof(f310,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,X19) = s(bool,happ(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))),s(X16,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f223,f170])).
fof(f309,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f223,f159])).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f223,f223])).
fof(f288,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23))),s(X22,X24)))) )),
  inference(superposition,[],[f223,f170])).
fof(f287,plain,(
  ( ! [X21,X20] : (s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20)))) )),
  inference(superposition,[],[f223,f159])).
fof(f228,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f143])).
fof(f225,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f142,f143])).
fof(f223,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f143,f143])).
fof(f203,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X1,fun(X4,X0)),'_dst_x00x11x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,fun(X4,X2)),happ(s(fun(X3,fun(X1,fun(X4,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x111x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,fun(X4,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X1,fun(X4,X0)),'_dst_x00x11x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,fun(X4,X2)),happ(s(fun(X3,fun(X1,fun(X4,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x111x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,fun(X4,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X1,fun(X4,X0)),'_dst_x00x11x01_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X1,fun(X4,X2)),happ(s(fun(X3,fun(X1,fun(X4,X2))),X7),s(X3,X9))),s(fun(X4,X5),X8))) = s(fun(X1,fun(X4,X0)),happ(s(fun(X3,fun(X1,fun(X4,X0))),'_dst_x000x111x001_3'(s(fun(X2,fun(X5,X0)),X6),s(fun(X3,fun(X1,fun(X4,X2))),X7),s(fun(X4,X5),X8))),s(X3,X9)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X36,X37,X42,X34,X38,X43] : ! [X39,X35,X40,X41] : s(fun(X37,fun(X38,X36)),'_dst_x00x11x01_3'(s(fun(X42,fun(X43,X36)),X39),s(fun(X37,fun(X38,X42)),happ(s(fun(X34,fun(X37,fun(X38,X42))),X35),s(X34,X41))),s(fun(X38,X43),X40))) = s(fun(X37,fun(X38,X36)),happ(s(fun(X34,fun(X37,fun(X38,X36))),'_dst_x000x111x001_3'(s(fun(X42,fun(X43,X36)),X39),s(fun(X34,fun(X37,fun(X38,X42))),X35),s(fun(X38,X43),X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x000x111x001_ax')).
fof(f202,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X36,X42,X34,X37,X38,X43] : ! [X39,X35,X40,X41] : s(fun(X37,fun(X38,X36)),'_dst_x00x00x11_3'(s(fun(X42,fun(X43,X36)),X39),s(X42,happ(s(fun(X34,X42),X35),s(X34,X41))),s(fun(X37,fun(X38,X43)),X40))) = s(fun(X37,fun(X38,X36)),happ(s(fun(X34,fun(X37,fun(X38,X36))),'_dst_x000x100x011_3'(s(fun(X42,fun(X43,X36)),X39),s(fun(X34,X42),X35),s(fun(X37,fun(X38,X43)),X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x000x100x011_ax')).
fof(f201,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(X4,happ(s(fun(X5,X4),X8),s(X5,X9))))) = s(fun(X0,fun(X2,X1)),happ(s(fun(X5,fun(X0,fun(X2,X1))),'_dst_x000x001x100_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(fun(X5,X4),X8))),s(X5,X9)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(X4,happ(s(fun(X5,X4),X8),s(X5,X9))))) = s(fun(X0,fun(X2,X1)),happ(s(fun(X5,fun(X0,fun(X2,X1))),'_dst_x000x001x100_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(fun(X5,X4),X8))),s(X5,X9)))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(X4,happ(s(fun(X5,X4),X8),s(X5,X9))))) = s(fun(X0,fun(X2,X1)),happ(s(fun(X5,fun(X0,fun(X2,X1))),'_dst_x000x001x100_3'(s(fun(X3,fun(X4,X1)),X6),s(fun(X2,X3),X7),s(fun(X5,X4),X8))),s(X5,X9)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X37,X36,X38,X42,X43,X34] : ! [X39,X35,X40,X41] : s(fun(X37,fun(X38,X36)),'_dst_x00x01x00_3'(s(fun(X42,fun(X43,X36)),X39),s(fun(X38,X42),X35),s(X43,happ(s(fun(X34,X43),X40),s(X34,X41))))) = s(fun(X37,fun(X38,X36)),happ(s(fun(X34,fun(X37,fun(X38,X36))),'_dst_x000x001x100_3'(s(fun(X42,fun(X43,X36)),X39),s(fun(X38,X42),X35),s(fun(X34,X43),X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x000x001x100_ax')).
fof(f200,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X3,X1),happ(s(fun(X2,fun(X3,X1)),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x11x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,fun(X3,X1)),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X3,X1),happ(s(fun(X2,fun(X3,X1)),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x11x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,fun(X3,X1)),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X3,X1),happ(s(fun(X2,fun(X3,X1)),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x11x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,fun(X3,X1)),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X36,X38,X34,X37,X42] : ! [X39,X35,X40,X41] : s(fun(X37,X36),'_dst_x0x1x1_3'(s(fun(X38,fun(X42,X36)),X39),s(fun(X37,X38),happ(s(fun(X34,fun(X37,X38)),X35),s(X34,X41))),s(fun(X37,X42),X40))) = s(fun(X37,X36),happ(s(fun(X34,fun(X37,X36)),'_dst_x00x11x01_3'(s(fun(X38,fun(X42,X36)),X39),s(fun(X34,fun(X37,X38)),X35),s(fun(X37,X42),X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x00x11x01_ax')).
fof(f199,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X36,X38,X37,X42,X34] : ! [X39,X35,X40,X41] : s(fun(X37,X36),'_dst_x0x0x1_3'(s(fun(X38,fun(X42,X36)),X39),s(X38,X35),s(fun(X37,X42),happ(s(fun(X34,fun(X37,X42)),X40),s(X34,X41))))) = s(fun(X37,X36),happ(s(fun(X34,fun(X37,X36)),'_dst_x00x00x11_3'(s(fun(X38,fun(X42,X36)),X39),s(X38,X35),s(fun(X34,fun(X37,X42)),X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x00x00x11_ax')).
fof(f198,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))) = s(fun(X2,X1),happ(s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))) = s(fun(X2,X1),happ(s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))),s(X0,X8)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X1),'_dst_x0x1x0_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))) = s(fun(X2,X1),happ(s(fun(X0,fun(X2,X1)),'_dst_x00x01x00_3'(s(fun(X3,fun(X4,X1)),X5),s(fun(X2,X3),X6),s(X4,X7))),s(X0,X8)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X34,X36,X37,X38,X42] : ! [X39,X35,X40,X41] : s(fun(X37,X36),'_dst_x0x1x0_3'(s(fun(X38,fun(X42,X36)),X39),s(fun(X37,X38),X35),s(X42,X40))) = s(fun(X37,X36),happ(s(fun(X34,fun(X37,X36)),'_dst_x00x01x00_3'(s(fun(X38,fun(X42,X36)),X39),s(fun(X37,X38),X35),s(X42,X40))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x00x01x00_ax')).
fof(f197,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7] : s(fun(X1,fun(X2,X0)),'_dst_x00x11_2'(s(fun(X3,X0),X5),s(fun(X1,fun(X2,X3)),happ(s(fun(X4,fun(X1,fun(X2,X3))),X6),s(X4,X7))))) = s(fun(X1,fun(X2,X0)),happ(s(fun(X4,fun(X1,fun(X2,X0))),'_dst_x000x111_2'(s(fun(X3,X0),X5),s(fun(X4,fun(X1,fun(X2,X3))),X6))),s(X4,X7)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X36,X37,X38,X42,X34] : ! [X39,X35,X41] : s(fun(X37,fun(X38,X36)),'_dst_x00x11_2'(s(fun(X42,X36),X39),s(fun(X37,fun(X38,X42)),happ(s(fun(X34,fun(X37,fun(X38,X42))),X35),s(X34,X41))))) = s(fun(X37,fun(X38,X36)),happ(s(fun(X34,fun(X37,fun(X38,X36))),'_dst_x000x111_2'(s(fun(X42,X36),X39),s(fun(X34,fun(X37,fun(X38,X42))),X35))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x000x111_ax')).
fof(f196,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X36,X37,X38,X34] : ! [X39,X35,X40,X41] : s(X36,happ(s(fun(X34,X36),'_dst_x0x1x1_3'(s(fun(X37,fun(X38,X36)),X39),s(fun(X34,X37),X35),s(fun(X34,X38),X40))),s(X34,X41))) = s(X36,happ(s(fun(X38,X36),happ(s(fun(X37,fun(X38,X36)),X39),s(X37,happ(s(fun(X34,X37),X35),s(X34,X41))))),s(X38,happ(s(fun(X34,X38),X40),s(X34,X41)))))),
  file('thm.closProps.evaluate_app_const','_dst_x0x1x1_ax')).
fof(f195,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X36,X37,X34,X38] : ! [X39,X35,X40,X41] : s(X36,happ(s(fun(X34,X36),'_dst_x0x1x0_3'(s(fun(X37,fun(X38,X36)),X39),s(fun(X34,X37),X35),s(X38,X40))),s(X34,X41))) = s(X36,happ(s(fun(X38,X36),happ(s(fun(X37,fun(X38,X36)),X39),s(X37,happ(s(fun(X34,X37),X35),s(X34,X41))))),s(X38,X40)))),
  file('thm.closProps.evaluate_app_const','_dst_x0x1x0_ax')).
fof(f194,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X36,X37,X38,X34] : ! [X39,X35,X40,X41] : s(X36,happ(s(fun(X34,X36),'_dst_x0x0x1_3'(s(fun(X37,fun(X38,X36)),X39),s(X37,X35),s(fun(X34,X38),X40))),s(X34,X41))) = s(X36,happ(s(fun(X38,X36),happ(s(fun(X37,fun(X38,X36)),X39),s(X37,X35))),s(X38,happ(s(fun(X34,X38),X40),s(X34,X41)))))),
  file('thm.closProps.evaluate_app_const','_dst_x0x0x1_ax')).
fof(f193,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X36,X37,X38,X34] : ! [X39,X35,X41] : s(fun(X37,X36),'_dst_x0x1_2'(s(fun(X38,X36),X39),s(fun(X37,X38),happ(s(fun(X34,fun(X37,X38)),X35),s(X34,X41))))) = s(fun(X37,X36),happ(s(fun(X34,fun(X37,X36)),'_dst_x00x11_2'(s(fun(X38,X36),X39),s(fun(X34,fun(X37,X38)),X35))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x00x11_ax')).
fof(f215,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f214])).
fof(f214,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f192])).
fof(f192,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14,X19,X15,X20] : ((s(bool,X14) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X15) = s(bool,X20))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.closProps.evaluate_app_const','thm.bool.IMP_CONG')).
fof(f184,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X5),s(X1,X4))),s(X2,X3))) = s(X0,'const.pair.pair_CASE_2'(s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X3))),s(fun(X1,fun(X2,X0)),X5)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X27,X26] : ! [X15,X14,X28] : s(X6,happ(s(fun(X26,X6),happ(s(fun(X27,fun(X26,X6)),X28),s(X27,X14))),s(X26,X15))) = s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X27,X26),'const.pair.,_2'(s(X27,X14),s(X26,X15))),s(fun(X27,fun(X26,X6)),X28)))),
  file('thm.closProps.evaluate_app_const','thm.pair.pair_case_thm')).
fof(f183,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X26,X6,X27] : ! [X28,X14,X15] : s(X26,happ(s(fun(X27,X26),happ(s(fun(X6,fun(X27,X26)),X28),s(X6,X14))),s(X27,X15))) = s(X26,'const.pair.UNCURRY_2'(s(fun(X6,fun(X27,X26)),X28),s('type.pair.prod'(X6,X27),'const.pair.,_2'(s(X6,X14),s(X27,X15)))))),
  file('thm.closProps.evaluate_app_const','thm.pair.UNCURRY_DEF')).
fof(f182,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X36,X37,X34] : ! [X39,X35,X41] : s(X36,happ(s(fun(X37,X36),X39),s(X37,happ(s(fun(X34,X37),X35),s(X34,X41))))) = s(X36,happ(s(fun(X34,X36),'_dst_x0x1_2'(s(fun(X37,X36),X39),s(fun(X34,X37),X35))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x0x1_ax')).
fof(f181,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X2,'const.pair.pair_CASE_2'(s('type.pair.prod'(X0,X1),X3),s(fun(X0,fun(X1,X2)),X4))) = s(X2,happ(s(fun(fun(X0,fun(X1,X2)),X2),happ(s(fun('type.pair.prod'(X0,X1),fun(fun(X0,fun(X1,X2)),X2)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X0,X1),X3))),s(fun(X0,fun(X1,X2)),X4)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X27,X26,X6] : ! [X33,X47] : s(X6,'const.pair.pair_CASE_2'(s('type.pair.prod'(X27,X26),X33),s(fun(X27,fun(X26,X6)),X47))) = s(X6,happ(s(fun(fun(X27,fun(X26,X6)),X6),happ(s(fun('type.pair.prod'(X27,X26),fun(fun(X27,fun(X26,X6)),X6)),'const.pair.pair_CASE_0'),s('type.pair.prod'(X27,X26),X33))),s(fun(X27,fun(X26,X6)),X47)))),
  file('thm.closProps.evaluate_app_const','ARITY_0_2_const.pair.pair_CASE')).
fof(f180,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),happ(s(fun(fun(X1,fun(X2,X0)),fun('type.pair.prod'(X1,X2),X0)),'const.pair.UNCURRY_0'),s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X26,X6,X27] : ! [X33,X47] : s(X26,'const.pair.UNCURRY_2'(s(fun(X6,fun(X27,X26)),X33),s('type.pair.prod'(X6,X27),X47))) = s(X26,happ(s(fun('type.pair.prod'(X6,X27),X26),happ(s(fun(fun(X6,fun(X27,X26)),fun('type.pair.prod'(X6,X27),X26)),'const.pair.UNCURRY_0'),s(fun(X6,fun(X27,X26)),X33))),s('type.pair.prod'(X6,X27),X47)))),
  file('thm.closProps.evaluate_app_const','ARITY_0_2_const.pair.UNCURRY')).
fof(f171,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK10(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK10(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.closProps.evaluate_app_const',aHL_EXT)).
fof(f170,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X34,X36] : ! [X39,X41] : s(X36,X39) = s(X36,happ(s(fun(X34,X36),'_dst_x0_1'(s(X36,X39))),s(X34,X41)))),
  file('thm.closProps.evaluate_app_const','_dst_x0_ax')).
fof(f162,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool),happ(s(fun(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'_dst_x0x1_2'(s(fun(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool)),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x111x001_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),bool)),'const.pair.pair_CASE_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))))),'_dst_x000x111_2'(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))))),'_dst_x000x100x011_3'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))))),'const.pair.,_0'),s(fun('type.list.list'('type.closLang.exp'),'type.list.list'('type.closLang.exp')),'_i_0_0'),s(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))),'const.pair.,_0'))))),s(fun('type.closSem.state'(X0),fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x001x100_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))))))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool),happ(s(fun(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'_dst_x0x1_2'(s(fun(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool)),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x111x001_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),bool)),'const.pair.pair_CASE_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))))),'_dst_x000x111_2'(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))))),'_dst_x000x100x011_3'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))))),'const.pair.,_0'),s(fun('type.list.list'('type.closLang.exp'),'type.list.list'('type.closLang.exp')),'_i_0_0'),s(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))),'const.pair.,_0'))))),s(fun('type.closSem.state'(X0),fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x001x100_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))))))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool),happ(s(fun(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'_dst_x0x1_2'(s(fun(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool)),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x111x001_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),fun(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool)),bool)),'const.pair.pair_CASE_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))))),'_dst_x000x111_2'(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))))),'_dst_x000x100x011_3'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))))),'const.pair.,_0'),s(fun('type.list.list'('type.closLang.exp'),'type.list.list'('type.closLang.exp')),'_i_0_0'),s(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X0),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0)))),'const.pair.,_0'))))),s(fun('type.closSem.state'(X0),fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X0),bool))),'_dst_x000x001x100_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))))))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X21] : ! [X22] : p(s(bool,happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),bool),happ(s(fun(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)),bool)),fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)),bool)),'_dst_x0x1_2'(s(fun(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X21),bool)),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)),bool)),'const.pair.UNCURRY_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X21),bool))),'_dst_x000x111x001_3'(s(fun('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),fun(fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X21),bool)),bool)),'const.pair.pair_CASE_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X21),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21))))),'_dst_x000x111_2'(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21))),'const.closSem.evaluate_0'),s(fun('type.list.list'('type.closLang.exp'),fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X21),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)))))),'_dst_x000x100x011_3'(s(fun('type.list.list'('type.closLang.exp'),fun('type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)),'type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))))),'const.pair.,_0'),s(fun('type.list.list'('type.closLang.exp'),'type.list.list'('type.closLang.exp')),'_i_0_0'),s(fun('type.list.list'('type.closSem.v'),fun('type.closSem.state'(X21),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21)))),'const.pair.,_0'))))),s(fun('type.closSem.state'(X21),fun('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),fun('type.closSem.state'(X21),bool))),'_dst_x000x001x100_3'(s(fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),fun('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),bool)),'=_0'),s(fun('type.closSem.state'(X21),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s(fun('type.closSem.state'(X21),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'))))))))),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),X22))))),
  file('thm.closProps.evaluate_app_const','thm.closProps.closProps_59_prove_ltb|split|0')).
fof(f161,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X0))),X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X21] : ! [X33] : s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),'const.closSem.evaluate_1'(s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),X33))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21)),happ(s(fun('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),'type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X21))),'const.closSem.evaluate_0'),s('type.pair.prod'('type.list.list'('type.closLang.exp'),'type.pair.prod'('type.list.list'('type.closSem.v'),'type.closSem.state'(X21))),X33)))),
  file('thm.closProps.evaluate_app_const','ARITY_0_1_const.closSem.evaluate')).
fof(f160,plain,(
  ( ! [X0,X1] : (s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X0),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X0),X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X21] : ! [X33] : s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X21),X33))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),happ(s(fun('type.closSem.state'(X21),'type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp'))),'const.closSem.state_code_0'),s('type.closSem.state'(X21),X33)))),
  file('thm.closProps.evaluate_app_const','ARITY_0_1_const.closSem.state_code')).
fof(f159,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X34] : ! [X35] : s(X34,X35) = s(X34,happ(s(fun(X34,X34),'_i_0_0'),s(X34,X35)))),
  file('thm.closProps.evaluate_app_const','_i0_ax')).
fof(f143,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.closProps.evaluate_app_const',aHL_BOOL_CASES)).
fof(f140,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.closProps.evaluate_app_const',aHL_TRUTH)).
fof(f138,plain,(
  s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),sK4),s('type.closSem.v',sK3),s('type.list.list'('type.closSem.v'),sK2),s('type.closSem.state'(sK0),sK1)))),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),sK6),s('type.closSem.state'(sK0),sK5))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),sK4),s('type.closSem.v',sK3),s('type.list.list'('type.closSem.v'),sK2),s('type.closSem.state'(sK0),sK1))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),X6),s('type.closSem.state'(X0),X5))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X4),s('type.closSem.v',X3),s('type.list.list'('type.closSem.v'),X2),s('type.closSem.state'(X0),X1))) & s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X5))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),X6),s('type.closSem.state'(X0),X5))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X4),s('type.closSem.v',X3),s('type.list.list'('type.closSem.v'),X2),s('type.closSem.state'(X0),X1))) => s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X5))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),X6),s('type.closSem.state'(X0),X5))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X4),s('type.closSem.v',X3),s('type.list.list'('type.closSem.v'),X2),s('type.closSem.state'(X0),X1))) => s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X1))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X0),X5))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),X12),s('type.closSem.state'(X6),X11))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X10),s('type.closSem.v',X9),s('type.list.list'('type.closSem.v'),X8),s('type.closSem.state'(X6),X7))) => s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X7))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X11))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),X12),s('type.closSem.state'(X6),X11))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.closSem.evaluate_app_4'(s('type.option.option'('type.num.num'),X10),s('type.closSem.v',X9),s('type.list.list'('type.closSem.v'),X8),s('type.closSem.state'(X6),X7))) => s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X7))) = s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(X6),X11))))),
  file('thm.closProps.evaluate_app_const','thm.closProps.evaluate_app_const')).
fof(f139,plain,(
  s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK1))) != s('type.finite_map.fmap'('type.num.num','type.pair.prod'('type.num.num','type.closLang.exp')),'const.closSem.state_code_1'(s('type.closSem.state'(sK0),sK5)))),
  inference(cnf_transformation,[],[f116])).
% SZS output end Proof for thm.closProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 95
% Passive clauses: 1296
% Generated clauses: 12395
% Final active clauses: 50
% Final passive clauses: 822
% Input formulas: 39
% Initial clauses: 66
% 
% Duplicate literals: 191
% Trivial inequalities: 58
% Fw subsumption resolutions: 123
% Fw demodulations: 2403
% Bw demodulations: 289
% Condensations: 363
% Global subsumptions: 1
% 
% Simple tautologies: 115
% Equational tautologies: 222
% Forward subsumptions: 7491
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 2189
% Forward superposition: 3398
% Backward superposition: 2745
% Self superposition: 173
% Equality factoring: 30
% Equality resolution: 4
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 1323
% SAT solver unit clauses: 489
% SAT solver binary clauses: 411
% 
% Sat splits: 5
% Sat splitting refutations: 14
% 
% Memory used [KB]: 22643
% Time elapsed: 0.742 s
% ------------------------------
% ------------------------------
% Success in time 2.12 s
