Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.stack_removeProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stack_removeProof
% SZS output start Proof for thm.stack_removeProof
fof(f2405,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f2404])).
fof(f2404,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))))))),
  inference(forward_demodulation,[],[f2403,f46])).
fof(f46,plain,(
  ( ! [X2,X0,X1] : (s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1,X2] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X17,X18] : s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X17),s(fun(fun(X6,bool),bool),X18))) = s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X18),s(fun(fun(X6,bool),bool),X17)))),
  file('thm.stack_removeProof.stack_removeProof_88_prove_ltb','thm.set_sep.STAR_COMM')).
fof(f2403,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3)))))))),
  inference(forward_demodulation,[],[f2402,f806])).
fof(f806,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(fun(X4,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X4,bool),bool),X5),s(fun(fun(X4,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X4,bool),bool),X6),s(fun(fun(X4,bool),bool),X7))))) = s(fun(fun(X4,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X4,bool),bool),X7),s(fun(fun(X4,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X4,bool),bool),X5),s(fun(fun(X4,bool),bool),X6)))))) )),
  inference(superposition,[],[f47,f46])).
fof(f47,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X3))))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X3)))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X3))))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X3)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : ! [X1,X2,X3] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X3))))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X1),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X3)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X17,X18,X19] : s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X17),s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X18),s(fun(fun(X6,bool),bool),X19))))) = s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X17),s(fun(fun(X6,bool),bool),X18))),s(fun(fun(X6,bool),bool),X19)))),
  file('thm.stack_removeProof.stack_removeProof_88_prove_ltb','thm.set_sep.STAR_ASSOC')).
fof(f2402,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))),
  inference(forward_demodulation,[],[f2401,f46])).
fof(f2401,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),sK4))),s(fun(fun(sK0,bool),bool),sK3))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4)))))))),
  inference(superposition,[],[f2167,f806])).
fof(f2167,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),sK4)))))))),
  inference(forward_demodulation,[],[f2098,f806])).
fof(f2098,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK1)))))))),
  inference(backward_demodulation,[],[f806,f885])).
fof(f885,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),sK4))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),sK2)))))))),
  inference(forward_demodulation,[],[f884,f46])).
fof(f884,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3))))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),sK2)))))))),
  inference(forward_demodulation,[],[f883,f46])).
fof(f883,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3))),s(fun(fun(sK0,bool),bool),sK2))))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),sK2)))))))),
  inference(forward_demodulation,[],[f882,f46])).
fof(f882,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3))),s(fun(fun(sK0,bool),bool),sK2))),s(fun(fun(sK0,bool),bool),sK1))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK1),s(fun(fun(sK0,bool),bool),sK2)))))))),
  inference(forward_demodulation,[],[f38,f46])).
fof(f38,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3))),s(fun(fun(sK0,bool),bool),sK2))),s(fun(fun(sK0,bool),bool),sK1))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),sK1)))))))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),sK3))),s(fun(fun(sK0,bool),bool),sK2))),s(fun(fun(sK0,bool),bool),sK1))) != s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK3),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK4),s(fun(fun(sK0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(sK0,bool),bool),sK2),s(fun(fun(sK0,bool),bool),sK1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f27])).
fof(f27,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1))))))) != s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),X3))),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X1)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1))))))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),X3))),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X1)))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X3),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X2),s(fun(fun(X0,bool),bool),X1))))))) = s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X0,bool),bool),X4),s(fun(fun(X0,bool),bool),X3))),s(fun(fun(X0,bool),bool),X2))),s(fun(fun(X0,bool),bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X9),s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X10),s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X8),s(fun(fun(X6,bool),bool),X7))))))) = s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X10),s(fun(fun(X6,bool),bool),X9))),s(fun(fun(X6,bool),bool),X8))),s(fun(fun(X6,bool),bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X9),s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X10),s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X8),s(fun(fun(X6,bool),bool),X7))))))) = s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),'const.set_sep.STAR_2'(s(fun(fun(X6,bool),bool),X10),s(fun(fun(X6,bool),bool),X9))),s(fun(fun(X6,bool),bool),X8))),s(fun(fun(X6,bool),bool),X7)))),
  file('thm.stack_removeProof.stack_removeProof_88_prove_ltb','thm.stack_removeProof.stack_removeProof_88_prove_ltb')).
% SZS output end Proof for thm.stack_removeProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 58
% Passive clauses: 314
% Generated clauses: 2307
% Final active clauses: 47
% Final passive clauses: 194
% Input formulas: 11
% Initial clauses: 14
% 
% Duplicate literals: 193
% Trivial inequalities: 18
% Fw subsumption resolutions: 149
% Fw demodulations: 172
% Bw demodulations: 8
% 
% Simple tautologies: 62
% Equational tautologies: 20
% Forward subsumptions: 1316
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 59
% Forward superposition: 899
% Backward superposition: 649
% Self superposition: 72
% Equality factoring: 32
% Equality resolution: 1
% 
% Split clauses: 84
% Split components: 174
% SAT solver clauses: 115
% SAT solver unit clauses: 4
% SAT solver binary clauses: 29
% 
% Sat splits: 84
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 115
% TWLsolver calls for satisfiability: 27
% 
% Memory used [KB]: 2046
% Time elapsed: 0.057 s
% ------------------------------
% ------------------------------
% Success in time 0.059 s
