Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (110533)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 2422
% Generated clauses: 13537
% Final active clauses: 132
% Final passive clauses: 2123
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 482
% Trivial inequalities: 250
% Fw subsumption resolutions: 1244
% Fw demodulations: 1812
% Bw demodulations: 14
% 
% Simple tautologies: 244
% Equational tautologies: 44
% Forward subsumptions: 6674
% 
% Binary resolution: 414
% Factoring: 1
% Forward superposition: 4392
% Backward superposition: 4644
% Self superposition: 116
% Equality factoring: 68
% Equality resolution: 28
% 
% Split clauses: 205
% Split components: 412
% SAT solver clauses: 518
% SAT solver unit clauses: 9
% SAT solver binary clauses: 138
% 
% Sat splits: 205
% 
% TWLsolver clauses: 518
% TWLsolver calls for satisfiability: 142
% 
% Memory used [KB]: 10874
% Time elapsed: 0.606 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% (110553)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1100
% Generated clauses: 16003
% Final active clauses: 140
% Final passive clauses: 869
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 268
% Trivial inequalities: 306
% Fw subsumption resolutions: 846
% Fw demodulations: 3470
% Bw demodulations: 83
% Condensations: 799
% Global subsumptions: 8
% 
% Simple tautologies: 115
% Equational tautologies: 41
% Forward subsumptions: 6808
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 145
% Factoring: 1
% Forward superposition: 4267
% Backward superposition: 4314
% Self superposition: 148
% Equality factoring: 72
% Equality resolution: 27
% 
% Split clauses: 82
% Split components: 164
% SAT solver clauses: 2202
% SAT solver unit clauses: 152
% SAT solver binary clauses: 53
% SAT solver learnt clauses: 16
% 
% Sat splits: 82
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2170
% TWLsolver calls for satisfiability: 383
% 
% Memory used [KB]: 18421
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f32359,plain,(
  $false),
  inference(subsumption_resolution,[],[f32358,f32357])).
fof(f32357,plain,(
  ( ! [X368] : (~p(s(bool,X368))) )),
  inference(forward_demodulation,[],[f32356,f320])).
fof(f320,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X16] : ! [X22,X31,X32,X33,X34] : s(X16,X34) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),X16),X22),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|6','thm.stackSem.result_case_def|split|4')).
fof(f32356,plain,(
  ( ! [X364,X366,X368,X365,X367] : (~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X364),s(fun('type.wordSem.word_loc'(sK0),bool),X365),s(fun('type.wordSem.word_loc'(sK0),bool),X366),s(bool,X367),s(bool,X368))))) )),
  inference(subsumption_resolution,[],[f32178,f32113])).
fof(f32113,plain,(
  ( ! [X6,X8,X11] : (s(X6,X11) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),X8),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f20961,f320])).
fof(f20961,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),X1),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X1),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f324,f231])).
fof(f231,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Error_0') = s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.Error_0') = s('type.stackSem.result'(sK0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f185])).
fof(f185,plain,(
  ? [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.Error_0') = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.stackSem.result'(X0),'const.stackSem.Error_0') != s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : ~s('type.stackSem.result'(X0),'const.stackSem.Error_0') = s('type.stackSem.result'(X0),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.Error_0') = s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ~s('type.stackSem.result'(X6),'const.stackSem.Error_0') = s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.stackSem.result_distinct|split|6','thm.stackSem.result_distinct|split|6')).
fof(f324,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),X4),s('type.wordSem.word_loc'(X1),X2))) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X1),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X1),X2))),s(fun('type.wordSem.word_loc'(X1),X0),X3),s(fun('type.wordSem.word_loc'(X1),X0),X4),s(fun('type.wordSem.word_loc'(X1),X0),X5),s(X0,X6),s(X0,X7)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X16,X6] : ! [X7,X22,X31,X32,X33,X34] : s(X16,happ(s(fun('type.wordSem.word_loc'(X6),X16),X31),s('type.wordSem.word_loc'(X6),X7))) = s(X16,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Exception_1'(s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),X16),X22),s(fun('type.wordSem.word_loc'(X6),X16),X31),s(fun('type.wordSem.word_loc'(X6),X16),X32),s(X16,X33),s(X16,X34)))),
  file('thm.stackSem.result_distinct|split|6','thm.stackSem.result_case_def|split|1')).
fof(f32178,plain,(
  ( ! [X364,X366,X368,X365,X367] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X365),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X364),s(fun('type.wordSem.word_loc'(sK0),bool),X365),s(fun('type.wordSem.word_loc'(sK0),bool),X366),s(bool,X367),s(bool,X368))))) )),
  inference(superposition,[],[f246,f20961])).
fof(f246,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X28] : (s(bool,'F_0') = s(bool,X28) <=> ~p(s(bool,X28)))),
  file('thm.stackSem.result_distinct|split|6','thm.bool.EQ_CLAUSES|split|2')).
fof(f32358,plain,(
  ( ! [X423,X427] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X423),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X427))))) )),
  inference(subsumption_resolution,[],[f32188,f32357])).
fof(f32188,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),X423),s('type.wordSem.word_loc'(sK0),sK1))))),s(bool,X427)))) | p(s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X422),s(fun('type.wordSem.word_loc'(sK0),bool),X423),s(fun('type.wordSem.word_loc'(sK0),bool),X424),s(bool,X425),s(bool,X426))))) )),
  inference(superposition,[],[f313,f20961])).
fof(f313,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X41))),s(bool,X42)))) <=> (p(s(bool,X41)) => p(s(bool,X42))))),
  file('thm.stackSem.result_distinct|split|6','==>_HL_0')).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 251
% Passive clauses: 2450
% Generated clauses: 32617
% Final active clauses: 163
% Final passive clauses: 1435
% Input formulas: 69
% Initial clauses: 103
% 
% Duplicate literals: 914
% Trivial inequalities: 151
% Fw subsumption resolutions: 1565
% Fw demodulations: 5702
% Bw demodulations: 780
% Condensations: 1520
% Global subsumptions: 13
% 
% Simple tautologies: 226
% Equational tautologies: 279
% Forward subsumptions: 19400
% Fw demodulations to eq. taut.: 78
% Bw demodulations to eq. taut.: 51
% 
% Unit resulting resolution: 2467
% Factoring: 1
% Forward superposition: 9746
% Backward superposition: 8742
% Self superposition: 198
% Equality factoring: 86
% Equality resolution: 27
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 2657
% SAT solver unit clauses: 866
% SAT solver binary clauses: 988
% 
% Sat splits: 14
% Sat splitting refutations: 164
% 
% Memory used [KB]: 29167
% Time elapsed: 1.336 s
% ------------------------------
% ------------------------------
% Success in time 2.67 s
