Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.stackSem
% (39257)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 174
% Passive clauses: 3370
% Generated clauses: 19220
% Final active clauses: 155
% Final passive clauses: 3008
% Input formulas: 69
% Initial clauses: 101
% 
% Duplicate literals: 553
% Trivial inequalities: 325
% Fw subsumption resolutions: 1524
% Fw demodulations: 2615
% Bw demodulations: 19
% 
% Simple tautologies: 269
% Equational tautologies: 45
% Forward subsumptions: 10075
% 
% Binary resolution: 513
% Factoring: 1
% Forward superposition: 6628
% Backward superposition: 6641
% Self superposition: 174
% Equality factoring: 70
% Equality resolution: 31
% 
% Split clauses: 229
% Split components: 460
% SAT solver clauses: 632
% SAT solver unit clauses: 9
% SAT solver binary clauses: 133
% 
% Sat splits: 229
% 
% TWLsolver clauses: 632
% TWLsolver calls for satisfiability: 172
% 
% Memory used [KB]: 19701
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.stackSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.stackSem
% SZS output start Proof for thm.stackSem
fof(f10426,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f10359,f3255,f4620,f241,f344,f318,f399,f4596,f230,f336,f337,f339,f9514,f9621,f9894,f229,f334,f317,f398])).
fof(f398,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) ) | $spl108),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X15] : ! [X23,X31,X32,X33,X34] : s(X15,X34) = s(X15,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),X15),X23),s(fun('type.wordSem.word_loc'(X6),X15),X31),s(fun('type.wordSem.word_loc'(X6),X15),X32),s(X15,X33),s(X15,X34)))),
  file('thm.stackSem.result_distinct|split|9','thm.stackSem.result_case_def|split|4')).
fof(f334,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Error_0') | $spl0),
  inference(cnf_transformation,[],[f334_D])).
fof(f334_D,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Error_0') <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f229,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Error_0')),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(sK0),'const.stackSem.Error_0')),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f186])).
fof(f186,plain,(
  ? [X0] : s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X0),'const.stackSem.Error_0')),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0] : s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') != s('type.stackSem.result'(X0),'const.stackSem.Error_0')),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ~s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X0),'const.stackSem.Error_0')),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X6),'const.stackSem.Error_0')),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0') = s('type.stackSem.result'(X6),'const.stackSem.Error_0')),
  file('thm.stackSem.result_distinct|split|9','thm.stackSem.result_distinct|split|9')).
fof(f9894,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(fun('type.wordSem.word_loc'(sK0),X0),X4),s(X0,X5),s(X0,X1)))) ) | $spl1196),
  inference(cnf_transformation,[],[f9894_D])).
fof(f9894_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(fun('type.wordSem.word_loc'(sK0),X0),X4),s(X0,X5),s(X0,X1)))) ) <=> ~$spl1196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1196])])).
fof(f9621,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),X0),X2),s(fun('type.wordSem.word_loc'(sK0),X0),X3),s(fun('type.wordSem.word_loc'(sK0),X0),X4),s(X0,X5),s(X0,X1)))) ) | ($spl0 | $spl108)),
  inference(superposition,[],[f398,f334])).
fof(f9514,plain,(
  $false | ($spl2 | $spl574)),
  inference(resolution,[],[f4620,f336])).
fof(f339,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f337,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f233,f230])).
fof(f233,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f37])).
fof(f37,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.stackSem.result_distinct|split|9','thm.bool.NOT_CLAUSES|split|1')).
fof(f336,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f336_D])).
fof(f336_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f230,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.stackSem.result_distinct|split|9',aHL_TRUTH)).
fof(f4596,plain,(
  $false | ($spl2 | $spl5 | $spl426)),
  inference(subsumption_resolution,[],[f4511,f336])).
fof(f4511,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl426)),
  inference(backward_demodulation,[],[f3255,f339])).
fof(f399,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) ) | $spl110),
  inference(cnf_transformation,[],[f399_D])).
fof(f399_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f318,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),X1),X2),s(fun('type.wordSem.word_loc'(X0),X1),X3),s(fun('type.wordSem.word_loc'(X0),X1),X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6,X15] : ! [X23,X31,X32,X33,X34] : s(X15,X33) = s(X15,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(X6),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),X15),X23),s(fun('type.wordSem.word_loc'(X6),X15),X31),s(fun('type.wordSem.word_loc'(X6),X15),X32),s(X15,X33),s(X15,X34)))),
  file('thm.stackSem.result_distinct|split|9','thm.stackSem.result_case_def|split|3')).
fof(f344,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f241,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.stackSem.result_distinct|split|9','thm.bool.EQ_CLAUSES|split|0')).
fof(f4620,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | $spl574),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f3255,plain,(
  ( ! [X72] : (s(bool,'T_0') = s(bool,X72)) ) | $spl426),
  inference(cnf_transformation,[],[f3255_D])).
fof(f3255_D,plain,(
  ( ! [X72] : (s(bool,'T_0') = s(bool,X72)) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f10359,plain,(
  ( ! [X187,X188] : (s(bool,'T_0') = s(bool,X187) | ~p(s(bool,X188))) ) | ($spl8 | $spl110 | $spl1196)),
  inference(forward_demodulation,[],[f10230,f399])).
fof(f10230,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,X188)) | s(bool,'T_0') = s(bool,'const.stackSem.result_CASE_6'(s('type.stackSem.result'(sK0),'const.stackSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK0),bool),X184),s(fun('type.wordSem.word_loc'(sK0),bool),X185),s(fun('type.wordSem.word_loc'(sK0),bool),X186),s(bool,X187),s(bool,X188)))) ) | ($spl8 | $spl1196)),
  inference(superposition,[],[f344,f9894])).
% SZS output end Proof for thm.stackSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 109
% Passive clauses: 725
% Generated clauses: 10166
% Final active clauses: 91
% Final passive clauses: 470
% Input formulas: 69
% Initial clauses: 101
% 
% Duplicate literals: 413
% Trivial inequalities: 162
% Fw subsumption resolutions: 749
% Fw demodulations: 1007
% Bw demodulations: 148
% Condensations: 974
% Global subsumptions: 8
% 
% Simple tautologies: 163
% Equational tautologies: 43
% Forward subsumptions: 4629
% Bw demodulations to eq. taut.: 11
% 
% Binary resolution: 200
% Factoring: 1
% Forward superposition: 2917
% Backward superposition: 2528
% Self superposition: 126
% Equality factoring: 74
% Equality resolution: 20
% 
% Split clauses: 83
% Split components: 166
% SAT solver clauses: 1408
% SAT solver unit clauses: 151
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 17
% 
% Sat splits: 83
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1317
% TWLsolver calls for satisfiability: 337
% 
% Memory used [KB]: 10362
% Time elapsed: 0.300 s
% ------------------------------
% ------------------------------
% Success in time 0.91 s
