Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relation
% (94825)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 560
% Passive clauses: 3529
% Generated clauses: 14639
% Final active clauses: 553
% Final passive clauses: 2246
% Input formulas: 31
% Initial clauses: 146
% 
% Duplicate literals: 619
% Trivial inequalities: 39
% Fw subsumption resolutions: 1940
% Fw demodulations: 281
% Bw demodulations: 6
% 
% Simple tautologies: 220
% Equational tautologies: 59
% Forward subsumptions: 5911
% 
% Binary resolution: 4188
% Factoring: 1068
% Forward superposition: 1950
% Backward superposition: 3234
% Self superposition: 108
% Equality factoring: 108
% Equality resolution: 6
% 
% Split clauses: 1616
% Split components: 3402
% SAT solver clauses: 1967
% SAT solver unit clauses: 34
% SAT solver binary clauses: 202
% 
% Sat splits: 1616
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1960
% TWLsolver calls for satisfiability: 371
% 
% Memory used [KB]: 6524
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relation
% SZS output start Proof for thm.clos_relation
fof(f4567,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4564,f187,f345,f453,f457,f185,f341,f201,f353,f455,f459,f184,f343,f1616,f1617,f346,f348,f227,f361,f474,f479,f186,f339,f205,f355])).
fof(f355,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f355_D])).
fof(f355_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f205,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.bool.EQ_CLAUSES|split|2')).
fof(f339,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f186,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))))),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))) & ? [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X1)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X4))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))) => ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X1)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X4))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2)))) => ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X1)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X4))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X7))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X8)))) => ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X7)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X10))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X7))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X8)))) => ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X7)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X10))),s('type.list.list'('type.closSem.v'),X9),s('type.list.list'('type.closSem.v'),X8))))))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.clos_relation.val_rel_mono_list|39|')).
fof(f479,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) | $spl188),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f474,plain,(
  s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK16))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) | ($spl1 | $spl16)),
  inference(resolution,[],[f355,f339])).
fof(f361,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1))))) ) | $spl20),
  inference(cnf_transformation,[],[f361_D])).
fof(f361_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f227,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num',X2))),s('type.list.list'('type.closSem.v'),X3),s('type.list.list'('type.closSem.v'),X4)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X10,X12,X13,X14] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X10)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X10))),s('type.list.list'('type.closSem.v'),X13),s('type.list.list'('type.closSem.v'),X14))))) => p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X6),'const.bool.the_value_0'),s('type.num.num',X12))),s('type.list.list'('type.closSem.v'),X13),s('type.list.list'('type.closSem.v'),X14)))))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.clos_relation.val_rel_mono_list')).
fof(f348,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f346,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f190,f187])).
fof(f190,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.bool.NOT_CLAUSES|split|2')).
fof(f1617,plain,(
  ( ! [X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',X1))))) ) | $spl554),
  inference(cnf_transformation,[],[f1617_D])).
fof(f1617_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',X1))))) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f1616,plain,(
  ( ! [X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',X1))))) ) | ($spl9 | $spl20 | $spl188)),
  inference(subsumption_resolution,[],[f1605,f348])).
fof(f1605,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',X1))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',X1))))) ) | ($spl20 | $spl188)),
  inference(superposition,[],[f361,f479])).
fof(f343,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) | $spl4),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f184,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))))),
  inference(cnf_transformation,[],[f119])).
fof(f459,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) | $spl182),
  inference(cnf_transformation,[],[f459_D])).
fof(f459_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f455,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK12),'const.bool.the_value_0'),s('type.num.num',sK13))),s('type.list.list'('type.closSem.v'),sK15),s('type.list.list'('type.closSem.v'),sK14))) | ($spl4 | $spl12)),
  inference(resolution,[],[f353,f343])).
fof(f353,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f201,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.bool.EQ_CLAUSES|split|0')).
fof(f341,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13)))) | $spl2),
  inference(cnf_transformation,[],[f341_D])).
fof(f341_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f185,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13))))),
  inference(cnf_transformation,[],[f119])).
fof(f457,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13))) | $spl180),
  inference(cnf_transformation,[],[f457_D])).
fof(f457_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f453,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13))) | ($spl2 | $spl12)),
  inference(resolution,[],[f353,f341])).
fof(f345,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f345_D])).
fof(f345_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f187,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relation.val_rel_mono_list|39|','thm.bool.TRUTH')).
fof(f4564,plain,(
  $false | ($spl6 | $spl180 | $spl182 | $spl554)),
  inference(subsumption_resolution,[],[f4563,f345])).
fof(f4563,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl180 | $spl182 | $spl554)),
  inference(forward_demodulation,[],[f4562,f457])).
fof(f4562,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13)))) | ($spl6 | $spl182 | $spl554)),
  inference(subsumption_resolution,[],[f4531,f345])).
fof(f4531,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK13)))) | ($spl182 | $spl554)),
  inference(superposition,[],[f1617,f459])).
% SZS output end Proof for thm.clos_relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 275
% Passive clauses: 865
% Generated clauses: 4288
% Final active clauses: 270
% Final passive clauses: 590
% Input formulas: 31
% Initial clauses: 146
% 
% Duplicate literals: 153
% Trivial inequalities: 25
% Fw subsumption resolutions: 461
% Fw demodulations: 76
% Bw demodulations: 4
% Condensations: 99
% Global subsumptions: 13
% 
% Simple tautologies: 207
% Equational tautologies: 15
% Forward subsumptions: 1274
% 
% Binary resolution: 1408
% Factoring: 66
% Forward superposition: 465
% Backward superposition: 451
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 5
% 
% Split clauses: 57
% Split components: 119
% SAT solver clauses: 1270
% SAT solver unit clauses: 167
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 57
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1217
% TWLsolver calls for satisfiability: 1540
% 
% Memory used [KB]: 4349
% Time elapsed: 0.160 s
% ------------------------------
% ------------------------------
% Success in time 0.78 s
