Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f8837,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f8803])).
fof(f8803,plain,(
  s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),sK1))) != s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),sK1)))),
  inference(superposition,[],[f131,f8586])).
fof(f8586,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X2))) = s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X1),s('type.closSem.state'(X0),X2)))))) )),
  inference(forward_demodulation,[],[f8547,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) = s(fun('type.num.num',bool),'const.pred_set.GSPEC_1'(s(fun('type.num.num','type.pair.prod'('type.num.num',bool)),'_dst_x1x1_2'(s(fun('type.num.num',fun(bool,'type.pair.prod'('type.num.num',bool))),'const.pair.,_0'),s(fun('type.num.num',bool),'_dst_x0x1_2'(s(fun(fun('type.closSem.v',bool),bool),'?_0'),s(fun('type.num.num',fun('type.closSem.v',bool)),'_dst_x00x10x01_3'(s(fun('type.option.option'('type.option.option'('type.closSem.v')),fun('type.option.option'('type.option.option'('type.closSem.v')),bool)),'=_0'),s(fun('type.num.num','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.list.list'('type.option.option'('type.closSem.v')),'type.option.option'('type.option.option'('type.closSem.v')))),'const.closSem.get_global_0'),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))),s(fun('type.closSem.v','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x0x1_2'(s(fun('type.option.option'('type.closSem.v'),'type.option.option'('type.option.option'('type.closSem.v'))),'const.option.SOME_0'),s(fun('type.closSem.v','type.option.option'('type.closSem.v')),'const.option.SOME_0')))))))))))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) = s(fun('type.num.num',bool),'const.pred_set.GSPEC_1'(s(fun('type.num.num','type.pair.prod'('type.num.num',bool)),'_dst_x1x1_2'(s(fun('type.num.num',fun(bool,'type.pair.prod'('type.num.num',bool))),'const.pair.,_0'),s(fun('type.num.num',bool),'_dst_x0x1_2'(s(fun(fun('type.closSem.v',bool),bool),'?_0'),s(fun('type.num.num',fun('type.closSem.v',bool)),'_dst_x00x10x01_3'(s(fun('type.option.option'('type.option.option'('type.closSem.v')),fun('type.option.option'('type.option.option'('type.closSem.v')),bool)),'=_0'),s(fun('type.num.num','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.list.list'('type.option.option'('type.closSem.v')),'type.option.option'('type.option.option'('type.closSem.v')))),'const.closSem.get_global_0'),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))),s(fun('type.closSem.v','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x0x1_2'(s(fun('type.option.option'('type.closSem.v'),'type.option.option'('type.option.option'('type.closSem.v'))),'const.option.SOME_0'),s(fun('type.closSem.v','type.option.option'('type.closSem.v')),'const.option.SOME_0')))))))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) = s(fun('type.num.num',bool),'const.pred_set.GSPEC_1'(s(fun('type.num.num','type.pair.prod'('type.num.num',bool)),'_dst_x1x1_2'(s(fun('type.num.num',fun(bool,'type.pair.prod'('type.num.num',bool))),'const.pair.,_0'),s(fun('type.num.num',bool),'_dst_x0x1_2'(s(fun(fun('type.closSem.v',bool),bool),'?_0'),s(fun('type.num.num',fun('type.closSem.v',bool)),'_dst_x00x10x01_3'(s(fun('type.option.option'('type.option.option'('type.closSem.v')),fun('type.option.option'('type.option.option'('type.closSem.v')),bool)),'=_0'),s(fun('type.num.num','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.list.list'('type.option.option'('type.closSem.v')),'type.option.option'('type.option.option'('type.closSem.v')))),'const.closSem.get_global_0'),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))))),s(fun('type.closSem.v','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x0x1_2'(s(fun('type.option.option'('type.closSem.v'),'type.option.option'('type.option.option'('type.closSem.v'))),'const.option.SOME_0'),s(fun('type.closSem.v','type.option.option'('type.closSem.v')),'const.option.SOME_0')))))))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X6),X7))) = s(fun('type.num.num',bool),'const.pred_set.GSPEC_1'(s(fun('type.num.num','type.pair.prod'('type.num.num',bool)),'_dst_x1x1_2'(s(fun('type.num.num',fun(bool,'type.pair.prod'('type.num.num',bool))),'const.pair.,_0'),s(fun('type.num.num',bool),'_dst_x0x1_2'(s(fun(fun('type.closSem.v',bool),bool),'?_0'),s(fun('type.num.num',fun('type.closSem.v',bool)),'_dst_x00x10x01_3'(s(fun('type.option.option'('type.option.option'('type.closSem.v')),fun('type.option.option'('type.option.option'('type.closSem.v')),bool)),'=_0'),s(fun('type.num.num','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.list.list'('type.option.option'('type.closSem.v')),'type.option.option'('type.option.option'('type.closSem.v')))),'const.closSem.get_global_0'),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X6),X7))))),s(fun('type.closSem.v','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x0x1_2'(s(fun('type.option.option'('type.closSem.v'),'type.option.option'('type.option.option'('type.closSem.v'))),'const.option.SOME_0'),s(fun('type.closSem.v','type.option.option'('type.closSem.v')),'const.option.SOME_0')))))))))))),
  file('thm.clos_knownProof.mapped_globals_dec_clock','thm.clos_knownProof.mapped_globals_def')).
fof(f8547,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X1),s('type.closSem.state'(X0),X2))))) = s(fun('type.num.num',bool),'const.pred_set.GSPEC_1'(s(fun('type.num.num','type.pair.prod'('type.num.num',bool)),'_dst_x1x1_2'(s(fun('type.num.num',fun(bool,'type.pair.prod'('type.num.num',bool))),'const.pair.,_0'),s(fun('type.num.num',bool),'_dst_x0x1_2'(s(fun(fun('type.closSem.v',bool),bool),'?_0'),s(fun('type.num.num',fun('type.closSem.v',bool)),'_dst_x00x10x01_3'(s(fun('type.option.option'('type.option.option'('type.closSem.v')),fun('type.option.option'('type.option.option'('type.closSem.v')),bool)),'=_0'),s(fun('type.num.num','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.list.list'('type.option.option'('type.closSem.v')),'type.option.option'('type.option.option'('type.closSem.v')))),'const.closSem.get_global_0'),s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X2))))),s(fun('type.closSem.v','type.option.option'('type.option.option'('type.closSem.v'))),'_dst_x0x1_2'(s(fun('type.option.option'('type.closSem.v'),'type.option.option'('type.option.option'('type.closSem.v'))),'const.option.SOME_0'),s(fun('type.closSem.v','type.option.option'('type.closSem.v')),'const.option.SOME_0')))))))))))) )),
  inference(superposition,[],[f137,f1695])).
fof(f1695,plain,(
  ( ! [X10,X11,X9] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),X10))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X9),'const.closSem.dec_clock_2'(s('type.num.num',X11),s('type.closSem.state'(X9),X10)))))) )),
  inference(superposition,[],[f141,f156])).
fof(f156,plain,(
  ( ! [X2,X0,X1] : (s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X1),s('type.closSem.state'(X0),X2))) = s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.num.num',X1))))),s('type.closSem.state'(X0),X2)))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X1),s('type.closSem.state'(X0),X2))) = s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.num.num',X1))))),s('type.closSem.state'(X0),X2)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X1),s('type.closSem.state'(X0),X2))) = s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.num.num',X1))))),s('type.closSem.state'(X0),X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : ! [X8,X7] : s('type.closSem.state'(X10),'const.closSem.dec_clock_2'(s('type.num.num',X8),s('type.closSem.state'(X10),X7))) = s('type.closSem.state'(X10),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),'const.combin.K_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X10),X7))),s('type.num.num',X8))))),s('type.closSem.state'(X10),X7)))),
  file('thm.clos_knownProof.mapped_globals_dec_clock','thm.closSem.dec_clock_def')).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),X1))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X0),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ! [X7,X11] : s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X10),X7))) = s('type.list.list'('type.option.option'('type.closSem.v')),'const.closSem.state_globals_1'(s('type.closSem.state'(X10),'const.closSem.state_clock_fupd_2'(s(fun('type.num.num','type.num.num'),X11),s('type.closSem.state'(X10),X7)))))),
  file('thm.clos_knownProof.mapped_globals_dec_clock','thm.closSem.closSem_65_prove_ltb')).
fof(f131,plain,(
  s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),sK1))) != s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),'const.closSem.dec_clock_2'(s('type.num.num',sK2),s('type.closSem.state'(sK0),sK1)))))),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),sK1))) != s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(sK0),'const.closSem.dec_clock_2'(s('type.num.num',sK2),s('type.closSem.state'(sK0),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f122])).
fof(f122,plain,(
  ? [X0,X1,X2] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) != s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X2),s('type.closSem.state'(X0),X1)))))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1,X2] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) = s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X2),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ~! [X0] : ! [X1,X2] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),X1))) = s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X0),'const.closSem.dec_clock_2'(s('type.num.num',X2),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X6),X7))) = s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X6),'const.closSem.dec_clock_2'(s('type.num.num',X8),s('type.closSem.state'(X6),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X6),X7))) = s(fun('type.num.num',bool),'const.clos_knownProof.mapped_globals_1'(s('type.closSem.state'(X6),'const.closSem.dec_clock_2'(s('type.num.num',X8),s('type.closSem.state'(X6),X7)))))),
  file('thm.clos_knownProof.mapped_globals_dec_clock','thm.clos_knownProof.mapped_globals_dec_clock')).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 124
% Passive clauses: 1790
% Generated clauses: 8498
% Final active clauses: 88
% Final passive clauses: 1446
% Input formulas: 43
% Initial clauses: 46
% 
% Duplicate literals: 437
% Trivial inequalities: 60
% Fw subsumption resolutions: 423
% Fw demodulations: 915
% Bw demodulations: 25
% 
% Simple tautologies: 74
% Equational tautologies: 22
% Forward subsumptions: 4439
% 
% Binary resolution: 106
% Forward superposition: 3315
% Backward superposition: 2905
% Self superposition: 132
% Equality factoring: 58
% Equality resolution: 10
% 
% Split clauses: 231
% Split components: 471
% SAT solver clauses: 372
% SAT solver unit clauses: 13
% SAT solver binary clauses: 90
% 
% Sat splits: 231
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 370
% TWLsolver calls for satisfiability: 81
% 
% Memory used [KB]: 11257
% Time elapsed: 0.481 s
% ------------------------------
% ------------------------------
% Success in time 0.49 s
