Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordProps
% SZS output start Proof for thm.wordProps
fof(f872,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f868,f80,f118,f612,f611,f591,f593,f489,f547])).
fof(f547,plain,(
  $false | $spl10),
  inference(trivial_inequality_removal,[],[f546])).
fof(f546,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl10),
  inference(superposition,[],[f128,f489])).
fof(f128,plain,(
  s(bool,'T_0') != s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))),
  inference(resolution,[],[f73,f82])).
fof(f82,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f73,plain,(
  ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))))),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(sK0),sK1),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2))),s('type.list.list'('type.wordSem.stack_frame'(sK0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(sK0),sK3),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))))) & ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X1),s('type.list.list'('type.wordSem.stack_frame'(X0)),X2))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))) & ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X1),s('type.list.list'('type.wordSem.stack_frame'(X0)),X2))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))) => p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X1),s('type.list.list'('type.wordSem.stack_frame'(X0)),X2))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))) => p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X7),s('type.list.list'('type.wordSem.stack_frame'(X6)),X8))),s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X9),s('type.list.list'('type.wordSem.stack_frame'(X6)),X10)))))) => p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),X8),s('type.list.list'('type.wordSem.stack_frame'(X6)),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X7),s('type.list.list'('type.wordSem.stack_frame'(X6)),X8))),s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X9),s('type.list.list'('type.wordSem.stack_frame'(X6)),X10)))))) => p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),X8),s('type.list.list'('type.wordSem.stack_frame'(X6)),X10)))))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb','thm.wordProps.wordProps_|49|67_prove_ltb')).
fof(f489,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl10),
  inference(cnf_transformation,[],[f489_D])).
fof(f489_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f593,plain,(
  ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))) | $spl15),
  inference(cnf_transformation,[],[f593_D])).
fof(f593_D,plain,(
  ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f591,plain,(
  ( ! [X13] : (~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f585,f81])).
fof(f81,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb',aHL_BOOL_CASES)).
fof(f585,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | ~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4)))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f86,f303])).
fof(f303,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(trivial_inequality_removal,[],[f271])).
fof(f271,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X20) = s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f128,f129])).
fof(f129,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f81,f81])).
fof(f86,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f611,plain,(
  s(bool,'F_0') = s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))) | $spl16),
  inference(cnf_transformation,[],[f611_D])).
fof(f611_D,plain,(
  s(bool,'F_0') = s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f612,plain,(
  s(bool,'F_0') = s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))) | $spl15),
  inference(resolution,[],[f593,f87])).
fof(f87,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f118,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f80,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb',aHL_FALSITY)).
fof(f868,plain,(
  $false | ($spl3 | $spl16)),
  inference(subsumption_resolution,[],[f867,f118])).
fof(f867,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(forward_demodulation,[],[f857,f611])).
fof(f857,plain,(
  p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))))),
  inference(resolution,[],[f109,f72])).
fof(f72,plain,(
  p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(sK0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(sK0),sK1),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK2))),s('type.list.list'('type.wordSem.stack_frame'(sK0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(sK0),sK3),s('type.list.list'('type.wordSem.stack_frame'(sK0)),sK4))))))),
  inference(cnf_transformation,[],[f50])).
fof(f109,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))))) | p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))))) | (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) & p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2)))))) & (~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) | ~p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2)))) | p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1))))))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))))) | (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) & p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2)))))) & ((~p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) | ~p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2))))) | p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1))))))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))))) <=> (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) & p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.list.list'('type.wordSem.stack_frame'(X0)),X3))),s('type.list.list'('type.wordSem.stack_frame'(X0)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X0),X2),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))))) <=> (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X0)),X3),s('type.list.list'('type.wordSem.stack_frame'(X0)),X1)))) & p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X0),X4),s('type.wordSem.stack_frame'(X0),X2))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X11,X12,X13,X14] : (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X14),s('type.list.list'('type.wordSem.stack_frame'(X6)),X13))),s('type.list.list'('type.wordSem.stack_frame'(X6)),'const.list.CONS_2'(s('type.wordSem.stack_frame'(X6),X12),s('type.list.list'('type.wordSem.stack_frame'(X6)),X11)))))) <=> (p(s(bool,'const.wordProps.s_val_eq_2'(s('type.list.list'('type.wordSem.stack_frame'(X6)),X13),s('type.list.list'('type.wordSem.stack_frame'(X6)),X11)))) & p(s(bool,'const.wordProps.s_frame_val_eq_2'(s('type.wordSem.stack_frame'(X6),X14),s('type.wordSem.stack_frame'(X6),X12))))))),
  file('thm.wordProps.wordProps_|49|67_prove_ltb','thm.wordProps.s_val_eq_def|split|1')).
% SZS output end Proof for thm.wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 36
% Passive clauses: 100
% Generated clauses: 770
% Final active clauses: 27
% Final passive clauses: 33
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 80
% Trivial inequalities: 29
% Fw subsumption resolutions: 69
% Fw demodulations: 16
% Bw demodulations: 8
% 
% Simple tautologies: 67
% Equational tautologies: 17
% Forward subsumptions: 383
% 
% Binary resolution: 35
% Forward superposition: 242
% Backward superposition: 167
% Self superposition: 30
% Equality factoring: 28
% Equality resolution: 9
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 40
% SAT solver unit clauses: 7
% SAT solver binary clauses: 26
% SAT solver learnt clauses: 1
% 
% Sat splits: 19
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 39
% TWLsolver calls for satisfiability: 15
% 
% Memory used [KB]: 895
% Time elapsed: 0.024 s
% ------------------------------
% ------------------------------
% Success in time 0.018 s
