Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pegComplete
% (131708)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 845
% Passive clauses: 12443
% Generated clauses: 27068
% Final active clauses: 835
% Final passive clauses: 10949
% Input formulas: 42
% Initial clauses: 180
% 
% Duplicate literals: 713
% Trivial inequalities: 102
% Fw subsumption resolutions: 2494
% Fw demodulations: 574
% Bw demodulations: 7
% 
% Simple tautologies: 282
% Equational tautologies: 60
% Forward subsumptions: 8879
% 
% Binary resolution: 13270
% Factoring: 1499
% Forward superposition: 4387
% Backward superposition: 3453
% Self superposition: 38
% Equality factoring: 38
% Equality resolution: 38
% 
% Split clauses: 931
% Split components: 1873
% SAT solver clauses: 1454
% SAT solver unit clauses: 20
% SAT solver binary clauses: 217
% 
% Sat splits: 931
% 
% TWLsolver clauses: 1454
% TWLsolver calls for satisfiability: 491
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pegComplete
% (131741)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 381
% Passive clauses: 2509
% Generated clauses: 9809
% Final active clauses: 376
% Final passive clauses: 2105
% Input formulas: 42
% Initial clauses: 180
% 
% Duplicate literals: 189
% Trivial inequalities: 11
% Fw subsumption resolutions: 689
% Fw demodulations: 63
% Bw demodulations: 24
% Condensations: 249
% Global subsumptions: 13
% 
% Simple tautologies: 220
% Equational tautologies: 24
% Forward subsumptions: 3033
% 
% Binary resolution: 3861
% Factoring: 348
% Forward superposition: 825
% Backward superposition: 782
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 2959
% SAT solver unit clauses: 183
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 32
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2907
% TWLsolver calls for satisfiability: 1883
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pegComplete
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pegComplete
% SZS output start Proof for thm.pegComplete
fof(f2455,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f245,f433,f432,f244,f435,f243,f437,f242,f439,f431,f1201,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1241,f1244,f1247,f1250,f1280,f1202,f1282,f1206,f1284,f1209,f1286,f1212,f1288,f1215,f1290,f1218,f1292,f1221,f1294,f1224,f1296,f1227,f1298,f1230,f1300,f1233,f1302,f1236,f1304,f1239,f1306,f1242,f1308,f1245,f1310,f1248,f1312,f1251,f1325,f1326,f1658,f1661,f1659,f1678,f1681,f1203,f1679,f1814,f1817,f1815,f1819,f1820,f1911,f1912,f2248,f2250,f2249,f2252,f1816,f2253,f2348,f2357,f2358,f2359,f2360,f2396,f2401,f2402,f2403,f2450])).
fof(f2450,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2435,f2404])).
fof(f2404,plain,(
  s('type.list.list'(sK13),sK16) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK26(sK18,sK17,sK16,sK15,sK13,sK14)),s('type.list.list'(sK13),sK27(sK18,sK17,sK16,sK15,sK13,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f431,f334])).
fof(f334,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK26(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK27(X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK26(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK27(X5,X4,X3,X2,X1,X0)))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,sK26(X5,X4,X3,X2,X1,X0)))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),sK27(X5,X4,X3,X2,X1,X0)))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | ? [X8,X9] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X8),s('type.list.list'(X1),X9))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X8))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X9))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7))))) & (! [X6,X7] : (s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) <=> ? [X6,X7] : (s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7))) & s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) & s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X11),s('type.list.list'(X7),X9))) <=> ? [X21,X13] : (s('type.list.list'(X7),X9) = s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X21),s('type.list.list'(X7),X13))) & s(X6,X10) = s(X6,happ(s(fun(X7,X6),X11),s(X7,X21))) & s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X11),s('type.list.list'(X7),X13)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.list.MAP_EQ_CONS')).
fof(f2435,plain,(
  s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK26(sK18,sK17,sK16,sK15,sK13,sK14)),s('type.list.list'(sK13),sK27(sK18,sK17,sK16,sK15,sK13,sK14)))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f2406,f2405,f433])).
fof(f2405,plain,(
  s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK27(sK18,sK17,sK16,sK15,sK13,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f431,f336])).
fof(f336,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),sK27(X5,X4,X3,X2,X1,X0)))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f2406,plain,(
  s(sK14,sK17) = s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK26(sK18,sK17,sK16,sK15,sK13,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f431,f335])).
fof(f335,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X5),s(X1,sK26(X5,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f201])).
fof(f2403,plain,(
  $false | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2392,f435])).
fof(f2392,plain,(
  s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))) | ($spl2 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f439,f437,f433])).
fof(f2402,plain,(
  $false | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2393,f439])).
fof(f2393,plain,(
  s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) | ($spl2 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f437,f435,f433])).
fof(f2401,plain,(
  $false | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2395,f437])).
fof(f2395,plain,(
  s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) | ($spl2 | $spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f439,f435,f433])).
fof(f2396,plain,(
  $false | ($spl2 | $spl4 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f437,f439,f435,f433])).
fof(f2360,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2344,f432])).
fof(f2344,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) | ($spl4 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f437,f439,f435,f337])).
fof(f337,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X7))) | s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X1,X0),X5),s('type.list.list'(X1),X3))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X5),s(X1,X6))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X7)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f2359,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2345,f435])).
fof(f2345,plain,(
  s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))) | ($spl1 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f439,f437,f432,f337])).
fof(f2358,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2346,f439])).
fof(f2346,plain,(
  s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) | ($spl1 | $spl4 | $spl6)),
  inference(unit_resulting_resolution,[],[f437,f435,f432,f337])).
fof(f2357,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2347,f437])).
fof(f2347,plain,(
  s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) | ($spl1 | $spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f439,f435,f432,f337])).
fof(f2348,plain,(
  $false | ($spl1 | $spl4 | $spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f437,f439,f435,f432,f337])).
fof(f2253,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl58),
  inference(cnf_transformation,[],[f2253_D])).
fof(f2253_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1816,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl52),
  inference(cnf_transformation,[],[f1816_D])).
fof(f1816_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2252,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f2246,f260])).
fof(f260,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f2246,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1257,f253])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb',aHL_BOOL_CASES)).
fof(f1257,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f981,f246])).
fof(f246,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb',aHL_TRUTH)).
fof(f981,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f328,f565])).
fof(f565,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f253,f253])).
fof(f328,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f195])).
fof(f195,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2249,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl56),
  inference(cnf_transformation,[],[f2249_D])).
fof(f2249_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2250,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f2242,f260])).
fof(f2242,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1257,f565])).
fof(f2248,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2241,f260])).
fof(f2241,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1257,f565])).
fof(f1912,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f1845,f246])).
fof(f1845,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(backward_demodulation,[],[f1842,f440])).
fof(f440,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f249,f246])).
fof(f249,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f1842,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,X27)) ) | $spl52),
  inference(global_subsumption,[],[f245,f244,f243,f242,f246,f253,f261,f260,f303,f300,f310,f309,f442,f328,f325,f324,f323,f333,f337,f336,f335,f334,f343,f342,f341,f425,f339,f338,f354,f353,f352,f351,f350,f349,f348,f358,f357,f356,f355,f362,f361,f360,f359,f366,f365,f364,f363,f370,f369,f368,f367,f424,f374,f373,f372,f371,f381,f380,f379,f378,f387,f384,f383,f382,f394,f393,f392,f423,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f411,f410,f409,f422,f407,f406,f405,f415,f414,f413,f412,f419,f418,f417,f416,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f463,f464,f471,f472,f474,f480,f485,f487,f492,f500,f503,f504,f506,f513,f517,f518,f521,f527,f528,f532,f535,f538,f565,f617,f626,f625,f636,f634,f648,f658,f664,f668,f673,f677,f684,f693,f711,f718,f727,f747,f754,f763,f786,f795,f844,f938,f1201,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1241,f1244,f1247,f1250,f1252,f1253,f1254,f1255,f1256,f1257,f1258,f1259,f1260,f1261,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1280,f1282,f1284,f1286,f1288,f1290,f1292,f1294,f1296,f1298,f1300,f1302,f1304,f1306,f1308,f1310,f1312,f288,f441,f294,f292,f299,f297,f1274,f440,f265,f264,f1586,f1585,f1584,f1583,f1587,f1589,f1591,f1592,f1594,f1596,f1606,f1607,f1608,f1609,f1611,f1612,f1613,f1614,f1615,f1616,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1656,f1658,f1661,f1655,f1676,f1678,f1681,f1675,f1688,f1690,f1692,f1698,f1705,f1707,f1709,f1715,f1722,f1724,f1726,f1727,f1729,f1731,f1741,f1743,f1745,f1746,f1748,f1758,f1766,f1767,f1768,f1769,f1772,f1784,f1814,f1817,f1818,f1819,f1813,f1816,f1824,f1825])).
fof(f1825,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X23))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X24,bool),X25),s('type.list.list'(X24),X26))) | s('type.list.list'(X24),X26) = s('type.list.list'(X24),'const.list.CONS_2'(s(X24,sK26(X25,X22,X26,X23,X24,bool)),s('type.list.list'(X24),sK27(X25,X22,X26,X23,X24,bool)))) | s(bool,X22) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f334,f565])).
fof(f1824,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (s('type.list.list'('type.list.list'(X12)),'const.list.CONS_2'(s('type.list.list'(X12),X17),s('type.list.list'('type.list.list'(X12)),X18))) != s('type.list.list'('type.list.list'(X12)),'const.list.MAP_2'(s(fun(X19,'type.list.list'(X12)),X20),s('type.list.list'(X19),X21))) | s('type.list.list'(X19),X21) = s('type.list.list'(X19),'const.list.CONS_2'(s(X19,sK26(X20,'const.list.MAP_2'(s(fun(X13,X12),X14),s('type.list.list'(X13),sK27(X14,X15,X16,X17,X13,X12))),X21,X18,X19,'type.list.list'(X12))),s('type.list.list'(X19),sK27(X20,'const.list.MAP_2'(s(fun(X13,X12),X14),s('type.list.list'(X13),sK27(X14,X15,X16,X17,X13,X12))),X21,X18,X19,'type.list.list'(X12))))) | s('type.list.list'(X12),'const.list.CONS_2'(s(X12,X15),s('type.list.list'(X12),X17))) != s('type.list.list'(X12),'const.list.MAP_2'(s(fun(X13,X12),X14),s('type.list.list'(X13),X16)))) )),
  inference(superposition,[],[f334,f336])).
fof(f1813,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X30,bool),X31)))) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31))))) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1256,f288])).
fof(f1818,plain,(
  ( ! [X17,X18,X16] : (sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f1809,f440])).
fof(f1809,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | sP1(X17,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1256,f253])).
fof(f1784,plain,(
  ( ! [X19,X20,X18] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X18,bool),X19))) | s(bool,X20) = s(bool,happ(s(fun(X18,bool),X19),s(X18,'const.min.@_1'(s(fun(X18,bool),X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f1252,f288])).
fof(f1772,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1252,f253])).
fof(f1769,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.list.list'('type.list.list'(X34)),'const.list.CONS_2'(s('type.list.list'(X34),X39),s('type.list.list'('type.list.list'(X34)),X40))) != s('type.list.list'('type.list.list'(X34)),'const.list.MAP_2'(s(fun(X41,'type.list.list'(X34)),X42),s('type.list.list'(X41),X43))) | s('type.list.list'(X34),'const.list.MAP_2'(s(fun(X35,X34),X36),s('type.list.list'(X35),sK27(X36,X37,X38,X39,X35,X34)))) = s('type.list.list'(X34),happ(s(fun(X41,'type.list.list'(X34)),X42),s(X41,sK26(X42,'const.list.MAP_2'(s(fun(X35,X34),X36),s('type.list.list'(X35),sK27(X36,X37,X38,X39,X35,X34))),X43,X40,X41,'type.list.list'(X34))))) | s('type.list.list'(X34),'const.list.CONS_2'(s(X34,X37),s('type.list.list'(X34),X39))) != s('type.list.list'(X34),'const.list.MAP_2'(s(fun(X35,X34),X36),s('type.list.list'(X35),X38)))) )),
  inference(superposition,[],[f335,f336])).
fof(f1768,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.list.list'(X24),'const.list.CONS_2'(s(X24,X30),s('type.list.list'(X24),X29))) != s('type.list.list'(X24),'const.list.MAP_2'(s(fun(X31,X24),X32),s('type.list.list'(X31),X33))) | s(X24,X30) = s(X24,happ(s(fun(X31,X24),X32),s(X31,sK26(X32,X30,X33,'const.list.MAP_2'(s(fun(X25,X24),X26),s('type.list.list'(X25),sK27(X26,X27,X28,X29,X25,X24))),X31,X24)))) | s('type.list.list'(X24),'const.list.CONS_2'(s(X24,X27),s('type.list.list'(X24),X29))) != s('type.list.list'(X24),'const.list.MAP_2'(s(fun(X25,X24),X26),s('type.list.list'(X25),X28)))) )),
  inference(superposition,[],[f335,f336])).
fof(f1767,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.list.list'(X20),'const.list.CONS_2'(s(X20,X21),s('type.list.list'(X20),X22))) != s('type.list.list'(X20),'const.list.MAP_2'(s(fun(X14,X20),X23),s('type.list.list'(X14),X19))) | s(X20,X21) = s(X20,happ(s(fun(X14,X20),X23),s(X14,sK26(X23,X21,'const.list.MAP_2'(s(fun(X15,X14),X16),s('type.list.list'(X15),sK27(X16,X17,X18,X19,X15,X14))),X22,X14,X20)))) | s('type.list.list'(X14),'const.list.CONS_2'(s(X14,X17),s('type.list.list'(X14),X19))) != s('type.list.list'(X14),'const.list.MAP_2'(s(fun(X15,X14),X16),s('type.list.list'(X15),X18)))) )),
  inference(superposition,[],[f335,f336])).
fof(f1766,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (s('type.list.list'(X6),X11) != s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) | s(X6,X12) = s(X6,happ(s(fun(X7,X6),X8),s(X7,sK26(X8,X12,sK27(X8,X9,X10,X11,X7,X6),X13,X7,X6)))) | s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),X11))) != s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X7,X6),X8),s('type.list.list'(X7),X10)))) )),
  inference(superposition,[],[f335,f336])).
fof(f1758,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s(bool,'|dollar|exists_1'(s(fun(X58,bool),X59))) = s(bool,happ(s(fun(X61,bool),X62),s(X61,sK26(X62,happ(s(fun(X58,bool),X59),s(X58,'const.min.@_1'(s(fun(X58,bool),X59)))),X63,X60,X61,bool)))) | s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X61,bool),X62),s('type.list.list'(X61),X63))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'|dollar|exists_1'(s(fun(X58,bool),X59))),s('type.list.list'(bool),X60)))) )),
  inference(forward_demodulation,[],[f1749,f288])).
fof(f1749,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X61,bool),X62),s('type.list.list'(X61),X63))) != s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'|dollar|exists_1'(s(fun(X58,bool),X59))),s('type.list.list'(bool),X60))) | s(bool,happ(s(fun(X58,bool),X59),s(X58,'const.min.@_1'(s(fun(X58,bool),X59))))) = s(bool,happ(s(fun(X61,bool),X62),s(X61,sK26(X62,happ(s(fun(X58,bool),X59),s(X58,'const.min.@_1'(s(fun(X58,bool),X59)))),X63,X60,X61,bool))))) )),
  inference(superposition,[],[f335,f288])).
fof(f1748,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X53),s('type.list.list'(bool),X54))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X55,bool),X56),s('type.list.list'(X55),X57))) | s(bool,'F_0') = s(bool,happ(s(fun(X55,bool),X56),s(X55,sK26(X56,'F_0',X57,X54,X55,bool)))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f335,f253])).
fof(f1746,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X41),s('type.list.list'(bool),X42))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),X45))) | s(bool,'T_0') = s(bool,happ(s(fun(X43,bool),X44),s(X43,sK26(X44,'T_0',X45,X42,X43,bool)))) | s(bool,X41) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f335,f565])).
fof(f1745,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X37))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X38,bool),X39),s('type.list.list'(X38),X40))) | s(bool,X36) = s(bool,happ(s(fun(X38,bool),X39),s(X38,sK26(X39,X36,X40,X37,X38,bool)))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f335,f253])).
fof(f1743,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X25),s('type.list.list'(bool),X26))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X27,bool),X28),s('type.list.list'(X27),X29))) | s(bool,X24) = s(bool,happ(s(fun(X27,bool),X28),s(X27,sK26(X28,X24,X29,X26,X27,bool)))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f335,f565])).
fof(f1741,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X13))) != s('type.list.list'(bool),'const.list.MAP_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),X16))) | s(bool,X12) = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK26(X15,X12,X16,X13,X14,bool)))) | s(bool,X12) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f335,f565])).
fof(f1731,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK25(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK25(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f333,f253])).
fof(f1729,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK25(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK25(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f333,f565])).
fof(f1727,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK25(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK25(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f333,f565])).
fof(f1726,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK25(X16,X17,X18,bool))) )),
  inference(superposition,[],[f333,f253])).
fof(f1724,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK25(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f333,f565])).
fof(f1722,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK25(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f333,f565])).
fof(f1715,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1710,f246])).
fof(f1710,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f299,f565])).
fof(f1709,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | p(s(bool,X17)) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK24(X16,X17,bool))) )),
  inference(superposition,[],[f299,f253])).
fof(f1707,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X10)))) | p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | s(bool,'T_0') = s(bool,sK24(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f299,f565])).
fof(f1705,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,X3) = s(bool,sK24(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f299,f565])).
fof(f1698,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1693,f246])).
fof(f1693,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X1,X2,X0)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f297,f565])).
fof(f1692,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,'F_0')))) | p(s(bool,X17)) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X18)))) | s(bool,'T_0') = s(bool,sK23(X16,X17,bool))) )),
  inference(superposition,[],[f297,f253])).
fof(f1690,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X10)))) | p(s(bool,X9)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X11)))) | s(bool,'T_0') = s(bool,sK23(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f297,f565])).
fof(f1688,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,X3) = s(bool,sK23(X0,X1,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f297,f565])).
fof(f1675,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39)))) | p(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) | ~p(s(bool,happ(s(fun(X40,bool),X41),s(X40,sK22(happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39)))),X41,X40)))))) )),
  inference(superposition,[],[f294,f288])).
fof(f1676,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22(X0,X2,X1))))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1667,f246])).
fof(f1667,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22(X0,X2,X1))))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f294,f565])).
fof(f1655,plain,(
  ( ! [X39,X41,X38,X42,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X38,bool),X39)))) | p(s(bool,happ(s(fun(X40,bool),X41),s(X40,X42)))) | ~p(s(bool,happ(s(fun(X40,bool),X41),s(X40,sK21(happ(s(fun(X38,bool),X39),s(X38,'const.min.@_1'(s(fun(X38,bool),X39)))),X41,X40)))))) )),
  inference(superposition,[],[f292,f288])).
fof(f1656,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21(X0,X2,X1))))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1647,f246])).
fof(f1647,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21(X0,X2,X1))))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f292,f565])).
fof(f1626,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98)))) | ~sP9(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | sP10(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100)) )),
  inference(superposition,[],[f423,f288])).
fof(f1625,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'|dollar|exists_1'(s(fun(X93,bool),X94)))) | ~sP11(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X95,X96) | sP12(happ(s(fun(X93,bool),X94),s(X93,'const.min.@_1'(s(fun(X93,bool),X94)))),X95,X96)) )),
  inference(superposition,[],[f422,f288])).
fof(f1624,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90)))) | ~p(s(bool,X91)) | p(s(bool,X92)) | ~sP12(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X92,X91)) )),
  inference(superposition,[],[f419,f288])).
fof(f1623,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),X86)))) | ~p(s(bool,X87)) | sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88)) )),
  inference(superposition,[],[f418,f288])).
fof(f1622,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),X82)))) | p(s(bool,X83)) | sP12(happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))),X84,X83)) )),
  inference(superposition,[],[f417,f288])).
fof(f1621,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'|dollar|exists_1'(s(fun(X77,bool),X78)))) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~sP11(X80,X79,happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))))) )),
  inference(superposition,[],[f412,f288])).
fof(f1620,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74)))) | p(s(bool,X75)) | p(s(bool,X76)) | ~sP10(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X76,X75)) )),
  inference(superposition,[],[f402,f288])).
fof(f1619,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70)))) | ~p(s(bool,X71)) | sP10(happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))),X71,X72)) )),
  inference(superposition,[],[f401,f288])).
fof(f1618,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66)))) | ~p(s(bool,X67)) | sP10(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X68,X67)) )),
  inference(superposition,[],[f400,f288])).
fof(f1617,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62)))) | p(s(bool,X63)) | p(s(bool,X64)) | ~sP9(happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))),X63,X64)) )),
  inference(superposition,[],[f395,f288])).
fof(f1616,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58)))) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP8(X60,X59,happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))))) )),
  inference(superposition,[],[f387,f288])).
fof(f1615,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X55,happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X56)) )),
  inference(superposition,[],[f378,f288])).
fof(f1614,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52)) )),
  inference(superposition,[],[f359,f288])).
fof(f1613,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'|dollar|exists_1'(s(fun(X45,bool),X46)))) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48)) )),
  inference(superposition,[],[f355,f288])).
fof(f1612,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42)))) | ~p(s(bool,X43)) | sP2(X43,X44,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))))) )),
  inference(superposition,[],[f343,f288])).
fof(f1611,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38)))) | ~p(s(bool,X39)) | sP2(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) )),
  inference(superposition,[],[f342,f288])).
fof(f1609,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'|dollar|exists_1'(s(fun(X30,bool),X31)))) | p(s(bool,X32)) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31)))))) )),
  inference(superposition,[],[f328,f288])).
fof(f1608,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28)))) | p(s(bool,X29)) | sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29)) )),
  inference(superposition,[],[f325,f288])).
fof(f1607,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25)))) | sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26)) )),
  inference(superposition,[],[f310,f288])).
fof(f1606,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X22)))) | sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22)))))) )),
  inference(superposition,[],[f309,f288])).
fof(f1596,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(bool,'T_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,'const.min.@_1'(s(fun(X14,bool),X15)))))) )),
  inference(superposition,[],[f288,f253])).
fof(f1594,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,'const.min.@_1'(s(fun(X8,bool),X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f288,f565])).
fof(f1592,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f288,f565])).
fof(f1591,plain,(
  ( ! [X8] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8))) = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f288,f253])).
fof(f1589,plain,(
  ( ! [X4,X5] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f288,f565])).
fof(f1587,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f288,f565])).
fof(f1583,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f440,f288,f370])).
fof(f1584,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f246,f288,f368])).
fof(f1585,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f440,f288,f370])).
fof(f1586,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f246,f288,f368])).
fof(f264,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f265,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f1274,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'|dollar|exists_1'(s(fun(X109,bool),X110)))) | s(bool,X112) = s(bool,happ(s(fun(X109,bool),X110),s(X109,X111))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f1005,f246])).
fof(f1005,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X109,bool),X110)))) | s(bool,X112) = s(bool,happ(s(fun(X109,bool),X110),s(X109,X111))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f441,f565])).
fof(f297,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7] : ! [X14,X19] : (! [X15] : (p(s(bool,X14)) | p(s(bool,happ(s(fun(X7,bool),X19),s(X7,X15))))) <=> (p(s(bool,X14)) | ! [X15] : p(s(bool,happ(s(fun(X7,bool),X19),s(X7,X15))))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f299,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f292,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2)) | ! [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2)) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (((? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : ((! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7] : ! [X14,X19] : ((! [X15] : p(s(bool,happ(s(fun(X7,bool),X14),s(X7,X15)))) & p(s(bool,X19))) <=> ! [X15] : (p(s(bool,happ(s(fun(X7,bool),X14),s(X7,X15)))) & p(s(bool,X19))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.LEFT_AND_FORALL_THM')).
fof(f294,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f183])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f289,f288])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : ! [X14,X15] : (p(s(bool,happ(s(fun(X7,bool),X14),s(X7,X15)))) => p(s(bool,happ(s(fun(X7,bool),X14),s(X7,'const.min.@_1'(s(fun(X7,bool),X14)))))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.SELECT_AX')).
fof(f288,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.EXISTS_DEF')).
fof(f1273,plain,(
  ( ! [X97,X95,X98,X96] : (~sP9(X95,X96,X97) | sP10(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f998,f246])).
fof(f998,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~sP9(X95,X96,X97) | sP10(X95,X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f423,f565])).
fof(f1272,plain,(
  ( ! [X94,X92,X93,X91] : (~sP11(X91,X92,X93) | sP12(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f997,f246])).
fof(f997,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~sP11(X91,X92,X93) | sP12(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f422,f565])).
fof(f1271,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | ~sP12(X87,X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f996,f246])).
fof(f996,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X88)) | p(s(bool,X89)) | ~sP12(X87,X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f419,f565])).
fof(f1270,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | sP12(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f995,f246])).
fof(f995,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X84)) | sP12(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f418,f565])).
fof(f1269,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f994,f246])).
fof(f994,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f417,f565])).
fof(f1268,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | ~p(s(bool,X77)) | ~sP11(X77,X76,X75) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f993,f246])).
fof(f993,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X76)) | ~p(s(bool,X77)) | ~sP11(X77,X76,X75) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f412,f565])).
fof(f1267,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | p(s(bool,X73)) | ~sP10(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f992,f246])).
fof(f992,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | p(s(bool,X73)) | ~sP10(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f402,f565])).
fof(f1266,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP10(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f991,f246])).
fof(f991,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP10(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f401,f565])).
fof(f1265,plain,(
  ( ! [X66,X64,X65,X63] : (sP10(X63,X65,X64) | ~p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f990,f246])).
fof(f990,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP10(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f400,f565])).
fof(f1264,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f989,f246])).
fof(f989,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f395,f565])).
fof(f1263,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X57,X56,X55) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f988,f246])).
fof(f988,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X57,X56,X55) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f387,f565])).
fof(f1262,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f987,f246])).
fof(f987,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f378,f565])).
fof(f1261,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP4(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f986,f246])).
fof(f986,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP4(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f359,f565])).
fof(f1260,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP5(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f985,f246])).
fof(f985,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP5(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f355,f565])).
fof(f1259,plain,(
  ( ! [X39,X41,X42,X40] : (sP2(X40,X41,X39) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f984,f246])).
fof(f984,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f343,f565])).
fof(f1258,plain,(
  ( ! [X37,X35,X38,X36] : (sP2(X36,X35,X37) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f983,f246])).
fof(f983,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f342,f565])).
fof(f1256,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f980,f246])).
fof(f980,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f325,f565])).
fof(f1255,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f979,f246])).
fof(f979,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f310,f565])).
fof(f1254,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f978,f246])).
fof(f978,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f309,f565])).
fof(f1253,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f977,f246])).
fof(f977,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f303,f565])).
fof(f1252,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f973,f246])).
fof(f973,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f264,f565])).
fof(f938,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f260,f565])).
fof(f844,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f565,f565])).
fof(f795,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f451,f448,f446,f440,f353])).
fof(f786,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f453,f448,f445,f246,f352])).
fof(f763,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f754,f407])).
fof(f754,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f440,f246,f412])).
fof(f747,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f462,f246,f440,f411])).
fof(f727,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f718,f390])).
fof(f718,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f440,f246,f395])).
fof(f711,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f459,f440,f440,f394])).
fof(f693,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f684,f371])).
fof(f684,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f246,f246,f378])).
fof(f677,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f454,f246,f246,f374])).
fof(f673,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f668,f348])).
fof(f668,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f440,f440,f363])).
fof(f664,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f658,f349])).
fof(f658,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f246,f246,f359])).
fof(f648,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f246,f246,f355])).
fof(f634,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f464,f369])).
fof(f636,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f246,f464,f369])).
fof(f625,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f440,f464,f367])).
fof(f626,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f464,f367])).
fof(f617,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f614,f440])).
fof(f614,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f300,f253])).
fof(f538,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f246,f418])).
fof(f535,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f440,f246,f417])).
fof(f532,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f440,f406])).
fof(f528,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f440,f440,f405])).
fof(f527,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f246,f401])).
fof(f521,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f246,f400])).
fof(f518,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f440,f389])).
fof(f517,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f440,f388])).
fof(f513,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f440,f440,f383])).
fof(f506,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f440,f440,f382])).
fof(f504,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f440,f373])).
fof(f503,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f246,f440,f372])).
fof(f500,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f246,f343])).
fof(f492,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f440,f440,f339])).
fof(f487,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f440,f440,f338])).
fof(f485,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f440,f442])).
fof(f480,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f440,f246,f328])).
fof(f474,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f246,f325])).
fof(f472,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f246,f246,f324])).
fof(f471,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f440,f440,f323])).
fof(f464,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f440,f260])).
fof(f463,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f425])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f440,f415])).
fof(f461,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f414])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f440,f413])).
fof(f459,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f440,f398])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f397])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f396])).
fof(f456,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f440,f381])).
fof(f455,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f440,f380])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f379])).
fof(f453,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f246,f366])).
fof(f452,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f246,f365])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f364])).
fof(f450,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f440,f362])).
fof(f449,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f440,f361])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f360])).
fof(f447,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f440,f358])).
fof(f446,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f440,f357])).
fof(f445,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f356])).
fof(f444,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f246,f310])).
fof(f443,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f246,f309])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f137,f156,f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) => p(s(bool,X22)))) <=> ((p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (~p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.sat.sat_27_prove_ltb')).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f107,f153,f152])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) | p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.sat.sat_25_prove_ltb')).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f231])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f228])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f105,f150,f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> (p(s(bool,X23)) & p(s(bool,X22)))) <=> ((p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X22))) & (p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X24)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.sat.sat_23_prove_ltb')).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f225])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f149])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f223])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f223])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f221])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f103,f147,f146,f145,f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f145,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f146,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f147,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.sat.sat_2|49|_prove_ltb')).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f144])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f217])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f214])).
fof(f214,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f145])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f215])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f212])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f212])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | ~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f203])).
fof(f203,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f204])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f204])).
fof(f333,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb',aHL_EXT)).
fof(f323,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f196])).
fof(f324,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f325,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f442,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f308,f440])).
fof(f308,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f309,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f310,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f300,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X16,X17] : ((p(s(bool,X16)) => p(s(bool,X17))) => ((p(s(bool,X17)) => p(s(bool,X16))) => s(bool,X16) = s(bool,X17)))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f303,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f123])).
fof(f261,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f1911,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f1910,f1816])).
fof(f1910,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(subsumption_resolution,[],[f1843,f1842])).
fof(f1843,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl52),
  inference(backward_demodulation,[],[f1842,f264])).
fof(f1820,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl54),
  inference(cnf_transformation,[],[f1820_D])).
fof(f1820_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1819,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f1812,f260])).
fof(f1812,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1256,f253])).
fof(f1815,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl50),
  inference(cnf_transformation,[],[f1815_D])).
fof(f1815_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1817,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1808,f260])).
fof(f1808,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1256,f565])).
fof(f1814,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1807,f260])).
fof(f1807,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1256,f565])).
fof(f1679,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1)))))) ) | $spl48),
  inference(cnf_transformation,[],[f1679_D])).
fof(f1679_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1203,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f1203_D])).
fof(f1203_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1681,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1)))))) )),
  inference(condensation,[],[f1680])).
fof(f1680,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK22('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1673,f617])).
fof(f1673,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK22('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f294,f565])).
fof(f1678,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1)))))) )),
  inference(condensation,[],[f1677])).
fof(f1677,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK22('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f1672,f617])).
fof(f1672,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK22('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f294,f565])).
fof(f1659,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1)))))) ) | $spl46),
  inference(cnf_transformation,[],[f1659_D])).
fof(f1659_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1661,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1)))))) )),
  inference(condensation,[],[f1660])).
fof(f1660,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK21('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1653,f617])).
fof(f1653,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,happ(s(fun(X30,bool),X31),s(X30,X32)))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,sK21('T_0',X31,X30))))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f292,f565])).
fof(f1658,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1)))))) )),
  inference(condensation,[],[f1657])).
fof(f1657,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK21('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f1652,f617])).
fof(f1652,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(X25,bool),X26),s(X25,X27)))) | ~p(s(bool,happ(s(fun(X25,bool),X26),s(X25,sK21('T_0',X26,X25))))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f292,f565])).
fof(f1326,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1315,f246])).
fof(f1315,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f1203,f440])).
fof(f1325,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f245,f244,f243,f242,f246,f440,f253,f261,f260,f265,f264,f288,f441,f294,f292,f299,f297,f303,f300,f310,f309,f442,f328,f325,f324,f323,f333,f337,f336,f335,f334,f343,f342,f341,f425,f339,f338,f354,f353,f352,f351,f350,f349,f348,f358,f357,f356,f355,f362,f361,f360,f359,f366,f365,f364,f363,f370,f369,f368,f367,f424,f374,f373,f372,f371,f381,f380,f379,f378,f387,f384,f383,f382,f394,f393,f392,f423,f390,f389,f388,f398,f397,f396,f395,f402,f401,f400,f399,f411,f410,f409,f422,f407,f406,f405,f415,f414,f413,f412,f419,f418,f417,f416,f443,f444,f445,f446,f447,f448,f449,f450,f451,f452,f453,f454,f455,f456,f457,f458,f459,f460,f461,f462,f463,f464,f471,f472,f474,f480,f485,f487,f492,f500,f503,f504,f506,f513,f517,f518,f521,f527,f528,f532,f535,f538,f565,f617,f626,f625,f636,f634,f648,f658,f664,f668,f673,f677,f684,f693,f711,f718,f727,f747,f754,f763,f786,f795,f844,f938,f1201,f1205,f1208,f1211,f1214,f1217,f1220,f1223,f1226,f1229,f1232,f1235,f1238,f1241,f1244,f1247,f1250,f1252,f1253,f1254,f1255,f1256,f1257,f1258,f1259,f1260,f1261,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1280,f1282,f1284,f1286,f1288,f1290,f1292,f1294,f1296,f1298,f1300,f1302,f1304,f1306,f1308,f1310,f1312,f1203,f1324])).
fof(f1324,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f1313,f1203])).
fof(f1313,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl12),
  inference(backward_demodulation,[],[f1203,f264])).
fof(f1251,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f1251_D])).
fof(f1251_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1312,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1311])).
fof(f1311,plain,(
  ( ! [X101,X99,X102,X100] : (~sP9('T_0',X100,X101) | sP10('T_0',X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f1036,f617])).
fof(f1036,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | ~sP9('T_0',X100,X101) | sP10('T_0',X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f423,f565])).
fof(f1248,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f1248_D])).
fof(f1248_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1310,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1309])).
fof(f1309,plain,(
  ( ! [X97,X95,X98,X96] : (~sP11('T_0',X96,X97) | sP12('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f1035,f617])).
fof(f1035,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP11('T_0',X96,X97) | sP12('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f422,f565])).
fof(f1245,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f1245_D])).
fof(f1245_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1308,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1307])).
fof(f1307,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f1034,f617])).
fof(f1034,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f419,f565])).
fof(f1242,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f1242_D])).
fof(f1242_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1305])).
fof(f1305,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | sP12('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1033,f617])).
fof(f1033,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | sP12('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f418,f565])).
fof(f1239,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f1239_D])).
fof(f1239_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1303])).
fof(f1303,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | sP12('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1032,f617])).
fof(f1032,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | p(s(bool,X84)) | sP12('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f417,f565])).
fof(f1236,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1236_D])).
fof(f1236_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1302,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1301])).
fof(f1301,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,'T_0') | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1031,f617])).
fof(f1031,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,'T_0') | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f412,f565])).
fof(f1233,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f1233_D])).
fof(f1233_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1300,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1299])).
fof(f1299,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP10('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1030,f617])).
fof(f1030,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP10('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f402,f565])).
fof(f1230,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1230_D])).
fof(f1230_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1298,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1297])).
fof(f1297,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1029,f617])).
fof(f1029,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f401,f565])).
fof(f1227,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1227_D])).
fof(f1227_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1296,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1295])).
fof(f1295,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP10('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1028,f617])).
fof(f1028,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | sP10('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f400,f565])).
fof(f1224,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1294,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1293])).
fof(f1293,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1027,f617])).
fof(f1027,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f395,f565])).
fof(f1221,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1221_D])).
fof(f1221_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1292,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1291])).
fof(f1291,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1026,f617])).
fof(f1026,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f387,f565])).
fof(f1218,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f1218_D])).
fof(f1218_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1290,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1289])).
fof(f1289,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1025,f617])).
fof(f1025,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f378,f565])).
fof(f1215,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1287])).
fof(f1287,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,'T_0',X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1024,f617])).
fof(f1024,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,'T_0',X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f359,f565])).
fof(f1212,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1285])).
fof(f1285,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1023,f617])).
fof(f1023,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f355,f565])).
fof(f1209,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f1209_D])).
fof(f1209_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1284,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1283])).
fof(f1283,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1022,f617])).
fof(f1022,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f343,f565])).
fof(f1206,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f1206_D])).
fof(f1206_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1282,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1281])).
fof(f1281,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1019,f617])).
fof(f1019,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f328,f565])).
fof(f1202,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl10),
  inference(cnf_transformation,[],[f1202_D])).
fof(f1202_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1280,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1279])).
fof(f1279,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1018,f617])).
fof(f1018,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f325,f565])).
fof(f1250,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1249])).
fof(f1249,plain,(
  ( ! [X101,X99,X102,X100] : (~sP9('T_0',X100,X101) | sP10('T_0',X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f967,f617])).
fof(f967,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | ~sP9('T_0',X100,X101) | sP10('T_0',X100,X101) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f423,f565])).
fof(f1247,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1246])).
fof(f1246,plain,(
  ( ! [X97,X95,X98,X96] : (~sP11('T_0',X96,X97) | sP12('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f966,f617])).
fof(f966,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~sP11('T_0',X96,X97) | sP12('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f422,f565])).
fof(f1244,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1243])).
fof(f1243,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f965,f617])).
fof(f965,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~p(s(bool,X92)) | p(s(bool,X93)) | ~sP12('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f419,f565])).
fof(f1241,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1240])).
fof(f1240,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | sP12('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f964,f617])).
fof(f964,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | sP12('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f418,f565])).
fof(f1238,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1237])).
fof(f1237,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | sP12('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f963,f617])).
fof(f963,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | p(s(bool,X84)) | sP12('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f417,f565])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1234])).
fof(f1234,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,'T_0') | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f962,f617])).
fof(f962,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X80)) | ~p(s(bool,X81)) | ~sP11(X81,X80,'T_0') | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f412,f565])).
fof(f1232,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1231])).
fof(f1231,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | p(s(bool,X77)) | ~sP10('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f961,f617])).
fof(f961,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | p(s(bool,X77)) | ~sP10('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f402,f565])).
fof(f1229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1228])).
fof(f1228,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f960,f617])).
fof(f960,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | sP10('T_0',X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f401,f565])).
fof(f1226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1225])).
fof(f1225,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP10('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f959,f617])).
fof(f959,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | sP10('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f400,f565])).
fof(f1223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1222])).
fof(f1222,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f958,f617])).
fof(f958,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f395,f565])).
fof(f1220,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1219])).
fof(f1219,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f957,f617])).
fof(f957,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP8(X61,X60,'T_0') | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f387,f565])).
fof(f1217,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1216])).
fof(f1216,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f956,f617])).
fof(f956,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X56,'T_0',X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f378,f565])).
fof(f1214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1213])).
fof(f1213,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,'T_0',X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f955,f617])).
fof(f955,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,'T_0',X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f359,f565])).
fof(f1211,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1210])).
fof(f1210,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f954,f617])).
fof(f954,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP5(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f355,f565])).
fof(f1208,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1207])).
fof(f1207,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f953,f617])).
fof(f953,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | sP2(X44,X45,'T_0') | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f343,f565])).
fof(f1205,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1204])).
fof(f1204,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f950,f617])).
fof(f950,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f328,f565])).
fof(f1201,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1200])).
fof(f1200,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f949,f617])).
fof(f949,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f325,f565])).
fof(f431,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) | $spl0),
  inference(cnf_transformation,[],[f431_D])).
fof(f431_D,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f439,plain,(
  s('type.list.list'(sK13),sK16) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) | $spl8),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  s('type.list.list'(sK13),sK16) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f242,plain,(
  s('type.list.list'(sK13),sK16) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) | s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16)))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  (s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) | (s('type.list.list'(sK13),sK16) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK19),s('type.list.list'(sK13),sK20))) & s(sK14,sK17) = s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) & s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))))) & (s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) | ! [X6,X7] : (s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X6),s('type.list.list'(sK13),X7))) | s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,X6))) | s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),X7)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17,sK18,sK19,sK20])],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) | ? [X8,X9] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X8),s('type.list.list'(X0),X9))) & s(X1,X4) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X8))) & s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X9))))) & (s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) != s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) | ! [X6,X7] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | s(X1,X4) != s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) | s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7))))))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) | ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X4) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) & s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7))))) & (s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) != s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) | ! [X6,X7] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | s(X1,X4) != s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) | s('type.list.list'(X1),X2) != s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7))))))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) <~> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X4) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) & s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7)))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) <=> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X4) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) & s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7)))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X2))) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X3))) <=> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & s(X1,X4) = s(X1,happ(s(fun(X0,X1),X5),s(X0,X6))) & s('type.list.list'(X1),X2) = s('type.list.list'(X1),'const.list.MAP_2'(s(fun(X0,X1),X5),s('type.list.list'(X0),X7)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X8))) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X11),s('type.list.list'(X6),X9))) <=> ? [X12,X13] : (s('type.list.list'(X6),X9) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & s(X7,X10) = s(X7,happ(s(fun(X6,X7),X11),s(X6,X12))) & s('type.list.list'(X7),X8) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X11),s('type.list.list'(X6),X13)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X8))) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X11),s('type.list.list'(X6),X9))) <=> ? [X12,X13] : (s('type.list.list'(X6),X9) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & s(X7,X10) = s(X7,happ(s(fun(X6,X7),X11),s(X6,X12))) & s('type.list.list'(X7),X8) = s('type.list.list'(X7),'const.list.MAP_2'(s(fun(X6,X7),X11),s('type.list.list'(X6),X13)))))),
  file('thm.pegComplete.pegComplete_|48|_prove_ltb','thm.pegComplete.pegComplete_|48|_prove_ltb')).
fof(f437,plain,(
  s(sK14,sK17) = s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) | $spl6),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  s(sK14,sK17) = s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f243,plain,(
  s(sK14,sK17) = s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,sK19))) | s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16)))),
  inference(cnf_transformation,[],[f160])).
fof(f435,plain,(
  s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))) | $spl4),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f244,plain,(
  s('type.list.list'(sK14),sK15) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK20))) | s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) = s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16)))),
  inference(cnf_transformation,[],[f160])).
fof(f432,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) | $spl1),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f433,plain,(
  ( ! [X6,X7] : (s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,X6))) | s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),X7))) | s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X6),s('type.list.list'(sK13),X7)))) ) | $spl2),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  ( ! [X6,X7] : (s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,X6))) | s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),X7))) | s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X6),s('type.list.list'(sK13),X7)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f245,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK14),sK15) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),X7))) | s(sK14,sK17) != s(sK14,happ(s(fun(sK13,sK14),sK18),s(sK13,X6))) | s('type.list.list'(sK13),sK16) != s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X6),s('type.list.list'(sK13),X7))) | s('type.list.list'(sK14),'const.list.CONS_2'(s(sK14,sK17),s('type.list.list'(sK14),sK15))) != s('type.list.list'(sK14),'const.list.MAP_2'(s(fun(sK13,sK14),sK18),s('type.list.list'(sK13),sK16)))) )),
  inference(cnf_transformation,[],[f160])).
% SZS output end Proof for thm.pegComplete
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 223
% Passive clauses: 469
% Generated clauses: 2276
% Final active clauses: 206
% Final passive clauses: 154
% Input formulas: 42
% Initial clauses: 180
% 
% Duplicate literals: 51
% Trivial inequalities: 15
% Fw subsumption resolutions: 124
% Fw demodulations: 44
% Bw demodulations: 70
% Condensations: 110
% Global subsumptions: 4
% 
% Simple tautologies: 60
% Equational tautologies: 67
% Forward subsumptions: 1270
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 792
% Factoring: 1
% Forward superposition: 338
% Backward superposition: 382
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 3
% 
% Split clauses: 48
% Split components: 96
% SAT solver clauses: 585
% SAT solver unit clauses: 79
% SAT solver binary clauses: 128
% 
% Sat splits: 48
% Sat splitting refutations: 14
% 
% Memory used [KB]: 2686
% Time elapsed: 0.177 s
% ------------------------------
% ------------------------------
% Success in time 1.595 s
