Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (3007)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 161
% Passive clauses: 3348
% Generated clauses: 19438
% Final active clauses: 142
% Final passive clauses: 2948
% Input formulas: 68
% Initial clauses: 99
% 
% Duplicate literals: 549
% Trivial inequalities: 309
% Fw subsumption resolutions: 1342
% Fw demodulations: 2503
% Bw demodulations: 17
% 
% Simple tautologies: 255
% Equational tautologies: 44
% Forward subsumptions: 10167
% 
% Binary resolution: 335
% Factoring: 1
% Forward superposition: 6935
% Backward superposition: 6990
% Self superposition: 189
% Equality factoring: 72
% Equality resolution: 31
% 
% Split clauses: 190
% Split components: 382
% SAT solver clauses: 535
% SAT solver unit clauses: 9
% SAT solver binary clauses: 94
% 
% Sat splits: 190
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 534
% TWLsolver calls for satisfiability: 150
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (3029)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 116
% Passive clauses: 858
% Generated clauses: 9884
% Final active clauses: 92
% Final passive clauses: 520
% Input formulas: 68
% Initial clauses: 99
% 
% Duplicate literals: 354
% Trivial inequalities: 116
% Fw subsumption resolutions: 433
% Fw demodulations: 1068
% Bw demodulations: 240
% Condensations: 605
% Global subsumptions: 8
% 
% Simple tautologies: 159
% Equational tautologies: 40
% Forward subsumptions: 4902
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 210
% Factoring: 1
% Forward superposition: 3105
% Backward superposition: 2589
% Self superposition: 107
% Equality factoring: 62
% Equality resolution: 18
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 1345
% SAT solver unit clauses: 150
% SAT solver binary clauses: 44
% SAT solver learnt clauses: 16
% 
% Sat splits: 29
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1313
% TWLsolver calls for satisfiability: 338
% 
% Memory used [KB]: 10746
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f59569,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f227,f59561])).
fof(f59561,plain,(
  ( ! [X578] : (~p(s(bool,X578))) )),
  inference(forward_demodulation,[],[f59560,f311])).
fof(f311,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X17] : ! [X21,X35,X36,X37,X38] : s(X17,X38) = s(X17,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X21),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s(X17,X36),s(X17,X37),s(X17,X38)))),
  file('thm.wordSem.result_distinct|split|6','thm.wordSem.result_case_def|split|4')).
fof(f59560,plain,(
  ( ! [X575,X576,X578,X574,X577] : (~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X574),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s(bool,X576),s(bool,X577),s(bool,X578))))) )),
  inference(subsumption_resolution,[],[f58924,f58832])).
fof(f58832,plain,(
  ( ! [X6,X8,X11] : (s(X6,X11) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X6)),X8),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f12794,f311])).
fof(f12794,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(X0,X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f313,f226])).
fof(f226,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Error_0') = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.Error_0') = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.Error_0') != s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.Error_0') = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.Error_0') = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.wordSem.result_distinct|split|6','thm.wordSem.result_distinct|split|6')).
fof(f313,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X17,X6] : ! [X8,X7,X21,X35,X36,X37,X38] : s(X17,happ(s(fun('type.wordSem.word_loc'(X6),X17),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s('type.wordSem.word_loc'(X6),X8))),s('type.wordSem.word_loc'(X6),X7))) = s(X17,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X21),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s(X17,X36),s(X17,X37),s(X17,X38)))),
  file('thm.wordSem.result_distinct|split|6','thm.wordSem.result_case_def|split|1')).
fof(f58924,plain,(
  ( ! [X575,X576,X578,X574,X577] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X574),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s(bool,X576),s(bool,X577),s(bool,X578))))) )),
  inference(superposition,[],[f241,f12794])).
fof(f241,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X30] : (s(bool,'F_0') = s(bool,X30) <=> ~p(s(bool,X30)))),
  file('thm.wordSem.result_distinct|split|6','thm.bool.EQ_CLAUSES|split|2')).
fof(f227,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.wordSem.result_distinct|split|6',aHL_TRUTH)).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 255
% Passive clauses: 3828
% Generated clauses: 60747
% Final active clauses: 165
% Final passive clauses: 2030
% Input formulas: 68
% Initial clauses: 99
% 
% Duplicate literals: 1054
% Trivial inequalities: 362
% Fw subsumption resolutions: 2911
% Fw demodulations: 13973
% Bw demodulations: 1474
% Condensations: 2836
% Global subsumptions: 110
% 
% Simple tautologies: 407
% Equational tautologies: 332
% Forward subsumptions: 32320
% Fw demodulations to eq. taut.: 2089
% Bw demodulations to eq. taut.: 67
% 
% Unit resulting resolution: 9024
% Factoring: 1
% Forward superposition: 13756
% Backward superposition: 13456
% Self superposition: 320
% Equality factoring: 140
% Equality resolution: 38
% 
% Split clauses: 29
% Split components: 58
% SAT solver clauses: 4415
% SAT solver unit clauses: 1206
% SAT solver binary clauses: 1676
% 
% Sat splits: 29
% Sat splitting refutations: 417
% 
% Memory used [KB]: 53730
% Time elapsed: 2.968 s
% ------------------------------
% ------------------------------
% Success in time 4.35 s
