Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f386,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f384,f343,f297,f105,f175,f104,f177,f171,f351,f107,f178,f173,f108,f176,f174,f172])).
fof(f172,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) | $spl1),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f174,plain,(
  p(s(bool,sK4)) | $spl2),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  p(s(bool,sK4)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f176,plain,(
  s(sK0,sK1) != s(sK0,sK2) | $spl5),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s(sK0,sK1) != s(sK0,sK2) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f108,plain,(
  s(sK0,sK1) != s(sK0,sK2) | p(s(bool,sK4)) | s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2)))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  (s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) | (p(s(bool,sK4)) & s(sK0,sK1) = s(sK0,sK3)) | (~p(s(bool,sK4)) & s(sK0,sK1) = s(sK0,sK2))) & (s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) | ((~p(s(bool,sK4)) | s(sK0,sK1) != s(sK0,sK3)) & (p(s(bool,sK4)) | s(sK0,sK1) != s(sK0,sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) | (p(s(bool,X4)) & s(X0,X1) = s(X0,X3)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X2))) & (s(X0,X1) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) | ((~p(s(bool,X4)) | s(X0,X1) != s(X0,X3)) & (p(s(bool,X4)) | s(X0,X1) != s(X0,X2)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2,X3,X4] : ((s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) | ((p(s(bool,X4)) & s(X0,X1) = s(X0,X3)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X2)))) & (s(X0,X1) != s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) | ((~p(s(bool,X4)) | s(X0,X1) != s(X0,X3)) & (p(s(bool,X4)) | s(X0,X1) != s(X0,X2)))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) <~> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X3)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X2))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) <=> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X3)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X2))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,X4),s(X0,X3),s(X0,X2))) <=> ((p(s(bool,X4)) & s(X0,X1) = s(X0,X3)) | (~p(s(bool,X4)) & s(X0,X1) = s(X0,X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X9),s(X6,X8))) <=> ((p(s(bool,X10)) & s(X6,X7) = s(X6,X9)) | (~p(s(bool,X10)) & s(X6,X7) = s(X6,X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (s(X6,X7) = s(X6,'const.bool.COND_3'(s(bool,X10),s(X6,X9),s(X6,X8))) <=> ((p(s(bool,X10)) & s(X6,X7) = s(X6,X9)) | (~p(s(bool,X10)) & s(X6,X7) = s(X6,X8))))),
  file('thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb','thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb')).
fof(f173,plain,(
  ~p(s(bool,sK4)) | $spl3),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ~p(s(bool,sK4)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f178,plain,(
  s(sK0,sK1) != s(sK0,sK3) | $spl7),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(sK0,sK1) != s(sK0,sK3) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f107,plain,(
  s(sK0,sK1) != s(sK0,sK3) | ~p(s(bool,sK4)) | s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2)))),
  inference(cnf_transformation,[],[f73])).
fof(f351,plain,(
  $false | ($spl0 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f345,f178])).
fof(f345,plain,(
  s(sK0,sK1) = s(sK0,sK3) | ($spl0 | $spl2)),
  inference(superposition,[],[f344,f161])).
fof(f161,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f344,plain,(
  s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK3),s(sK0,sK2))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f171,f307])).
fof(f307,plain,(
  s(bool,'T_0') = s(bool,sK4) | $spl2),
  inference(resolution,[],[f174,f120])).
fof(f120,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f171,plain,(
  s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) | $spl0),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f177,plain,(
  s(sK0,sK1) = s(sK0,sK3) | $spl6),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  s(sK0,sK1) = s(sK0,sK3) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f104,plain,(
  ~p(s(bool,sK4)) | s(sK0,sK1) = s(sK0,sK3) | s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2)))),
  inference(cnf_transformation,[],[f73])).
fof(f175,plain,(
  s(sK0,sK1) = s(sK0,sK2) | $spl4),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s(sK0,sK1) = s(sK0,sK2) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f105,plain,(
  s(sK0,sK1) = s(sK0,sK2) | p(s(bool,sK4)) | s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK2)))),
  inference(cnf_transformation,[],[f73])).
fof(f297,plain,(
  $false | ($spl1 | $spl3 | $spl4)),
  inference(trivial_inequality_removal,[],[f294])).
fof(f294,plain,(
  s(sK0,sK1) != s(sK0,sK1) | ($spl1 | $spl3 | $spl4)),
  inference(superposition,[],[f242,f162])).
fof(f162,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X11),s(X6,X12)))),
  file('thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f242,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,'F_0'),s(sK0,sK3),s(sK0,sK1))) | ($spl1 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f241,f188])).
fof(f188,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl3),
  inference(resolution,[],[f124,f173])).
fof(f124,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.clos_relationProps.clos_relationProps_|49||49|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f241,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK3),s(sK0,sK1))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f172,f175])).
fof(f343,plain,(
  $false | ($spl1 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f342,f161])).
fof(f342,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,'T_0'),s(sK0,sK1),s(sK0,sK2))) | ($spl1 | $spl2 | $spl6)),
  inference(forward_demodulation,[],[f341,f307])).
fof(f341,plain,(
  s(sK0,sK1) != s(sK0,'const.bool.COND_3'(s(bool,sK4),s(sK0,sK1),s(sK0,sK2))) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f172,f177])).
fof(f384,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f378,f176])).
fof(f378,plain,(
  s(sK0,sK1) = s(sK0,sK2) | ($spl0 | $spl3)),
  inference(superposition,[],[f375,f162])).
fof(f375,plain,(
  s(sK0,sK1) = s(sK0,'const.bool.COND_3'(s(bool,'F_0'),s(sK0,sK3),s(sK0,sK2))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f171,f355])).
fof(f355,plain,(
  s(bool,'F_0') = s(bool,sK4) | $spl3),
  inference(resolution,[],[f173,f124])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 26
% Passive clauses: 61
% Generated clauses: 282
% Final active clauses: 16
% Final passive clauses: 19
% Input formulas: 29
% Initial clauses: 65
% 
% Duplicate literals: 3
% Trivial inequalities: 5
% Fw subsumption resolutions: 15
% Fw demodulations: 32
% Bw demodulations: 4
% 
% Simple tautologies: 32
% Equational tautologies: 5
% Forward subsumptions: 115
% 
% Binary resolution: 22
% Factoring: 1
% Forward superposition: 41
% Backward superposition: 71
% Self superposition: 6
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 18
% Split components: 41
% SAT solver clauses: 33
% SAT solver unit clauses: 6
% SAT solver binary clauses: 14
% SAT solver learnt clauses: 4
% SAT solver learnt literals: 4
% 
% Sat splits: 18
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 28
% TWLsolver calls for satisfiability: 7
% 
% Memory used [KB]: 767
% Time elapsed: 0.044 s
% ------------------------------
% ------------------------------
% Success in time 0.034 s
