Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvlSem
% (36200)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 176
% Passive clauses: 3859
% Generated clauses: 21350
% Final active clauses: 167
% Final passive clauses: 3493
% Input formulas: 69
% Initial clauses: 141
% 
% Duplicate literals: 714
% Trivial inequalities: 602
% Fw subsumption resolutions: 2277
% Fw demodulations: 3250
% Bw demodulations: 6
% 
% Simple tautologies: 478
% Equational tautologies: 53
% Forward subsumptions: 9650
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 440
% Factoring: 1
% Forward superposition: 7371
% Backward superposition: 6169
% Self superposition: 154
% Equality factoring: 88
% Equality resolution: 43
% 
% Split clauses: 231
% Split components: 470
% SAT solver clauses: 539
% SAT solver unit clauses: 14
% SAT solver binary clauses: 106
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 231
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 538
% TWLsolver calls for satisfiability: 158
% 
% Memory used [KB]: 14583
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvlSem
% (36215)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 1297
% Generated clauses: 14604
% Final active clauses: 129
% Final passive clauses: 1122
% Input formulas: 69
% Initial clauses: 141
% 
% Duplicate literals: 429
% Trivial inequalities: 396
% Fw subsumption resolutions: 949
% Fw demodulations: 2069
% Bw demodulations: 38
% Condensations: 728
% Global subsumptions: 8
% 
% Simple tautologies: 255
% Equational tautologies: 44
% Forward subsumptions: 6393
% 
% Binary resolution: 172
% Factoring: 1
% Forward superposition: 4310
% Backward superposition: 3807
% Self superposition: 125
% Equality factoring: 60
% Equality resolution: 32
% 
% Split clauses: 47
% Split components: 94
% SAT solver clauses: 2134
% SAT solver unit clauses: 151
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 2
% 
% Sat splits: 47
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2100
% TWLsolver calls for satisfiability: 574
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvlSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvlSem
% SZS output start Proof for thm.bvlSem
fof(f27057,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f255,f417,f415,f254,f416,f414,f419,f1372,f1383,f2062,f2068,f2065,f5182,f5189,f5190,f5192,f5194,f5198,f5200,f5202,f5204,f5208,f5210,f5212,f5214,f5217,f5219,f5221,f5223,f5226,f5228,f5230,f5232,f5234,f5236,f5238,f5262,f5264,f5296,f5300,f5302,f5547,f5552,f5553,f5554,f5555,f6069,f6070,f6077,f6078,f7301,f7309,f7302,f7475,f10111,f10250,f10252,f10290,f10291,f10292,f10293,f7185,f10251,f10294,f10295,f10296,f10297,f6957,f10253,f6011,f6010,f10306,f10305,f5944,f5967,f4611,f4631,f10321,f1382,f10322,f10795,f10797,f10806,f10809,f10810,f10811,f11013,f11015,f11017,f11019,f11023,f11025,f11027,f11029,f11031,f11033,f11035,f11037,f11039,f11041,f11043,f11045,f11047,f11049,f11051,f11053,f15417,f10180,f15554,f14910,f7493,f15589,f7491,f15556,f18441,f18524,f20392,f20484,f10179,f26420,f2063,f26419,f2069,f26421,f26943,f26958])).
fof(f26958,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f26957,f417])).
fof(f26957,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | $spl0),
  inference(forward_demodulation,[],[f26742,f305])).
fof(f305,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X18] : ! [X20] : s(X18,X20) = s(X18,happ(s(fun(X18,X18),'_i_0_0'),s(X18,X20)))),
  file('thm.bvlSem.v_|49||49||split|1','_i0_ax')).
fof(f26742,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),'_i_0_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) | $spl0),
  inference(superposition,[],[f305,f24589])).
fof(f24589,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X1),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) = s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X1),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f24291,f326])).
fof(f326,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_case_def|split|1')).
fof(f24291,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X1),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X1),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X3),s(fun('type.num.num',X0),X4),s(fun('type.num.num',X0),X5)))) ) | $spl0),
  inference(superposition,[],[f326,f414])).
fof(f26943,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f26942,f417])).
fof(f26942,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | $spl0),
  inference(forward_demodulation,[],[f26719,f305])).
fof(f26719,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),'_i_0_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) | $spl0),
  inference(superposition,[],[f24589,f305])).
fof(f26421,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.min.@_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'))))))) ) | $spl46),
  inference(cnf_transformation,[],[f26421_D])).
fof(f26421_D,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.min.@_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'))))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2069,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) | $spl12),
  inference(cnf_transformation,[],[f2069_D])).
fof(f2069_D,plain,(
  ( ! [X24] : (~p(s(bool,X24))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f26419,plain,(
  ( ! [X56,X55] : (~p(s(bool,X56)) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.min.@_1'(s(fun('type.pair.prod'(bool,X55),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f2070,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X23,X9] : ! [X24,X25] : s(X9,X24) = s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X23),'const.pair.,_2'(s(X9,X24),s(X23,X25)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.pair.FST')).
fof(f2070,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(forward_demodulation,[],[f2048,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9,X23] : ! [X8] : s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,X23),X8))) = s(X9,happ(s(fun('type.pair.prod'(X9,X23),X9),'const.pair.FST_0'),s('type.pair.prod'(X9,X23),X8)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_1_const.pair.FST')).
fof(f2048,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),X29)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X28),'const.min.@_1'(s(fun('type.pair.prod'(bool,X28),bool),'const.pair.FST_0'))))))) )),
  inference(superposition,[],[f313,f341])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X9] : ! [X26,X24] : (p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24)))) => p(s(bool,happ(s(fun(X9,bool),X26),s(X9,'const.min.@_1'(s(fun(X9,bool),X26)))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.SELECT_AX')).
fof(f2063,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f2063_D])).
fof(f2063_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f26420,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.min.@_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f26399,f256])).
fof(f256,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bvlSem.v_|49||49||split|1',aHL_TRUTH)).
fof(f26399,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.min.@_1'(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0')))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f2070,f684])).
fof(f684,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f343,f266])).
fof(f266,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvlSem.v_|49||49||split|1',aHL_BOOL_CASES)).
fof(f10179,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) | $spl26),
  inference(cnf_transformation,[],[f10179_D])).
fof(f10179_D,plain,(
  ( ! [X71] : (s(bool,'T_0') = s(bool,X71)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f20484,plain,(
  ( ! [X77] : (s(bool,X77) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X77)) ) | $spl44),
  inference(cnf_transformation,[],[f20484_D])).
fof(f20484_D,plain,(
  ( ! [X77] : (s(bool,X77) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X77)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f20392,plain,(
  ( ! [X76,X77] : (s(bool,X77) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X76)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f20226])).
fof(f20226,plain,(
  ( ! [X76,X77] : (s(bool,X77) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) ) | $spl10),
  inference(superposition,[],[f15551,f19441])).
fof(f19441,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f19426])).
fof(f19426,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f19003,f433])).
fof(f433,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f266,f266])).
fof(f19003,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl10),
  inference(forward_demodulation,[],[f18913,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X23,X9] : ! [X24] : s(X9,X24) = s(X9,'const.sum.OUTL_1'(s('type.sum.sum'(X9,X23),'const.sum.INL_1'(s(X9,X24)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.sum.OUTL')).
fof(f18913,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7110,f9561])).
fof(f9561,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f9473,f258])).
fof(f258,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bvlSem.v_|49||49||split|1',aHL_FALSITY)).
fof(f9473,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl10),
  inference(superposition,[],[f267,f2716])).
fof(f2716,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'T_0') = s(bool,X21)) ) | $spl10),
  inference(forward_demodulation,[],[f2595,f2093])).
fof(f2093,plain,(
  s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2065,f268])).
fof(f268,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f2595,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21)))) ) | $spl10),
  inference(superposition,[],[f2336,f305])).
fof(f2336,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(subsumption_resolution,[],[f2313,f256])).
fof(f2313,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X30)) ) | $spl10),
  inference(superposition,[],[f1999,f2093])).
fof(f1999,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f1998,f305])).
fof(f1998,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f1966,f256])).
fof(f1966,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f332,f460])).
fof(f460,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f305,f266])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X33,X34] : ((p(s(bool,X33)) => p(s(bool,X34))) => ((p(s(bool,X34)) => p(s(bool,X33))) => s(bool,X33) = s(bool,X34)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f267,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f7110,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'F_0')))))) )),
  inference(unit_resulting_resolution,[],[f420,f6868])).
fof(f6868,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6797,f258])).
fof(f6797,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f267,f1997])).
fof(f1997,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(forward_demodulation,[],[f1952,f305])).
fof(f1952,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X37),'const.sum.INL_1'(s(bool,X38)))))) )),
  inference(superposition,[],[f460,f339])).
fof(f420,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f258,f267])).
fof(f15551,plain,(
  ( ! [X59,X57] : (s(bool,X57) != s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f15550,f258])).
fof(f15550,plain,(
  ( ! [X59,X57] : (p(s(bool,'F_0')) | s(bool,X57) != s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57)) )),
  inference(forward_demodulation,[],[f15262,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9,X23] : ! [X24] : s(X23,X24) = s(X23,'const.sum.OUTR_1'(s('type.sum.sum'(X9,X23),'const.sum.INR_1'(s(X23,X24)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.sum.OUTR')).
fof(f15262,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,'F_0')))))) | s(bool,X57) != s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f2001,f11390])).
fof(f11390,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(trivial_inequality_removal,[],[f11380])).
fof(f11380,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f6883,f433])).
fof(f6883,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) )),
  inference(unit_resulting_resolution,[],[f420,f6656])).
fof(f6656,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6586,f258])).
fof(f6586,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f267,f1996])).
fof(f1996,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,X36))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(forward_demodulation,[],[f1951,f305])).
fof(f1951,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X35,bool),'const.sum.INR_1'(s(bool,X36)))))) )),
  inference(superposition,[],[f460,f340])).
fof(f2001,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,'F_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f2000,f305])).
fof(f2000,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1967,f256])).
fof(f1967,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | s(bool,X13) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X12))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12)) )),
  inference(superposition,[],[f406,f460])).
fof(f406,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f376])).
fof(f376,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X24,X50,X25,X51] : ((s(bool,X24) = s(bool,X50) & (p(s(bool,X50)) => s(bool,X25) = s(bool,X51))) => ((p(s(bool,X24)) => p(s(bool,X25))) <=> (p(s(bool,X50)) => p(s(bool,X51)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.IMP_CONG')).
fof(f18524,plain,(
  ( ! [X114,X113] : (s(bool,X114) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114)) ) | $spl42),
  inference(cnf_transformation,[],[f18524_D])).
fof(f18524_D,plain,(
  ( ! [X114,X113] : (s(bool,X114) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f18441,plain,(
  ( ! [X114,X112,X113] : (s(bool,X114) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114) | s(bool,'T_0') = s(bool,X112)) )),
  inference(duplicate_literal_removal,[],[f18291])).
fof(f18291,plain,(
  ( ! [X114,X112,X113] : (s(bool,X114) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f15551,f11512])).
fof(f11512,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(trivial_inequality_removal,[],[f11502])).
fof(f11502,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f7110,f433])).
fof(f15556,plain,(
  s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | $spl40),
  inference(cnf_transformation,[],[f15556_D])).
fof(f15556_D,plain,(
  s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7491,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(cnf_transformation,[],[f7491_D])).
fof(f7491_D,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f15589,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f9468,f258])).
fof(f9468,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | ($spl10 | $spl20)),
  inference(superposition,[],[f7302,f2716])).
fof(f7493,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl24),
  inference(cnf_transformation,[],[f7493_D])).
fof(f7493_D,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f14910,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f7302,f305])).
fof(f15554,plain,(
  p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f15003,f258])).
fof(f15003,plain,(
  p(s(bool,'F_0')) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | ($spl10 | $spl20)),
  inference(superposition,[],[f7302,f2716])).
fof(f10180,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) | $spl28),
  inference(cnf_transformation,[],[f10180_D])).
fof(f10180_D,plain,(
  ( ! [X72,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f15417,plain,(
  ( ! [X109,X110,X108] : (s(bool,X110) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X108) | s(bool,'F_0') = s(bool,X110)) ) | ($spl10 | $spl22)),
  inference(duplicate_literal_removal,[],[f15281])).
fof(f15281,plain,(
  ( ! [X109,X110,X108] : (s(bool,X110) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X109,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X108) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X108)) ) | ($spl10 | $spl22)),
  inference(superposition,[],[f9004,f11390])).
fof(f9004,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl10 | $spl22)),
  inference(superposition,[],[f7988,f7910])).
fof(f7910,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,sK6('_i_0_0','_i_0_0',bool)) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl22)),
  inference(subsumption_resolution,[],[f7871,f258])).
fof(f7871,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,sK6('_i_0_0','_i_0_0',bool)) | s(bool,'F_0') = s(bool,X5)) ) | ($spl10 | $spl22)),
  inference(superposition,[],[f7718,f2604])).
fof(f2604,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f2336,f2336])).
fof(f7718,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | $spl22),
  inference(forward_demodulation,[],[f7717,f305])).
fof(f7717,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0','_i_0_0',bool))))) | $spl22),
  inference(subsumption_resolution,[],[f7716,f256])).
fof(f7716,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0','_i_0_0',bool))))) | $spl22),
  inference(forward_demodulation,[],[f7697,f305])).
fof(f7697,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0','_i_0_0',bool))))) | $spl22),
  inference(condensation,[],[f7683])).
fof(f7683,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0','_i_0_0',bool)))))) ) | $spl22),
  inference(superposition,[],[f317,f7588])).
fof(f7588,plain,(
  s(bool,'T_0') = s(bool,sK5('_i_0_0','_i_0_0',bool)) | $spl22),
  inference(unit_resulting_resolution,[],[f420,f7491,f1999])).
fof(f317,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : ! [X53,X26] : (! [X24] : (p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24)))) => p(s(bool,happ(s(fun(X9,bool),X53),s(X9,X24))))) => (? [X24] : p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24)))) => ? [X24] : p(s(bool,happ(s(fun(X9,bool),X53),s(X9,X24))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.MONO_EXISTS')).
fof(f7988,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,sK6('_i_0_0','_i_0_0',bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl22),
  inference(subsumption_resolution,[],[f7937,f256])).
fof(f7937,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,sK6('_i_0_0','_i_0_0',bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl22),
  inference(superposition,[],[f435,f7862])).
fof(f7862,plain,(
  s(bool,'T_0') = s(bool,sK6('_i_0_0','_i_0_0',bool)) | $spl22),
  inference(unit_resulting_resolution,[],[f420,f7718,f1999])).
fof(f435,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f271,f266])).
fof(f271,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f11053,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042,f11044,f11046,f11048,f11050,f11052])).
fof(f11052,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10776,f256])).
fof(f10776,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8533])).
fof(f11050,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10775,f256])).
fof(f10775,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f7992])).
fof(f11048,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (s(X393,X401) = s(X393,happ(s(fun(X394,X393),X400),s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),X396),s(X395,happ(s(fun(X397,X395),X398),s(X397,X399))))),s(X393,X401)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10774,f256])).
fof(f10774,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (~p(s(bool,'T_0')) | s(X393,X401) = s(X393,happ(s(fun(X394,X393),X400),s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),X396),s(X395,happ(s(fun(X397,X395),X398),s(X397,X399))))),s(X393,X401)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5038])).
fof(f11046,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (s(X376,X383) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),X378),s(X377,happ(s(fun(X379,X377),X380),s(X379,X381))))),s(X375,happ(s(fun(X376,X375),X382),s(X376,X383)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10773,f256])).
fof(f10773,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (~p(s(bool,'T_0')) | s(X376,X383) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),X378),s(X377,happ(s(fun(X379,X377),X380),s(X379,X381))))),s(X375,happ(s(fun(X376,X375),X382),s(X376,X383)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5036])).
fof(f11044,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10772,f256])).
fof(f10772,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8534])).
fof(f11042,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10771,f256])).
fof(f10771,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f7993])).
fof(f11040,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (s(X380,X388) = s(X380,happ(s(fun(X381,X380),X387),s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10770,f256])).
fof(f10770,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (~p(s(bool,'T_0')) | s(X380,X388) = s(X380,happ(s(fun(X381,X380),X387),s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5115])).
fof(f11038,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (s(X363,X370) = s(X363,happ(s(fun(X362,X363),happ(s(fun(X364,fun(X362,X363)),'_dst_x0x1_2'(s(fun(X365,fun(X362,X363)),X366),s(fun(X364,X365),X367))),s(X364,X368))),s(X362,happ(s(fun(X363,X362),X369),s(X363,X370)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10769,f256])).
fof(f10769,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (~p(s(bool,'T_0')) | s(X363,X370) = s(X363,happ(s(fun(X362,X363),happ(s(fun(X364,fun(X362,X363)),'_dst_x0x1_2'(s(fun(X365,fun(X362,X363)),X366),s(fun(X364,X365),X367))),s(X364,X368))),s(X362,happ(s(fun(X363,X362),X369),s(X363,X370)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5113])).
fof(f11036,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10768,f256])).
fof(f10768,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8546])).
fof(f11034,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10767,f256])).
fof(f10767,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8005])).
fof(f11032,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (s(X403,X410) = s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),X405),s(X404,happ(s(fun(X406,X404),X407),s(X406,X408))))),s(X402,happ(s(fun(X403,X402),X409),s(X403,X410)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10766,f256])).
fof(f10766,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (~p(s(bool,'T_0')) | s(X403,X410) = s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),X405),s(X404,happ(s(fun(X406,X404),X407),s(X406,X408))))),s(X402,happ(s(fun(X403,X402),X409),s(X403,X410)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5039])).
fof(f11030,plain,(
  ( ! [X389,X391,X385,X387,X388,X390,X384,X386,X392] : (s(X384,X392) = s(X384,happ(s(fun(X385,X384),X391),s(X385,happ(s(fun(X384,X385),happ(s(fun(X386,fun(X384,X385)),X387),s(X386,happ(s(fun(X388,X386),X389),s(X388,X390))))),s(X384,X392)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10765,f256])).
fof(f10765,plain,(
  ( ! [X389,X391,X385,X387,X388,X390,X384,X386,X392] : (~p(s(bool,'T_0')) | s(X384,X392) = s(X384,happ(s(fun(X385,X384),X391),s(X385,happ(s(fun(X384,X385),happ(s(fun(X386,fun(X384,X385)),X387),s(X386,happ(s(fun(X388,X386),X389),s(X388,X390))))),s(X384,X392)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5037])).
fof(f11028,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10764,f256])).
fof(f10764,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8547])).
fof(f11026,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10763,f256])).
fof(f10763,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'T_0')) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f8006])).
fof(f11024,plain,(
  ( ! [X389,X391,X397,X393,X395,X390,X396,X392,X394] : (s(X390,X397) = s(X390,happ(s(fun(X389,X390),happ(s(fun(X391,fun(X389,X390)),'_dst_x0x1_2'(s(fun(X392,fun(X389,X390)),X393),s(fun(X391,X392),X394))),s(X391,X395))),s(X389,happ(s(fun(X390,X389),X396),s(X390,X397)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10762,f256])).
fof(f10762,plain,(
  ( ! [X389,X391,X397,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'T_0')) | s(X390,X397) = s(X390,happ(s(fun(X389,X390),happ(s(fun(X391,fun(X389,X390)),'_dst_x0x1_2'(s(fun(X392,fun(X389,X390)),X393),s(fun(X391,X392),X394))),s(X391,X395))),s(X389,happ(s(fun(X390,X389),X396),s(X390,X397)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5116])).
fof(f11022,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (s(X371,X379) = s(X371,happ(s(fun(X372,X371),X378),s(X372,happ(s(fun(X371,X372),happ(s(fun(X373,fun(X371,X372)),'_dst_x0x1_2'(s(fun(X374,fun(X371,X372)),X375),s(fun(X373,X374),X376))),s(X373,X377))),s(X371,X379)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10761,f256])).
fof(f10761,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (~p(s(bool,'T_0')) | s(X371,X379) = s(X371,happ(s(fun(X372,X371),X378),s(X372,happ(s(fun(X371,X372),happ(s(fun(X373,fun(X371,X372)),'_dst_x0x1_2'(s(fun(X374,fun(X371,X372)),X375),s(fun(X373,X374),X376))),s(X373,X377))),s(X371,X379)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5114])).
fof(f11018,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10756,f256])).
fof(f10756,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f357])).
fof(f11016,plain,(
  ( ! [X2,X0,X7,X3,X1] : (s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10755,f256])).
fof(f10755,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f356])).
fof(f11014,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10754,f256])).
fof(f10754,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f352])).
fof(f11012,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl26),
  inference(subsumption_resolution,[],[f10753,f256])).
fof(f10753,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f351])).
fof(f10796,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl26),
  inference(subsumption_resolution,[],[f10324,f10179])).
fof(f10324,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f271])).
fof(f2243,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X49),s('type.num.num',X50))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X48),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50)))))) )),
  inference(superposition,[],[f266,f320])).
fof(f2575,plain,(
  ( ! [X21,X22,X20] : (s(X22,f(s(bool,'F_0'))) != s(X22,g(s(bool,'F_0'))) | s(fun(bool,X22),X20) = s(fun(bool,X22),X21) | s(bool,'T_0') = s(bool,sK12(X20,X21,X22,bool))) )),
  inference(superposition,[],[f353,f266])).
fof(f2581,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,f(s(X20,sK12(X21,X22,bool,X20)))) | s(fun(X20,bool),X21) = s(fun(X20,bool),X22) | s(bool,'T_0') = s(bool,g(s(X20,sK12(X21,X22,bool,X20))))) )),
  inference(superposition,[],[f353,f266])).
fof(f10307,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(subsumption_resolution,[],[f5987,f258])).
fof(f5987,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) )),
  inference(superposition,[],[f408,f266])).
fof(f9729,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK10(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9703,f256])).
fof(f9703,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK10(X13,X11,X12,X10,X9))))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f349,f433])).
fof(f9728,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f9727,f305])).
fof(f9727,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(subsumption_resolution,[],[f9701,f256])).
fof(f9701,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK10(X1,'_i_0_0','_i_0_0',X0,X0)))))))) )),
  inference(superposition,[],[f349,f428])).
fof(f9700,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | ~p(s(bool,happ(s(fun(X95,bool),X99),s(X95,X100)))) | p(s(bool,happ(s(fun(X95,bool),X99),s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,sK10(X99,X98,sK4(X96,X97,fun(X94,X95)),X94,X95))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f819])).
fof(f9726,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),X90),s(X87,sK10(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | ~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,X93))))) )),
  inference(forward_demodulation,[],[f9699,f339])).
fof(f9699,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | ~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,X93)))) | p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90))))),s(X87,sK10(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88)))))))) )),
  inference(superposition,[],[f349,f339])).
fof(f9725,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),X83),s(X80,sK10(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | ~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,X86))))) )),
  inference(forward_demodulation,[],[f9698,f340])).
fof(f9698,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | ~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,X86)))) | p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83))))),s(X80,sK10(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81)))))))) )),
  inference(superposition,[],[f349,f340])).
fof(f9724,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),X75),s(X72,sK10(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | ~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,X79))))) )),
  inference(forward_demodulation,[],[f9697,f343])).
fof(f9697,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | ~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,X79)))) | p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76))))),s(X72,sK10(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73)))))))) )),
  inference(superposition,[],[f349,f343])).
fof(f9723,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(X64,sK10(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | ~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,X71))))) )),
  inference(forward_demodulation,[],[f9696,f320])).
fof(f9696,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | ~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,X71)))) | p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))),s(X64,sK10(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65)))))))) )),
  inference(superposition,[],[f349,f320])).
fof(f9722,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),X59),s(X57,sK10(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | ~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,X63))))) )),
  inference(forward_demodulation,[],[f9695,f312])).
fof(f9695,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | ~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,X63)))) | p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0'))),s(X57,sK10(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58)))))))) )),
  inference(superposition,[],[f349,f312])).
fof(f9721,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),X53),s(X49,sK10(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | ~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,X56))))) )),
  inference(forward_demodulation,[],[f9694,f344])).
fof(f9694,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | ~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,X56)))) | p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53))))),s(X49,sK10(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50)))))))) )),
  inference(superposition,[],[f349,f344])).
fof(f9720,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f9693,f341])).
fof(f9693,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK10(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))))) )),
  inference(superposition,[],[f349,f341])).
fof(f9719,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f9692,f359])).
fof(f9692,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK10(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))))) )),
  inference(superposition,[],[f349,f359])).
fof(f9691,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33)))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK10(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f9690,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23)))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK10(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f9718,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f9689,f305])).
fof(f9689,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK10(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))))) )),
  inference(superposition,[],[f349,f305])).
fof(f9717,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f9688,f367])).
fof(f9688,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK10(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))))) )),
  inference(superposition,[],[f349,f367])).
fof(f9687,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | ~p(s(bool,happ(s(fun(X94,bool),X99),s(X94,X100)))) | p(s(bool,happ(s(fun(X94,bool),X99),s(X94,happ(s(fun(X95,X94),X98),s(X95,sK10(X99,sK4(X96,X97,fun(X94,X95)),X98,X95,X94))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f349,f819])).
fof(f9686,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X87,X88),X90),s(fun(X88,X87),X91)))) | ~p(s(bool,happ(s(fun(X87,bool),X92),s(X87,X93)))) | p(s(bool,happ(s(fun(X87,bool),X92),s(X87,happ(s(fun(X88,X87),X91),s(X88,sK10(X92,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X91,X88,X87)))))))) )),
  inference(superposition,[],[f349,f339])).
fof(f9685,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X80,X81),X83),s(fun(X81,X80),X84)))) | ~p(s(bool,happ(s(fun(X80,bool),X85),s(X80,X86)))) | p(s(bool,happ(s(fun(X80,bool),X85),s(X80,happ(s(fun(X81,X80),X84),s(X81,sK10(X85,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X84,X81,X80)))))))) )),
  inference(superposition,[],[f349,f340])).
fof(f9684,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X72,X73),X75),s(fun(X73,X72),X77)))) | ~p(s(bool,happ(s(fun(X72,bool),X78),s(X72,X79)))) | p(s(bool,happ(s(fun(X72,bool),X78),s(X72,happ(s(fun(X73,X72),X77),s(X73,sK10(X78,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X77,X73,X72)))))))) )),
  inference(superposition,[],[f349,f343])).
fof(f9683,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(fun(X65,X64),X69)))) | ~p(s(bool,happ(s(fun(X64,bool),X70),s(X64,X71)))) | p(s(bool,happ(s(fun(X64,bool),X70),s(X64,happ(s(fun(X65,X64),X69),s(X65,sK10(X70,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X69,X65,X64)))))))) )),
  inference(superposition,[],[f349,f320])).
fof(f9682,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | ~p(s(bool,happ(s(fun(X57,bool),X62),s(X57,X63)))) | p(s(bool,happ(s(fun(X57,bool),X62),s(X57,happ(s(fun(X58,X57),X61),s(X58,sK10(X62,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X61,X58,X57)))))))) )),
  inference(superposition,[],[f349,f312])).
fof(f9681,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),X53),s(fun(X50,X49),X54)))) | ~p(s(bool,happ(s(fun(X49,bool),X55),s(X49,X56)))) | p(s(bool,happ(s(fun(X49,bool),X55),s(X49,happ(s(fun(X50,X49),X54),s(X50,sK10(X55,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X54,X50,X49)))))))) )),
  inference(superposition,[],[f349,f344])).
fof(f9680,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48)))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK10(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42)))))))) )),
  inference(superposition,[],[f349,f341])).
fof(f9679,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41)))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK10(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34)))))))) )),
  inference(superposition,[],[f349,f359])).
fof(f9678,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33)))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK10(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f9677,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23)))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK10(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14)))))))) )),
  inference(superposition,[],[f349,f368])).
fof(f9676,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK10(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8)))))))) )),
  inference(superposition,[],[f349,f305])).
fof(f9675,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK10(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0)))))))) )),
  inference(superposition,[],[f349,f367])).
fof(f9386,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK9(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9360,f256])).
fof(f9360,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,sK9(X13,X11,X12,X10,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X14)))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f348,f433])).
fof(f9385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK9(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f9384,f305])).
fof(f9384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK9(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(subsumption_resolution,[],[f9358,f256])).
fof(f9358,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,sK9(X1,'_i_0_0','_i_0_0',X0,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(superposition,[],[f348,f428])).
fof(f9357,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | ~p(s(bool,happ(s(fun(X95,bool),X99),s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,sK9(X99,X98,sK4(X96,X97,fun(X94,X95)),X94,X95))))))) | p(s(bool,happ(s(fun(X95,bool),X99),s(X95,X100)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f348,f819])).
fof(f9383,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),X90),s(X87,sK9(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | p(s(bool,happ(s(fun(X88,bool),X92),s(X88,X93))))) )),
  inference(forward_demodulation,[],[f9356,f339])).
fof(f9356,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | ~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90))))),s(X87,sK9(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88))))))) | p(s(bool,happ(s(fun(X88,bool),X92),s(X88,X93))))) )),
  inference(superposition,[],[f348,f339])).
fof(f9382,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),X83),s(X80,sK9(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | p(s(bool,happ(s(fun(X81,bool),X85),s(X81,X86))))) )),
  inference(forward_demodulation,[],[f9355,f340])).
fof(f9355,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | ~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83))))),s(X80,sK9(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81))))))) | p(s(bool,happ(s(fun(X81,bool),X85),s(X81,X86))))) )),
  inference(superposition,[],[f348,f340])).
fof(f9381,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),X75),s(X72,sK9(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | p(s(bool,happ(s(fun(X73,bool),X78),s(X73,X79))))) )),
  inference(forward_demodulation,[],[f9354,f343])).
fof(f9354,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | ~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76))))),s(X72,sK9(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73))))))) | p(s(bool,happ(s(fun(X73,bool),X78),s(X73,X79))))) )),
  inference(superposition,[],[f348,f343])).
fof(f9380,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(X64,sK9(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | p(s(bool,happ(s(fun(X65,bool),X70),s(X65,X71))))) )),
  inference(forward_demodulation,[],[f9353,f320])).
fof(f9353,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | ~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))),s(X64,sK9(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65))))))) | p(s(bool,happ(s(fun(X65,bool),X70),s(X65,X71))))) )),
  inference(superposition,[],[f348,f320])).
fof(f9379,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),X59),s(X57,sK9(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | p(s(bool,happ(s(fun(X58,bool),X62),s(X58,X63))))) )),
  inference(forward_demodulation,[],[f9352,f312])).
fof(f9352,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | ~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0'))),s(X57,sK9(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58))))))) | p(s(bool,happ(s(fun(X58,bool),X62),s(X58,X63))))) )),
  inference(superposition,[],[f348,f312])).
fof(f9378,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),X53),s(X49,sK9(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | p(s(bool,happ(s(fun(X50,bool),X55),s(X50,X56))))) )),
  inference(forward_demodulation,[],[f9351,f344])).
fof(f9351,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | ~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53))))),s(X49,sK9(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50))))))) | p(s(bool,happ(s(fun(X50,bool),X55),s(X50,X56))))) )),
  inference(superposition,[],[f348,f344])).
fof(f9377,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK9(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(forward_demodulation,[],[f9350,f341])).
fof(f9350,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,sK9(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,X48))))) )),
  inference(superposition,[],[f348,f341])).
fof(f9376,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,sK9(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(forward_demodulation,[],[f9349,f359])).
fof(f9349,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,sK9(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,X41))))) )),
  inference(superposition,[],[f348,f359])).
fof(f9348,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,sK9(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,X33))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9347,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,sK9(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,X23))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9375,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,sK9(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(forward_demodulation,[],[f9346,f305])).
fof(f9346,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,sK9(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,X13))))) )),
  inference(superposition,[],[f348,f305])).
fof(f9374,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK9(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(forward_demodulation,[],[f9345,f367])).
fof(f9345,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,sK9(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,X7))))) )),
  inference(superposition,[],[f348,f367])).
fof(f9344,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | ~p(s(bool,happ(s(fun(X94,bool),X99),s(X94,happ(s(fun(X95,X94),X98),s(X95,sK9(X99,sK4(X96,X97,fun(X94,X95)),X98,X95,X94))))))) | p(s(bool,happ(s(fun(X94,bool),X99),s(X94,X100)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f348,f819])).
fof(f9343,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X87,X88),X90),s(fun(X88,X87),X91)))) | ~p(s(bool,happ(s(fun(X87,bool),X92),s(X87,happ(s(fun(X88,X87),X91),s(X88,sK9(X92,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X91,X88,X87))))))) | p(s(bool,happ(s(fun(X87,bool),X92),s(X87,X93))))) )),
  inference(superposition,[],[f348,f339])).
fof(f9342,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X80,X81),X83),s(fun(X81,X80),X84)))) | ~p(s(bool,happ(s(fun(X80,bool),X85),s(X80,happ(s(fun(X81,X80),X84),s(X81,sK9(X85,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X84,X81,X80))))))) | p(s(bool,happ(s(fun(X80,bool),X85),s(X80,X86))))) )),
  inference(superposition,[],[f348,f340])).
fof(f9341,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X72,X73),X75),s(fun(X73,X72),X77)))) | ~p(s(bool,happ(s(fun(X72,bool),X78),s(X72,happ(s(fun(X73,X72),X77),s(X73,sK9(X78,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X77,X73,X72))))))) | p(s(bool,happ(s(fun(X72,bool),X78),s(X72,X79))))) )),
  inference(superposition,[],[f348,f343])).
fof(f9340,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(fun(X65,X64),X69)))) | ~p(s(bool,happ(s(fun(X64,bool),X70),s(X64,happ(s(fun(X65,X64),X69),s(X65,sK9(X70,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X69,X65,X64))))))) | p(s(bool,happ(s(fun(X64,bool),X70),s(X64,X71))))) )),
  inference(superposition,[],[f348,f320])).
fof(f9339,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | ~p(s(bool,happ(s(fun(X57,bool),X62),s(X57,happ(s(fun(X58,X57),X61),s(X58,sK9(X62,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X61,X58,X57))))))) | p(s(bool,happ(s(fun(X57,bool),X62),s(X57,X63))))) )),
  inference(superposition,[],[f348,f312])).
fof(f9338,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),X53),s(fun(X50,X49),X54)))) | ~p(s(bool,happ(s(fun(X49,bool),X55),s(X49,happ(s(fun(X50,X49),X54),s(X50,sK9(X55,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X54,X50,X49))))))) | p(s(bool,happ(s(fun(X49,bool),X55),s(X49,X56))))) )),
  inference(superposition,[],[f348,f344])).
fof(f9337,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,sK9(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,X48))))) )),
  inference(superposition,[],[f348,f341])).
fof(f9336,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,sK9(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,X41))))) )),
  inference(superposition,[],[f348,f359])).
fof(f9335,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,sK9(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,X33))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9334,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,sK9(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,X23))))) )),
  inference(superposition,[],[f348,f368])).
fof(f9333,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,sK9(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X13))))) )),
  inference(superposition,[],[f348,f305])).
fof(f9332,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK9(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7))))) )),
  inference(superposition,[],[f348,f367])).
fof(f9132,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f9106,f256])).
fof(f9106,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,sK8(X13,X11,X12,X10,X9))))) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,happ(s(fun(X10,X9),X12),s(X10,X14)))))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f347,f433])).
fof(f9131,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,'_i_0_0','_i_0_0',X0,X0)))))) )),
  inference(forward_demodulation,[],[f9130,f305])).
fof(f9130,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f9104,f256])).
fof(f9104,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X1,'_i_0_0','_i_0_0',X0,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f347,f428])).
fof(f9103,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | ~p(s(bool,happ(s(fun(X95,bool),X99),s(X95,sK8(X99,X98,sK4(X96,X97,fun(X94,X95)),X94,X95))))) | p(s(bool,happ(s(fun(X95,bool),X99),s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f347,f819])).
fof(f9129,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),X90),s(X87,X93)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | ~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,sK8(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88)))))) )),
  inference(forward_demodulation,[],[f9102,f339])).
fof(f9102,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X88,X87),X91),s(fun(X87,X88),X90)))) | ~p(s(bool,happ(s(fun(X88,bool),X92),s(X88,sK8(X92,X91,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X87,X88))))) | p(s(bool,happ(s(fun(X88,bool),X92),s(X88,happ(s(fun(X87,X88),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90))))),s(X87,X93))))))) )),
  inference(superposition,[],[f347,f339])).
fof(f9128,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),X83),s(X80,X86)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | ~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,sK8(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81)))))) )),
  inference(forward_demodulation,[],[f9101,f340])).
fof(f9101,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X81,X80),X84),s(fun(X80,X81),X83)))) | ~p(s(bool,happ(s(fun(X81,bool),X85),s(X81,sK8(X85,X84,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X80,X81))))) | p(s(bool,happ(s(fun(X81,bool),X85),s(X81,happ(s(fun(X80,X81),'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83))))),s(X80,X86))))))) )),
  inference(superposition,[],[f347,f340])).
fof(f9127,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),X75),s(X72,X79)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | ~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,sK8(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73)))))) )),
  inference(forward_demodulation,[],[f9100,f343])).
fof(f9100,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X73,X72),X77),s(fun(X72,X73),X75)))) | ~p(s(bool,happ(s(fun(X73,bool),X78),s(X73,sK8(X78,X77,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X72,X73))))) | p(s(bool,happ(s(fun(X73,bool),X78),s(X73,happ(s(fun(X72,X73),'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76))))),s(X72,X79))))))) )),
  inference(superposition,[],[f347,f343])).
fof(f9126,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(X64,X71)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | ~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,sK8(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65)))))) )),
  inference(forward_demodulation,[],[f9099,f320])).
fof(f9099,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X65,X64),X69),s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68)))))) | ~p(s(bool,happ(s(fun(X65,bool),X70),s(X65,sK8(X70,X69,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X64,X65))))) | p(s(bool,happ(s(fun(X65,bool),X70),s(X65,happ(s(fun(X64,X65),'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))),s(X64,X71))))))) )),
  inference(superposition,[],[f347,f320])).
fof(f9125,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),X59),s(X57,X63)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | ~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,sK8(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58)))))) )),
  inference(forward_demodulation,[],[f9098,f312])).
fof(f9098,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | ~p(s(bool,happ(s(fun(X58,bool),X62),s(X58,sK8(X62,X61,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X57,X58))))) | p(s(bool,happ(s(fun(X58,bool),X62),s(X58,happ(s(fun(X57,X58),'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0'))),s(X57,X63))))))) )),
  inference(superposition,[],[f347,f312])).
fof(f9124,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),X53),s(X49,X56)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | ~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,sK8(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50)))))) )),
  inference(forward_demodulation,[],[f9097,f344])).
fof(f9097,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X50,X49),X54),s(fun(X49,X50),X53)))) | ~p(s(bool,happ(s(fun(X50,bool),X55),s(X50,sK8(X55,X54,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X49,X50))))) | p(s(bool,happ(s(fun(X50,bool),X55),s(X50,happ(s(fun(X49,X50),'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53))))),s(X49,X56))))))) )),
  inference(superposition,[],[f347,f344])).
fof(f9123,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK8(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43)))))) )),
  inference(forward_demodulation,[],[f9096,f341])).
fof(f9096,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X43,X42),X46),s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45)))))) | ~p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK8(X47,X46,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X42,X43))))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,happ(s(fun(X42,X43),happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45))),s(X42,X48))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f9122,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),X38),s(X34,X41)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK8(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35)))))) )),
  inference(forward_demodulation,[],[f9095,f359])).
fof(f9095,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X35,X34),X39),s(fun(X34,X35),X38)))) | ~p(s(bool,happ(s(fun(X35,bool),X40),s(X35,sK8(X40,X39,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X34,X35))))) | p(s(bool,happ(s(fun(X35,bool),X40),s(X35,happ(s(fun(X34,X35),happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0'))),s(X34,X41))))))) )),
  inference(superposition,[],[f347,f359])).
fof(f9094,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | ~p(s(bool,happ(s(fun(X25,bool),X32),s(X25,sK8(X32,X31,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X24,X25))))) | p(s(bool,happ(s(fun(X25,bool),X32),s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9093,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | ~p(s(bool,happ(s(fun(X15,bool),X22),s(X15,sK8(X22,X21,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X14,X15))))) | p(s(bool,happ(s(fun(X15,bool),X22),s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9121,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),X10),s(X8,X13)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK8(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9)))))) )),
  inference(forward_demodulation,[],[f9092,f305])).
fof(f9092,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X8),X11),s(fun(X8,X9),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK8(X12,X11,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X8,X9))))) | p(s(bool,happ(s(fun(X9,bool),X12),s(X9,happ(s(fun(X8,X9),happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10))),s(X8,X13))))))) )),
  inference(superposition,[],[f347,f305])).
fof(f9120,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,X7)))))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK8(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))))) )),
  inference(forward_demodulation,[],[f9091,f367])).
fof(f9091,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X5),s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3)))))) | ~p(s(bool,happ(s(fun(X1,bool),X6),s(X1,sK8(X6,X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) | p(s(bool,happ(s(fun(X1,bool),X6),s(X1,happ(s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4))),s(X0,X7))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f9090,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | ~p(s(bool,happ(s(fun(X94,bool),X99),s(X94,sK8(X99,sK4(X96,X97,fun(X94,X95)),X98,X95,X94))))) | p(s(bool,happ(s(fun(X94,bool),X99),s(X94,happ(s(fun(X95,X94),X98),s(X95,X100)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f347,f819])).
fof(f9089,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X87,X88),X90),s(fun(X88,X87),X91)))) | ~p(s(bool,happ(s(fun(X87,bool),X92),s(X87,sK8(X92,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X87,X88),X89),'const.sum.INL_1'(s(fun(X87,X88),X90)))),X91,X88,X87))))) | p(s(bool,happ(s(fun(X87,bool),X92),s(X87,happ(s(fun(X88,X87),X91),s(X88,X93))))))) )),
  inference(superposition,[],[f347,f339])).
fof(f9088,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X80,X81),X83),s(fun(X81,X80),X84)))) | ~p(s(bool,happ(s(fun(X80,bool),X85),s(X80,sK8(X85,'const.sum.OUTR_1'(s('type.sum.sum'(X82,fun(X80,X81)),'const.sum.INR_1'(s(fun(X80,X81),X83)))),X84,X81,X80))))) | p(s(bool,happ(s(fun(X80,bool),X85),s(X80,happ(s(fun(X81,X80),X84),s(X81,X86))))))) )),
  inference(superposition,[],[f347,f340])).
fof(f9087,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X72,X73),X75),s(fun(X73,X72),X77)))) | ~p(s(bool,happ(s(fun(X72,bool),X78),s(X72,sK8(X78,'const.pair.FST_1'(s('type.pair.prod'(fun(X72,X73),X74),'const.pair.,_2'(s(fun(X72,X73),X75),s(X74,X76)))),X77,X73,X72))))) | p(s(bool,happ(s(fun(X72,bool),X78),s(X72,happ(s(fun(X73,X72),X77),s(X73,X79))))))) )),
  inference(superposition,[],[f347,f343])).
fof(f9086,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X64,X65),happ(s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num',X68))),s(fun(X65,X64),X69)))) | ~p(s(bool,happ(s(fun(X64,bool),X70),s(X64,sK8(X70,'const.ind_type.FCONS_3'(s(fun(X64,X65),X66),s(fun('type.num.num',fun(X64,X65)),X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68)))),X69,X65,X64))))) | p(s(bool,happ(s(fun(X64,bool),X70),s(X64,happ(s(fun(X65,X64),X69),s(X65,X71))))))) )),
  inference(superposition,[],[f347,f320])).
fof(f9085,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | ~p(s(bool,happ(s(fun(X57,bool),X62),s(X57,sK8(X62,'const.ind_type.FCONS_3'(s(fun(X57,X58),X59),s(fun('type.num.num',fun(X57,X58)),X60),s('type.num.num','const.num.|48|_0')),X61,X58,X57))))) | p(s(bool,happ(s(fun(X57,bool),X62),s(X57,happ(s(fun(X58,X57),X61),s(X58,X63))))))) )),
  inference(superposition,[],[f347,f312])).
fof(f9084,plain,(
  ( ! [X54,X52,X50,X56,X55,X53,X51,X49] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X49,X50),X53),s(fun(X50,X49),X54)))) | ~p(s(bool,happ(s(fun(X49,bool),X55),s(X49,sK8(X55,'const.pair.SND_1'(s('type.pair.prod'(X51,fun(X49,X50)),'const.pair.,_2'(s(X51,X52),s(fun(X49,X50),X53)))),X54,X50,X49))))) | p(s(bool,happ(s(fun(X49,bool),X55),s(X49,happ(s(fun(X50,X49),X54),s(X50,X56))))))) )),
  inference(superposition,[],[f347,f344])).
fof(f9083,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X42,X43),'const.pair.FST_1'(s('type.pair.prod'(fun(X42,X43),X44),X45))),s(fun(X43,X42),X46)))) | ~p(s(bool,happ(s(fun(X42,bool),X47),s(X42,sK8(X47,happ(s(fun('type.pair.prod'(fun(X42,X43),X44),fun(X42,X43)),'const.pair.FST_0'),s('type.pair.prod'(fun(X42,X43),X44),X45)),X46,X43,X42))))) | p(s(bool,happ(s(fun(X42,bool),X47),s(X42,happ(s(fun(X43,X42),X46),s(X43,X48))))))) )),
  inference(superposition,[],[f347,f341])).
fof(f9082,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X34,X35),X38),s(fun(X35,X34),X39)))) | ~p(s(bool,happ(s(fun(X34,bool),X40),s(X34,sK8(X40,happ(s(fun('type.list.list'(X36),fun(X34,X35)),sK16(X37,X38,X36,fun(X34,X35))),s('type.list.list'(X36),'const.list.NIL_0')),X39,X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X40),s(X34,happ(s(fun(X35,X34),X39),s(X35,X41))))))) )),
  inference(superposition,[],[f347,f359])).
fof(f9081,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | ~p(s(bool,happ(s(fun(X24,bool),X32),s(X24,sK8(X32,happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30)),X31,X25,X24))))) | p(s(bool,happ(s(fun(X24,bool),X32),s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9080,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | ~p(s(bool,happ(s(fun(X14,bool),X22),s(X14,sK8(X22,happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20)))),X21,X15,X14))))) | p(s(bool,happ(s(fun(X14,bool),X22),s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))))))) )),
  inference(superposition,[],[f347,f368])).
fof(f9079,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X8,X9),X10),s(fun(X9,X8),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,sK8(X12,happ(s(fun(fun(X8,X9),fun(X8,X9)),'_i_0_0'),s(fun(X8,X9),X10)),X11,X9,X8))))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,happ(s(fun(X9,X8),X11),s(X9,X13))))))) )),
  inference(superposition,[],[f347,f305])).
fof(f9078,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK8(X6,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) )),
  inference(superposition,[],[f347,f367])).
fof(f8592,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f8559,f256])).
fof(f8559,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X10,X13) != s(X10,happ(s(fun(X9,X10),X11),s(X9,X14))) | s(X9,X14) = s(X9,happ(s(fun(X10,X9),X12),s(X10,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f352,f433])).
fof(f8556,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f352,f819])).
fof(f8547,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f352,f368])).
fof(f8546,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f352,f368])).
fof(f8543,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f352,f819])).
fof(f8534,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f352,f368])).
fof(f8533,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f352,f368])).
fof(f8051,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f8018,f256])).
fof(f8018,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,X13) != s(X9,happ(s(fun(X10,X9),X12),s(X10,X14))) | s(X10,X14) = s(X10,happ(s(fun(X9,X10),X11),s(X9,X13))) | s(bool,X15) = s(bool,'const.ind_type.ISO_2'(s(fun(X9,X10),X11),s(fun(X10,X9),X12))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f351,f433])).
fof(f8050,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(subsumption_resolution,[],[f8016,f256])).
fof(f8016,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X2))) | s(X0,X2) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f351,f428])).
fof(f8015,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X95,X94),X98),s(fun(X94,X95),X96)))) | s(X95,X99) != s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X100))) | s(X94,X100) = s(X94,happ(s(fun(X95,X94),X98),s(X95,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f351,f819])).
fof(f8006,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X25,X24),X31),s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30)))))))) | s(X25,X32) != s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X33))) | s(X24,X33) = s(X24,happ(s(fun(X25,X24),X31),s(X25,X32)))) )),
  inference(superposition,[],[f351,f368])).
fof(f8005,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X15,X14),X21),s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20)))))) | s(X15,X22) != s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X23))) | s(X14,X23) = s(X14,happ(s(fun(X15,X14),X21),s(X15,X22)))) )),
  inference(superposition,[],[f351,f368])).
fof(f8002,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X94,X95),X96),s(fun(X95,X94),X98)))) | s(X94,X99) != s(X94,happ(s(fun(X95,X94),X98),s(X95,X100))) | s(X95,X100) = s(X95,happ(s(fun(X94,X95),sK4(X96,X97,fun(X94,X95))),s(X94,X99))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X94,X95),bool),X97)))) )),
  inference(superposition,[],[f351,f819])).
fof(f7993,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X24,X25),happ(s(fun(X27,fun(X24,X25)),X28),s(X27,happ(s(fun(X26,X27),X29),s(X26,X30))))),s(fun(X25,X24),X31)))) | s(X24,X32) != s(X24,happ(s(fun(X25,X24),X31),s(X25,X33))) | s(X25,X33) = s(X25,happ(s(fun(X24,X25),happ(s(fun(X26,fun(X24,X25)),'_dst_x0x1_2'(s(fun(X27,fun(X24,X25)),X28),s(fun(X26,X27),X29))),s(X26,X30))),s(X24,X32)))) )),
  inference(superposition,[],[f351,f368])).
fof(f7992,plain,(
  ( ! [X14,X23,X21,X19,X17,X15,X22,X20,X18,X16] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X14,X15),happ(s(fun(X18,fun(X14,X15)),'_dst_x0x1_2'(s(fun(X16,fun(X14,X15)),X17),s(fun(X18,X16),X19))),s(X18,X20))),s(fun(X15,X14),X21)))) | s(X14,X22) != s(X14,happ(s(fun(X15,X14),X21),s(X15,X23))) | s(X15,X23) = s(X15,happ(s(fun(X14,X15),happ(s(fun(X16,fun(X14,X15)),X17),s(X16,happ(s(fun(X18,X16),X19),s(X18,X20))))),s(X14,X22)))) )),
  inference(superposition,[],[f351,f368])).
fof(f7803,plain,(
  ( ! [X70,X71] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X70),s('type.list.list'('type.bvlSem.v'),sK13(happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)),X71,'type.list.list'('type.bvlSem.v'),'type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),X71),s('type.list.list'('type.bvlSem.v'),sK13(happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)),X71,'type.list.list'('type.bvlSem.v'),'type.bvlSem.v')))) | s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),X71) = s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)))) )),
  inference(superposition,[],[f355,f328])).
fof(f7802,plain,(
  ( ! [X69] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK13('const.bvlSem.Word64_0',X69,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),X69),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK13('const.bvlSem.Word64_0',X69,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v')))) | s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0') = s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),X69)) )),
  inference(superposition,[],[f355,f265])).
fof(f7801,plain,(
  ( ! [X68] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK13('const.bvlSem.Number_0',X68,'type.integer.int','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),X68),s('type.integer.int',sK13('const.bvlSem.Number_0',X68,'type.integer.int','type.bvlSem.v')))) | s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0') = s(fun('type.integer.int','type.bvlSem.v'),X68)) )),
  inference(superposition,[],[f355,f262])).
fof(f7800,plain,(
  ( ! [X67] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK13('const.bvlSem.CodePtr_0',X67,'type.num.num','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),X67),s('type.num.num',sK13('const.bvlSem.CodePtr_0',X67,'type.num.num','type.bvlSem.v')))) | s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0') = s(fun('type.num.num','type.bvlSem.v'),X67)) )),
  inference(superposition,[],[f355,f261])).
fof(f7799,plain,(
  ( ! [X66] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK13('const.bvlSem.RefPtr_0',X66,'type.num.num','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),X66),s('type.num.num',sK13('const.bvlSem.RefPtr_0',X66,'type.num.num','type.bvlSem.v')))) | s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0') = s(fun('type.num.num','type.bvlSem.v'),X66)) )),
  inference(superposition,[],[f355,f260])).
fof(f7798,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK13('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK13('_dstsel_0_x01_1'(s(fun(X61,X63),X64)),X65,fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f355,f367])).
fof(f7787,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK13(X20,X21,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK13(X20,X21,X19,bool))))) )),
  inference(superposition,[],[f355,f433])).
fof(f7786,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK13(X16,X17,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK13(X16,X17,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f355,f433])).
fof(f7785,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK13(X13,X14,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK13(X13,X14,X12,bool))))) )),
  inference(superposition,[],[f355,f266])).
fof(f7784,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK13('_i_0_0',X11,bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK13('_i_0_0',X11,bool,bool))) )),
  inference(superposition,[],[f355,f460])).
fof(f7783,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK13('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK13('const.pair.FST_0',X10,'type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f355,f341])).
fof(f7782,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK13('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK13('_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X7,X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f355,f368])).
fof(f7781,plain,(
  ( ! [X0,X1] : (s(X0,sK13('_i_0_0',X1,X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK13('_i_0_0',X1,X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f355,f305])).
fof(f7770,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,happ(s(fun(bool,X9),X7),s(bool,X10))) != s(X9,happ(s(fun(bool,X9),X8),s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK13(X7,X8,bool,X9))) )),
  inference(superposition,[],[f355,f433])).
fof(f7769,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK13(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f355,f433])).
fof(f7768,plain,(
  ( ! [X2,X0,X1] : (s(X2,happ(s(fun(bool,X2),X0),s(bool,'F_0'))) != s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,X2))) )),
  inference(superposition,[],[f355,f266])).
fof(f7767,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK13(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK13(sK4(X72,X73,fun(X70,X71)),X74,X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f355,f819])).
fof(f7828,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(fun(X65,X66),X68) = s(fun(X65,X66),X69) | s(X66,happ(s(fun(X65,X66),X68),s(X65,sK13('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X69,X65,X66)))) != s(X66,happ(s(fun(X65,X66),X69),s(X65,sK13('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X69,X65,X66))))) )),
  inference(forward_demodulation,[],[f7766,f339])).
fof(f7766,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(X66,happ(s(fun(X65,X66),X68),s(X65,sK13('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X69,X65,X66)))) != s(X66,happ(s(fun(X65,X66),X69),s(X65,sK13('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X69,X65,X66)))) | s(fun(X65,X66),X69) = s(fun(X65,X66),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))))) )),
  inference(superposition,[],[f355,f339])).
fof(f7827,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(fun(X60,X61),X63) = s(fun(X60,X61),X64) | s(X61,happ(s(fun(X60,X61),X63),s(X60,sK13('const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X64,X60,X61)))) != s(X61,happ(s(fun(X60,X61),X64),s(X60,sK13('const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X64,X60,X61))))) )),
  inference(forward_demodulation,[],[f7765,f340])).
fof(f7765,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(X61,happ(s(fun(X60,X61),X63),s(X60,sK13('const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X64,X60,X61)))) != s(X61,happ(s(fun(X60,X61),X64),s(X60,sK13('const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X64,X60,X61)))) | s(fun(X60,X61),X64) = s(fun(X60,X61),'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))))) )),
  inference(superposition,[],[f355,f340])).
fof(f7826,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(fun(X54,X55),X57) = s(fun(X54,X55),X59) | s(X55,happ(s(fun(X54,X55),X57),s(X54,sK13('const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X59,X54,X55)))) != s(X55,happ(s(fun(X54,X55),X59),s(X54,sK13('const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X59,X54,X55))))) )),
  inference(forward_demodulation,[],[f7764,f343])).
fof(f7764,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(X55,happ(s(fun(X54,X55),X57),s(X54,sK13('const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X59,X54,X55)))) != s(X55,happ(s(fun(X54,X55),X59),s(X54,sK13('const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X59,X54,X55)))) | s(fun(X54,X55),X59) = s(fun(X54,X55),'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))))) )),
  inference(superposition,[],[f355,f343])).
fof(f7825,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(fun(X48,X49),X53) = s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))) | s(X49,happ(s(fun(X48,X49),X53),s(X48,sK13('const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,X48,X49)))) != s(X49,happ(s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))),s(X48,sK13('const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,X48,X49))))) )),
  inference(forward_demodulation,[],[f7763,f320])).
fof(f7763,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X49,happ(s(fun(X48,X49),X53),s(X48,sK13('const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,X48,X49)))) != s(X49,happ(s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))),s(X48,sK13('const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X53,X48,X49)))) | s(fun(X48,X49),X53) = s(fun(X48,X49),'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))))) )),
  inference(superposition,[],[f355,f320])).
fof(f7824,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(fun(X43,X44),X45) = s(fun(X43,X44),X47) | s(X44,happ(s(fun(X43,X44),X45),s(X43,sK13('const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X47,X43,X44)))) != s(X44,happ(s(fun(X43,X44),X47),s(X43,sK13('const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X47,X43,X44))))) )),
  inference(forward_demodulation,[],[f7762,f312])).
fof(f7762,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(X44,happ(s(fun(X43,X44),X45),s(X43,sK13('const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X47,X43,X44)))) != s(X44,happ(s(fun(X43,X44),X47),s(X43,sK13('const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X47,X43,X44)))) | s(fun(X43,X44),X47) = s(fun(X43,X44),'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f355,f312])).
fof(f7823,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(fun(X37,X38),X41) = s(fun(X37,X38),X42) | s(X38,happ(s(fun(X37,X38),X41),s(X37,sK13('const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X42,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X42),s(X37,sK13('const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X42,X37,X38))))) )),
  inference(forward_demodulation,[],[f7761,f344])).
fof(f7761,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(X38,happ(s(fun(X37,X38),X41),s(X37,sK13('const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X42,X37,X38)))) != s(X38,happ(s(fun(X37,X38),X42),s(X37,sK13('const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X42,X37,X38)))) | s(fun(X37,X38),X42) = s(fun(X37,X38),'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))))) )),
  inference(superposition,[],[f355,f344])).
fof(f7822,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK13(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK13(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33))))) )),
  inference(forward_demodulation,[],[f7760,f341])).
fof(f7760,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK13(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK13(happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X36,X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f355,f341])).
fof(f7821,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK13(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK13(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27))))) )),
  inference(forward_demodulation,[],[f7759,f359])).
fof(f7759,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK13(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK13(happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X31,X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f355,f359])).
fof(f7758,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK13(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK13(happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X25,X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f355,f368])).
fof(f7757,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK13(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK13(happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X17,X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f355,f368])).
fof(f7820,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK13(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK13(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7))))) )),
  inference(forward_demodulation,[],[f7756,f305])).
fof(f7756,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK13(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK13(happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X9,X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f355,f305])).
fof(f7819,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK13(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK13(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1))))) )),
  inference(forward_demodulation,[],[f7755,f367])).
fof(f7755,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK13(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK13(happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X5,X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f355,f367])).
fof(f7754,plain,(
  ( ! [X70,X71] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X70),s('type.list.list'('type.bvlSem.v'),sK13(X71,happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)),'type.list.list'('type.bvlSem.v'),'type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),X71),s('type.list.list'('type.bvlSem.v'),sK13(X71,happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)),'type.list.list'('type.bvlSem.v'),'type.bvlSem.v')))) | s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),X71) = s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X70)))) )),
  inference(superposition,[],[f355,f328])).
fof(f7753,plain,(
  ( ! [X69] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK13(X69,'const.bvlSem.Word64_0','type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),X69),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK13(X69,'const.bvlSem.Word64_0','type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v')))) | s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0') = s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),X69)) )),
  inference(superposition,[],[f355,f265])).
fof(f7752,plain,(
  ( ! [X68] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK13(X68,'const.bvlSem.Number_0','type.integer.int','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),X68),s('type.integer.int',sK13(X68,'const.bvlSem.Number_0','type.integer.int','type.bvlSem.v')))) | s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0') = s(fun('type.integer.int','type.bvlSem.v'),X68)) )),
  inference(superposition,[],[f355,f262])).
fof(f7751,plain,(
  ( ! [X67] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',sK13(X67,'const.bvlSem.CodePtr_0','type.num.num','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),X67),s('type.num.num',sK13(X67,'const.bvlSem.CodePtr_0','type.num.num','type.bvlSem.v')))) | s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0') = s(fun('type.num.num','type.bvlSem.v'),X67)) )),
  inference(superposition,[],[f355,f261])).
fof(f7750,plain,(
  ( ! [X66] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',sK13(X66,'const.bvlSem.RefPtr_0','type.num.num','type.bvlSem.v')))) != s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),X66),s('type.num.num',sK13(X66,'const.bvlSem.RefPtr_0','type.num.num','type.bvlSem.v')))) | s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0') = s(fun('type.num.num','type.bvlSem.v'),X66)) )),
  inference(superposition,[],[f355,f260])).
fof(f7749,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(fun(X61,X62),'_dst_x0x1_2'(s(fun(X63,X62),sK13(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))),s(fun(X61,X63),X64))) != s(fun(X61,X62),happ(s(fun(fun(X63,X62),fun(X61,X62)),X65),s(fun(X63,X62),sK13(X65,'_dstsel_0_x01_1'(s(fun(X61,X63),X64)),fun(X63,X62),fun(X61,X62))))) | s(fun(fun(X63,X62),fun(X61,X62)),X65) = s(fun(fun(X63,X62),fun(X61,X62)),'_dstsel_0_x01_1'(s(fun(X61,X63),X64)))) )),
  inference(superposition,[],[f355,f367])).
fof(f7738,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X22) != s(bool,happ(s(fun(X19,bool),X21),s(X19,sK13(X21,X20,X19,bool)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X20),s(X19,sK13(X21,X20,X19,bool))))) )),
  inference(superposition,[],[f355,f433])).
fof(f7737,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,happ(s(fun(X15,bool),X17),s(X15,sK13(X17,X16,X15,bool)))) | s(fun(X15,bool),X16) = s(fun(X15,bool),X17) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,sK13(X17,X16,X15,bool)))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f355,f433])).
fof(f7736,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(X12,bool),X14),s(X12,sK13(X14,X13,X12,bool)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK13(X14,X13,X12,bool))))) )),
  inference(superposition,[],[f355,f266])).
fof(f7735,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X11),s(bool,sK13(X11,'_i_0_0',bool,bool)))) | s(fun(bool,bool),'_i_0_0') = s(fun(bool,bool),X11) | s(bool,'F_0') = s(bool,sK13(X11,'_i_0_0',bool,bool))) )),
  inference(superposition,[],[f355,f460])).
fof(f7734,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X9),sK13(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) != s(X8,happ(s(fun('type.pair.prod'(X8,X9),X8),X10),s('type.pair.prod'(X8,X9),sK13(X10,'const.pair.FST_0','type.pair.prod'(X8,X9),X8)))) | s(fun('type.pair.prod'(X8,X9),X8),'const.pair.FST_0') = s(fun('type.pair.prod'(X8,X9),X8),X10)) )),
  inference(superposition,[],[f355,f341])).
fof(f7733,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X2,happ(s(fun(X3,X2),X7),s(X3,sK13(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))) != s(X2,happ(s(fun(X4,X2),X5),s(X4,happ(s(fun(X3,X4),X6),s(X3,sK13(X7,'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)),X3,X2)))))) | s(fun(X3,X2),X7) = s(fun(X3,X2),'_dst_x0x1_2'(s(fun(X4,X2),X5),s(fun(X3,X4),X6)))) )),
  inference(superposition,[],[f355,f368])).
fof(f7732,plain,(
  ( ! [X0,X1] : (s(X0,sK13(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK13(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f355,f305])).
fof(f7731,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(X71,happ(s(fun(X70,X71),X72),s(X70,sK13(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) != s(X71,happ(s(fun(X70,X71),X74),s(X70,sK13(X74,sK4(X72,X73,fun(X70,X71)),X70,X71)))) | s(fun(X70,X71),X74) = s(fun(X70,X71),sK4(X72,X73,fun(X70,X71))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,X71),bool),X73)))) )),
  inference(superposition,[],[f355,f819])).
fof(f7814,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(fun(X65,X66),X68) = s(fun(X65,X66),X69) | s(X66,happ(s(fun(X65,X66),X68),s(X65,sK13(X69,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X65,X66)))) != s(X66,happ(s(fun(X65,X66),X69),s(X65,sK13(X69,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X65,X66))))) )),
  inference(forward_demodulation,[],[f7730,f339])).
fof(f7730,plain,(
  ( ! [X68,X66,X69,X67,X65] : (s(X66,happ(s(fun(X65,X66),X68),s(X65,sK13(X69,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X65,X66)))) != s(X66,happ(s(fun(X65,X66),X69),s(X65,sK13(X69,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))),X65,X66)))) | s(fun(X65,X66),X69) = s(fun(X65,X66),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,X66),X67),'const.sum.INL_1'(s(fun(X65,X66),X68)))))) )),
  inference(superposition,[],[f355,f339])).
fof(f7813,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(fun(X60,X61),X63) = s(fun(X60,X61),X64) | s(X61,happ(s(fun(X60,X61),X63),s(X60,sK13(X64,'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X60,X61)))) != s(X61,happ(s(fun(X60,X61),X64),s(X60,sK13(X64,'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X60,X61))))) )),
  inference(forward_demodulation,[],[f7729,f340])).
fof(f7729,plain,(
  ( ! [X61,X64,X62,X60,X63] : (s(X61,happ(s(fun(X60,X61),X63),s(X60,sK13(X64,'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X60,X61)))) != s(X61,happ(s(fun(X60,X61),X64),s(X60,sK13(X64,'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))),X60,X61)))) | s(fun(X60,X61),X64) = s(fun(X60,X61),'const.sum.OUTR_1'(s('type.sum.sum'(X62,fun(X60,X61)),'const.sum.INR_1'(s(fun(X60,X61),X63)))))) )),
  inference(superposition,[],[f355,f340])).
fof(f7812,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(fun(X54,X55),X57) = s(fun(X54,X55),X59) | s(X55,happ(s(fun(X54,X55),X57),s(X54,sK13(X59,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X54,X55)))) != s(X55,happ(s(fun(X54,X55),X59),s(X54,sK13(X59,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X54,X55))))) )),
  inference(forward_demodulation,[],[f7728,f343])).
fof(f7728,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(X55,happ(s(fun(X54,X55),X57),s(X54,sK13(X59,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X54,X55)))) != s(X55,happ(s(fun(X54,X55),X59),s(X54,sK13(X59,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))),X54,X55)))) | s(fun(X54,X55),X59) = s(fun(X54,X55),'const.pair.FST_1'(s('type.pair.prod'(fun(X54,X55),X56),'const.pair.,_2'(s(fun(X54,X55),X57),s(X56,X58)))))) )),
  inference(superposition,[],[f355,f343])).
fof(f7811,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(fun(X48,X49),X53) = s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))) | s(X49,happ(s(fun(X48,X49),X53),s(X48,sK13(X53,'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))),s(X48,sK13(X53,'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X48,X49))))) )),
  inference(forward_demodulation,[],[f7727,f320])).
fof(f7727,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X49,happ(s(fun(X48,X49),X53),s(X48,sK13(X53,'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X48,X49)))) != s(X49,happ(s(fun(X48,X49),happ(s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num',X52))),s(X48,sK13(X53,'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))),X48,X49)))) | s(fun(X48,X49),X53) = s(fun(X48,X49),'const.ind_type.FCONS_3'(s(fun(X48,X49),X50),s(fun('type.num.num',fun(X48,X49)),X51),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))))) )),
  inference(superposition,[],[f355,f320])).
fof(f7810,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(fun(X43,X44),X45) = s(fun(X43,X44),X47) | s(X44,happ(s(fun(X43,X44),X45),s(X43,sK13(X47,'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X43,X44)))) != s(X44,happ(s(fun(X43,X44),X47),s(X43,sK13(X47,'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X43,X44))))) )),
  inference(forward_demodulation,[],[f7726,f312])).
fof(f7726,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(X44,happ(s(fun(X43,X44),X45),s(X43,sK13(X47,'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X43,X44)))) != s(X44,happ(s(fun(X43,X44),X47),s(X43,sK13(X47,'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')),X43,X44)))) | s(fun(X43,X44),X47) = s(fun(X43,X44),'const.ind_type.FCONS_3'(s(fun(X43,X44),X45),s(fun('type.num.num',fun(X43,X44)),X46),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f355,f312])).
fof(f7809,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(fun(X37,X38),X41) = s(fun(X37,X38),X42) | s(X38,happ(s(fun(X37,X38),X41),s(X37,sK13(X42,'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X42),s(X37,sK13(X42,'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X37,X38))))) )),
  inference(forward_demodulation,[],[f7725,f344])).
fof(f7725,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(X38,happ(s(fun(X37,X38),X41),s(X37,sK13(X42,'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X37,X38)))) != s(X38,happ(s(fun(X37,X38),X42),s(X37,sK13(X42,'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))),X37,X38)))) | s(fun(X37,X38),X42) = s(fun(X37,X38),'const.pair.SND_1'(s('type.pair.prod'(X39,fun(X37,X38)),'const.pair.,_2'(s(X39,X40),s(fun(X37,X38),X41)))))) )),
  inference(superposition,[],[f355,f344])).
fof(f7808,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(fun(X32,X33),X36) = s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))) | s(X33,happ(s(fun(X32,X33),X36),s(X32,sK13(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK13(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33))))) )),
  inference(forward_demodulation,[],[f7724,f341])).
fof(f7724,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(X33,happ(s(fun(X32,X33),X36),s(X32,sK13(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) != s(X33,happ(s(fun(X32,X33),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,X33),X34),X35))),s(X32,sK13(X36,happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)),X32,X33)))) | s(fun(X32,X33),X36) = s(fun(X32,X33),happ(s(fun('type.pair.prod'(fun(X32,X33),X34),fun(X32,X33)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,X33),X34),X35)))) )),
  inference(superposition,[],[f355,f341])).
fof(f7807,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(fun(X26,X27),X30) = s(fun(X26,X27),X31) | s(X27,happ(s(fun(X26,X27),X30),s(X26,sK13(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK13(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27))))) )),
  inference(forward_demodulation,[],[f7723,f359])).
fof(f7723,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(X27,happ(s(fun(X26,X27),X30),s(X26,sK13(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) != s(X27,happ(s(fun(X26,X27),X31),s(X26,sK13(X31,happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')),X26,X27)))) | s(fun(X26,X27),X31) = s(fun(X26,X27),happ(s(fun('type.list.list'(X28),fun(X26,X27)),sK16(X29,X30,X28,fun(X26,X27))),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f355,f359])).
fof(f7722,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (s(X19,happ(s(fun(X18,X19),X25),s(X18,sK13(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) != s(X19,happ(s(fun(X18,X19),happ(s(fun(X21,fun(X18,X19)),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24))))),s(X18,sK13(X25,happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)),X18,X19)))) | s(fun(X18,X19),X25) = s(fun(X18,X19),happ(s(fun(X20,fun(X18,X19)),'_dst_x0x1_2'(s(fun(X21,fun(X18,X19)),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f355,f368])).
fof(f7721,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,happ(s(fun(X10,X11),X17),s(X10,sK13(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) != s(X11,happ(s(fun(X10,X11),happ(s(fun(X14,fun(X10,X11)),'_dst_x0x1_2'(s(fun(X12,fun(X10,X11)),X13),s(fun(X14,X12),X15))),s(X14,X16))),s(X10,sK13(X17,happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))),X10,X11)))) | s(fun(X10,X11),X17) = s(fun(X10,X11),happ(s(fun(X12,fun(X10,X11)),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16)))))) )),
  inference(superposition,[],[f355,f368])).
fof(f7806,plain,(
  ( ! [X6,X8,X7,X9] : (s(fun(X6,X7),X8) = s(fun(X6,X7),X9) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK13(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK13(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7))))) )),
  inference(forward_demodulation,[],[f7720,f305])).
fof(f7720,plain,(
  ( ! [X6,X8,X7,X9] : (s(X7,happ(s(fun(X6,X7),X8),s(X6,sK13(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) != s(X7,happ(s(fun(X6,X7),X9),s(X6,sK13(X9,happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)),X6,X7)))) | s(fun(X6,X7),X9) = s(fun(X6,X7),happ(s(fun(fun(X6,X7),fun(X6,X7)),'_i_0_0'),s(fun(X6,X7),X8)))) )),
  inference(superposition,[],[f355,f305])).
fof(f7805,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X1),X5) = s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))) | s(X1,happ(s(fun(X0,X1),X5),s(X0,sK13(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK13(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1))))) )),
  inference(forward_demodulation,[],[f7719,f367])).
fof(f7719,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X5),s(X0,sK13(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) != s(X1,happ(s(fun(X0,X1),'_dst_x0x1_2'(s(fun(X2,X1),X4),s(fun(X0,X2),X3))),s(X0,sK13(X5,happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)),X0,X1)))) | s(fun(X0,X1),X5) = s(fun(X0,X1),happ(s(fun(fun(X2,X1),fun(X0,X1)),'_dstsel_0_x01_1'(s(fun(X0,X2),X3))),s(fun(X2,X1),X4)))) )),
  inference(superposition,[],[f355,f367])).
fof(f7562,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74))))) )),
  inference(forward_demodulation,[],[f7536,f341])).
fof(f7536,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),sK5(X73,'const.pair.FST_0','type.pair.prod'(bool,X72)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),X73),s('type.pair.prod'(bool,X72),X74)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),sK6(X73,'const.pair.FST_0','type.pair.prod'(bool,X72))))))) )),
  inference(superposition,[],[f317,f341])).
fof(f7535,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X67,bool),X68),s(X67,happ(s(fun(X66,X67),X69),s(X66,sK5(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66))))))) | ~p(s(bool,happ(s(fun(X66,bool),X70),s(X66,X71)))) | p(s(bool,happ(s(fun(X66,bool),'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69))),s(X66,sK6(X70,'_dst_x0x1_2'(s(fun(X67,bool),X68),s(fun(X66,X67),X69)),X66)))))) )),
  inference(superposition,[],[f317,f368])).
fof(f7561,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | p(s(bool,sK6(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65))))) )),
  inference(forward_demodulation,[],[f7534,f305])).
fof(f7534,plain,(
  ( ! [X64,X65] : (~p(s(bool,sK5(X64,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X64),s(bool,X65)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X64,'_i_0_0',bool)))))) )),
  inference(superposition,[],[f317,f305])).
fof(f7553,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7522,f256])).
fof(f7522,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,X7,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,X7,X6)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f317,f433])).
fof(f7552,plain,(
  ( ! [X0,X1] : (p(s(bool,sK6(X0,'_i_0_0',bool))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(forward_demodulation,[],[f7551,f305])).
fof(f7551,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f7520,f256])).
fof(f7520,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6(X0,'_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK5(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f317,f460])).
fof(f7509,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X9)))) | ~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X10)))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK6(X7,X8,bool))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool))) )),
  inference(superposition,[],[f317,f433])).
fof(f7508,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,sK6(X3,X4,bool))))) | s(bool,X6) = s(bool,sK5(X3,X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f317,f433])).
fof(f7507,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6(X0,X1,bool))))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool))) )),
  inference(superposition,[],[f317,f266])).
fof(f7506,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,sK5(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | ~p(s(bool,happ(s(fun(X70,bool),X73),s(X70,X74)))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK6(X73,sK4(X71,X72,fun(X70,bool)),X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f317,f819])).
fof(f7546,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),X67),s(X65,sK6(X68,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X65))))) | ~p(s(bool,happ(s(fun(X65,bool),X67),s(X65,sK5(X68,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X65))))) | ~p(s(bool,happ(s(fun(X65,bool),X68),s(X65,X69))))) )),
  inference(forward_demodulation,[],[f7505,f339])).
fof(f7505,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),X67),s(X65,sK5(X68,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X65))))) | ~p(s(bool,happ(s(fun(X65,bool),X68),s(X65,X69)))) | p(s(bool,happ(s(fun(X65,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67))))),s(X65,sK6(X68,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X65)))))) )),
  inference(superposition,[],[f317,f339])).
fof(f7545,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun(X60,bool),X62),s(X60,sK6(X63,'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X60))))) | ~p(s(bool,happ(s(fun(X60,bool),X62),s(X60,sK5(X63,'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X60))))) | ~p(s(bool,happ(s(fun(X60,bool),X63),s(X60,X64))))) )),
  inference(forward_demodulation,[],[f7504,f340])).
fof(f7504,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun(X60,bool),X62),s(X60,sK5(X63,'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X60))))) | ~p(s(bool,happ(s(fun(X60,bool),X63),s(X60,X64)))) | p(s(bool,happ(s(fun(X60,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62))))),s(X60,sK6(X63,'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X60)))))) )),
  inference(superposition,[],[f317,f340])).
fof(f7544,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,happ(s(fun(X54,bool),X56),s(X54,sK6(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X54))))) | ~p(s(bool,happ(s(fun(X54,bool),X56),s(X54,sK5(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X54))))) | ~p(s(bool,happ(s(fun(X54,bool),X58),s(X54,X59))))) )),
  inference(forward_demodulation,[],[f7503,f343])).
fof(f7503,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),X56),s(X54,sK5(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X54))))) | ~p(s(bool,happ(s(fun(X54,bool),X58),s(X54,X59)))) | p(s(bool,happ(s(fun(X54,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57))))),s(X54,sK6(X58,'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X54)))))) )),
  inference(superposition,[],[f317,f343])).
fof(f7543,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,sK6(X52,'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,sK5(X52,'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X52),s(X48,X53))))) )),
  inference(forward_demodulation,[],[f7502,f320])).
fof(f7502,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,sK5(X52,'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X48))))) | ~p(s(bool,happ(s(fun(X48,bool),X52),s(X48,X53)))) | p(s(bool,happ(s(fun(X48,bool),'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))),s(X48,sK6(X52,'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X48)))))) )),
  inference(superposition,[],[f317,f320])).
fof(f7542,plain,(
  ( ! [X47,X45,X43,X46,X44] : (p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X43))))) | ~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X43))))) | ~p(s(bool,happ(s(fun(X43,bool),X46),s(X43,X47))))) )),
  inference(forward_demodulation,[],[f7501,f312])).
fof(f7501,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK5(X46,'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X43))))) | ~p(s(bool,happ(s(fun(X43,bool),X46),s(X43,X47)))) | p(s(bool,happ(s(fun(X43,bool),'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0'))),s(X43,sK6(X46,'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X43)))))) )),
  inference(superposition,[],[f317,f312])).
fof(f7541,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X40),s(X37,sK6(X41,'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,sK5(X41,'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X41),s(X37,X42))))) )),
  inference(forward_demodulation,[],[f7500,f344])).
fof(f7500,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,sK5(X41,'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X41),s(X37,X42)))) | p(s(bool,happ(s(fun(X37,bool),'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40))))),s(X37,sK6(X41,'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X37)))))) )),
  inference(superposition,[],[f317,f344])).
fof(f7540,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36))))) )),
  inference(forward_demodulation,[],[f7499,f341])).
fof(f7499,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32))))) | ~p(s(bool,happ(s(fun(X32,bool),X35),s(X32,X36)))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK6(X35,happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X32)))))) )),
  inference(superposition,[],[f317,f341])).
fof(f7539,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31))))) )),
  inference(forward_demodulation,[],[f7498,f359])).
fof(f7498,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X30),s(X26,X31)))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK6(X30,happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X26)))))) )),
  inference(superposition,[],[f317,f359])).
fof(f7497,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,sK5(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18))))) | ~p(s(bool,happ(s(fun(X18,bool),X24),s(X18,X25)))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK6(X24,happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X18)))))) )),
  inference(superposition,[],[f317,f368])).
fof(f7496,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,sK5(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10))))) | ~p(s(bool,happ(s(fun(X10,bool),X16),s(X10,X17)))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK6(X16,happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X10)))))) )),
  inference(superposition,[],[f317,f368])).
fof(f7538,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))))) )),
  inference(forward_demodulation,[],[f7495,f305])).
fof(f7495,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK6(X8,happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X6)))))) )),
  inference(superposition,[],[f317,f305])).
fof(f7537,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f7494,f367])).
fof(f7494,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK6(X4,happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X0)))))) )),
  inference(superposition,[],[f317,f367])).
fof(f7310,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | ~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(forward_demodulation,[],[f7281,f341])).
fof(f7281,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),X83)))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),X84),s('type.pair.prod'(bool,X82),sK6('const.pair.FST_0',X84,'type.pair.prod'(bool,X82)))))) | p(s(bool,happ(s(fun('type.pair.prod'(bool,X82),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X82),sK5('const.pair.FST_0',X84,'type.pair.prod'(bool,X82))))))) )),
  inference(superposition,[],[f316,f341])).
fof(f7280,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X80)) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),X81),s('type.list.list'(X78),sK6(sK16(X79,X80,X78,bool),X81,'type.list.list'(X78)))))) | p(s(bool,happ(s(fun('type.list.list'(X78),bool),sK16(X79,X80,X78,bool)),s('type.list.list'(X78),sK5(sK16(X79,X80,X78,bool),X81,'type.list.list'(X78))))))) )),
  inference(superposition,[],[f316,f359])).
fof(f7279,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | p(s(bool,happ(s(fun(X72,bool),X77),s(X72,sK6('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72))))) | p(s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,sK5('_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75)),X77,X72)))))) )),
  inference(superposition,[],[f316,f368])).
fof(f7278,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (~p(s(bool,happ(s(fun(X68,bool),'_dst_x0x1_2'(s(fun(X66,bool),X67),s(fun(X68,X66),X69))),s(X68,X70)))) | p(s(bool,happ(s(fun(X66,bool),X71),s(X66,sK6(X67,X71,X66))))) | p(s(bool,happ(s(fun(X66,bool),X67),s(X66,sK5(X67,X71,X66)))))) )),
  inference(superposition,[],[f316,f368])).
fof(f7303,plain,(
  ( ! [X6,X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7265,f256])).
fof(f7265,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(X7,X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(X7,X9,X6))))) | s(bool,X10) = s(bool,happ(s(fun(X6,bool),X7),s(X6,X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f316,f433])).
fof(f7207,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,happ(s(fun(X70,bool),X71),s(X70,X73)))) | p(s(bool,happ(s(fun(X70,bool),X74),s(X70,sK6(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | p(s(bool,happ(s(fun(X70,bool),sK4(X71,X72,fun(X70,bool))),s(X70,sK5(sK4(X71,X72,fun(X70,bool)),X74,X70))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X70,bool),bool),X72)))) )),
  inference(superposition,[],[f316,f819])).
fof(f7299,plain,(
  ( ! [X68,X66,X69,X67,X65] : (p(s(bool,happ(s(fun(X65,bool),X67),s(X65,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X69,X65))))) | ~p(s(bool,happ(s(fun(X65,bool),X67),s(X65,X68)))) | p(s(bool,happ(s(fun(X65,bool),X69),s(X65,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X69,X65)))))) )),
  inference(forward_demodulation,[],[f7206,f339])).
fof(f7206,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,happ(s(fun(X65,bool),X67),s(X65,X68)))) | p(s(bool,happ(s(fun(X65,bool),X69),s(X65,sK6('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X69,X65))))) | p(s(bool,happ(s(fun(X65,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67))))),s(X65,sK5('const.sum.OUTL_1'(s('type.sum.sum'(fun(X65,bool),X66),'const.sum.INL_1'(s(fun(X65,bool),X67)))),X69,X65)))))) )),
  inference(superposition,[],[f316,f339])).
fof(f7298,plain,(
  ( ! [X61,X64,X62,X60,X63] : (p(s(bool,happ(s(fun(X60,bool),X62),s(X60,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X64,X60))))) | ~p(s(bool,happ(s(fun(X60,bool),X62),s(X60,X63)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X64,X60)))))) )),
  inference(forward_demodulation,[],[f7205,f340])).
fof(f7205,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,happ(s(fun(X60,bool),X62),s(X60,X63)))) | p(s(bool,happ(s(fun(X60,bool),X64),s(X60,sK6('const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X64,X60))))) | p(s(bool,happ(s(fun(X60,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62))))),s(X60,sK5('const.sum.OUTR_1'(s('type.sum.sum'(X61,fun(X60,bool)),'const.sum.INR_1'(s(fun(X60,bool),X62)))),X64,X60)))))) )),
  inference(superposition,[],[f316,f340])).
fof(f7297,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,happ(s(fun(X54,bool),X56),s(X54,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X59,X54))))) | ~p(s(bool,happ(s(fun(X54,bool),X56),s(X54,X58)))) | p(s(bool,happ(s(fun(X54,bool),X59),s(X54,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X59,X54)))))) )),
  inference(forward_demodulation,[],[f7204,f343])).
fof(f7204,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),X56),s(X54,X58)))) | p(s(bool,happ(s(fun(X54,bool),X59),s(X54,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X59,X54))))) | p(s(bool,happ(s(fun(X54,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57))))),s(X54,sK5('const.pair.FST_1'(s('type.pair.prod'(fun(X54,bool),X55),'const.pair.,_2'(s(fun(X54,bool),X56),s(X55,X57)))),X59,X54)))))) )),
  inference(superposition,[],[f316,f343])).
fof(f7296,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,sK5('const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X48))))) | ~p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),X53),s(X48,sK6('const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X48)))))) )),
  inference(forward_demodulation,[],[f7203,f320])).
fof(f7203,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (~p(s(bool,happ(s(fun(X48,bool),happ(s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num',X51))),s(X48,X52)))) | p(s(bool,happ(s(fun(X48,bool),X53),s(X48,sK6('const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X48))))) | p(s(bool,happ(s(fun(X48,bool),'const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))),s(X48,sK5('const.ind_type.FCONS_3'(s(fun(X48,bool),X49),s(fun('type.num.num',fun(X48,bool)),X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))),X53,X48)))))) )),
  inference(superposition,[],[f316,f320])).
fof(f7295,plain,(
  ( ! [X47,X45,X43,X46,X44] : (p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK5('const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X47,X43))))) | ~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X46)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK6('const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X47,X43)))))) )),
  inference(forward_demodulation,[],[f7202,f312])).
fof(f7202,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X46)))) | p(s(bool,happ(s(fun(X43,bool),X47),s(X43,sK6('const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X47,X43))))) | p(s(bool,happ(s(fun(X43,bool),'const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0'))),s(X43,sK5('const.ind_type.FCONS_3'(s(fun(X43,bool),X44),s(fun('type.num.num',fun(X43,bool)),X45),s('type.num.num','const.num.|48|_0')),X47,X43)))))) )),
  inference(superposition,[],[f316,f312])).
fof(f7294,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(X37,bool),X40),s(X37,sK5('const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X42,X37))))) | ~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),X42),s(X37,sK6('const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X42,X37)))))) )),
  inference(forward_demodulation,[],[f7201,f344])).
fof(f7201,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~p(s(bool,happ(s(fun(X37,bool),X40),s(X37,X41)))) | p(s(bool,happ(s(fun(X37,bool),X42),s(X37,sK6('const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X42,X37))))) | p(s(bool,happ(s(fun(X37,bool),'const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40))))),s(X37,sK5('const.pair.SND_1'(s('type.pair.prod'(X38,fun(X37,bool)),'const.pair.,_2'(s(X38,X39),s(fun(X37,bool),X40)))),X42,X37)))))) )),
  inference(superposition,[],[f316,f344])).
fof(f7293,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | ~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(forward_demodulation,[],[f7200,f341])).
fof(f7200,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,happ(s(fun(X32,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,X35)))) | p(s(bool,happ(s(fun(X32,bool),X36),s(X32,sK6(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32))))) | p(s(bool,happ(s(fun(X32,bool),happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34))),s(X32,sK5(happ(s(fun('type.pair.prod'(fun(X32,bool),X33),fun(X32,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X32,bool),X33),X34)),X36,X32)))))) )),
  inference(superposition,[],[f316,f341])).
fof(f7292,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (p(s(bool,happ(s(fun(X26,bool),X29),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(forward_demodulation,[],[f7199,f359])).
fof(f7199,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X29),s(X26,X30)))) | p(s(bool,happ(s(fun(X26,bool),X31),s(X26,sK6(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26))))) | p(s(bool,happ(s(fun(X26,bool),happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0'))),s(X26,sK5(happ(s(fun('type.list.list'(X27),fun(X26,bool)),sK16(X28,X29,X27,fun(X26,bool))),s('type.list.list'(X27),'const.list.NIL_0')),X31,X26)))))) )),
  inference(superposition,[],[f316,f359])).
fof(f7198,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X20,fun(X18,bool)),X21),s(X20,happ(s(fun(X19,X20),X22),s(X19,X23))))),s(X18,X24)))) | p(s(bool,happ(s(fun(X18,bool),X25),s(X18,sK6(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18))))) | p(s(bool,happ(s(fun(X18,bool),happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23))),s(X18,sK5(happ(s(fun(X19,fun(X18,bool)),'_dst_x0x1_2'(s(fun(X20,fun(X18,bool)),X21),s(fun(X19,X20),X22))),s(X19,X23)),X25,X18)))))) )),
  inference(superposition,[],[f316,f368])).
fof(f7197,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X13,fun(X10,bool)),'_dst_x0x1_2'(s(fun(X11,fun(X10,bool)),X12),s(fun(X13,X11),X14))),s(X13,X15))),s(X10,X16)))) | p(s(bool,happ(s(fun(X10,bool),X17),s(X10,sK6(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10))))) | p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15))))),s(X10,sK5(happ(s(fun(X11,fun(X10,bool)),X12),s(X11,happ(s(fun(X13,X11),X14),s(X13,X15)))),X17,X10)))))) )),
  inference(superposition,[],[f316,f368])).
fof(f7291,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(forward_demodulation,[],[f7196,f305])).
fof(f7196,plain,(
  ( ! [X6,X8,X7,X9] : (~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))) | p(s(bool,happ(s(fun(X6,bool),X9),s(X6,sK6(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7))),s(X6,sK5(happ(s(fun(fun(X6,bool),fun(X6,bool)),'_i_0_0'),s(fun(X6,bool),X7)),X9,X6)))))) )),
  inference(superposition,[],[f316,f305])).
fof(f7290,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f7195,f367])).
fof(f7195,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X3),s(fun(X0,X1),X2))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK6(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3))),s(X0,sK5(happ(s(fun(fun(X1,bool),fun(X0,bool)),'_dstsel_0_x01_1'(s(fun(X0,X1),X2))),s(fun(X1,bool),X3)),X5,X0)))))) )),
  inference(superposition,[],[f316,f367])).
fof(f7159,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f6868,f819])).
fof(f7158,plain,(
  ( ! [X127,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X128),'const.sum.INL_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X126),'const.sum.INL_1'(s(bool,X127)))))) )),
  inference(superposition,[],[f6868,f339])).
fof(f7157,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X125),'const.sum.INL_1'(s(bool,X124))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X123,bool),'const.sum.INR_1'(s(bool,X124)))))) )),
  inference(superposition,[],[f6868,f340])).
fof(f7156,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X122),'const.sum.INL_1'(s(bool,X120))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X120),s(X119,X121)))))) )),
  inference(superposition,[],[f6868,f343])).
fof(f7155,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X116),s('type.num.num',X117))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X115),s(fun('type.num.num',bool),X116),s('type.num.num','const.num.SUC_1'(s('type.num.num',X117)))))) )),
  inference(superposition,[],[f6868,f320])).
fof(f7154,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X112),s(fun('type.num.num',bool),X113),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6868,f312])).
fof(f7153,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X111),'const.sum.INL_1'(s(bool,X110))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,X110)))))) )),
  inference(superposition,[],[f6868,f344])).
fof(f7152,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f6868,f341])).
fof(f7151,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X104),'const.sum.INL_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK16(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6868,f359])).
fof(f7150,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f6868,f368])).
fof(f7149,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f6868,f368])).
fof(f6978,plain,(
  ( ! [X2,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),happ(s(fun(X1,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),X2),s(X1,X3))))) = s('type.bvlSem.v',happ(s(fun(X1,'type.bvlSem.v'),'_dst_x0x1_2'(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s(fun(X1,'type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one')))))))),X2))),s(X1,X3)))) )),
  inference(superposition,[],[f265,f368])).
fof(f6931,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f6656,f819])).
fof(f6930,plain,(
  ( ! [X127,X128,X126] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X127))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X126),'const.sum.INL_1'(s(bool,X127)))))) )),
  inference(superposition,[],[f6656,f339])).
fof(f6929,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X125,bool),'const.sum.INR_1'(s(bool,X124))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X123,bool),'const.sum.INR_1'(s(bool,X124)))))) )),
  inference(superposition,[],[f6656,f340])).
fof(f6928,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X122,bool),'const.sum.INR_1'(s(bool,X120))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X120),s(X119,X121)))))) )),
  inference(superposition,[],[f6656,f343])).
fof(f6927,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X118,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X116),s('type.num.num',X117))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X115),s(fun('type.num.num',bool),X116),s('type.num.num','const.num.SUC_1'(s('type.num.num',X117)))))) )),
  inference(superposition,[],[f6656,f320])).
fof(f6926,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X114,bool),'const.sum.INR_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X112),s(fun('type.num.num',bool),X113),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f6656,f312])).
fof(f6925,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,X110))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,X110)))))) )),
  inference(superposition,[],[f6656,f344])).
fof(f6924,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X107,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f6656,f341])).
fof(f6923,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK16(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f6656,f359])).
fof(f6922,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f6656,f368])).
fof(f6921,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X94,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f6656,f368])).
fof(f6876,plain,(
  ( ! [X61,X62,X63] : (s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f6818,f258])).
fof(f6818,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'F_0')) | s(bool,X63) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X61),'const.sum.INL_1'(s(bool,X62))))) | s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f2001,f1997])).
fof(f6783,plain,(
  ( ! [X130,X131,X129] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X131),'const.sum.INL_1'(s(bool,X129))))) | s(bool,'T_0') = s(bool,sK4(X129,X130,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X130)))) )),
  inference(superposition,[],[f1997,f819])).
fof(f6780,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X122),'const.sum.INL_1'(s(bool,X120))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X120),s(X119,X121)))))) )),
  inference(superposition,[],[f1997,f343])).
fof(f6779,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X118),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X116),s('type.num.num',X117))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X115),s(fun('type.num.num',bool),X116),s('type.num.num','const.num.SUC_1'(s('type.num.num',X117)))))) )),
  inference(superposition,[],[f1997,f320])).
fof(f6778,plain,(
  ( ! [X114,X112,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X114),'const.sum.INL_1'(s(bool,X112))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X112),s(fun('type.num.num',bool),X113),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1997,f312])).
fof(f6777,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X111),'const.sum.INL_1'(s(bool,X110))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,X110)))))) )),
  inference(superposition,[],[f1997,f344])).
fof(f6776,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X107),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X105),X106))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X105),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X105),X106)))) )),
  inference(superposition,[],[f1997,f341])).
fof(f6775,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X104),'const.sum.INL_1'(s(bool,X103))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X101),bool),sK16(X102,X103,X101,bool)),s('type.list.list'(X101),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1997,f359])).
fof(f6774,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X100),'const.sum.INL_1'(s(bool,happ(s(fun(X96,bool),X97),s(X96,happ(s(fun(X95,X96),X98),s(X95,X99))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'_dst_x0x1_2'(s(fun(X96,bool),X97),s(fun(X95,X96),X98))),s(X95,X99)))) )),
  inference(superposition,[],[f1997,f368])).
fof(f6773,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X94),'const.sum.INL_1'(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X89,bool),X90),s(fun(X91,X89),X92))),s(X91,X93))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X89,bool),X90),s(X89,happ(s(fun(X91,X89),X92),s(X91,X93)))))) )),
  inference(superposition,[],[f1997,f368])).
fof(f6664,plain,(
  ( ! [X59,X60,X58] : (s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f6606,f258])).
fof(f6606,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,'F_0')) | s(bool,X60) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X58,bool),'const.sum.INR_1'(s(bool,X59))))) | s(bool,'F_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f2001,f1996])).
fof(f6572,plain,(
  ( ! [X127,X128,X126] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X128,bool),'const.sum.INR_1'(s(bool,X126))))) | s(bool,'T_0') = s(bool,sK4(X126,X127,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X127)))) )),
  inference(superposition,[],[f1996,f819])).
fof(f6569,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X119,bool),'const.sum.INR_1'(s(bool,X117))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X116),'const.pair.,_2'(s(bool,X117),s(X116,X118)))))) )),
  inference(superposition,[],[f1996,f343])).
fof(f6568,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X115,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X113),s('type.num.num',X114))))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X112),s(fun('type.num.num',bool),X113),s('type.num.num','const.num.SUC_1'(s('type.num.num',X114)))))) )),
  inference(superposition,[],[f1996,f320])).
fof(f6567,plain,(
  ( ! [X111,X109,X110] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,X109))))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X109),s(fun('type.num.num',bool),X110),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1996,f312])).
fof(f6566,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X108,bool),'const.sum.INR_1'(s(bool,X107))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X105,bool),'const.pair.,_2'(s(X105,X106),s(bool,X107)))))) )),
  inference(superposition,[],[f1996,f344])).
fof(f6565,plain,(
  ( ! [X103,X102,X104] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X104,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),X103))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X102),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X102),X103)))) )),
  inference(superposition,[],[f1996,f341])).
fof(f6564,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X101,bool),'const.sum.INR_1'(s(bool,X100))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X98),bool),sK16(X99,X100,X98,bool)),s('type.list.list'(X98),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1996,f359])).
fof(f6563,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X97,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X93,bool),X94),s(X93,happ(s(fun(X92,X93),X95),s(X92,X96))))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X93,bool),X94),s(fun(X92,X93),X95))),s(X92,X96)))) )),
  inference(superposition,[],[f1996,f368])).
fof(f6562,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X91,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X88,bool),'_dst_x0x1_2'(s(fun(X86,bool),X87),s(fun(X88,X86),X89))),s(X88,X90))))))) | s(bool,'T_0') = s(bool,happ(s(fun(X86,bool),X87),s(X86,happ(s(fun(X88,X86),X89),s(X88,X90)))))) )),
  inference(superposition,[],[f1996,f368])).
fof(f6056,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,X160)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X162),bool),sK16(X163,sK4(X160,X161,bool),X162,bool)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161)))) )),
  inference(superposition,[],[f4654,f819])).
fof(f6055,plain,(
  ( ! [X158,X156,X159,X157] : (~p(s(bool,X157)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X158),bool),sK16(X159,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X156),'const.sum.INL_1'(s(bool,X157)))),X158,bool))))) )),
  inference(superposition,[],[f4654,f339])).
fof(f6054,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,X153)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X154),bool),sK16(X155,'const.sum.OUTR_1'(s('type.sum.sum'(X152,bool),'const.sum.INR_1'(s(bool,X153)))),X154,bool))))) )),
  inference(superposition,[],[f4654,f340])).
fof(f6053,plain,(
  ( ! [X151,X149,X147,X150,X148] : (~p(s(bool,X148)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X150),bool),sK16(X151,'const.pair.FST_1'(s('type.pair.prod'(bool,X147),'const.pair.,_2'(s(bool,X148),s(X147,X149)))),X150,bool))))) )),
  inference(superposition,[],[f4654,f343])).
fof(f6052,plain,(
  ( ! [X146,X144,X142,X145,X143] : (~p(s(bool,happ(s(fun('type.num.num',bool),X143),s('type.num.num',X144)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X145),bool),sK16(X146,'const.ind_type.FCONS_3'(s(bool,X142),s(fun('type.num.num',bool),X143),s('type.num.num','const.num.SUC_1'(s('type.num.num',X144)))),X145,bool))))) )),
  inference(superposition,[],[f4654,f320])).
fof(f6051,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,X138)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X140),bool),sK16(X141,'const.ind_type.FCONS_3'(s(bool,X138),s(fun('type.num.num',bool),X139),s('type.num.num','const.num.|48|_0')),X140,bool))))) )),
  inference(superposition,[],[f4654,f312])).
fof(f6050,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,X135)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X136),bool),sK16(X137,'const.pair.SND_1'(s('type.pair.prod'(X133,bool),'const.pair.,_2'(s(X133,X134),s(bool,X135)))),X136,bool))))) )),
  inference(superposition,[],[f4654,f344])).
fof(f6049,plain,(
  ( ! [X132,X130,X131,X129] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X131),bool),sK16(X132,happ(s(fun('type.pair.prod'(bool,X129),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X129),X130)),X131,bool))))) )),
  inference(superposition,[],[f4654,f341])).
fof(f6048,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X126)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X127),bool),sK16(X128,happ(s(fun('type.list.list'(X124),bool),sK16(X125,X126,X124,bool)),s('type.list.list'(X124),'const.list.NIL_0')),X127,bool))))) )),
  inference(superposition,[],[f4654,f359])).
fof(f6047,plain,(
  ( ! [X123,X121,X118,X122,X120,X119,X117] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X122),bool),sK16(X123,happ(s(fun(X117,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X117,X118),X120))),s(X117,X121)),X122,bool))))) )),
  inference(superposition,[],[f4654,f368])).
fof(f6046,plain,(
  ( ! [X111,X116,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(X112,bool),'_dst_x0x1_2'(s(fun(X110,bool),X111),s(fun(X112,X110),X113))),s(X112,X114)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X115),bool),sK16(X116,happ(s(fun(X110,bool),X111),s(X110,happ(s(fun(X112,X110),X113),s(X112,X114)))),X115,bool))))) )),
  inference(superposition,[],[f4654,f368])).
fof(f6045,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X108),bool),sK16(X109,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X107)),X108,bool))))) )),
  inference(superposition,[],[f4654,f305])).
fof(f6080,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f6042,f256])).
fof(f6042,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X98),bool),sK16(X99,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X96),'const.sum.INL_1'(s(bool,X97)))),X98,bool)))) | s(bool,'F_0') = s(bool,X97)) )),
  inference(superposition,[],[f4654,f501])).
fof(f6076,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'const.ind_type.ISO_2'(s(fun(X88,X88),'_i_0_0'),s(fun(X88,X88),'_i_0_0')),X89,bool))))) )),
  inference(subsumption_resolution,[],[f6039,f256])).
fof(f6039,plain,(
  ( ! [X90,X88,X89] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X89),bool),sK16(X90,'const.ind_type.ISO_2'(s(fun(X88,X88),'_i_0_0'),s(fun(X88,X88),'_i_0_0')),X89,bool))))) )),
  inference(superposition,[],[f4654,f428])).
fof(f6075,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X86),bool),sK16(X87,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X85)))),X86,bool)))) | s(bool,'F_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f6038,f256])).
fof(f6038,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X86),bool),sK16(X87,'const.sum.OUTR_1'(s('type.sum.sum'(X84,bool),'const.sum.INR_1'(s(bool,X85)))),X86,bool)))) | s(bool,'F_0') = s(bool,X85)) )),
  inference(superposition,[],[f4654,f528])).
fof(f6072,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X79),bool),sK16(X80,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'T_0',X77,bool))),X79,bool))))) )),
  inference(subsumption_resolution,[],[f6036,f256])).
fof(f6036,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X79),bool),sK16(X80,'|dollar|exists_1'(s(fun('type.list.list'(X77),bool),sK16(X78,'T_0',X77,bool))),X79,bool))))) )),
  inference(superposition,[],[f4654,f4742])).
fof(f6068,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'|dollar|exists_1'(s(fun(X67,bool),X68)),X69,bool)))) | s(X67,X71) = s(X67,sK4(X71,X68,X67))) )),
  inference(subsumption_resolution,[],[f6033,f256])).
fof(f6033,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X69),bool),sK16(X70,'|dollar|exists_1'(s(fun(X67,bool),X68)),X69,bool)))) | s(X67,X71) = s(X67,sK4(X71,X68,X67))) )),
  inference(superposition,[],[f4654,f819])).
fof(f6067,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f6032,f256])).
fof(f6032,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X65),bool),sK16(X66,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64)),X65,bool)))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f4654,f460])).
fof(f6057,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6016,f256])).
fof(f6016,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X3,X4,bool)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4654,f433])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : ! [X26,X22] : (! [X24] : (s(X9,X22) = s(X9,X24) => p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24))))) => p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X26)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.JRH_INDUCT_UTIL')).
fof(f4480,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f4431,f256])).
fof(f4431,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK4(X34,X32,X31))) )),
  inference(superposition,[],[f332,f819])).
fof(f4481,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f4432,f256])).
fof(f4432,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) != s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | p(s(bool,X37)) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f406,f819])).
fof(f5032,plain,(
  ( ! [X350,X344,X346,X349,X345,X347,X348] : (p(s(bool,'|dollar|exists_1'(s(fun(X344,bool),happ(s(fun(X347,fun(X344,bool)),'_dst_x0x1_2'(s(fun(X345,fun(X344,bool)),X346),s(fun(X347,X345),X348))),s(X347,X349)))))) | s(X344,X350) = s(X344,sK4(X350,happ(s(fun(X345,fun(X344,bool)),X346),s(X345,happ(s(fun(X347,X345),X348),s(X347,X349)))),X344))) )),
  inference(superposition,[],[f314,f368])).
fof(f5109,plain,(
  ( ! [X333,X335,X331,X337,X332,X334,X336] : (p(s(bool,'|dollar|exists_1'(s(fun(X331,bool),happ(s(fun(X333,fun(X331,bool)),X334),s(X333,happ(s(fun(X332,X333),X335),s(X332,X336)))))))) | s(X331,X337) = s(X331,sK4(X337,happ(s(fun(X332,fun(X331,bool)),'_dst_x0x1_2'(s(fun(X333,fun(X331,bool)),X334),s(fun(X332,X333),X335))),s(X332,X336)),X331))) )),
  inference(superposition,[],[f314,f368])).
fof(f812,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f314,f305])).
fof(f4690,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'|dollar|exists_1'(s(fun(X31,bool),'_dst_x0x1_2'(s(fun(X32,bool),X34),s(fun(X31,X32),X33)))))) | s(X31,X35) = s(X31,sK4(X35,happ(s(fun(fun(X32,bool),fun(X31,bool)),'_dstsel_0_x01_1'(s(fun(X31,X32),X33))),s(fun(X32,bool),X34)),X31))) )),
  inference(superposition,[],[f314,f367])).
fof(f918,plain,(
  ( ! [X90,X88,X87,X91,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X90)))) | s(X87,X91) = s(X87,sK4(X91,happ(s(fun('type.list.list'(X88),fun(X87,bool)),sK16(X89,X90,X88,fun(X87,bool))),s('type.list.list'(X88),'const.list.NIL_0')),X87))) )),
  inference(superposition,[],[f314,f359])).
fof(f1503,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X85,bool),X86),X87)))))) | s(X85,X88) = s(X85,sK4(X88,happ(s(fun('type.pair.prod'(fun(X85,bool),X86),fun(X85,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X85,bool),X86),X87)),X85))) )),
  inference(superposition,[],[f314,f341])).
fof(f813,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | s(X3,X7) = s(X3,sK4(X7,'const.pair.SND_1'(s('type.pair.prod'(X4,fun(X3,bool)),'const.pair.,_2'(s(X4,X5),s(fun(X3,bool),X6)))),X3))) )),
  inference(superposition,[],[f314,f344])).
fof(f814,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | s(X8,X11) = s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) )),
  inference(superposition,[],[f314,f312])).
fof(f2264,plain,(
  ( ! [X132,X130,X128,X131,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X128,bool),happ(s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num',X131)))))) | s(X128,X132) = s(X128,sK4(X132,'const.ind_type.FCONS_3'(s(fun(X128,bool),X129),s(fun('type.num.num',fun(X128,bool)),X130),s('type.num.num','const.num.SUC_1'(s('type.num.num',X131)))),X128))) )),
  inference(superposition,[],[f314,f320])).
fof(f815,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | s(X12,X16) = s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) )),
  inference(superposition,[],[f314,f343])).
fof(f816,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | s(X17,X20) = s(X17,sK4(X20,'const.sum.OUTR_1'(s('type.sum.sum'(X18,fun(X17,bool)),'const.sum.INR_1'(s(fun(X17,bool),X19)))),X17))) )),
  inference(superposition,[],[f314,f340])).
fof(f817,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X23)))) | s(X21,X24) = s(X21,sK4(X24,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X21,bool),X22),'const.sum.INL_1'(s(fun(X21,bool),X23)))),X21))) )),
  inference(superposition,[],[f314,f339])).
fof(f4402,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,'|dollar|exists_1'(s(fun(X159,bool),X160)))) | s(X159,X162) = s(X159,sK4(X162,sK4(X160,X161,fun(X159,bool)),X159)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X159,bool),bool),X161)))) )),
  inference(superposition,[],[f314,f819])).
fof(f4366,plain,(
  ( ! [X56,X55] : (p(s(bool,X55)) | s(bool,'T_0') != s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f267,f819])).
fof(f4367,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,sK4(X57,X58,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f268,f819])).
fof(f4368,plain,(
  ( ! [X59,X60] : (s(bool,'F_0') != s(bool,X59) | ~p(s(bool,sK4(X59,X60,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f271,f819])).
fof(f4369,plain,(
  ( ! [X61,X62] : (p(s(bool,X61)) | s(bool,'F_0') = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f272,f819])).
fof(f4370,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,X63)) | p(s(bool,X65)) | s(bool,X65) = s(bool,sK4(X63,X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f329,f819])).
fof(f4372,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | s(bool,X71) = s(bool,sK4(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f332,f819])).
fof(f4373,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,X74) != s(bool,sK4(X72,X73,bool)) | p(s(bool,X74)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X73)))) )),
  inference(superposition,[],[f406,f819])).
fof(f4378,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X87)) | s(bool,X89) != s(bool,sK4(X87,X88,bool)) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f435,f819])).
fof(f4380,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | s(bool,X94) = s(bool,sK4(X92,X93,bool)) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f482,f819])).
fof(f4385,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f1999,f819])).
fof(f4386,plain,(
  ( ! [X111,X109,X110] : (p(s(bool,X109)) | s(bool,X111) != s(bool,sK4(X109,X110,bool)) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f2001,f819])).
fof(f4552,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,happ(s(fun(X39,bool),X40),s(X39,X42)))) | p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),sK4(X40,X41,fun(X39,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f4479,f819])).
fof(f4400,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X154)))) | p(s(bool,happ(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool))),s(X151,'const.min.@_1'(s(fun(X151,bool),sK4(X152,X153,fun(X151,bool)))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f313,f819])).
fof(f4403,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun('type.list.list'(X163),bool),X164),s('type.list.list'(X163),sK2(sK4(X164,X165,fun('type.list.list'(X163),bool)),X163))))) | ~p(s(bool,happ(s(fun('type.list.list'(X163),bool),sK4(X164,X165,fun('type.list.list'(X163),bool))),s('type.list.list'(X163),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.list.list'(X163),bool),bool),X165)))) )),
  inference(superposition,[],[f408,f819])).
fof(f4503,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X58,X57),X61),s(fun(X57,X58),X59)))) | s(X58,X62) = s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,happ(s(fun(X58,X57),X61),s(X58,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f357,f819])).
fof(f4493,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X57,X58),X59),s(fun(X58,X57),X61)))) | s(X57,X62) = s(X57,happ(s(fun(X58,X57),X61),s(X58,happ(s(fun(X57,X58),sK4(X59,X60,fun(X57,X58))),s(X57,X62))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X57,X58),bool),X60)))) )),
  inference(superposition,[],[f357,f819])).
fof(f4406,plain,(
  ( ! [X182,X180,X178,X183,X181,X179] : (s(fun(X178,X179),X180) != s(fun(X178,X179),X182) | s(X179,happ(s(fun(X178,X179),X182),s(X178,X183))) = s(X179,happ(s(fun(X178,X179),sK4(X180,X181,fun(X178,X179))),s(X178,X183))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X178,X179),bool),X181)))) )),
  inference(superposition,[],[f354,f819])).
fof(f4405,plain,(
  ( ! [X177,X175,X173,X176,X174,X172] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X172,X173),X174),s(fun(X173,X172),X176)))) | s(X173,X177) = s(X173,happ(s(fun(X172,X173),sK4(X174,X175,fun(X172,X173))),s(X172,happ(s(fun(X173,X172),X176),s(X173,X177))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X172,X173),bool),X175)))) )),
  inference(superposition,[],[f356,f819])).
fof(f4404,plain,(
  ( ! [X167,X171,X169,X166,X170,X168] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X167,X166),X170),s(fun(X166,X167),X168)))) | s(X166,X171) = s(X166,happ(s(fun(X167,X166),X170),s(X167,happ(s(fun(X166,X167),sK4(X168,X169,fun(X166,X167))),s(X166,X171))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X166,X167),bool),X169)))) )),
  inference(superposition,[],[f356,f819])).
fof(f4414,plain,(
  ( ! [X208,X206,X207] : (s('type.integer.int',X206) != s('type.integer.int',X208) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X208))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',sK4(X206,X207,'type.integer.int')))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X207)))) )),
  inference(superposition,[],[f334,f819])).
fof(f5069,plain,(
  ( ! [X99,X97,X98,X96] : (p(s(bool,happ(s(fun(X97,bool),X98),s(X97,happ(s(fun('type.list.list'(X96),X97),X99),s('type.list.list'(X96),sK2('_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99)),X96))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X96),bool),'_dst_x0x1_2'(s(fun(X97,bool),X98),s(fun('type.list.list'(X96),X97),X99))),s('type.list.list'(X96),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f368])).
fof(f4712,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36)))))) )),
  inference(forward_demodulation,[],[f4691,f367])).
fof(f4691,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun('type.list.list'(X36),bool),'_dst_x0x1_2'(s(fun(X37,bool),X39),s(fun('type.list.list'(X36),X37),X38))),s('type.list.list'(X36),sK2(happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39)),X36))))) | ~p(s(bool,happ(s(fun('type.list.list'(X36),bool),happ(s(fun(fun(X37,bool),fun('type.list.list'(X36),bool)),'_dstsel_0_x01_1'(s(fun('type.list.list'(X36),X37),X38))),s(fun(X37,bool),X39))),s('type.list.list'(X36),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f367])).
fof(f4990,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),'_dst_x0x1_2'(s(fun(X90,bool),X91),s(fun(X92,X90),X93))),s(X92,X94)))) | p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f4479,f368])).
fof(f4991,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X97,bool),'_dst_x0x1_2'(s(fun(X95,bool),X96),s(fun(X97,X95),X98))),s(X97,X99)))) | p(s(bool,happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96))))))) )),
  inference(superposition,[],[f313,f368])).
fof(f4993,plain,(
  ( ! [X111,X109,X112,X110,X108] : (p(s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X108,bool),X109),s(fun(X110,X108),X111))),s(X110,X112)))) | s(bool,'T_0') != s(bool,happ(s(fun(X108,bool),X109),s(X108,happ(s(fun(X110,X108),X111),s(X110,X112)))))) )),
  inference(superposition,[],[f267,f368])).
fof(f4994,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X113,bool),X114),s(fun(X115,X113),X116))),s(X115,X117)))) | s(bool,'T_0') = s(bool,happ(s(fun(X113,bool),X114),s(X113,happ(s(fun(X115,X113),X116),s(X115,X117)))))) )),
  inference(superposition,[],[f268,f368])).
fof(f4995,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'F_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) )),
  inference(superposition,[],[f271,f368])).
fof(f4996,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X123,bool),X124),s(fun(X125,X123),X126))),s(X125,X127)))) | s(bool,'F_0') = s(bool,happ(s(fun(X123,bool),X124),s(X123,happ(s(fun(X125,X123),X126),s(X125,X127)))))) )),
  inference(superposition,[],[f272,f368])).
fof(f4997,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,happ(s(fun(X130,bool),'_dst_x0x1_2'(s(fun(X128,bool),X129),s(fun(X130,X128),X131))),s(X130,X132)))) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) )),
  inference(superposition,[],[f329,f368])).
fof(f4999,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (~p(s(bool,happ(s(fun(X142,bool),'_dst_x0x1_2'(s(fun(X140,bool),X141),s(fun(X142,X140),X143))),s(X142,X144)))) | ~p(s(bool,X145)) | s(bool,X145) = s(bool,happ(s(fun(X140,bool),X141),s(X140,happ(s(fun(X142,X140),X143),s(X142,X144)))))) )),
  inference(superposition,[],[f332,f368])).
fof(f5000,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X148,bool),'_dst_x0x1_2'(s(fun(X146,bool),X147),s(fun(X148,X146),X149))),s(X148,X150)))) | s(bool,X151) != s(bool,happ(s(fun(X146,bool),X147),s(X146,happ(s(fun(X148,X146),X149),s(X148,X150))))) | p(s(bool,X151))) )),
  inference(superposition,[],[f406,f368])).
fof(f5005,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (~p(s(bool,happ(s(fun(X178,bool),'_dst_x0x1_2'(s(fun(X176,bool),X177),s(fun(X178,X176),X179))),s(X178,X180)))) | s(bool,X181) != s(bool,happ(s(fun(X176,bool),X177),s(X176,happ(s(fun(X178,X176),X179),s(X178,X180))))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f435,f368])).
fof(f5006,plain,(
  ( ! [X185,X182,X186,X184,X183] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X182,bool),X183),s(fun(X184,X182),X185))),s(X184,X186))))) | s(bool,'F_0') = s(bool,happ(s(fun(X182,bool),X183),s(X182,happ(s(fun(X184,X182),X185),s(X184,X186)))))) )),
  inference(superposition,[],[f460,f368])).
fof(f5007,plain,(
  ( ! [X191,X189,X187,X192,X190,X188] : (p(s(bool,happ(s(fun(X189,bool),'_dst_x0x1_2'(s(fun(X187,bool),X188),s(fun(X189,X187),X190))),s(X189,X191)))) | s(bool,X192) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,X191))))) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f482,f368])).
fof(f5008,plain,(
  ( ! [X198,X196,X194,X197,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X198),'const.sum.INL_1'(s(bool,happ(s(fun(X195,bool),'_dst_x0x1_2'(s(fun(X193,bool),X194),s(fun(X195,X193),X196))),s(X195,X197))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X193,bool),X194),s(X193,happ(s(fun(X195,X193),X196),s(X195,X197)))))) )),
  inference(superposition,[],[f501,f368])).
fof(f5009,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X201,bool),'_dst_x0x1_2'(s(fun(X199,bool),X200),s(fun(X201,X199),X202))),s(X201,X203))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X199,bool),X200),s(X199,happ(s(fun(X201,X199),X202),s(X201,X203)))))) )),
  inference(superposition,[],[f528,f368])).
fof(f5012,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (~p(s(bool,happ(s(fun(X218,bool),'_dst_x0x1_2'(s(fun(X216,bool),X217),s(fun(X218,X216),X219))),s(X218,X220)))) | s(bool,X221) = s(bool,happ(s(fun(X216,bool),X217),s(X216,happ(s(fun(X218,X216),X219),s(X218,X220))))) | s(bool,'F_0') = s(bool,X221)) )),
  inference(superposition,[],[f1999,f368])).
fof(f5013,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (p(s(bool,happ(s(fun(X224,bool),'_dst_x0x1_2'(s(fun(X222,bool),X223),s(fun(X224,X222),X225))),s(X224,X226)))) | s(bool,X227) != s(bool,happ(s(fun(X222,bool),X223),s(X222,happ(s(fun(X224,X222),X225),s(X224,X226))))) | s(bool,'F_0') = s(bool,X227)) )),
  inference(superposition,[],[f2001,f368])).
fof(f5071,plain,(
  ( ! [X109,X107,X105,X108,X106] : (p(s(bool,happ(s(fun(X106,bool),X107),s(X106,happ(s(fun(X105,X106),X108),s(X105,X109)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X105,bool),'_dst_x0x1_2'(s(fun(X106,bool),X107),s(fun(X105,X106),X108))),s(X105,X109)))) )),
  inference(superposition,[],[f267,f368])).
fof(f5072,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,happ(s(fun(X111,bool),X112),s(X111,happ(s(fun(X110,X111),X113),s(X110,X114)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X110,bool),'_dst_x0x1_2'(s(fun(X111,bool),X112),s(fun(X110,X111),X113))),s(X110,X114)))) )),
  inference(superposition,[],[f268,f368])).
fof(f5073,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'F_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) )),
  inference(superposition,[],[f271,f368])).
fof(f5074,plain,(
  ( ! [X123,X121,X124,X122,X120] : (p(s(bool,happ(s(fun(X121,bool),X122),s(X121,happ(s(fun(X120,X121),X123),s(X120,X124)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X121,bool),X122),s(fun(X120,X121),X123))),s(X120,X124)))) )),
  inference(superposition,[],[f272,f368])).
fof(f5075,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,happ(s(fun(X126,bool),X127),s(X126,happ(s(fun(X125,X126),X128),s(X125,X129)))))) | p(s(bool,X130)) | s(bool,X130) = s(bool,happ(s(fun(X125,bool),'_dst_x0x1_2'(s(fun(X126,bool),X127),s(fun(X125,X126),X128))),s(X125,X129)))) )),
  inference(superposition,[],[f329,f368])).
fof(f5077,plain,(
  ( ! [X142,X140,X138,X141,X139,X137] : (~p(s(bool,happ(s(fun(X138,bool),X139),s(X138,happ(s(fun(X137,X138),X140),s(X137,X141)))))) | ~p(s(bool,X142)) | s(bool,X142) = s(bool,happ(s(fun(X137,bool),'_dst_x0x1_2'(s(fun(X138,bool),X139),s(fun(X137,X138),X140))),s(X137,X141)))) )),
  inference(superposition,[],[f332,f368])).
fof(f5078,plain,(
  ( ! [X146,X144,X147,X145,X143,X148] : (~p(s(bool,happ(s(fun(X144,bool),X145),s(X144,happ(s(fun(X143,X144),X146),s(X143,X147)))))) | s(bool,X148) != s(bool,happ(s(fun(X143,bool),'_dst_x0x1_2'(s(fun(X144,bool),X145),s(fun(X143,X144),X146))),s(X143,X147))) | p(s(bool,X148))) )),
  inference(superposition,[],[f406,f368])).
fof(f5083,plain,(
  ( ! [X177,X175,X173,X178,X176,X174] : (~p(s(bool,happ(s(fun(X174,bool),X175),s(X174,happ(s(fun(X173,X174),X176),s(X173,X177)))))) | s(bool,X178) != s(bool,happ(s(fun(X173,bool),'_dst_x0x1_2'(s(fun(X174,bool),X175),s(fun(X173,X174),X176))),s(X173,X177))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f435,f368])).
fof(f5085,plain,(
  ( ! [X189,X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X185,bool),X186),s(X185,happ(s(fun(X184,X185),X187),s(X184,X188)))))) | s(bool,X189) = s(bool,happ(s(fun(X184,bool),'_dst_x0x1_2'(s(fun(X185,bool),X186),s(fun(X184,X185),X187))),s(X184,X188))) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f482,f368])).
fof(f5090,plain,(
  ( ! [X218,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(X214,bool),X215),s(X214,happ(s(fun(X213,X214),X216),s(X213,X217)))))) | s(bool,X218) = s(bool,happ(s(fun(X213,bool),'_dst_x0x1_2'(s(fun(X214,bool),X215),s(fun(X213,X214),X216))),s(X213,X217))) | s(bool,'F_0') = s(bool,X218)) )),
  inference(superposition,[],[f1999,f368])).
fof(f5091,plain,(
  ( ! [X222,X220,X219,X223,X221,X224] : (p(s(bool,happ(s(fun(X220,bool),X221),s(X220,happ(s(fun(X219,X220),X222),s(X219,X223)))))) | s(bool,X224) != s(bool,happ(s(fun(X219,bool),'_dst_x0x1_2'(s(fun(X220,bool),X221),s(fun(X219,X220),X222))),s(X219,X223))) | s(bool,'F_0') = s(bool,X224)) )),
  inference(superposition,[],[f2001,f368])).
fof(f5068,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,happ(s(fun(X92,bool),X93),s(X92,happ(s(fun(X91,X92),X94),s(X91,X95)))))) | p(s(bool,happ(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))),s(X91,'const.min.@_1'(s(fun(X91,bool),'_dst_x0x1_2'(s(fun(X92,bool),X93),s(fun(X91,X92),X94))))))))) )),
  inference(superposition,[],[f313,f368])).
fof(f347,plain,(
  ( ! [X2,X0,X14,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,sK8(X11,X3,X2,X1,X0))))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0))))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X11] : ((? [X15] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X15)))) | ! [X14] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X14))))))) & (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,X13)))))) | ! [X12] : p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12)))))) & ! [X6] : ((! [X10] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | ? [X9] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X9))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : ((? [X10] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) | ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & (? [X11] : ~p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11)))))) | ! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))))) & ! [X6] : ((! [X7] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) | ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & (! [X8] : ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))))) & ! [X4,X5] : ((s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) & (s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(ennf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) => (! [X9] : (! [X10] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,X10)))) <=> ! [X11] : p(s(bool,happ(s(fun(X0,bool),X9),s(X0,happ(s(fun(X1,X0),X2),s(X1,X11))))))) & ! [X6] : (? [X7] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X7)))) <=> ? [X8] : p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8))))))) & ! [X4,X5] : (s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) <=> s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9,X23] : ! [X37,X10] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X23),X10),s(fun(X23,X9),X37)))) => (! [X26] : (! [X24] : p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24)))) <=> ! [X24] : p(s(bool,happ(s(fun(X9,bool),X26),s(X9,happ(s(fun(X23,X9),X37),s(X23,X24))))))) & ! [X26] : (? [X24] : p(s(bool,happ(s(fun(X9,bool),X26),s(X9,X24)))) <=> ? [X24] : p(s(bool,happ(s(fun(X9,bool),X26),s(X9,happ(s(fun(X23,X9),X37),s(X23,X24))))))) & ! [X6,X46] : (s(X9,X6) = s(X9,happ(s(fun(X23,X9),X37),s(X23,X46))) <=> s(X23,X46) = s(X23,happ(s(fun(X9,X23),X10),s(X9,X6))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.ISO_USAGE')).
fof(f350,plain,(
  ( ! [X6,X2,X0,X8,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,X8)))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK11(X6,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f5067,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,happ(s(fun(X87,bool),X88),s(X87,happ(s(fun(X86,X87),X89),s(X86,X90)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),'_dst_x0x1_2'(s(fun(X87,bool),X88),s(fun(X86,X87),X89))))))) )),
  inference(superposition,[],[f4479,f368])).
fof(f5084,plain,(
  ( ! [X182,X180,X183,X181,X179] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(X180,bool),X181),s(X180,happ(s(fun(X179,X180),X182),s(X179,X183))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X179,bool),'_dst_x0x1_2'(s(fun(X180,bool),X181),s(fun(X179,X180),X182))),s(X179,X183)))) )),
  inference(superposition,[],[f460,f368])).
fof(f5086,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X195),'const.sum.INL_1'(s(bool,happ(s(fun(X191,bool),X192),s(X191,happ(s(fun(X190,X191),X193),s(X190,X194))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X190,bool),'_dst_x0x1_2'(s(fun(X191,bool),X192),s(fun(X190,X191),X193))),s(X190,X194)))) )),
  inference(superposition,[],[f501,f368])).
fof(f5087,plain,(
  ( ! [X198,X196,X200,X199,X197,X201] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X201,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X197,bool),X198),s(X197,happ(s(fun(X196,X197),X199),s(X196,X200))))))))) | s(bool,'F_0') = s(bool,happ(s(fun(X196,bool),'_dst_x0x1_2'(s(fun(X197,bool),X198),s(fun(X196,X197),X199))),s(X196,X200)))) )),
  inference(superposition,[],[f528,f368])).
fof(f348,plain,(
  ( ! [X2,X0,X12,X3,X1,X11] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X11),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X11,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X11),s(X0,X12))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f349,plain,(
  ( ! [X6,X2,X0,X10,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X10)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f240])).
fof(f2579,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) != s(bool,f(s(X12,sK12(X13,X14,bool,X12)))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,g(s(X12,sK12(X13,X14,bool,X12)))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f353,f433])).
fof(f2577,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,f(s(X4,sK12(X5,X6,bool,X4)))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,g(s(X4,sK12(X5,X6,bool,X4)))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f353,f433])).
fof(f2229,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X9),s('type.num.num',X10))) | s(bool,X11) = s(bool,'const.ind_type.FCONS_3'(s(bool,X8),s(fun('type.num.num',bool),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f320,f433])).
fof(f2231,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',X18))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f320,f433])).
fof(f2257,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X100),s('type.num.num',X101))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X99),s(fun('type.num.num',bool),X100),s('type.num.num','const.num.SUC_1'(s('type.num.num',X101)))))) )),
  inference(superposition,[],[f460,f320])).
fof(f4053,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X82),'const.sum.INL_1'(s(bool,happ(s(fun('type.num.num',bool),X80),s('type.num.num',X81))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X79),s(fun('type.num.num',bool),X80),s('type.num.num','const.num.SUC_1'(s('type.num.num',X81)))))) )),
  inference(superposition,[],[f501,f320])).
fof(f4228,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X85,bool),'const.sum.INR_1'(s(bool,happ(s(fun('type.num.num',bool),X83),s('type.num.num',X84))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X82),s(fun('type.num.num',bool),X83),s('type.num.num','const.num.SUC_1'(s('type.num.num',X84)))))) )),
  inference(superposition,[],[f528,f320])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK2(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X9] : ! [X26] : ((p(s(bool,happ(s(fun('type.list.list'(X9),bool),X26),s('type.list.list'(X9),'const.list.NIL_0')))) & ! [X22] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X26),s('type.list.list'(X9),X22)))) => ! [X35] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X26),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X35),s('type.list.list'(X9),X22)))))))) => ! [X36] : p(s(bool,happ(s(fun('type.list.list'(X9),bool),X26),s('type.list.list'(X9),X36)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.list.list_INDUCT')).
fof(f408,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(condensation,[],[f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f2432,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f2417,f305])).
fof(f2417,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK2(happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1)),X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),happ(s(fun(fun('type.list.list'(X0),bool),fun('type.list.list'(X0),bool)),'_i_0_0'),s(fun('type.list.list'(X0),bool),X1))),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f305])).
fof(f2433,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2)))))) )),
  inference(forward_demodulation,[],[f2418,f359])).
fof(f2418,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),X5),s('type.list.list'(X2),sK2(happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0')),X2))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),happ(s(fun('type.list.list'(X3),fun('type.list.list'(X2),bool)),sK16(X4,X5,X3,fun('type.list.list'(X2),bool))),s('type.list.list'(X3),'const.list.NIL_0'))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f359])).
fof(f2435,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X9),bool),X12),s('type.list.list'(X9),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X9),bool),X12),s('type.list.list'(X9),sK2('const.pair.SND_1'(s('type.pair.prod'(X10,fun('type.list.list'(X9),bool)),'const.pair.,_2'(s(X10,X11),s(fun('type.list.list'(X9),bool),X12)))),X9)))))) )),
  inference(forward_demodulation,[],[f2420,f344])).
fof(f2420,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,happ(s(fun('type.list.list'(X9),bool),X12),s('type.list.list'(X9),sK2('const.pair.SND_1'(s('type.pair.prod'(X10,fun('type.list.list'(X9),bool)),'const.pair.,_2'(s(X10,X11),s(fun('type.list.list'(X9),bool),X12)))),X9))))) | ~p(s(bool,happ(s(fun('type.list.list'(X9),bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun('type.list.list'(X9),bool)),'const.pair.,_2'(s(X10,X11),s(fun('type.list.list'(X9),bool),X12))))),s('type.list.list'(X9),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f344])).
fof(f2436,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun('type.list.list'(X13),bool),X14),s('type.list.list'(X13),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X13),bool),X14),s('type.list.list'(X13),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X13),bool),X14),s(fun('type.num.num',fun('type.list.list'(X13),bool)),X15),s('type.num.num','const.num.|48|_0')),X13)))))) )),
  inference(forward_demodulation,[],[f2421,f312])).
fof(f2421,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun('type.list.list'(X13),bool),X14),s('type.list.list'(X13),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X13),bool),X14),s(fun('type.num.num',fun('type.list.list'(X13),bool)),X15),s('type.num.num','const.num.|48|_0')),X13))))) | ~p(s(bool,happ(s(fun('type.list.list'(X13),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X13),bool),X14),s(fun('type.num.num',fun('type.list.list'(X13),bool)),X15),s('type.num.num','const.num.|48|_0'))),s('type.list.list'(X13),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f312])).
fof(f2438,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X20),bool),X21),'const.pair.,_2'(s(fun('type.list.list'(X20),bool),X22),s(X21,X23)))),X20)))))) )),
  inference(forward_demodulation,[],[f2423,f343])).
fof(f2423,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.list.list'(X20),bool),X22),s('type.list.list'(X20),sK2('const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X20),bool),X21),'const.pair.,_2'(s(fun('type.list.list'(X20),bool),X22),s(X21,X23)))),X20))))) | ~p(s(bool,happ(s(fun('type.list.list'(X20),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X20),bool),X21),'const.pair.,_2'(s(fun('type.list.list'(X20),bool),X22),s(X21,X23))))),s('type.list.list'(X20),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f343])).
fof(f2439,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,happ(s(fun('type.list.list'(X24),bool),X26),s('type.list.list'(X24),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X24),bool),X26),s('type.list.list'(X24),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X25,fun('type.list.list'(X24),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X24),bool),X26)))),X24)))))) )),
  inference(forward_demodulation,[],[f2424,f340])).
fof(f2424,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,happ(s(fun('type.list.list'(X24),bool),X26),s('type.list.list'(X24),sK2('const.sum.OUTR_1'(s('type.sum.sum'(X25,fun('type.list.list'(X24),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X24),bool),X26)))),X24))))) | ~p(s(bool,happ(s(fun('type.list.list'(X24),bool),'const.sum.OUTR_1'(s('type.sum.sum'(X25,fun('type.list.list'(X24),bool)),'const.sum.INR_1'(s(fun('type.list.list'(X24),bool),X26))))),s('type.list.list'(X24),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f340])).
fof(f2440,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27)))))) )),
  inference(forward_demodulation,[],[f2425,f339])).
fof(f2425,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,happ(s(fun('type.list.list'(X27),bool),X29),s('type.list.list'(X27),sK2('const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29)))),X27))))) | ~p(s(bool,happ(s(fun('type.list.list'(X27),bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun('type.list.list'(X27),bool),X28),'const.sum.INL_1'(s(fun('type.list.list'(X27),bool),X29))))),s('type.list.list'(X27),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f339])).
fof(f316,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK5(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f2067,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2041,f256])).
fof(f2041,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.min.@_1'(s(fun(X5,bool),X6)))))) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f313,f433])).
fof(f4479,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f4350,f267])).
fof(f4350,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f315,f819])).
fof(f4633,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4600,f256])).
fof(f4600,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f4479,f433])).
fof(f2437,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun('type.list.list'(X16),bool),happ(s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num',X19))),s('type.list.list'(X16),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X16),bool),happ(s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num',X19))),s('type.list.list'(X16),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X16),bool),X17),s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f2422,f320])).
fof(f2422,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun('type.list.list'(X16),bool),happ(s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num',X19))),s('type.list.list'(X16),sK2('const.ind_type.FCONS_3'(s(fun('type.list.list'(X16),bool),X17),s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))),X16))))) | ~p(s(bool,happ(s(fun('type.list.list'(X16),bool),'const.ind_type.FCONS_3'(s(fun('type.list.list'(X16),bool),X17),s(fun('type.num.num',fun('type.list.list'(X16),bool)),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))),s('type.list.list'(X16),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f320])).
fof(f5030,plain,(
  ( ! [X333,X335,X331,X332,X334,X330,X336] : (~p(s(bool,happ(s(fun(X330,bool),happ(s(fun(X333,fun(X330,bool)),'_dst_x0x1_2'(s(fun(X331,fun(X330,bool)),X332),s(fun(X333,X331),X334))),s(X333,X335))),s(X330,X336)))) | p(s(bool,happ(s(fun(X330,bool),happ(s(fun(X331,fun(X330,bool)),X332),s(X331,happ(s(fun(X333,X331),X334),s(X333,X335))))),s(X330,'const.min.@_1'(s(fun(X330,bool),happ(s(fun(X331,fun(X330,bool)),X332),s(X331,happ(s(fun(X333,X331),X334),s(X333,X335))))))))))) )),
  inference(superposition,[],[f313,f368])).
fof(f5028,plain,(
  ( ! [X321,X316,X318,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(X316,bool),happ(s(fun(X319,fun(X316,bool)),'_dst_x0x1_2'(s(fun(X317,fun(X316,bool)),X318),s(fun(X319,X317),X320))),s(X319,X321))),s(X316,X322)))) | p(s(bool,'|dollar|exists_1'(s(fun(X316,bool),happ(s(fun(X317,fun(X316,bool)),X318),s(X317,happ(s(fun(X319,X317),X320),s(X319,X321))))))))) )),
  inference(superposition,[],[f4479,f368])).
fof(f5107,plain,(
  ( ! [X321,X323,X318,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(X317,bool),happ(s(fun(X319,fun(X317,bool)),X320),s(X319,happ(s(fun(X318,X319),X321),s(X318,X322))))),s(X317,X323)))) | p(s(bool,happ(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))),s(X317,'const.min.@_1'(s(fun(X317,bool),happ(s(fun(X318,fun(X317,bool)),'_dst_x0x1_2'(s(fun(X319,fun(X317,bool)),X320),s(fun(X318,X319),X321))),s(X318,X322))))))))) )),
  inference(superposition,[],[f313,f368])).
fof(f5105,plain,(
  ( ! [X308,X304,X306,X303,X309,X305,X307] : (~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X305,fun(X303,bool)),X306),s(X305,happ(s(fun(X304,X305),X307),s(X304,X308))))),s(X303,X309)))) | p(s(bool,'|dollar|exists_1'(s(fun(X303,bool),happ(s(fun(X304,fun(X303,bool)),'_dst_x0x1_2'(s(fun(X305,fun(X303,bool)),X306),s(fun(X304,X305),X307))),s(X304,X308))))))) )),
  inference(superposition,[],[f4479,f368])).
fof(f5033,plain,(
  ( ! [X356,X352,X354,X351,X353,X355] : (p(s(bool,happ(s(fun('type.list.list'(X351),bool),happ(s(fun(X354,fun('type.list.list'(X351),bool)),'_dst_x0x1_2'(s(fun(X352,fun('type.list.list'(X351),bool)),X353),s(fun(X354,X352),X355))),s(X354,X356))),s('type.list.list'(X351),sK2(happ(s(fun(X352,fun('type.list.list'(X351),bool)),X353),s(X352,happ(s(fun(X354,X352),X355),s(X354,X356)))),X351))))) | ~p(s(bool,happ(s(fun('type.list.list'(X351),bool),happ(s(fun(X352,fun('type.list.list'(X351),bool)),X353),s(X352,happ(s(fun(X354,X352),X355),s(X354,X356))))),s('type.list.list'(X351),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f368])).
fof(f5110,plain,(
  ( ! [X341,X343,X339,X340,X342,X338] : (p(s(bool,happ(s(fun('type.list.list'(X338),bool),happ(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(X340,happ(s(fun(X339,X340),X342),s(X339,X343))))),s('type.list.list'(X338),sK2(happ(s(fun(X339,fun('type.list.list'(X338),bool)),'_dst_x0x1_2'(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(fun(X339,X340),X342))),s(X339,X343)),X338))))) | ~p(s(bool,happ(s(fun('type.list.list'(X338),bool),happ(s(fun(X339,fun('type.list.list'(X338),bool)),'_dst_x0x1_2'(s(fun(X340,fun('type.list.list'(X338),bool)),X341),s(fun(X339,X340),X342))),s(X339,X343))),s('type.list.list'(X338),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f368])).
fof(f2434,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f2419,f341])).
fof(f2419,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),'const.pair.FST_1'(s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),sK2(happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun('type.list.list'(X6),bool),happ(s(fun('type.pair.prod'(fun('type.list.list'(X6),bool),X7),fun('type.list.list'(X6),bool)),'const.pair.FST_0'),s('type.pair.prod'(fun('type.list.list'(X6),bool),X7),X8))),s('type.list.list'(X6),'const.list.NIL_0'))))) )),
  inference(superposition,[],[f408,f341])).
fof(f2047,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | p(s(bool,happ(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)),s('type.list.list'(X25),'const.min.@_1'(s(fun('type.list.list'(X25),bool),sK16(X26,X27,X25,bool)))))))) )),
  inference(superposition,[],[f313,f359])).
fof(f4329,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f819,f305])).
fof(f4330,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))) )),
  inference(superposition,[],[f819,f359])).
fof(f4332,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X15))) | s(X12,X16) = s(X12,sK4(X16,'const.pair.SND_1'(s('type.pair.prod'(X13,fun(X12,bool)),'const.pair.,_2'(s(X13,X14),s(fun(X12,bool),X15)))),X12))) )),
  inference(superposition,[],[f819,f344])).
fof(f4333,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),X18))) | s(X17,X20) = s(X17,sK4(X20,'const.ind_type.FCONS_3'(s(fun(X17,bool),X18),s(fun('type.num.num',fun(X17,bool)),X19),s('type.num.num','const.num.|48|_0')),X17))) )),
  inference(superposition,[],[f819,f312])).
fof(f4335,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28))) | s(X26,X30) = s(X26,sK4(X30,'const.pair.FST_1'(s('type.pair.prod'(fun(X26,bool),X27),'const.pair.,_2'(s(fun(X26,bool),X28),s(X27,X29)))),X26))) )),
  inference(superposition,[],[f819,f343])).
fof(f4336,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X33))) | s(X31,X34) = s(X31,sK4(X34,'const.sum.OUTR_1'(s('type.sum.sum'(X32,fun(X31,bool)),'const.sum.INR_1'(s(fun(X31,bool),X33)))),X31))) )),
  inference(superposition,[],[f819,f340])).
fof(f4337,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X37))) | s(X35,X38) = s(X35,sK4(X38,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X35,bool),X36),'const.sum.INL_1'(s(fun(X35,bool),X37)))),X35))) )),
  inference(superposition,[],[f819,f339])).
fof(f4338,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(X39,X42) = s(X39,sK4(X42,sK4(X40,X41,fun(X39,bool)),X39)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X39,bool),bool),X41)))) )),
  inference(superposition,[],[f819,f819])).
fof(f4334,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),happ(s(fun('type.num.num',fun(X21,bool)),X23),s('type.num.num',X24))))) | s(X21,X25) = s(X21,sK4(X25,'const.ind_type.FCONS_3'(s(fun(X21,bool),X22),s(fun('type.num.num',fun(X21,bool)),X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))),X21))) )),
  inference(superposition,[],[f819,f320])).
fof(f4687,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),'_dst_x0x1_2'(s(fun(X17,bool),X19),s(fun(X16,X17),X18))))) | s(X16,X20) = s(X16,sK4(X20,happ(s(fun(fun(X17,bool),fun(X16,bool)),'_dstsel_0_x01_1'(s(fun(X16,X17),X18))),s(fun(X17,bool),X19)),X16))) )),
  inference(superposition,[],[f819,f367])).
fof(f4331,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f819,f341])).
fof(f4654,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(subsumption_resolution,[],[f4642,f258])).
fof(f4642,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X4),bool),sK16(X5,X6,X4,bool))))) )),
  inference(superposition,[],[f4612,f266])).
fof(f4612,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X47),bool),sK16(X48,X49,X47,bool))))) | ~p(s(bool,X49))) )),
  inference(superposition,[],[f4479,f359])).
fof(f4743,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f4640,f4612])).
fof(f5556,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3)))) )),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f307,f408,f312,f313,f315,f314,f317,f316,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f350,f349,f348,f347,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f812,f813,f814,f815,f816,f817,f819,f839,f840,f841,f842,f843,f844,f897,f918,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1503,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2067,f2068,f2047,f2229,f2231,f2257,f2293,f2264,f2271,f2328,f2333,f2335,f2432,f2433,f2434,f2435,f2436,f2437,f2438,f2439,f2440,f2571,f2573,f2577,f2579,f4049,f4051,f4052,f4053,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4228,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4329,f4330,f4331,f4332,f4333,f4334,f4335,f4336,f4337,f4338,f4479,f4366,f4367,f4368,f4369,f4370,f4372,f4373,f4378,f4379,f4380,f4381,f4382,f4385,f4386,f4400,f4401,f4402,f4403,f4404,f4405,f4406,f4414,f4416,f4480,f4481,f4482,f4493,f4503,f4552,f4633,f4612,f4640,f4654,f4687,f4690,f4712,f4743,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f4990,f4991,f4993,f4994,f4995,f4996,f4997,f4999,f5000,f5005,f5006,f5007,f5008,f5009,f5012,f5013,f5028,f5029,f5030,f5032,f5033,f5034,f5048,f5050,f5066,f5067,f5068,f5069,f5071,f5072,f5073,f5074,f5075,f5077,f5078,f5083,f5084,f5085,f5086,f5087,f5090,f5091,f5105,f5106,f5107,f5109,f5110,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183])).
fof(f5183,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f5171,f256])).
fof(f5171,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4613,f433])).
fof(f4613,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X50),X51)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X50),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4479,f341])).
fof(f1472,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) | s(bool,X2) = s(bool,happ(s(fun('type.pair.prod'(bool,X0),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X0),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f341,f433])).
fof(f1474,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X6),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X6),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f341,f433])).
fof(f1476,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X12),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X12),X13)))) )),
  inference(superposition,[],[f341,f266])).
fof(f1947,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),X26))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X25),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X25),X26)))) )),
  inference(superposition,[],[f460,f341])).
fof(f4050,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X71),'const.sum.INL_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),X70))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X69),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X69),X70)))) )),
  inference(superposition,[],[f501,f341])).
fof(f4225,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X74,bool),'const.sum.INR_1'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),X73))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X72),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X72),X73)))) )),
  inference(superposition,[],[f528,f341])).
fof(f351,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | s(X1,X5) = s(X1,happ(s(fun(X0,X1),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f352,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X3),s(fun(X1,X0),X2)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X3),s(X0,X4))) | s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f356,plain,(
  ( ! [X2,X0,X7,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0)))))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X7] : s(X0,X7) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))))) & ! [X6] : s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & ((? [X5] : s(X0,X5) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) | ? [X4] : s(X1,X4) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) <=> (! [X5] : s(X0,X5) = s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,X5))))) & ! [X4] : s(X1,X4) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X23,X9] : ! [X10,X37] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X23),X10),s(fun(X23,X9),X37)))) <=> (! [X24] : s(X23,X24) = s(X23,happ(s(fun(X9,X23),X10),s(X9,happ(s(fun(X23,X9),X37),s(X23,X24))))) & ! [X25] : s(X9,X25) = s(X9,happ(s(fun(X23,X9),X37),s(X23,happ(s(fun(X9,X23),X10),s(X9,X25)))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.ISO')).
fof(f357,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X6)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X1,X0),X2),s(fun(X0,X1),X3)))) | s(X1,sK15(X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK15(X3,X2,X1,X0)))))) | s(X0,sK14(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK14(X3,X2,X1,X0))))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f394,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7))))))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4)))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) | ~p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) | p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(ennf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X2,X3),X7),s(fun(X3,X2),X6)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X1),X5),s(fun(X1,X0),X4))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X2,X0),fun(X3,X1)),'_dst_x00x11_2'(s(fun(X0,X1),X5),s(fun(fun(X2,X0),fun(X3,X0)),'_dstsel_0_x01_1'(s(fun(X3,X2),X6))))),s(fun(fun(X3,X1),fun(X2,X0)),'_dst_x00x11_2'(s(fun(X1,X0),X4),s(fun(fun(X3,X1),fun(X2,X1)),'_dstsel_0_x01_1'(s(fun(X2,X3),X7)))))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X23,X42,X9,X43] : ! [X44,X37,X45,X10] : ((p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X43),X10),s(fun(X43,X9),X45)))) & p(s(bool,'const.ind_type.ISO_2'(s(fun(X23,X42),X37),s(fun(X42,X23),X44))))) => p(s(bool,'const.ind_type.ISO_2'(s(fun(fun(X9,X23),fun(X43,X42)),'_dst_x00x11_2'(s(fun(X23,X42),X37),s(fun(fun(X9,X23),fun(X43,X23)),'_dstsel_0_x01_1'(s(fun(X43,X9),X45))))),s(fun(fun(X43,X42),fun(X9,X23)),'_dst_x00x11_2'(s(fun(X42,X23),X44),s(fun(fun(X43,X42),fun(X9,X42)),'_dstsel_0_x01_1'(s(fun(X9,X43),X10)))))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.ISO_FUN')).
fof(f4189,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4158,f256])).
fof(f4158,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X8,X11) = s(X8,happ(s(fun(X7,X8),X9),s(X7,happ(s(fun(X8,X7),X10),s(X8,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f356,f433])).
fof(f4537,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f4506,f256])).
fof(f4506,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | s(X7,X11) = s(X7,happ(s(fun(X8,X7),X10),s(X8,happ(s(fun(X7,X8),X9),s(X7,X11))))) | s(bool,X12) = s(bool,'const.ind_type.ISO_2'(s(fun(X7,X8),X9),s(fun(X8,X7),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f357,f433])).
fof(f5039,plain,(
  ( ! [X408,X410,X405,X407,X403,X409,X404,X406,X402] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X403,X402),X409),s(fun(X402,X403),happ(s(fun(X406,fun(X402,X403)),'_dst_x0x1_2'(s(fun(X404,fun(X402,X403)),X405),s(fun(X406,X404),X407))),s(X406,X408)))))) | s(X403,X410) = s(X403,happ(s(fun(X402,X403),happ(s(fun(X404,fun(X402,X403)),X405),s(X404,happ(s(fun(X406,X404),X407),s(X406,X408))))),s(X402,happ(s(fun(X403,X402),X409),s(X403,X410)))))) )),
  inference(superposition,[],[f357,f368])).
fof(f5037,plain,(
  ( ! [X389,X391,X385,X387,X388,X390,X384,X386,X392] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X385,X384),X391),s(fun(X384,X385),happ(s(fun(X388,fun(X384,X385)),'_dst_x0x1_2'(s(fun(X386,fun(X384,X385)),X387),s(fun(X388,X386),X389))),s(X388,X390)))))) | s(X384,X392) = s(X384,happ(s(fun(X385,X384),X391),s(X385,happ(s(fun(X384,X385),happ(s(fun(X386,fun(X384,X385)),X387),s(X386,happ(s(fun(X388,X386),X389),s(X388,X390))))),s(X384,X392)))))) )),
  inference(superposition,[],[f356,f368])).
fof(f5116,plain,(
  ( ! [X389,X391,X397,X393,X395,X390,X396,X392,X394] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X390,X389),X396),s(fun(X389,X390),happ(s(fun(X392,fun(X389,X390)),X393),s(X392,happ(s(fun(X391,X392),X394),s(X391,X395)))))))) | s(X390,X397) = s(X390,happ(s(fun(X389,X390),happ(s(fun(X391,fun(X389,X390)),'_dst_x0x1_2'(s(fun(X392,fun(X389,X390)),X393),s(fun(X391,X392),X394))),s(X391,X395))),s(X389,happ(s(fun(X390,X389),X396),s(X390,X397)))))) )),
  inference(superposition,[],[f357,f368])).
fof(f5114,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X372,X371),X378),s(fun(X371,X372),happ(s(fun(X374,fun(X371,X372)),X375),s(X374,happ(s(fun(X373,X374),X376),s(X373,X377)))))))) | s(X371,X379) = s(X371,happ(s(fun(X372,X371),X378),s(X372,happ(s(fun(X371,X372),happ(s(fun(X373,fun(X371,X372)),'_dst_x0x1_2'(s(fun(X374,fun(X371,X372)),X375),s(fun(X373,X374),X376))),s(X373,X377))),s(X371,X379)))))) )),
  inference(superposition,[],[f356,f368])).
fof(f5038,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X393,X394),happ(s(fun(X397,fun(X393,X394)),'_dst_x0x1_2'(s(fun(X395,fun(X393,X394)),X396),s(fun(X397,X395),X398))),s(X397,X399))),s(fun(X394,X393),X400)))) | s(X393,X401) = s(X393,happ(s(fun(X394,X393),X400),s(X394,happ(s(fun(X393,X394),happ(s(fun(X395,fun(X393,X394)),X396),s(X395,happ(s(fun(X397,X395),X398),s(X397,X399))))),s(X393,X401)))))) )),
  inference(superposition,[],[f357,f368])).
fof(f5036,plain,(
  ( ! [X383,X377,X379,X380,X382,X376,X378,X375,X381] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X375,X376),happ(s(fun(X379,fun(X375,X376)),'_dst_x0x1_2'(s(fun(X377,fun(X375,X376)),X378),s(fun(X379,X377),X380))),s(X379,X381))),s(fun(X376,X375),X382)))) | s(X376,X383) = s(X376,happ(s(fun(X375,X376),happ(s(fun(X377,fun(X375,X376)),X378),s(X377,happ(s(fun(X379,X377),X380),s(X379,X381))))),s(X375,happ(s(fun(X376,X375),X382),s(X376,X383)))))) )),
  inference(superposition,[],[f356,f368])).
fof(f5115,plain,(
  ( ! [X383,X385,X387,X380,X382,X388,X384,X386,X381] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X380,X381),happ(s(fun(X383,fun(X380,X381)),X384),s(X383,happ(s(fun(X382,X383),X385),s(X382,X386))))),s(fun(X381,X380),X387)))) | s(X380,X388) = s(X380,happ(s(fun(X381,X380),X387),s(X381,happ(s(fun(X380,X381),happ(s(fun(X382,fun(X380,X381)),'_dst_x0x1_2'(s(fun(X383,fun(X380,X381)),X384),s(fun(X382,X383),X385))),s(X382,X386))),s(X380,X388)))))) )),
  inference(superposition,[],[f357,f368])).
fof(f5113,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (~p(s(bool,'const.ind_type.ISO_2'(s(fun(X362,X363),happ(s(fun(X365,fun(X362,X363)),X366),s(X365,happ(s(fun(X364,X365),X367),s(X364,X368))))),s(fun(X363,X362),X369)))) | s(X363,X370) = s(X363,happ(s(fun(X362,X363),happ(s(fun(X364,fun(X362,X363)),'_dst_x0x1_2'(s(fun(X365,fun(X362,X363)),X366),s(fun(X364,X365),X367))),s(X364,X368))),s(X362,happ(s(fun(X363,X362),X369),s(X363,X370)))))) )),
  inference(superposition,[],[f356,f368])).
fof(f5127,plain,(
  ( ! [X466,X463,X465,X467,X462,X464] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X467))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X463,'type.integer.int'),X464),s(X463,happ(s(fun(X462,X463),X465),s(X462,X466))))))) | s('type.integer.int',X467) = s('type.integer.int',happ(s(fun(X462,'type.integer.int'),'_dst_x0x1_2'(s(fun(X463,'type.integer.int'),X464),s(fun(X462,X463),X465))),s(X462,X466)))) )),
  inference(superposition,[],[f333,f368])).
fof(f5125,plain,(
  ( ! [X453,X455,X451,X452,X454,X450] : (s('type.integer.int',X455) != s('type.integer.int',happ(s(fun(X451,'type.integer.int'),X452),s(X451,happ(s(fun(X450,X451),X453),s(X450,X454))))) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X455))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X450,'type.integer.int'),'_dst_x0x1_2'(s(fun(X451,'type.integer.int'),X452),s(fun(X450,X451),X453))),s(X450,X454)))))) )),
  inference(superposition,[],[f334,f368])).
fof(f5111,plain,(
  ( ! [X350,X344,X346,X352,X349,X351,X345,X347,X348] : (s(fun(X344,X345),X351) != s(fun(X344,X345),happ(s(fun(X347,fun(X344,X345)),X348),s(X347,happ(s(fun(X346,X347),X349),s(X346,X350))))) | s(X345,happ(s(fun(X344,X345),X351),s(X344,X352))) = s(X345,happ(s(fun(X344,X345),happ(s(fun(X346,fun(X344,X345)),'_dst_x0x1_2'(s(fun(X347,fun(X344,X345)),X348),s(fun(X346,X347),X349))),s(X346,X350))),s(X344,X352)))) )),
  inference(superposition,[],[f354,f368])).
fof(f5106,plain,(
  ( ! [X313,X315,X310,X316,X312,X314,X311] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X310,bool),happ(s(fun(X312,fun(X310,bool)),X313),s(X312,happ(s(fun(X311,X312),X314),s(X311,X315))))))) | s(X310,X316) = s(X310,sK4(X316,happ(s(fun(X311,fun(X310,bool)),'_dst_x0x1_2'(s(fun(X312,fun(X310,bool)),X313),s(fun(X311,X312),X314))),s(X311,X315)),X310))) )),
  inference(superposition,[],[f819,f368])).
fof(f5066,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X82,bool),X83),s(X82,happ(s(fun(X81,X82),X84),s(X81,sK4(X85,'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84)),X81))))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'_dst_x0x1_2'(s(fun(X82,bool),X83),s(fun(X81,X82),X84))))))) )),
  inference(superposition,[],[f315,f368])).
fof(f5050,plain,(
  ( ! [X492,X488,X490,X493,X489,X491] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X493))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X490,'type.integer.int'),'_dst_x0x1_2'(s(fun(X488,'type.integer.int'),X489),s(fun(X490,X488),X491))),s(X490,X492))))) | s('type.integer.int',X493) = s('type.integer.int',happ(s(fun(X488,'type.integer.int'),X489),s(X488,happ(s(fun(X490,X488),X491),s(X490,X492)))))) )),
  inference(superposition,[],[f333,f368])).
fof(f5048,plain,(
  ( ! [X476,X478,X480,X477,X479,X481] : (s('type.integer.int',X481) != s('type.integer.int',happ(s(fun(X478,'type.integer.int'),'_dst_x0x1_2'(s(fun(X476,'type.integer.int'),X477),s(fun(X478,X476),X479))),s(X478,X480))) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X481))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X476,'type.integer.int'),X477),s(X476,happ(s(fun(X478,X476),X479),s(X478,X480)))))))) )),
  inference(superposition,[],[f334,f368])).
fof(f5034,plain,(
  ( ! [X358,X364,X360,X362,X357,X359,X365,X361,X363] : (s(fun(X357,X358),X364) != s(fun(X357,X358),happ(s(fun(X361,fun(X357,X358)),'_dst_x0x1_2'(s(fun(X359,fun(X357,X358)),X360),s(fun(X361,X359),X362))),s(X361,X363))) | s(X358,happ(s(fun(X357,X358),X364),s(X357,X365))) = s(X358,happ(s(fun(X357,X358),happ(s(fun(X359,fun(X357,X358)),X360),s(X359,happ(s(fun(X361,X359),X362),s(X361,X363))))),s(X357,X365)))) )),
  inference(superposition,[],[f354,f368])).
fof(f5029,plain,(
  ( ! [X325,X327,X323,X329,X324,X326,X328] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X323,bool),happ(s(fun(X326,fun(X323,bool)),'_dst_x0x1_2'(s(fun(X324,fun(X323,bool)),X325),s(fun(X326,X324),X327))),s(X326,X328))))) | s(X323,X329) = s(X323,sK4(X329,happ(s(fun(X324,fun(X323,bool)),X325),s(X324,happ(s(fun(X326,X324),X327),s(X326,X328)))),X323))) )),
  inference(superposition,[],[f819,f368])).
fof(f4889,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s(fun(X95,X96),'_dst_x0x1_2'(s(fun(X97,X96),happ(s(fun(X99,fun(X97,X96)),X100),s(X99,X101))),s(fun(X95,X97),X98))) = s(fun(X95,X96),happ(s(fun(X99,fun(X95,X96)),'_dst_x0x1_2'(s(fun(fun(X97,X96),fun(X95,X96)),'_dstsel_0_x01_1'(s(fun(X95,X97),X98))),s(fun(X99,fun(X97,X96)),X100))),s(X99,X101)))) )),
  inference(superposition,[],[f368,f367])).
fof(f4879,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun(X39,bool),'_dst_x0x1_2'(s(fun(X37,bool),X38),s(fun(X39,X37),X40))),s(X39,X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,happ(s(fun(X37,bool),X38),s(X37,happ(s(fun(X39,X37),X40),s(X39,X41)))))) )),
  inference(superposition,[],[f368,f433])).
fof(f4878,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X33,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X33,X31),X34))),s(X33,X35))) | s(bool,X36) = s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X33,X31),X34),s(X33,X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f368,f433])).
fof(f4877,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,happ(s(fun(X28,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X28,X26),X29))),s(X28,X30))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X28,X26),X29),s(X28,X30)))))) )),
  inference(superposition,[],[f368,f266])).
fof(f4876,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,X19),happ(s(fun(X20,'type.pair.prod'(X18,X19)),X21),s(X20,X22))))) = s(X18,happ(s(fun(X20,X18),'_dst_x0x1_2'(s(fun('type.pair.prod'(X18,X19),X18),'const.pair.FST_0'),s(fun(X20,'type.pair.prod'(X18,X19)),X21))),s(X20,X22)))) )),
  inference(superposition,[],[f368,f341])).
fof(f4875,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,happ(s(fun(X12,X10),X13),s(X12,happ(s(fun(X11,X12),X14),s(X11,happ(s(fun(X15,X11),X16),s(X15,X17))))))) = s(X10,happ(s(fun(X15,X10),'_dst_x0x1_2'(s(fun(X11,X10),'_dst_x0x1_2'(s(fun(X12,X10),X13),s(fun(X11,X12),X14))),s(fun(X15,X11),X16))),s(X15,X17)))) )),
  inference(superposition,[],[f368,f368])).
fof(f4874,plain,(
  ( ! [X103,X101,X99,X105,X102,X100,X104] : (s(X104,happ(s(fun(fun(X99,X100),X104),X105),s(fun(X99,X100),'_dst_x0x1_2'(s(fun(X101,X100),X103),s(fun(X99,X101),X102))))) = s(X104,happ(s(fun(fun(X101,X100),X104),'_dst_x0x1_2'(s(fun(fun(X99,X100),X104),X105),s(fun(fun(X101,X100),fun(X99,X100)),'_dstsel_0_x01_1'(s(fun(X99,X101),X102))))),s(fun(X101,X100),X103)))) )),
  inference(superposition,[],[f368,f367])).
fof(f4873,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(X24,happ(s(fun(X20,X24),X25),s(X20,X23))) = s(X24,happ(s(fun('type.list.list'(X21),X24),'_dst_x0x1_2'(s(fun(X20,X24),X25),s(fun('type.list.list'(X21),X20),sK16(X22,X23,X21,X20)))),s('type.list.list'(X21),'const.list.NIL_0')))) )),
  inference(superposition,[],[f368,f359])).
fof(f4872,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X18,happ(s(fun(X12,X18),X19),s(X12,happ(s(fun(X14,X12),X15),s(X14,happ(s(fun(X13,X14),X16),s(X13,X17))))))) = s(X18,happ(s(fun(X13,X18),'_dst_x0x1_2'(s(fun(X12,X18),X19),s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X14,X12),X15),s(fun(X13,X14),X16))))),s(X13,X17)))) )),
  inference(superposition,[],[f368,f368])).
fof(f4871,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X11,X9] : (s(X10,happ(s(fun(X4,X10),X11),s(X4,happ(s(fun(X7,X4),'_dst_x0x1_2'(s(fun(X5,X4),X6),s(fun(X7,X5),X8))),s(X7,X9))))) = s(X10,happ(s(fun(X5,X10),'_dst_x0x1_2'(s(fun(X4,X10),X11),s(fun(X5,X4),X6))),s(X5,happ(s(fun(X7,X5),X8),s(X7,X9)))))) )),
  inference(superposition,[],[f368,f368])).
fof(f4869,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X2,fun(X1,X0)),X3),s(X2,happ(s(fun(X4,X2),X5),s(X4,X6))))),s(X1,X7))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X4,fun(X1,X0)),'_dst_x0x1_2'(s(fun(X2,fun(X1,X0)),X3),s(fun(X4,X2),X5))),s(X4,X6))),s(X1,X7)))) )),
  inference(unit_resulting_resolution,[],[f368,f354])).
fof(f4870,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X0,'type.integer.int'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun(X2,'type.integer.int'),'_dst_x0x1_2'(s(fun(X0,'type.integer.int'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) )),
  inference(unit_resulting_resolution,[],[f368,f334])).
fof(f4865,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4818,f256])).
fof(f4818,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) != s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X39),bool),sK16(X40,'T_0',X39,bool)))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f435,f4742])).
fof(f4736,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,'|dollar|exists_1'(s(fun('type.list.list'(X2),bool),sK16(X3,'T_0',X2,bool))),X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f4640,f435])).
fof(f4742,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool))))) )),
  inference(unit_resulting_resolution,[],[f420,f4640,f1999])).
fof(f4640,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X0),bool),sK16(X1,'T_0',X0,bool)))))) )),
  inference(unit_resulting_resolution,[],[f256,f4612])).
fof(f4482,plain,(
  ( ! [X54,X52,X53,X51] : (s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(subsumption_resolution,[],[f4435,f256])).
fof(f4435,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X53) != s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X53) | s(X51,X54) = s(X51,sK4(X54,X52,X51))) )),
  inference(superposition,[],[f435,f819])).
fof(f4416,plain,(
  ( ! [X212,X213,X214] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X212))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X214))) | s('type.integer.int',X214) = s('type.integer.int',sK4(X212,X213,'type.integer.int')) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.integer.int',bool),X213)))) )),
  inference(superposition,[],[f333,f819])).
fof(f4401,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK4(X158,sK4(X156,X157,fun(X155,bool)),X155))))) | p(s(bool,'|dollar|exists_1'(s(fun(X155,bool),sK4(X156,X157,fun(X155,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X155,bool),bool),X157)))) )),
  inference(superposition,[],[f315,f819])).
fof(f4382,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X98))))) | s(bool,'F_0') = s(bool,sK4(X98,X99,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f528,f819])).
fof(f4381,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X97),'const.sum.INL_1'(s(bool,X95))))) | s(bool,'F_0') = s(bool,sK4(X95,X96,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f501,f819])).
fof(f4379,plain,(
  ( ! [X90,X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X90))) | s(bool,'F_0') = s(bool,sK4(X90,X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f460,f819])).
fof(f4319,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool))) )),
  inference(superposition,[],[f819,f433])).
fof(f4318,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,X10) = s(bool,sK4(X8,X9,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f819,f433])).
fof(f4317,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool))) )),
  inference(superposition,[],[f819,f266])).
fof(f4311,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4254,f256])).
fof(f4254,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f435,f528])).
fof(f4231,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X95,bool),'const.sum.INR_1'(s(bool,X94))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,X94)))))) )),
  inference(superposition,[],[f528,f339])).
fof(f4230,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X92,bool),'const.sum.INR_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X90,bool),'const.sum.INR_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f528,f340])).
fof(f4229,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X89,bool),'const.sum.INR_1'(s(bool,X87))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X86),'const.pair.,_2'(s(bool,X87),s(X86,X88)))))) )),
  inference(superposition,[],[f528,f343])).
fof(f4227,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X81,bool),'const.sum.INR_1'(s(bool,X79))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X79),s(fun('type.num.num',bool),X80),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f528,f312])).
fof(f4226,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,X77))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X75,bool),'const.pair.,_2'(s(X75,X76),s(bool,X77)))))) )),
  inference(superposition,[],[f528,f344])).
fof(f4224,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X71,bool),'const.sum.INR_1'(s(bool,X70))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X68),bool),sK16(X69,X70,X68,bool)),s('type.list.list'(X68),'const.list.NIL_0')))) )),
  inference(superposition,[],[f528,f359])).
fof(f4135,plain,(
  ( ! [X33,X34,X32] : (s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4079,f256])).
fof(f4079,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X34) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X32),'const.sum.INL_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f435,f501])).
fof(f4056,plain,(
  ( ! [X92,X90,X91] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X92),'const.sum.INL_1'(s(bool,X91))))) | s(bool,'F_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X91)))))) )),
  inference(superposition,[],[f501,f339])).
fof(f4055,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X89),'const.sum.INL_1'(s(bool,X88))))) | s(bool,'F_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X87,bool),'const.sum.INR_1'(s(bool,X88)))))) )),
  inference(superposition,[],[f501,f340])).
fof(f4054,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X86),'const.sum.INL_1'(s(bool,X84))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X83),'const.pair.,_2'(s(bool,X84),s(X83,X85)))))) )),
  inference(superposition,[],[f501,f343])).
fof(f4052,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,X76))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f501,f312])).
fof(f4051,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,X74))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X72,bool),'const.pair.,_2'(s(X72,X73),s(bool,X74)))))) )),
  inference(superposition,[],[f501,f344])).
fof(f4049,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X68),'const.sum.INL_1'(s(bool,X67))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X65),bool),sK16(X66,X67,X65,bool)),s('type.list.list'(X65),'const.list.NIL_0')))) )),
  inference(superposition,[],[f501,f359])).
fof(f2573,plain,(
  ( ! [X14,X12,X15,X13] : (s(X14,f(s(bool,X15))) != s(X14,g(s(bool,X15))) | s(fun(bool,X14),X12) = s(fun(bool,X14),X13) | s(bool,'T_0') = s(bool,sK12(X12,X13,X14,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f353,f433])).
fof(f2571,plain,(
  ( ! [X6,X4,X7,X5] : (s(X6,f(s(bool,'T_0'))) != s(X6,g(s(bool,'T_0'))) | s(fun(bool,X6),X4) = s(fun(bool,X6),X5) | s(bool,X7) = s(bool,sK12(X4,X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f353,f433])).
fof(f2335,plain,(
  ( ! [X28,X29] : (s(bool,X29) = s(bool,'const.ind_type.ISO_2'(s(fun(X28,X28),'_i_0_0'),s(fun(X28,X28),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f2312,f256])).
fof(f2312,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,'const.ind_type.ISO_2'(s(fun(X28,X28),'_i_0_0'),s(fun(X28,X28),'_i_0_0'))) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f1999,f428])).
fof(f2333,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2310,f256])).
fof(f2310,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X25))) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f1999,f460])).
fof(f2328,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f2302,f256])).
fof(f2302,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1999,f433])).
fof(f2271,plain,(
  ( ! [X156,X154,X155,X153] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X156))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',happ(s(fun('type.num.num','type.integer.int'),X154),s('type.num.num',X155))))) | s('type.integer.int',X156) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X153),s(fun('type.num.num','type.integer.int'),X154),s('type.num.num','const.num.SUC_1'(s('type.num.num',X155)))))) )),
  inference(superposition,[],[f333,f320])).
fof(f2293,plain,(
  ( ! [X127,X125,X123,X126,X124] : (p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126)))))) | ~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123)))))) )),
  inference(forward_demodulation,[],[f2263,f320])).
fof(f2263,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,happ(s(fun(X123,bool),happ(s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num',X126))),s(X123,sK4(X127,'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126)))),X123))))) | p(s(bool,'|dollar|exists_1'(s(fun(X123,bool),'const.ind_type.FCONS_3'(s(fun(X123,bool),X124),s(fun('type.num.num',fun(X123,bool)),X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X126))))))))) )),
  inference(superposition,[],[f315,f320])).
fof(f2002,plain,(
  ( ! [X21,X20] : (s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1970,f256])).
fof(f1970,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | s(bool,X21) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X20))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X20)) )),
  inference(superposition,[],[f435,f460])).
fof(f1995,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X33),s(X32,X34)))))) )),
  inference(forward_demodulation,[],[f1950,f305])).
fof(f1950,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X33),s(X32,X34)))))) )),
  inference(superposition,[],[f460,f343])).
fof(f1994,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1949,f305])).
fof(f1949,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X30))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f460,f312])).
fof(f1993,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29)))))) )),
  inference(forward_demodulation,[],[f1948,f305])).
fof(f1948,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X29))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29)))))) )),
  inference(superposition,[],[f460,f344])).
fof(f1992,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f1946,f305])).
fof(f1946,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(X22),bool),sK16(X23,X24,X22,bool)),s('type.list.list'(X22),'const.list.NIL_0')))) )),
  inference(superposition,[],[f460,f359])).
fof(f1508,plain,(
  ( ! [X101,X99,X100] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X101))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X99),X100))))) | s('type.integer.int',X101) = s('type.integer.int',happ(s(fun('type.pair.prod'('type.integer.int',X99),'type.integer.int'),'const.pair.FST_0'),s('type.pair.prod'('type.integer.int',X99),X100)))) )),
  inference(superposition,[],[f333,f341])).
fof(f1521,plain,(
  ( ! [X83,X81,X84,X82] : (p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83)))))) | ~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81)))))) )),
  inference(forward_demodulation,[],[f1502,f341])).
fof(f1502,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,happ(s(fun(X81,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X81,bool),X82),X83))),s(X81,sK4(X84,happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83)),X81))))) | p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),happ(s(fun('type.pair.prod'(fun(X81,bool),X82),fun(X81,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X81,bool),X82),X83))))))) )),
  inference(superposition,[],[f315,f341])).
fof(f1485,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),sK4(X36,'const.pair.FST_0','type.pair.prod'(bool,X35)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X35),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f315,f341])).
fof(f1380,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1367,f256])).
fof(f1367,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f315,f433])).
fof(f1366,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f315,f266])).
fof(f1364,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f315,f433])).
fof(f1362,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f315,f433])).
fof(f1379,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),X28)))) | ~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26)))))) )),
  inference(forward_demodulation,[],[f1361,f339])).
fof(f1361,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X28),s(X26,sK4(X29,'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28)))),X26))))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),'const.sum.OUTL_1'(s('type.sum.sum'(fun(X26,bool),X27),'const.sum.INL_1'(s(fun(X26,bool),X28))))))))) )),
  inference(superposition,[],[f315,f339])).
fof(f1378,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X24)))) | ~p(s(bool,happ(s(fun(X22,bool),X24),s(X22,sK4(X25,'const.sum.OUTR_1'(s('type.sum.sum'(X23,fun(X22,bool)),'const.sum.INR_1'(s(fun(X22,bool),X24)))),X22)))))) )),
  inference(forward_demodulation,[],[f1360,f340])).
fof(f1360,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),X24),s(X22,sK4(X25,'const.sum.OUTR_1'(s('type.sum.sum'(X23,fun(X22,bool)),'const.sum.INR_1'(s(fun(X22,bool),X24)))),X22))))) | p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),'const.sum.OUTR_1'(s('type.sum.sum'(X23,fun(X22,bool)),'const.sum.INR_1'(s(fun(X22,bool),X24))))))))) )),
  inference(superposition,[],[f315,f340])).
fof(f1377,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X19)))) | ~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK4(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f1359,f343])).
fof(f1359,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X19),s(X17,sK4(X21,'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X17,bool),X18),'const.pair.,_2'(s(fun(X17,bool),X19),s(X18,X20))))))))) )),
  inference(superposition,[],[f315,f343])).
fof(f1376,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),X14)))) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK4(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13)))))) )),
  inference(forward_demodulation,[],[f1358,f312])).
fof(f1358,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK4(X16,'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0')),X13))))) | p(s(bool,'|dollar|exists_1'(s(fun(X13,bool),'const.ind_type.FCONS_3'(s(fun(X13,bool),X14),s(fun('type.num.num',fun(X13,bool)),X15),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f315,f312])).
fof(f1375,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X11)))) | ~p(s(bool,happ(s(fun(X8,bool),X11),s(X8,sK4(X12,'const.pair.SND_1'(s('type.pair.prod'(X9,fun(X8,bool)),'const.pair.,_2'(s(X9,X10),s(fun(X8,bool),X11)))),X8)))))) )),
  inference(forward_demodulation,[],[f1357,f344])).
fof(f1357,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X11),s(X8,sK4(X12,'const.pair.SND_1'(s('type.pair.prod'(X9,fun(X8,bool)),'const.pair.,_2'(s(X9,X10),s(fun(X8,bool),X11)))),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.SND_1'(s('type.pair.prod'(X9,fun(X8,bool)),'const.pair.,_2'(s(X9,X10),s(fun(X8,bool),X11))))))))) )),
  inference(superposition,[],[f315,f344])).
fof(f1374,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X6)))) | ~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3)))))) )),
  inference(forward_demodulation,[],[f1356,f359])).
fof(f1356,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,sK4(X7,happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0')),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun('type.list.list'(X4),fun(X3,bool)),sK16(X5,X6,X4,fun(X3,bool))),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(superposition,[],[f315,f359])).
fof(f1373,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1355,f305])).
fof(f1355,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f315,f305])).
fof(f1348,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1105,f256])).
fof(f1105,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f435,f433])).
fof(f1347,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1104,f256])).
fof(f1104,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f406,f433])).
fof(f1346,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1103,f256])).
fof(f1103,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f332,f433])).
fof(f994,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X35),'const.sum.INL_1'(s(bool,X36))))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f433,f339])).
fof(f993,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X32,bool),'const.sum.INR_1'(s(bool,X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f433,f340])).
fof(f992,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,X31) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X29),s(X28,X30))))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f433,f343])).
fof(f991,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,'const.ind_type.FCONS_3'(s(bool,X25),s(fun('type.num.num',bool),X26),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f433,f312])).
fof(f990,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X21,bool),'const.pair.,_2'(s(X21,X22),s(bool,X23))))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f433,f344])).
fof(f989,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,happ(s(fun('type.list.list'(X17),bool),sK16(X18,X19,X17,bool)),s('type.list.list'(X17),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f433,f359])).
fof(f988,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f433,f305])).
fof(f978,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f433,f433])).
fof(f961,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X42),'const.sum.INL_1'(s(bool,X43)))))) )),
  inference(superposition,[],[f433,f339])).
fof(f960,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X39,bool),'const.sum.INR_1'(s(bool,X40)))))) )),
  inference(superposition,[],[f433,f340])).
fof(f959,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,X36),s(X35,X37)))))) )),
  inference(superposition,[],[f433,f343])).
fof(f958,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X32),s(fun('type.num.num',bool),X33),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f433,f312])).
fof(f957,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X29),s(bool,X30)))))) )),
  inference(superposition,[],[f433,f344])).
fof(f956,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X24),bool),sK16(X25,X26,X24,bool)),s('type.list.list'(X24),'const.list.NIL_0')))) )),
  inference(superposition,[],[f433,f359])).
fof(f955,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))) )),
  inference(superposition,[],[f433,f305])).
fof(f923,plain,(
  ( ! [X109,X107,X108,X106] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X108))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X109))) | s('type.integer.int',X109) = s('type.integer.int',happ(s(fun('type.list.list'(X106),'type.integer.int'),sK16(X107,X108,X106,'type.integer.int')),s('type.list.list'(X106),'const.list.NIL_0')))) )),
  inference(superposition,[],[f333,f359])).
fof(f897,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X0),bool),sK16(X1,X2,X0,bool)),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f359,f266])).
fof(f844,plain,(
  ( ! [X17,X18,X16] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X17))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X18))) | s('type.integer.int',X18) = s('type.integer.int','const.sum.OUTL_1'(s('type.sum.sum'('type.integer.int',X16),'const.sum.INL_1'(s('type.integer.int',X17)))))) )),
  inference(superposition,[],[f333,f339])).
fof(f843,plain,(
  ( ! [X14,X15,X13] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X14))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X15))) | s('type.integer.int',X15) = s('type.integer.int','const.sum.OUTR_1'(s('type.sum.sum'(X13,'type.integer.int'),'const.sum.INR_1'(s('type.integer.int',X14)))))) )),
  inference(superposition,[],[f333,f340])).
fof(f842,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X10))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X12))) | s('type.integer.int',X12) = s('type.integer.int','const.pair.FST_1'(s('type.pair.prod'('type.integer.int',X9),'const.pair.,_2'(s('type.integer.int',X10),s(X9,X11)))))) )),
  inference(superposition,[],[f333,f343])).
fof(f841,plain,(
  ( ! [X6,X8,X7] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))) | s('type.integer.int',X8) = s('type.integer.int','const.ind_type.FCONS_3'(s('type.integer.int',X6),s(fun('type.num.num','type.integer.int'),X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f333,f312])).
fof(f840,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X4))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X5))) | s('type.integer.int',X5) = s('type.integer.int','const.pair.SND_1'(s('type.pair.prod'(X2,'type.integer.int'),'const.pair.,_2'(s(X2,X3),s('type.integer.int',X4)))))) )),
  inference(superposition,[],[f333,f344])).
fof(f839,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int',happ(s(fun('type.integer.int','type.integer.int'),'_i_0_0'),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f333,f305])).
fof(f819,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f818,f258])).
fof(f818,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f314,f266])).
fof(f805,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f792,f256])).
fof(f792,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.ind_type.ISO_2'(s(fun(X4,X4),'_i_0_0'),s(fun(X4,X4),'_i_0_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f435,f428])).
fof(f744,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f344,f266])).
fof(f628,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f312,f266])).
fof(f528,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f340,f266])).
fof(f501,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f339,f266])).
fof(f482,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f477,f258])).
fof(f477,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f329,f266])).
fof(f428,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f259,f268])).
fof(f393,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X16,X17,X47,X18] : ! [X19,X20,X21] : s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X47,X16),X19),s(fun(X17,X47),happ(s(fun(X18,fun(X17,X47)),X20),s(X18,X21))))) = s(fun(X17,X16),happ(s(fun(X18,fun(X17,X16)),'_dst_x00x11_2'(s(fun(X47,X16),X19),s(fun(X18,fun(X17,X47)),X20))),s(X18,X21)))),
  file('thm.bvlSem.v_|49||49||split|1','_dst_x00x11_ax')).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))) = s(X16,happ(s(fun(X18,X16),'_dst_x0x1_2'(s(fun(X17,X16),X19),s(fun(X18,X17),X20))),s(X18,X21)))),
  file('thm.bvlSem.v_|49||49||split|1','_dst_x0x1_ax')).
fof(f367,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),X3))) = s(fun(X1,X0),happ(s(fun(fun(X2,X0),fun(X1,X0)),'_dstsel_0_x01_1'(s(fun(X1,X2),X3))),s(fun(X2,X0),X4)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X16,X17,X18] : ! [X20,X21] : s(fun(X17,X16),'_dst_x0x1_2'(s(fun(X18,X16),X21),s(fun(X17,X18),X20))) = s(fun(X17,X16),happ(s(fun(fun(X18,X16),fun(X17,X16)),'_dstsel_0_x01_1'(s(fun(X17,X18),X20))),s(fun(X18,X16),X21)))),
  file('thm.bvlSem.v_|49||49||split|1','_dstsel_0_x01_ax')).
fof(f359,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X23,X9] : ! [X38,X11] : ? [X39] : (s(X23,X38) = s(X23,happ(s(fun('type.list.list'(X9),X23),X39),s('type.list.list'(X9),'const.list.NIL_0'))) & ! [X40,X41] : s(X23,happ(s(fun('type.list.list'(X9),X23),X39),s('type.list.list'(X9),'const.list.CONS_2'(s(X9,X40),s('type.list.list'(X9),X41))))) = s(X23,happ(s(fun(X23,X23),happ(s(fun('type.list.list'(X9),fun(X23,X23)),happ(s(fun(X9,fun('type.list.list'(X9),fun(X23,X23))),X11),s(X9,X40))),s('type.list.list'(X9),X41))),s(X23,happ(s(fun('type.list.list'(X9),X23),X39),s('type.list.list'(X9),X41))))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.list.list_Axiom')).
fof(f360,plain,(
  ( ! [X6,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK16(X3,X2,X1,X0)),s('type.list.list'(X1),X6)))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f354,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X23,X9] : ! [X10,X37] : (s(fun(X9,X23),X10) = s(fun(X9,X23),X37) <=> ! [X24] : s(X23,happ(s(fun(X9,X23),X10),s(X9,X24))) = s(X23,happ(s(fun(X9,X23),X37),s(X9,X24))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bool.FUN_EQ_THM')).
fof(f355,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK13(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK13(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f244])).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.bvlSem.v_|49||49||split|1',aHL_EXT)).
fof(f344,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9,X23] : ! [X24,X25] : s(X23,X25) = s(X23,'const.pair.SND_1'(s('type.pair.prod'(X9,X23),'const.pair.,_2'(s(X9,X24),s(X23,X25)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.pair.SND')).
fof(f342,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK7(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X23,X9] : ! [X27] : ? [X10] : ! [X28,X29,X30] : s(X23,happ(s(fun('type.ind_type.recspace'(X9),X23),X10),s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X28),s(X9,X29),s(fun('type.num.num','type.ind_type.recspace'(X9)),X30))))) = s(X23,happ(s(fun(fun('type.num.num',X23),X23),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X23),X23)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X23),X23))),happ(s(fun('type.num.num',fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),fun(fun('type.num.num',X23),X23)))),X27),s('type.num.num',X28))),s(X9,X29))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X30))),s(fun('type.num.num',X23),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X9),X23),X10),s(fun('type.num.num','type.ind_type.recspace'(X9)),X30)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.CONSTR_REC')).
fof(f337,plain,(
  ( ! [X0,X1] : (s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) & (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ast.lit_|49||49||split|4')).
fof(f338,plain,(
  ( ! [X0,X1] : (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.ast.lit','const.ast.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))) )),
  inference(cnf_transformation,[],[f236])).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) & (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  file('thm.bvlSem.v_|49||49||split|1','thm.closSem.v_|49||49||split|1')).
fof(f336,plain,(
  ( ! [X0,X1] : (s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1) | s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.closSem.v','const.closSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))) )),
  inference(cnf_transformation,[],[f235])).
fof(f333,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) != s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X7))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_|49||49||split|0')).
fof(f334,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f234])).
fof(f329,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f328,plain,(
  ( ! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.bvlSem.v'),X1)))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X0),s('type.list.list'('type.bvlSem.v'),X1))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X0))),s('type.list.list'('type.bvlSem.v'),X1)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8,X15] : s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X8),s('type.list.list'('type.bvlSem.v'),X15))) = s('type.bvlSem.v',happ(s(fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v'),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),'type.bvlSem.v')),'const.bvlSem.Block_0'),s('type.num.num',X8))),s('type.list.list'('type.bvlSem.v'),X15)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_2_const.bvlSem.Block')).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.bvlSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),X2))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.bvlSem.v'),X2))),s(fun('type.integer.int',X0),X3),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X4),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s(fun('type.num.num',X0),X6),s(fun('type.num.num',X0),X7)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.bvlSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),X2))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.bvlSem.v'),X2))),s(fun('type.integer.int',X0),X3),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X4),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s(fun('type.num.num',X0),X6),s(fun('type.num.num',X0),X7)))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.bvlSem.v'),X0),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s('type.num.num',X1))),s('type.list.list'('type.bvlSem.v'),X2))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X1),s('type.list.list'('type.bvlSem.v'),X2))),s(fun('type.integer.int',X0),X3),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X4),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X5),s(fun('type.num.num',X0),X6),s(fun('type.num.num',X0),X7)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X9] : ! [X40,X41,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.list.list'('type.bvlSem.v'),X9),happ(s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s('type.num.num',X40))),s('type.list.list'('type.bvlSem.v'),X41))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Block_2'(s('type.num.num',X40),s('type.list.list'('type.bvlSem.v'),X41))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_case_def|split|2')).
fof(f325,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X6),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.num.num',X9),X14),s('type.num.num',X6))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_case_def|split|4')).
fof(f324,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.num.num',X9),X13),s('type.num.num',X6))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_case_def|split|3')).
fof(f323,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.integer.int',X0),X2),s('type.integer.int',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.integer.int',X0),X2),s('type.integer.int',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.integer.int',X0),X2),s('type.integer.int',X1))) = s(X0,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X1))),s(fun('type.integer.int',X0),X2),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0),X3),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X0)),X4),s(fun('type.num.num',X0),X5),s(fun('type.num.num',X0),X6)))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14] : s(X9,happ(s(fun('type.integer.int',X9),X10),s('type.integer.int',X6))) = s(X9,'const.bvlSem.v_CASE_6'(s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X6))),s(fun('type.integer.int',X9),X10),s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X9),X11),s(fun('type.num.num',fun('type.list.list'('type.bvlSem.v'),X9)),X12),s(fun('type.num.num',X9),X13),s(fun('type.num.num',X9),X14)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_case_def|split|0')).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X8,X15,X32] : s(X9,'const.ind_type.FCONS_3'(s(X9,X8),s(fun('type.num.num',X9),X15),s('type.num.num',X32))) = s(X9,happ(s(fun('type.num.num',X9),happ(s(fun(fun('type.num.num',X9),fun('type.num.num',X9)),happ(s(fun(X9,fun(fun('type.num.num',X9),fun('type.num.num',X9))),'const.ind_type.FCONS_0'),s(X9,X8))),s(fun('type.num.num',X9),X15))),s('type.num.num',X32)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_3_const.ind_type.FCONS')).
fof(f321,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : ! [X8,X15,X32] : s('type.ind_type.recspace'(X9),'const.ind_type.CONSTR_3'(s('type.num.num',X8),s(X9,X15),s(fun('type.num.num','type.ind_type.recspace'(X9)),X32))) = s('type.ind_type.recspace'(X9),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9)),happ(s(fun(X9,fun(fun('type.num.num','type.ind_type.recspace'(X9)),'type.ind_type.recspace'(X9))),'const.ind_type.CONSTR_1'(s('type.num.num',X8))),s(X9,X15))),s(fun('type.num.num','type.ind_type.recspace'(X9)),X32)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_1_3_const.ind_type.CONSTR')).
fof(f320,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : ! [X6,X10,X31] : s(X9,happ(s(fun('type.num.num',X9),X10),s('type.num.num',X31))) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X31)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.FCONS|split|1')).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X10] : s(X9,X6) = s(X9,'const.ind_type.FCONS_3'(s(X9,X6),s(fun('type.num.num',X9),X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.FCONS|split|0')).
fof(f272,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f265,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8))) = s('type.bvlSem.v',happ(s(fun('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),'type.bvlSem.v'),'const.bvlSem.Word64_0'),s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X8)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_1_const.bvlSem.Word64')).
fof(f264,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6)))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_size_def|split|1')).
fof(f263,plain,(
  ( ! [X0] : (s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X6] : s('type.num.num','const.bvlSem.v_size_1'(s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X6))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X6)))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_size_def|split|3')).
fof(f262,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X0))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.Number_1'(s('type.integer.int',X8))) = s('type.bvlSem.v',happ(s(fun('type.integer.int','type.bvlSem.v'),'const.bvlSem.Number_0'),s('type.integer.int',X8)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_1_const.bvlSem.Number')).
fof(f261,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.CodePtr_1'(s('type.num.num',X8))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.CodePtr_0'),s('type.num.num',X8)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_1_const.bvlSem.CodePtr')).
fof(f260,plain,(
  ( ! [X0] : (s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X0))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : s('type.bvlSem.v','const.bvlSem.RefPtr_1'(s('type.num.num',X8))) = s('type.bvlSem.v',happ(s(fun('type.num.num','type.bvlSem.v'),'const.bvlSem.RefPtr_0'),s('type.num.num',X8)))),
  file('thm.bvlSem.v_|49||49||split|1','ARITY_0_1_const.bvlSem.RefPtr')).
fof(f259,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X9] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X9,X9),'_i_0_0'),s(fun(X9,X9),'_i_0_0'))))),
  file('thm.bvlSem.v_|49||49||split|1','thm.ind_type.ISO_REFL')).
fof(f11051,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042,f11044,f11046,f11048,f11050])).
fof(f11049,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042,f11044,f11046,f11048])).
fof(f11047,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042,f11044,f11046])).
fof(f11045,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042,f11044])).
fof(f11043,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040,f11042])).
fof(f11041,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038,f11040])).
fof(f11039,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036,f11038])).
fof(f11037,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034,f11036])).
fof(f11035,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032,f11034])).
fof(f11033,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030,f11032])).
fof(f11031,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028,f11030])).
fof(f11029,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026,f11028])).
fof(f11027,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024,f11026])).
fof(f11025,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022,f11024])).
fof(f11023,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018,f11022])).
fof(f11019,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016,f11018])).
fof(f11017,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014,f11016])).
fof(f11015,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012,f11014])).
fof(f11013,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796,f11012])).
fof(f10811,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10476,f10179])).
fof(f10476,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,'T_0')))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f7110])).
fof(f10810,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10474,f10179])).
fof(f10474,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0')))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f6883])).
fof(f10809,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10808,f256])).
fof(f10808,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(forward_demodulation,[],[f10807,f10179])).
fof(f10807,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10434,f10179])).
fof(f10434,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X116,bool),X117),s(X116,happ(s(fun(X115,X116),X118),s(X115,X119))))) | ~p(s(bool,happ(s(fun(X115,bool),'_dst_x0x1_2'(s(fun(X116,bool),X117),s(fun(X115,X116),X118))),s(X115,X119))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f5073])).
fof(f10806,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10805,f256])).
fof(f10805,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(forward_demodulation,[],[f10804,f10179])).
fof(f10804,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl26),
  inference(subsumption_resolution,[],[f10426,f10179])).
fof(f10426,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,'T_0') != s(bool,happ(s(fun(X120,bool),'_dst_x0x1_2'(s(fun(X118,bool),X119),s(fun(X120,X118),X121))),s(X120,X122))) | ~p(s(bool,happ(s(fun(X118,bool),X119),s(X118,happ(s(fun(X120,X118),X121),s(X120,X122))))))) ) | $spl26),
  inference(backward_demodulation,[],[f10179,f4995])).
fof(f10797,plain,(
  $false | $spl26),
  inference(global_subsumption,[],[f255,f254,f256,f258,f259,f260,f261,f262,f263,f264,f265,f266,f268,f267,f272,f271,f305,f312,f315,f320,f321,f322,f323,f324,f325,f326,f327,f328,f332,f329,f334,f333,f336,f335,f338,f337,f339,f340,f341,f342,f343,f344,f353,f355,f354,f360,f359,f367,f368,f406,f393,f420,f428,f433,f435,f460,f482,f501,f528,f628,f684,f744,f805,f819,f839,f840,f841,f842,f843,f844,f897,f923,f955,f956,f957,f958,f959,f960,f961,f978,f988,f989,f990,f991,f992,f993,f994,f1346,f1347,f1348,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1362,f1364,f1366,f1380,f1372,f1485,f1521,f1508,f1992,f1993,f1994,f1995,f1996,f1997,f1999,f2001,f2002,f2062,f2068,f2293,f2271,f2328,f2333,f2335,f2571,f2573,f4049,f4051,f4052,f4054,f4055,f4056,f4135,f4224,f4226,f4227,f4229,f4230,f4231,f4311,f4317,f4318,f4319,f4379,f4381,f4382,f4401,f4416,f4482,f4640,f4742,f4736,f4865,f4870,f4869,f4871,f4872,f4873,f4874,f4875,f4876,f4877,f4878,f4879,f4889,f5029,f5034,f5048,f5050,f5066,f5106,f5111,f5125,f5127,f5182,f5113,f5115,f5036,f5038,f5114,f5116,f5037,f5039,f4537,f4189,f394,f358,f357,f356,f352,f351,f4225,f4050,f2070,f1947,f1476,f1474,f1472,f4613,f5183,f5556,f4743,f4612,f4654,f4331,f4687,f4334,f4338,f4337,f4336,f4335,f4333,f4332,f4330,f4329,f2047,f2434,f5110,f5033,f5105,f5107,f5028,f5030,f2437,f4633,f4479,f2067,f317,f316,f313,f2440,f2439,f2438,f2436,f2435,f2433,f2432,f408,f307,f4228,f4053,f2257,f2231,f2229,f2577,f2579,f349,f348,f5087,f5086,f5084,f5067,f350,f347,f5068,f5091,f5090,f5085,f5083,f5078,f5077,f5075,f5074,f5073,f5072,f5071,f5013,f5012,f5009,f5008,f5007,f5006,f5005,f5000,f4999,f4997,f4996,f4995,f4994,f4993,f4991,f4990,f4712,f5069,f4414,f4404,f4405,f4406,f4493,f4503,f4403,f4400,f4552,f4386,f4385,f4380,f4378,f4373,f4372,f4370,f4369,f4368,f4367,f4366,f4402,f817,f816,f815,f2264,f814,f813,f1503,f918,f4690,f812,f5109,f5032,f4481,f4480,f314,f6057,f6067,f6068,f6072,f6075,f6076,f6080,f6045,f6046,f6047,f6048,f6049,f6050,f6051,f6052,f6053,f6054,f6055,f6056,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6572,f6656,f6664,f6773,f6774,f6775,f6776,f6777,f6778,f6779,f6780,f6783,f6868,f6876,f6921,f6922,f6923,f6924,f6925,f6926,f6927,f6928,f6929,f6930,f6931,f6978,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7290,f7291,f7197,f7198,f7292,f7293,f7294,f7295,f7296,f7297,f7298,f7299,f7207,f7301,f7303,f7309,f7278,f7279,f7280,f7310,f7537,f7538,f7496,f7497,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7506,f7507,f7508,f7509,f7552,f7553,f7561,f7535,f7562,f7805,f7806,f7721,f7722,f7807,f7808,f7809,f7810,f7811,f7812,f7813,f7814,f7731,f7732,f7733,f7734,f7735,f7736,f7737,f7738,f7749,f7750,f7751,f7752,f7753,f7754,f7819,f7820,f7757,f7758,f7821,f7822,f7823,f7824,f7825,f7826,f7827,f7828,f7767,f7768,f7769,f7770,f7781,f7782,f7783,f7784,f7785,f7786,f7787,f7798,f7799,f7800,f7801,f7802,f7803,f7992,f7993,f8002,f8005,f8006,f8015,f8050,f8051,f8533,f8534,f8543,f8546,f8547,f8556,f8592,f9078,f9079,f9080,f9081,f9082,f9083,f9084,f9085,f9086,f9087,f9088,f9089,f9090,f9120,f9121,f9093,f9094,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9103,f9131,f9132,f9332,f9333,f9334,f9335,f9336,f9337,f9338,f9339,f9340,f9341,f9342,f9343,f9344,f9374,f9375,f9347,f9348,f9376,f9377,f9378,f9379,f9380,f9381,f9382,f9383,f9357,f9385,f9386,f9675,f9676,f9677,f9678,f9679,f9680,f9681,f9682,f9683,f9684,f9685,f9686,f9687,f9717,f9718,f9690,f9691,f9719,f9720,f9721,f9722,f9723,f9724,f9725,f9726,f9700,f9728,f9729,f10290,f10291,f7185,f7110,f10294,f10295,f6957,f6883,f6011,f6010,f10307,f5944,f5967,f4611,f4631,f2581,f2575,f2243,f10321,f10179,f10796])).
fof(f10795,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f10323,f256])).
fof(f10323,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(backward_demodulation,[],[f10179,f258])).
fof(f10322,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl38),
  inference(cnf_transformation,[],[f10322_D])).
fof(f10322_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1382,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f1382_D])).
fof(f1382_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f10321,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1958,f256])).
fof(f1958,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f315,f460])).
fof(f4631,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4598,f256])).
fof(f4598,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4479,f460])).
fof(f4611,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4479,f305])).
fof(f5967,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f5931,f256])).
fof(f5931,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4479,f460])).
fof(f5944,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f4479,f305])).
fof(f10305,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8))))) ) | $spl34),
  inference(cnf_transformation,[],[f10305_D])).
fof(f10305_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f10306,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) | $spl36),
  inference(cnf_transformation,[],[f10306_D])).
fof(f10306_D,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6010,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),sK2(X6,X5))))) )),
  inference(subsumption_resolution,[],[f5989,f267])).
fof(f5989,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | ~p(s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X5),bool),X6),s('type.list.list'(X5),sK2(X6,X5))))) )),
  inference(superposition,[],[f408,f433])).
fof(f6011,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8))))) )),
  inference(subsumption_resolution,[],[f5990,f267])).
fof(f5990,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),'const.list.NIL_0')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(X8),bool),X9),s('type.list.list'(X8),sK2(X9,X8)))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f408,f433])).
fof(f10253,plain,(
  ( ! [X73] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X73,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl32),
  inference(cnf_transformation,[],[f10253_D])).
fof(f10253_D,plain,(
  ( ! [X73] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X73,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6957,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f6886])).
fof(f6886,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6656,f266])).
fof(f10297,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X22,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X23)) ) | ($spl10 | $spl22)),
  inference(subsumption_resolution,[],[f6894,f9004])).
fof(f6894,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X22,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X21) | s(bool,X21) = s(bool,X23) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f6656,f2604])).
fof(f10296,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X32)) ) | ($spl10 | $spl22)),
  inference(subsumption_resolution,[],[f6897,f9004])).
fof(f6897,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X31,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6656,f2604])).
fof(f10295,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f6911,f6656])).
fof(f6911,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X67,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X65,bool),'const.sum.INR_1'(s(bool,X66))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f6656,f1996])).
fof(f10294,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f6916,f6868])).
fof(f6916,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X78,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6656,f1997])).
fof(f10251,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl30),
  inference(cnf_transformation,[],[f10251_D])).
fof(f10251_D,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f7185,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7114])).
fof(f7114,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6868,f266])).
fof(f10293,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X22),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,X23)) ) | ($spl10 | $spl22)),
  inference(subsumption_resolution,[],[f7122,f9004])).
fof(f7122,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X22),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X21) | s(bool,X21) = s(bool,X23) | s(bool,'F_0') = s(bool,X23)) ) | $spl10),
  inference(superposition,[],[f6868,f2604])).
fof(f10292,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X32)) ) | ($spl10 | $spl22)),
  inference(subsumption_resolution,[],[f7125,f9004])).
fof(f7125,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X31),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X30) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f6868,f2604])).
fof(f10291,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X67),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f7139,f6656])).
fof(f7139,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X67),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X65,bool),'const.sum.INR_1'(s(bool,X66))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f6868,f1996])).
fof(f10290,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f7144,f6868])).
fof(f7144,plain,(
  ( ! [X78,X76,X77] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X78),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X76),'const.sum.INL_1'(s(bool,X77))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f6868,f1997])).
fof(f10252,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X73,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl10),
  inference(subsumption_resolution,[],[f9508,f9561])).
fof(f9508,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X73,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl10),
  inference(superposition,[],[f6656,f2716])).
fof(f10250,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X74)) ) | $spl10),
  inference(subsumption_resolution,[],[f9509,f9561])).
fof(f9509,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X75),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | $spl10),
  inference(superposition,[],[f6868,f2716])).
fof(f10111,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f9947])).
fof(f9947,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) ) | $spl6),
  inference(superposition,[],[f2400,f7069])).
fof(f7069,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f7055])).
fof(f7055,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f6958,f433])).
fof(f6958,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0')))))) ) | $spl6),
  inference(forward_demodulation,[],[f6884,f1397])).
fof(f1397,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1383,f272])).
fof(f6884,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,sK4(X1,'_i_0_0',bool))))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1398,f6656])).
fof(f1398,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1383,f267])).
fof(f2400,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(subsumption_resolution,[],[f2399,f258])).
fof(f2399,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(forward_demodulation,[],[f2365,f1397])).
fof(f2365,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK4(X1,'_i_0_0',bool))) | s(bool,X0) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f2001,f1421])).
fof(f1421,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f1399,f256])).
fof(f1399,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,sK4(X0,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X1)) ) | $spl6),
  inference(superposition,[],[f1383,f433])).
fof(f7475,plain,(
  p(s(bool,sK6('_i_0_0','_i_0_0',bool))) | p(s(bool,sK5('_i_0_0','_i_0_0',bool))) | $spl20),
  inference(superposition,[],[f7302,f305])).
fof(f7302,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) | $spl20),
  inference(cnf_transformation,[],[f7302_D])).
fof(f7302_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,sK5('_i_0_0',X1,bool)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7309,plain,(
  ( ! [X64,X65] : (p(s(bool,sK5('_i_0_0',X65,bool))) | ~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool)))))) )),
  inference(forward_demodulation,[],[f7277,f305])).
fof(f7277,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),X65),s(bool,sK6('_i_0_0',X65,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X65,bool)))))) )),
  inference(superposition,[],[f316,f305])).
fof(f7301,plain,(
  ( ! [X0,X1] : (p(s(bool,sK5('_i_0_0',X1,bool))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f7300,f305])).
fof(f7300,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7263,f256])).
fof(f7263,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK6('_i_0_0',X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5('_i_0_0',X1,bool))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f460])).
fof(f6078,plain,(
  ( ! [X92,X93] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK16(X93,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X92,bool))))) ) | $spl18),
  inference(cnf_transformation,[],[f6078_D])).
fof(f6078_D,plain,(
  ( ! [X92,X93] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK16(X93,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X92,bool))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6077,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK16(X93,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X92,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f6040,f271])).
fof(f6040,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,X91)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X92),bool),sK16(X93,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')),X92,bool)))) | s(bool,'F_0') = s(bool,X91)) ) | $spl10),
  inference(superposition,[],[f4654,f2336])).
fof(f6070,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X73,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f6070_D])).
fof(f6070_D,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X73,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f6069,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X73,bool))))) ) | $spl6),
  inference(subsumption_resolution,[],[f6034,f271])).
fof(f6034,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,X72)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.list.list'(X73),bool),sK16(X74,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')),X73,bool)))) | s(bool,'F_0') = s(bool,X72)) ) | $spl6),
  inference(superposition,[],[f4654,f2334])).
fof(f2334,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X27)) ) | $spl6),
  inference(subsumption_resolution,[],[f2311,f256])).
fof(f2311,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X27)) ) | $spl6),
  inference(superposition,[],[f1999,f1598])).
fof(f1598,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl6),
  inference(unit_resulting_resolution,[],[f1545,f268])).
fof(f1545,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl6),
  inference(unit_resulting_resolution,[],[f1398,f314])).
fof(f5555,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5201,f258])).
fof(f5201,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f420,f2069,f329])).
fof(f5554,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5203,f2069])).
fof(f5203,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1398,f2069,f329])).
fof(f5553,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5211,f258])).
fof(f5211,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(unit_resulting_resolution,[],[f420,f2069,f329])).
fof(f5552,plain,(
  $false | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5213,f2069])).
fof(f5213,plain,(
  ( ! [X0] : (p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1398,f2069,f329])).
fof(f5547,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f5502,f258])).
fof(f5502,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(backward_demodulation,[],[f5443,f4743])).
fof(f5443,plain,(
  ( ! [X227] : (s(bool,'F_0') = s(bool,X227)) ) | $spl12),
  inference(subsumption_resolution,[],[f5442,f5441])).
fof(f5441,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f5440,f2069])).
fof(f5440,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) ) | $spl12),
  inference(subsumption_resolution,[],[f5408,f258])).
fof(f5408,plain,(
  ( ! [X132,X130,X128,X133,X131,X129] : (p(s(bool,'F_0')) | p(s(bool,X133)) | s(bool,X133) = s(bool,happ(s(fun(X128,bool),X129),s(X128,happ(s(fun(X130,X128),X131),s(X130,X132)))))) ) | $spl12),
  inference(backward_demodulation,[],[f5294,f4997])).
fof(f5294,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(X1,bool),X2),s(fun(X0,X1),X3))),s(X0,X4)))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f368,f2069,f2001])).
fof(f5442,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (s(bool,X227) != s(bool,happ(s(fun(X222,bool),X223),s(X222,happ(s(fun(X224,X222),X225),s(X224,X226))))) | s(bool,'F_0') = s(bool,X227)) ) | $spl12),
  inference(subsumption_resolution,[],[f5417,f258])).
fof(f5417,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (p(s(bool,'F_0')) | s(bool,X227) != s(bool,happ(s(fun(X222,bool),X223),s(X222,happ(s(fun(X224,X222),X225),s(X224,X226))))) | s(bool,'F_0') = s(bool,X227)) ) | $spl12),
  inference(backward_demodulation,[],[f5294,f5013])).
fof(f5302,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f2069,f4612])).
fof(f5300,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1398,f2069,f314])).
fof(f5296,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f428,f420,f2069,f2001])).
fof(f5264,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f1398,f1398,f2069,f482])).
fof(f5262,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f420,f420,f2069,f482])).
fof(f5238,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f339,f2069,f406])).
fof(f5236,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f340,f2069,f406])).
fof(f5234,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f343,f2069,f406])).
fof(f5232,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f312,f2069,f406])).
fof(f5230,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f344,f2069,f406])).
fof(f5228,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f359,f2069,f406])).
fof(f5226,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f305,f2069,f406])).
fof(f5223,plain,(
  $false | ($spl10 | $spl12)),
  inference(unit_resulting_resolution,[],[f256,f2093,f2069,f406])).
fof(f5221,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f4742,f2069,f406])).
fof(f5219,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f256,f1598,f2069,f406])).
fof(f5217,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f428,f2069,f406])).
fof(f5214,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2069,f1398,f2069,f329])).
fof(f5212,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f420,f2069,f329])).
fof(f5210,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2069,f1398,f2069,f329])).
fof(f5208,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2069,f420,f2069,f329])).
fof(f5204,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2069,f1398,f2069,f329])).
fof(f5202,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f258,f420,f2069,f329])).
fof(f5200,plain,(
  $false | ($spl6 | $spl12)),
  inference(unit_resulting_resolution,[],[f2069,f1398,f2069,f329])).
fof(f5198,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f2069,f420,f2069,f329])).
fof(f5194,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f420,f2069,f272])).
fof(f5192,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f428,f2069,f267])).
fof(f5190,plain,(
  $false | $spl12),
  inference(unit_resulting_resolution,[],[f256,f2069])).
fof(f5189,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) | $spl14),
  inference(cnf_transformation,[],[f5189_D])).
fof(f5189_D,plain,(
  ( ! [X33] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5182,plain,(
  ( ! [X33,X34] : (~p(s(bool,X34)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X33),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f4613,f343])).
fof(f2065,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl10),
  inference(cnf_transformation,[],[f2065_D])).
fof(f2065_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2068,plain,(
  ( ! [X24] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f2046,f305])).
fof(f2046,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f313,f305])).
fof(f2062,plain,(
  ( ! [X0] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2061,f305])).
fof(f2061,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2039,f256])).
fof(f2039,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f313,f460])).
fof(f1383,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f1383_D])).
fof(f1383_D,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1372,plain,(
  ( ! [X19] : (~p(s(bool,sK4(X19,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f315,f305])).
fof(f419,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f418])).
fof(f418,plain,(
  s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f415,f416])).
fof(f414,plain,(
  s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) | $spl0),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f416,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | $spl2),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f254,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1)))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1)) & (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f199])).
fof(f199,plain,(
  ? [X0,X1] : ((s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)) & (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) | s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1)))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ? [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <~> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X0) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6))) = s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))) <=> s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X6) = s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),X7))),
  file('thm.bvlSem.v_|49||49||split|1','thm.bvlSem.v_|49||49||split|1')).
fof(f415,plain,(
  s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) | $spl1),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f417,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | $spl3),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f255,plain,(
  s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0) != s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1) | s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK0))) != s('type.bvlSem.v','const.bvlSem.Word64_1'(s('type.fcp.cart'(bool,'type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.fcp.bit|48|'('type.one.one'))))))),sK1)))),
  inference(cnf_transformation,[],[f200])).
% SZS output end Proof for thm.bvlSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 215
% Passive clauses: 2524
% Generated clauses: 27787
% Final active clauses: 140
% Final passive clauses: 1471
% Input formulas: 69
% Initial clauses: 141
% 
% Duplicate literals: 958
% Trivial inequalities: 221
% Fw subsumption resolutions: 1386
% Fw demodulations: 4244
% Bw demodulations: 647
% Condensations: 1253
% Global subsumptions: 23
% 
% Simple tautologies: 334
% Equational tautologies: 158
% Forward subsumptions: 16439
% Fw demodulations to eq. taut.: 147
% Bw demodulations to eq. taut.: 69
% 
% Unit resulting resolution: 1504
% Factoring: 1
% Forward superposition: 8566
% Backward superposition: 7356
% Self superposition: 190
% Equality factoring: 88
% Equality resolution: 33
% 
% Split clauses: 39
% Split components: 82
% SAT solver clauses: 2688
% SAT solver unit clauses: 175
% SAT solver binary clauses: 949
% 
% Sat splits: 39
% Sat splitting refutations: 61
% 
% Memory used [KB]: 30319
% Time elapsed: 1.477 s
% ------------------------------
% ------------------------------
% Success in time 2.84 s
