Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exhSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exhSem
% SZS output start Proof for thm.exhSem
fof(f4248,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4247,f152,f254,f252,f4187,f151,f253,f251])).
fof(f251,plain,(
  s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) | $spl0),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f253,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) = s('type.list.list'('type.exhSem.v'),sK1) | $spl2),
  inference(cnf_transformation,[],[f253_D])).
fof(f253_D,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) = s('type.list.list'('type.exhSem.v'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f151,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) = s('type.list.list'('type.exhSem.v'),sK1) | s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1)))),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) | s('type.list.list'('type.exhSem.v'),sK0) = s('type.list.list'('type.exhSem.v'),sK1)) & (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) | s('type.list.list'('type.exhSem.v'),sK0) != s('type.list.list'('type.exhSem.v'),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f110])).
fof(f110,plain,(
  ? [X0,X1] : ((s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))) | s('type.list.list'('type.exhSem.v'),X0) = s('type.list.list'('type.exhSem.v'),X1)) & (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))) | s('type.list.list'('type.exhSem.v'),X0) != s('type.list.list'('type.exhSem.v'),X1)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ? [X0,X1] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))) <~> s('type.list.list'('type.exhSem.v'),X0) = s('type.list.list'('type.exhSem.v'),X1))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X0))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))) <=> s('type.list.list'('type.exhSem.v'),X0) = s('type.list.list'('type.exhSem.v'),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X7))) <=> s('type.list.list'('type.exhSem.v'),X6) = s('type.list.list'('type.exhSem.v'),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))) = s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X7))) <=> s('type.list.list'('type.exhSem.v'),X6) = s('type.list.list'('type.exhSem.v'),X7))),
  file('thm.exhSem.v_|49||49||split|5','thm.exhSem.v_|49||49||split|5')).
fof(f4187,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4186,f254])).
fof(f4186,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) = s('type.list.list'('type.exhSem.v'),sK1) | $spl0),
  inference(forward_demodulation,[],[f4088,f195])).
fof(f195,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X28] : ! [X30] : s(X28,X30) = s(X28,happ(s(fun(X28,X28),'_i_0_0'),s(X28,X30)))),
  file('thm.exhSem.v_|49||49||split|5','_i0_ax')).
fof(f4088,plain,(
  s('type.list.list'('type.exhSem.v'),sK1) = s('type.list.list'('type.exhSem.v'),happ(s(fun('type.list.list'('type.exhSem.v'),'type.list.list'('type.exhSem.v')),'_i_0_0'),s('type.list.list'('type.exhSem.v'),sK0))) | $spl0),
  inference(superposition,[],[f3278,f195])).
fof(f3278,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X1),s('type.list.list'('type.exhSem.v'),sK0))) = s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X1),s('type.list.list'('type.exhSem.v'),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f3164,f207])).
fof(f207,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X7),s('type.list.list'('type.exhSem.v'),X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X7),s('type.list.list'('type.exhSem.v'),X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X7),s('type.list.list'('type.exhSem.v'),X1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X1))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X7)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9] : ! [X6,X10,X11,X12,X13,X14,X15] : s(X9,happ(s(fun('type.list.list'('type.exhSem.v'),X9),X15),s('type.list.list'('type.exhSem.v'),X6))) = s(X9,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),X6))),s(fun('type.ast.lit',X9),X10),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X9)),X11),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X9))),X12),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X9))),X13),s(fun('type.num.num',X9),X14),s(fun('type.list.list'('type.exhSem.v'),X9),X15)))),
  file('thm.exhSem.v_|49||49||split|5','thm.exhSem.v_case_def|split|5')).
fof(f3164,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'('type.exhSem.v'),X0),X1),s('type.list.list'('type.exhSem.v'),sK1))) = s(X0,'const.exhSem.v_CASE_7'(s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))),s(fun('type.ast.lit',X0),X2),s(fun('type.num.num',fun('type.list.list'('type.exhSem.v'),X0)),X3),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.string.char'),fun('type.exhLang.exp',X0))),X4),s(fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.exhSem.v')),fun('type.list.list'('type.pair.prod'('type.list.list'('type.string.char'),'type.pair.prod'('type.list.list'('type.string.char'),'type.exhLang.exp'))),fun('type.list.list'('type.string.char'),X0))),X5),s(fun('type.num.num',X0),X6),s(fun('type.list.list'('type.exhSem.v'),X0),X1)))) ) | $spl0),
  inference(superposition,[],[f207,f251])).
fof(f252,plain,(
  s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) | $spl1),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f254,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) != s('type.list.list'('type.exhSem.v'),sK1) | $spl3),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) != s('type.list.list'('type.exhSem.v'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f152,plain,(
  s('type.list.list'('type.exhSem.v'),sK0) != s('type.list.list'('type.exhSem.v'),sK1) | s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK1)))),
  inference(cnf_transformation,[],[f111])).
fof(f4247,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f4246])).
fof(f4246,plain,(
  s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) != s('type.exhSem.v','const.exhSem.Vectorv_1'(s('type.list.list'('type.exhSem.v'),sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f252,f253])).
% SZS output end Proof for thm.exhSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 67
% Passive clauses: 894
% Generated clauses: 4025
% Final active clauses: 60
% Final passive clauses: 709
% Input formulas: 37
% Initial clauses: 97
% 
% Duplicate literals: 207
% Trivial inequalities: 80
% Fw subsumption resolutions: 449
% Fw demodulations: 439
% Bw demodulations: 4
% 
% Simple tautologies: 155
% Equational tautologies: 36
% Forward subsumptions: 1685
% 
% Binary resolution: 149
% Factoring: 1
% Forward superposition: 1164
% Backward superposition: 1248
% Self superposition: 79
% Equality factoring: 58
% Equality resolution: 15
% 
% Split clauses: 84
% Split components: 174
% SAT solver clauses: 155
% SAT solver unit clauses: 7
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 84
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 153
% TWLsolver calls for satisfiability: 51
% 
% Memory used [KB]: 3965
% Time elapsed: 0.109 s
% ------------------------------
% ------------------------------
% Success in time 0.113 s
