Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ml_optimise
% (27799)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 163
% Passive clauses: 1910
% Generated clauses: 20267
% Final active clauses: 154
% Final passive clauses: 1530
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 1355
% Trivial inequalities: 145
% Fw subsumption resolutions: 2784
% Fw demodulations: 3301
% Bw demodulations: 7
% 
% Simple tautologies: 322
% Equational tautologies: 32
% Forward subsumptions: 9160
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 357
% Forward superposition: 7233
% Backward superposition: 4748
% Self superposition: 102
% Equality factoring: 42
% Equality resolution: 42
% 
% Split clauses: 698
% Split components: 1419
% SAT solver clauses: 1425
% SAT solver unit clauses: 7
% SAT solver binary clauses: 65
% 
% Sat splits: 698
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1414
% TWLsolver calls for satisfiability: 197
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ml_optimise
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ml_optimise
% SZS output start Proof for thm.ml_optimise
fof(f4295,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4294,f82,f125,f4182,f4217,f115,f148,f4169,f4173,f2413,f2441,f120,f152,f2319,f2316,f83,f131,f163,f165,f80,f129,f87,f135,f737,f738,f111,f143,f159,f161,f81,f127])).
fof(f127,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))))) | $spl2),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f81,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))))))) & p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))))) & s('type.semanticPrimitives.v',sK3) != s('type.semanticPrimitives.v',sK4)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) & p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & s('type.semanticPrimitives.v',X3) != s('type.semanticPrimitives.v',X4))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) & (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) & s('type.semanticPrimitives.v',X3) != s('type.semanticPrimitives.v',X4)))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) => (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) => s('type.semanticPrimitives.v',X3) = s('type.semanticPrimitives.v',X4)))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4)))))))) => (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X6),s('type.semanticPrimitives.state'(X0),X2),s('type.ast.exp',X5),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X0),X1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X3)))))))) => s('type.semanticPrimitives.v',X3) = s('type.semanticPrimitives.v',X4)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X6),X8),s('type.ast.exp',X11),s('type.pair.prod'('type.semanticPrimitives.state'(X6),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X6),X7),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X10)))))))) => (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X6),X8),s('type.ast.exp',X11),s('type.pair.prod'('type.semanticPrimitives.state'(X6),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X6),X7),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X9)))))))) => s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X10)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X6),X8),s('type.ast.exp',X11),s('type.pair.prod'('type.semanticPrimitives.state'(X6),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X6),X7),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X10)))))))) => (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X6),X8),s('type.ast.exp',X11),s('type.pair.prod'('type.semanticPrimitives.state'(X6),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(X6),X7),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X9)))))))) => s('type.semanticPrimitives.v',X9) = s('type.semanticPrimitives.v',X10)))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb','thm.ml_optimise.ml_optimise_3|48|_prove_ltb')).
fof(f161,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))))) | $spl42),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f159,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))))) | ($spl2 | $spl12)),
  inference(resolution,[],[f127,f135])).
fof(f143,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5)))) | s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)) ) | $spl20),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5)))) | s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f111,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5)))) | ! [X6] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) | s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5)))) => ! [X6] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) => s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5)))) => ! [X6] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X1),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X2),s('type.semanticPrimitives.state'(X0),X3),s('type.ast.exp',X4),s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))) => s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X5) = s('type.pair.prod'('type.semanticPrimitives.state'(X0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X6)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X15] : ! [X13,X12,X8,X16,X17] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X15),X8),s('type.ast.exp',X16),s('type.pair.prod'('type.semanticPrimitives.state'(X15),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X17)))) => ! [X18] : (p(s(bool,'const.bigStep.evaluate_5'(s(bool,X13),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),X12),s('type.semanticPrimitives.state'(X15),X8),s('type.ast.exp',X16),s('type.pair.prod'('type.semanticPrimitives.state'(X15),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X18)))) => s('type.pair.prod'('type.semanticPrimitives.state'(X15),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X17) = s('type.pair.prod'('type.semanticPrimitives.state'(X15),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X18)))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb','thm.determ.big_exp_determ|split|0')).
fof(f738,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0)))) | s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))) ) | $spl194),
  inference(cnf_transformation,[],[f738_D])).
fof(f738_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0)))) | s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f737,plain,(
  ( ! [X0] : (~p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0)))) | s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))) ) | ($spl6 | $spl20 | $spl42)),
  inference(subsumption_resolution,[],[f706,f131])).
fof(f706,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0)))) | s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),X0) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3)))))) ) | ($spl20 | $spl42)),
  inference(superposition,[],[f143,f161])).
fof(f135,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f87,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f129,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))))))) | $spl4),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f80,plain,(
  p(s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))))))),
  inference(cnf_transformation,[],[f59])).
fof(f165,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))))) | $spl44),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f163,plain,(
  s(bool,'T_0') = s(bool,'const.bigStep.evaluate_5'(s(bool,sK7),s('type.semanticPrimitives.environment'('type.semanticPrimitives.v'),sK6),s('type.semanticPrimitives.state'(sK0),sK2),s('type.ast.exp',sK5),s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))))) | ($spl4 | $spl12)),
  inference(resolution,[],[f129,f135])).
fof(f131,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f83,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb',aHL_TRUTH)).
fof(f2316,plain,(
  s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))) | $spl484),
  inference(cnf_transformation,[],[f2316_D])).
fof(f2316_D,plain,(
  s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f2319,plain,(
  s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))) | ($spl6 | $spl44 | $spl194)),
  inference(subsumption_resolution,[],[f2307,f131])).
fof(f2307,plain,(
  ~p(s(bool,'T_0')) | s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) = s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))))) | ($spl44 | $spl194)),
  inference(superposition,[],[f738,f165])).
fof(f152,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) ) | $spl38),
  inference(cnf_transformation,[],[f152_D])).
fof(f152_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f120,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X3) = s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X23] : ! [X14,X24,X25,X26] : (s('type.pair.prod'(X6,X23),'const.pair.,_2'(s(X6,X14),s(X23,X24))) = s('type.pair.prod'(X6,X23),'const.pair.,_2'(s(X6,X25),s(X23,X26))) <=> (s(X6,X14) = s(X6,X25) & s(X23,X24) = s(X23,X26)))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb','thm.pair.CLOSED_PAIR_EQ')).
fof(f2441,plain,(
  ( ! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),X4),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5))) != s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))) ) | $spl492),
  inference(cnf_transformation,[],[f2441_D])).
fof(f2441_D,plain,(
  ( ! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),X4),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5))) != s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f2413,plain,(
  ( ! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),X4),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5))) != s('type.pair.prod'('type.semanticPrimitives.state'(sK0),'type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v')),'const.pair.,_2'(s('type.semanticPrimitives.state'(sK0),sK1),s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))))) | s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),X5) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4)))) ) | ($spl38 | $spl484)),
  inference(superposition,[],[f152,f2316])).
fof(f4173,plain,(
  s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))) | $spl716),
  inference(cnf_transformation,[],[f4173_D])).
fof(f4173_D,plain,(
  s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))) <=> ~$spl716),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl716])])).
fof(f4169,plain,(
  s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) = s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK4))) | $spl492),
  inference(equality_resolution,[],[f2441])).
fof(f148,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) ) | $spl30),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,X3) | s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) != s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X2))) = s('type.semanticPrimitives.result'(X1,X0),'const.semanticPrimitives.Rval_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X23,X6] : ! [X25,X27] : (s('type.semanticPrimitives.result'(X6,X23),'const.semanticPrimitives.Rval_1'(s(X6,X25))) = s('type.semanticPrimitives.result'(X6,X23),'const.semanticPrimitives.Rval_1'(s(X6,X27))) <=> s(X6,X25) = s(X6,X27))),
  file('thm.ml_optimise.ml_optimise_3|48|_prove_ltb','thm.semanticPrimitives.result_|49||49||split|0')).
fof(f4217,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4))) | s('type.semanticPrimitives.v',sK4) = s('type.semanticPrimitives.v',X4)) ) | $spl724),
  inference(cnf_transformation,[],[f4217_D])).
fof(f4217_D,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4))) | s('type.semanticPrimitives.v',sK4) = s('type.semanticPrimitives.v',X4)) ) <=> ~$spl724),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl724])])).
fof(f4182,plain,(
  ( ! [X4] : (s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',sK3))) != s('type.semanticPrimitives.result'('type.semanticPrimitives.v','type.semanticPrimitives.v'),'const.semanticPrimitives.Rval_1'(s('type.semanticPrimitives.v',X4))) | s('type.semanticPrimitives.v',sK4) = s('type.semanticPrimitives.v',X4)) ) | ($spl30 | $spl716)),
  inference(superposition,[],[f148,f4173])).
fof(f125,plain,(
  s('type.semanticPrimitives.v',sK3) != s('type.semanticPrimitives.v',sK4) | $spl1),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  s('type.semanticPrimitives.v',sK3) != s('type.semanticPrimitives.v',sK4) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f82,plain,(
  s('type.semanticPrimitives.v',sK3) != s('type.semanticPrimitives.v',sK4)),
  inference(cnf_transformation,[],[f59])).
fof(f4294,plain,(
  $false | ($spl1 | $spl724)),
  inference(subsumption_resolution,[],[f4292,f125])).
fof(f4292,plain,(
  s('type.semanticPrimitives.v',sK3) = s('type.semanticPrimitives.v',sK4) | $spl724),
  inference(equality_resolution,[],[f4217])).
% SZS output end Proof for thm.ml_optimise
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 100
% Passive clauses: 398
% Generated clauses: 4139
% Final active clauses: 75
% Final passive clauses: 199
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 245
% Trivial inequalities: 57
% Fw subsumption resolutions: 205
% Fw demodulations: 231
% Bw demodulations: 116
% Condensations: 323
% Global subsumptions: 4
% 
% Simple tautologies: 112
% Equational tautologies: 28
% Forward subsumptions: 1736
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 42
% Factoring: 3
% Forward superposition: 1382
% Backward superposition: 972
% Self superposition: 37
% Equality factoring: 28
% Equality resolution: 43
% 
% Split clauses: 71
% Split components: 147
% SAT solver clauses: 837
% SAT solver unit clauses: 48
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 9
% 
% Sat splits: 71
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 820
% TWLsolver calls for satisfiability: 232
% 
% Memory used [KB]: 4349
% Time elapsed: 0.181 s
% ------------------------------
% ------------------------------
% Success in time 0.8 s
