Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% (75471)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 626
% Passive clauses: 2822
% Generated clauses: 14700
% Final active clauses: 603
% Final passive clauses: 1992
% Input formulas: 132
% Initial clauses: 313
% 
% Duplicate literals: 179
% Trivial inequalities: 32
% Fw subsumption resolutions: 1035
% Fw demodulations: 742
% Bw demodulations: 7
% 
% Simple tautologies: 182
% Equational tautologies: 23
% Forward subsumptions: 9570
% 
% Binary resolution: 4012
% Factoring: 868
% Forward superposition: 3158
% Backward superposition: 4296
% Self superposition: 27
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 160
% Split components: 322
% SAT solver clauses: 277
% SAT solver unit clauses: 15
% SAT solver binary clauses: 57
% 
% Sat splits: 160
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 277
% TWLsolver calls for satisfiability: 238
% 
% Memory used [KB]: 7036
% Time elapsed: 0.615 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f14466,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14422,f451,f790,f14399,f14400,f797,f799,f580,f861,f1028,f1033,f450,f784,f510,f841,f11160,f11161,f448,f787,f10842,f10844,f612,f892,f1009,f1012,f449,f786,f506,f839,f7221,f7241,f6967,f6969,f651,f916,f1435,f1443,f1428,f1430,f447,f788])).
fof(f788,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X8))),s(sK12,X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X8))),s('type.num.num',sK14))))) ) | $spl6),
  inference(cnf_transformation,[],[f788_D])).
fof(f788_D,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X8))),s(sK12,X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X8))),s('type.num.num',sK14))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f447,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X8))),s(sK12,X9)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X8))),s('type.num.num',sK14))))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X8))),s('type.num.num',sK14)))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X8))),s(sK12,X9))))) & ! [X6,X7] : (s('type.num.num',sK14) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X6))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X6))),s(sK12,X7))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) & ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14,sK15,sK16])],[f329])).
fof(f329,plain,(
  ? [X0,X1,X2,X3] : (! [X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X8))),s('type.num.num',X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X8))),s(X1,X9))))) & ! [X6,X7] : (s('type.num.num',X3) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X6))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X6))),s(X1,X7))))) & ? [X4,X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X4))),s('type.num.num',X3)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X4))),s(X1,X5))))))),
  inference(rectify,[],[f286])).
fof(f286,plain,(
  ? [X0,X1,X2,X3] : (! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X6))),s('type.num.num',X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X6))),s(X1,X7))))) & ! [X4,X5] : (s('type.num.num',X3) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X4))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X4))),s(X1,X5))))) & ? [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X8))),s('type.num.num',X3)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X8))),s(X1,X9))))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ? [X0,X1,X2,X3] : ((! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X6))),s('type.num.num',X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X6))),s(X1,X7))))) & ! [X4,X5] : (s('type.num.num',X3) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X4))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X4))),s(X1,X5)))))) & ? [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X8))),s('type.num.num',X3)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X8))),s(X1,X9))))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2,X3] : ((! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X6))),s('type.num.num',X3)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X6))),s(X1,X7))))) & ! [X4,X5] : (s('type.num.num',X3) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X4))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X4))),s(X1,X5)))))) => ! [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X8))),s('type.num.num',X3)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X8))),s(X1,X9))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ~! [X0,X1] : ! [X2,X3] : ((! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X6))),s('type.num.num',X3)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X6))),s(X1,X7))))) & ! [X4,X5] : (s('type.num.num',X3) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X4))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X4))),s(X1,X5)))))) => ! [X8,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X8))),s('type.num.num',X3)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun('type.closSem.state'(X0),fun(X1,bool)),X2),s('type.closSem.state'(X0),X8))),s(X1,X9))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : ((! [X10,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))),s('type.num.num',X9)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11))))) & ! [X10,X11] : (s('type.num.num',X9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11)))))) => ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))),s('type.num.num',X9)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : ((! [X10,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))),s('type.num.num',X9)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11))))) & ! [X10,X11] : (s('type.num.num',X9) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11)))))) => ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X6),X10))),s('type.num.num',X9)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun('type.closSem.state'(X6),fun(X7,bool)),X8),s('type.closSem.state'(X6),X10))),s(X7,X11))))))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb')).
fof(f1430,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) | $spl517),
  inference(cnf_transformation,[],[f1430_D])).
fof(f1430_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) <=> ~$spl517),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl517])])).
fof(f1428,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) | ($spl6 | $spl17 | $spl356)),
  inference(subsumption_resolution,[],[f1418,f799])).
fof(f1418,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) | ($spl6 | $spl356)),
  inference(superposition,[],[f788,f1033])).
fof(f1443,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) | $spl526),
  inference(cnf_transformation,[],[f1443_D])).
fof(f1443_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f1435,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) | ($spl92 | $spl517)),
  inference(resolution,[],[f1430,f841])).
fof(f916,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl204),
  inference(cnf_transformation,[],[f916_D])).
fof(f916_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f651,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12,X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X13)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X12)))))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f6969,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) | $spl1922),
  inference(cnf_transformation,[],[f6969_D])).
fof(f6969_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) <=> ~$spl1922),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1922])])).
fof(f6967,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) | ($spl17 | $spl204 | $spl526)),
  inference(subsumption_resolution,[],[f6879,f799])).
fof(f6879,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) | ($spl204 | $spl526)),
  inference(superposition,[],[f916,f1443])).
fof(f7241,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))))) | $spl1966),
  inference(cnf_transformation,[],[f7241_D])).
fof(f7241_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))))) <=> ~$spl1966),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1966])])).
fof(f7221,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))))) | ($spl88 | $spl1922)),
  inference(resolution,[],[f6969,f839])).
fof(f839,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl88),
  inference(cnf_transformation,[],[f839_D])).
fof(f839_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f506,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f786,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) | $spl2),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f449,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f330])).
fof(f1012,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) | $spl352),
  inference(cnf_transformation,[],[f1012_D])).
fof(f1012_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f1009,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))),s('type.num.num',sK14))) | ($spl2 | $spl88)),
  inference(resolution,[],[f839,f786])).
fof(f892,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl166),
  inference(cnf_transformation,[],[f892_D])).
fof(f892_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f612,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(flattening,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num',X1) | (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12,X13] : (s('type.num.num',X12) = s('type.num.num',X13) <=> (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X13)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X12))))))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.arithmetic.EQ_LESS_EQ')).
fof(f10844,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) | $spl3072),
  inference(cnf_transformation,[],[f10844_D])).
fof(f10844_D,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) <=> ~$spl3072),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3072])])).
fof(f10842,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) | ($spl8 | $spl166 | $spl352 | $spl1966)),
  inference(subsumption_resolution,[],[f10841,f790])).
fof(f10841,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) | ($spl8 | $spl166 | $spl352 | $spl1966)),
  inference(forward_demodulation,[],[f10840,f7241])).
fof(f10840,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) | s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) | ($spl8 | $spl166 | $spl352)),
  inference(subsumption_resolution,[],[f10603,f790])).
fof(f10603,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15)))))) | s('type.num.num',sK14) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),sK15))) | ($spl166 | $spl352)),
  inference(superposition,[],[f892,f1012])).
fof(f787,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK14) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X6))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X6))),s(sK12,X7))))) ) | $spl4),
  inference(cnf_transformation,[],[f787_D])).
fof(f787_D,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK14) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X6))) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X6))),s(sK12,X7))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f448,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),X6))),s(sK12,X7)))) | s('type.num.num',sK14) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK11),X6)))) )),
  inference(cnf_transformation,[],[f330])).
fof(f11161,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,X0))))) ) | $spl3084),
  inference(cnf_transformation,[],[f11161_D])).
fof(f11161_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,X0))))) ) <=> ~$spl3084),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3084])])).
fof(f11160,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,X0))))) ) | ($spl4 | $spl3072)),
  inference(trivial_inequality_removal,[],[f11063])).
fof(f11063,plain,(
  ( ! [X0] : (s('type.num.num',sK14) != s('type.num.num',sK14) | p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,X0))))) ) | ($spl4 | $spl3072)),
  inference(superposition,[],[f787,f10844])).
fof(f841,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl92),
  inference(cnf_transformation,[],[f841_D])).
fof(f841_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f510,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f784,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16)))) | $spl1),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f450,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16))))),
  inference(cnf_transformation,[],[f330])).
fof(f1033,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16))) | $spl356),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f1028,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK16))) | ($spl1 | $spl92)),
  inference(resolution,[],[f841,f784])).
fof(f861,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) ) | $spl104),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f580,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK17(X2,X1,X0)))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : ((? [X6] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ~p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f379])).
fof(f379,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : (! [X3] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (~p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(ennf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) => ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X7] : ! [X8,X20] : (! [X11] : (p(s(bool,X8)) => p(s(bool,happ(s(fun(X7,bool),X20),s(X7,X11))))) <=> (p(s(bool,X8)) => ! [X11] : p(s(bool,happ(s(fun(X7,bool),X20),s(X7,X11))))))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.bool.RIGHT_FORALL_IMP_THM')).
fof(f799,plain,(
  ~p(s(bool,'F_0')) | $spl17),
  inference(cnf_transformation,[],[f799_D])).
fof(f799_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f797,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f457,f451])).
fof(f457,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f47])).
fof(f47,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f14400,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl3490),
  inference(cnf_transformation,[],[f14400_D])).
fof(f14400_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl3490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3490])])).
fof(f14399,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl17 | $spl104 | $spl356 | $spl3084)),
  inference(subsumption_resolution,[],[f14398,f11161])).
fof(f14398,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK18(happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15)),X0,sK12)))))) ) | ($spl17 | $spl104 | $spl356)),
  inference(subsumption_resolution,[],[f14380,f799])).
fof(f14380,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK12,bool),happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15))),s(sK12,sK18(happ(s(fun('type.closSem.state'(sK11),fun(sK12,bool)),sK13),s('type.closSem.state'(sK11),sK15)),X0,sK12)))))) ) | ($spl104 | $spl356)),
  inference(superposition,[],[f861,f1033])).
fof(f790,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f451,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.clos_relationProps_|49|7_prove_ltb',aHL_TRUTH)).
fof(f14422,plain,(
  $false | ($spl8 | $spl3490)),
  inference(resolution,[],[f14400,f790])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 426
% Passive clauses: 1877
% Generated clauses: 14131
% Final active clauses: 312
% Final passive clauses: 1133
% Input formulas: 132
% Initial clauses: 313
% 
% Duplicate literals: 79
% Trivial inequalities: 38
% Fw subsumption resolutions: 668
% Fw demodulations: 2377
% Bw demodulations: 403
% Condensations: 39
% Global subsumptions: 24
% 
% Simple tautologies: 219
% Equational tautologies: 8
% Forward subsumptions: 6524
% Fw demodulations to eq. taut.: 17
% 
% Binary resolution: 1583
% Factoring: 59
% Forward superposition: 3560
% Backward superposition: 2990
% Self superposition: 48
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 2530
% SAT solver unit clauses: 385
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 14
% 
% Sat splits: 5
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 2430
% TWLsolver calls for satisfiability: 1881
% 
% Memory used [KB]: 12153
% Time elapsed: 0.625 s
% ------------------------------
% ------------------------------
% Success in time 1.27 s
