Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 150 next slice time: 5
dis+11_7_3 on thm.evalProps
% (142670)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 229
% Passive clauses: 4405
% Generated clauses: 21603
% Final active clauses: 145
% Final passive clauses: 3373
% Input formulas: 15
% Initial clauses: 19
% 
% Duplicate literals: 1285
% Trivial inequalities: 308
% Fw subsumption resolutions: 2233
% Fw demodulations: 2034
% Bw demodulations: 24
% 
% Simple tautologies: 411
% Equational tautologies: 62
% Forward subsumptions: 9632
% 
% Binary resolution: 253
% Forward superposition: 8008
% Backward superposition: 6624
% Self superposition: 286
% Equality factoring: 127
% Equality resolution: 25
% 
% Split clauses: 825
% Split components: 1674
% SAT solver clauses: 1404
% SAT solver unit clauses: 12
% SAT solver binary clauses: 153
% 
% Sat splits: 825
% Sat splitting refutations: 43
% 
% TWLsolver clauses: 1385
% TWLsolver calls for satisfiability: 192
% 
% Memory used [KB]: 14711
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.evalProps
% (142688)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 167
% Passive clauses: 1286
% Generated clauses: 17131
% Final active clauses: 78
% Final passive clauses: 326
% Input formulas: 15
% Initial clauses: 19
% 
% Duplicate literals: 708
% Trivial inequalities: 112
% Fw subsumption resolutions: 610
% Fw demodulations: 699
% Bw demodulations: 739
% Condensations: 1243
% 
% Simple tautologies: 105
% Equational tautologies: 41
% Forward subsumptions: 9814
% Bw demodulations to eq. taut.: 76
% 
% Binary resolution: 52
% Forward superposition: 6569
% Backward superposition: 4316
% Self superposition: 318
% Equality factoring: 88
% Equality resolution: 9
% 
% Split clauses: 328
% Split components: 671
% SAT solver clauses: 2314
% SAT solver unit clauses: 27
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 5
% SAT solver learnt literals: 8
% 
% Sat splits: 328
% Sat splitting refutations: 23
% 
% TWLsolver clauses: 2309
% TWLsolver calls for satisfiability: 217
% 
% Memory used [KB]: 15863
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f1323,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f451,f469,f467,f498,f499,f507,f505,f500,f501,f510,f509,f678,f680,f681,f682,f685,f688,f690,f691,f1135,f1139,f1137,f1240,f1241,f1248,f1245,f1242,f1243,f1251,f1244,f1253,f1246,f1250,f1316,f1322])).
fof(f1322,plain,(
  $false | ($spl12 | $spl14)),
  inference(subsumption_resolution,[],[f1321,f1320])).
fof(f1320,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) | ($spl12 | $spl14)),
  inference(forward_demodulation,[],[f1318,f61])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X0),X1),s(X0,X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X11,X9] : s(X7,X9) = s(X7,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,X7),X11),s(X7,X9)))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.semanticPrimitives.eq_result_case_def|split|1')).
fof(f1318,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))))) | ($spl12 | $spl14)),
  inference(backward_demodulation,[],[f1137,f1317])).
fof(f1317,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))))) | $spl14),
  inference(backward_demodulation,[],[f1138,f46])).
fof(f46,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3)))))),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X1),X5),s(X1,X3))))) != s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X0),'_dst_x0x1_2'(s(fun(X1,X0),X4),s(fun(bool,X1),X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X3)))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X1),X5),s(X1,X3))))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X0),'_dst_x0x1_2'(s(fun(X1,X0),X4),s(fun(bool,X1),X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X3)))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X4),s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X1),X5),s(X1,X3))))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X2),s(fun(bool,X0),'_dst_x0x1_2'(s(fun(X1,X0),X4),s(fun(bool,X1),X5))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : s(X6,happ(s(fun(X7,X6),X10),s(X7,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X8),s(fun(bool,X7),X11),s(X7,X9))))) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X8),s(fun(bool,X6),'_dst_x0x1_2'(s(fun(X7,X6),X10),s(fun(bool,X7),X11))),s(X6,happ(s(fun(X7,X6),X10),s(X7,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : s(X6,happ(s(fun(X7,X6),X10),s(X7,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X8),s(fun(bool,X7),X11),s(X7,X9))))) = s(X6,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X8),s(fun(bool,X6),'_dst_x0x1_2'(s(fun(X7,X6),X10),s(fun(bool,X7),X11))),s(X6,happ(s(fun(X7,X6),X10),s(X7,X9)))))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.evalProps.evalProps_4|48|_prove_ltb')).
fof(f1138,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | $spl14),
  inference(cnf_transformation,[],[f1138_D])).
fof(f1138_D,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1321,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) | ($spl12 | $spl14)),
  inference(forward_demodulation,[],[f1319,f61])).
fof(f1319,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) = s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | ($spl12 | $spl14)),
  inference(backward_demodulation,[],[f1137,f1138])).
fof(f1316,plain,(
  $false | $spl15),
  inference(trivial_inequality_removal,[],[f1315])).
fof(f1315,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))) | $spl15),
  inference(forward_demodulation,[],[f1314,f61])).
fof(f1314,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | $spl15),
  inference(forward_demodulation,[],[f1312,f61])).
fof(f1312,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK1),sK5),s(sK1,sK3))))) != s(sK0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,sK3))))) | $spl15),
  inference(backward_demodulation,[],[f1310,f46])).
fof(f1310,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2) | $spl15),
  inference(trivial_inequality_removal,[],[f1309])).
fof(f1309,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2) | $spl15),
  inference(superposition,[],[f1139,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(bool,X1),X2),s(bool,sK6(X0)))) = s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X0),s(fun(bool,X1),X2),s(X1,X3))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f62,f51])).
fof(f51,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f20])).
fof(f20,plain,(
  ! [X0] : (? [X1] : s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : (? [X15] : s('type.semanticPrimitives.eq_result',X14) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X15))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X14))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.semanticPrimitives.eq_result_nchotomy')).
fof(f62,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun(bool,X0),X2),s(bool,X1))) = s(X0,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1))),s(fun(bool,X0),X2),s(X0,X3)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ! [X13,X11,X9] : s(X7,happ(s(fun(bool,X7),X11),s(bool,X13))) = s(X7,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X13))),s(fun(bool,X7),X11),s(X7,X9)))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.semanticPrimitives.eq_result_case_def|split|0')).
fof(f1250,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl22),
  inference(cnf_transformation,[],[f1250_D])).
fof(f1250_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1246,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool))) ) | $spl18),
  inference(cnf_transformation,[],[f1246_D])).
fof(f1246_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1253,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl27),
  inference(cnf_transformation,[],[f1253_D])).
fof(f1253_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1244,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK7(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f474,f50])).
fof(f50,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb',aHL_BOOL_CASES)).
fof(f474,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK7(X16,X17,X18,bool))) )),
  inference(superposition,[],[f63,f50])).
fof(f63,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb',aHL_EXT)).
fof(f1251,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl24),
  inference(cnf_transformation,[],[f1251_D])).
fof(f1251_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1243,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK7(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f474,f74])).
fof(f74,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f50,f50])).
fof(f1242,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK7(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f474,f74])).
fof(f1245,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f1245_D])).
fof(f1245_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1248,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl21),
  inference(cnf_transformation,[],[f1248_D])).
fof(f1248_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1241,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK7(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f474,f74])).
fof(f1240,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK7(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f474,f74])).
fof(f1137,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2) | $spl12),
  inference(cnf_transformation,[],[f1137_D])).
fof(f1137_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1139,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | $spl15),
  inference(cnf_transformation,[],[f1139_D])).
fof(f1139_D,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1135,plain,(
  s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,happ(s(fun(bool,sK1),sK5),s(bool,sK6(sK2)))))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2)),
  inference(forward_demodulation,[],[f1113,f64])).
fof(f64,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16,X17,X18] : ! [X19,X20,X21] : s(X16,happ(s(fun(X17,X16),X19),s(X17,happ(s(fun(X18,X17),X20),s(X18,X21))))) = s(X16,happ(s(fun(X18,X16),'_dst_x0x1_2'(s(fun(X17,X16),X19),s(fun(X18,X17),X20))),s(X18,X21)))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','_dst_x0x1_ax')).
fof(f1113,plain,(
  s(sK0,happ(s(fun(bool,sK0),'_dst_x0x1_2'(s(fun(sK1,sK0),sK4),s(fun(bool,sK1),sK5))),s(bool,sK6(sK2)))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',sK2),s(fun(bool,sK1),sK5),s(sK1,sK3))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK2)),
  inference(superposition,[],[f46,f343])).
fof(f691,plain,(
  $false | ($spl4 | $spl8)),
  inference(global_subsumption,[],[f46,f47,f49,f50,f51,f53,f52,f57,f56,f61,f62,f63,f64,f65,f74,f76,f95,f103,f134,f144,f154,f177,f194,f330,f211,f343,f357,f359,f361,f415,f462,f463,f424,f430,f435,f457,f466,f451,f470,f472,f474,f475,f476,f478,f480,f481,f509,f505,f521,f522,f523,f525,f527,f529,f558,f564,f565,f566,f567,f573,f574,f581,f582,f583,f584,f590,f591,f605,f638,f679,f677,f668,f669])).
fof(f669,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(bool,X1),X2),s(bool,'T_0'))) = s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X0),s(fun(bool,X1),X2),s(X1,X3))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f343])).
fof(f668,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f51])).
fof(f677,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f645])).
fof(f645,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f457])).
fof(f679,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f640,f638])).
fof(f640,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f56])).
fof(f638,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f637,f47])).
fof(f637,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f623,f505])).
fof(f623,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f76,f509])).
fof(f605,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f509,f50])).
fof(f591,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,happ(s(fun(X79,bool),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82)))))) | s(bool,X83) = s(bool,happ(s(fun(X78,bool),'_dst_x0x1_2'(s(fun(X79,bool),X80),s(fun(X78,X79),X81))),s(X78,X82))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f194,f64])).
fof(f590,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X73,bool),X74),s(X73,happ(s(fun(X72,X73),X75),s(X72,X76)))))) | s(bool,X77) != s(bool,happ(s(fun(X72,bool),'_dst_x0x1_2'(s(fun(X73,bool),X74),s(fun(X72,X73),X75))),s(X72,X76))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f76,f64])).
fof(f584,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,happ(s(fun(X36,bool),X37),s(X36,happ(s(fun(X35,X36),X38),s(X35,X39)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),'_dst_x0x1_2'(s(fun(X36,bool),X37),s(fun(X35,X36),X38))),s(X35,X39)))) )),
  inference(superposition,[],[f57,f64])).
fof(f583,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X30,X31),X33),s(X30,X34))))) | ~p(s(bool,happ(s(fun(X30,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X30,X31),X33))),s(X30,X34))))) )),
  inference(superposition,[],[f56,f64])).
fof(f582,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,happ(s(fun(X25,X26),X28),s(X25,X29)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'_dst_x0x1_2'(s(fun(X26,bool),X27),s(fun(X25,X26),X28))),s(X25,X29)))) )),
  inference(superposition,[],[f53,f64])).
fof(f581,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X21,bool),X22),s(X21,happ(s(fun(X20,X21),X23),s(X20,X24)))))) | s(bool,'T_0') != s(bool,happ(s(fun(X20,bool),'_dst_x0x1_2'(s(fun(X21,bool),X22),s(fun(X20,X21),X23))),s(X20,X24)))) )),
  inference(superposition,[],[f52,f64])).
fof(f574,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,happ(s(fun(X80,bool),'_dst_x0x1_2'(s(fun(X78,bool),X79),s(fun(X80,X78),X81))),s(X80,X82)))) | s(bool,X83) = s(bool,happ(s(fun(X78,bool),X79),s(X78,happ(s(fun(X80,X78),X81),s(X80,X82))))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f194,f64])).
fof(f573,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~p(s(bool,happ(s(fun(X74,bool),'_dst_x0x1_2'(s(fun(X72,bool),X73),s(fun(X74,X72),X75))),s(X74,X76)))) | s(bool,X77) != s(bool,happ(s(fun(X72,bool),X73),s(X72,happ(s(fun(X74,X72),X75),s(X74,X76))))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f76,f64])).
fof(f567,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,happ(s(fun(X37,bool),'_dst_x0x1_2'(s(fun(X35,bool),X36),s(fun(X37,X35),X38))),s(X37,X39)))) | s(bool,'F_0') = s(bool,happ(s(fun(X35,bool),X36),s(X35,happ(s(fun(X37,X35),X38),s(X37,X39)))))) )),
  inference(superposition,[],[f57,f64])).
fof(f566,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,happ(s(fun(X32,bool),'_dst_x0x1_2'(s(fun(X30,bool),X31),s(fun(X32,X30),X33))),s(X32,X34))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34))))))) )),
  inference(superposition,[],[f56,f64])).
fof(f565,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,happ(s(fun(X27,bool),'_dst_x0x1_2'(s(fun(X25,bool),X26),s(fun(X27,X25),X28))),s(X27,X29)))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,happ(s(fun(X27,X25),X28),s(X27,X29)))))) )),
  inference(superposition,[],[f53,f64])).
fof(f564,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X22,bool),'_dst_x0x1_2'(s(fun(X20,bool),X21),s(fun(X22,X20),X23))),s(X22,X24)))) | s(bool,'T_0') != s(bool,happ(s(fun(X20,bool),X21),s(X20,happ(s(fun(X22,X20),X23),s(X22,X24)))))) )),
  inference(superposition,[],[f52,f64])).
fof(f558,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,happ(s(fun(X24,X25),X27),s(X24,X28))))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),'_dst_x0x1_2'(s(fun(X25,bool),X26),s(fun(X24,X25),X27))),s(X24,X28)))) )),
  inference(superposition,[],[f64,f50])).
fof(f529,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(X40,happ(s(fun(X42,X40),X43),s(X42,happ(s(fun(X41,X42),X44),s(X41,happ(s(fun(X45,X41),X46),s(X45,X47))))))) = s(X40,happ(s(fun(X45,X40),'_dst_x0x1_2'(s(fun(X41,X40),'_dst_x0x1_2'(s(fun(X42,X40),X43),s(fun(X41,X42),X44))),s(fun(X45,X41),X46))),s(X45,X47)))) )),
  inference(superposition,[],[f64,f64])).
fof(f527,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'F_0') = s(bool,happ(s(fun(X26,bool),'_dst_x0x1_2'(s(fun(X24,bool),X25),s(fun(X26,X24),X27))),s(X26,X28))) | s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X25),s(X24,happ(s(fun(X26,X24),X27),s(X26,X28)))))) )),
  inference(superposition,[],[f64,f50])).
fof(f525,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X17) = s(bool,happ(s(fun(X14,bool),'_dst_x0x1_2'(s(fun(X12,bool),X13),s(fun(X14,X12),X15))),s(X14,X16))) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,happ(s(fun(X14,X12),X15),s(X14,X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f64,f74])).
fof(f523,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),'_dst_x0x1_2'(s(fun(X0,bool),X1),s(fun(X2,X0),X3))),s(X2,X4))) | s(bool,X5) = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f64,f74])).
fof(f522,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42,X40] : (s(X46,happ(s(fun(X40,X46),X47),s(X40,happ(s(fun(X42,X40),X43),s(X42,happ(s(fun(X41,X42),X44),s(X41,X45))))))) = s(X46,happ(s(fun(X41,X46),'_dst_x0x1_2'(s(fun(X40,X46),X47),s(fun(X41,X40),'_dst_x0x1_2'(s(fun(X42,X40),X43),s(fun(X41,X42),X44))))),s(X41,X45)))) )),
  inference(superposition,[],[f64,f64])).
fof(f521,plain,(
  ( ! [X39,X37,X35,X33,X38,X36,X34,X32] : (s(X38,happ(s(fun(X32,X38),X39),s(X32,happ(s(fun(X35,X32),'_dst_x0x1_2'(s(fun(X33,X32),X34),s(fun(X35,X33),X36))),s(X35,X37))))) = s(X38,happ(s(fun(X33,X38),'_dst_x0x1_2'(s(fun(X32,X38),X39),s(fun(X33,X32),X34))),s(X33,happ(s(fun(X35,X33),X36),s(X35,X37)))))) )),
  inference(superposition,[],[f64,f64])).
fof(f481,plain,(
  ( ! [X21,X19,X20] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) != s('type.semanticPrimitives.eq_result',f(s(X19,sK7(X20,X21,'type.semanticPrimitives.eq_result',X19)))) | s(fun(X19,'type.semanticPrimitives.eq_result'),X20) = s(fun(X19,'type.semanticPrimitives.eq_result'),X21) | s(bool,'T_0') = s(bool,sK6(g(s(X19,sK7(X20,X21,'type.semanticPrimitives.eq_result',X19))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',g(s(X19,sK7(X20,X21,'type.semanticPrimitives.eq_result',X19))))) )),
  inference(superposition,[],[f63,f95])).
fof(f480,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK7(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK7(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f63,f50])).
fof(f478,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK7(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK7(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f63,f74])).
fof(f476,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK7(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK7(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f63,f74])).
fof(f475,plain,(
  ( ! [X21,X19,X20] : (s(X21,f(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))))) != s(X21,g(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))))) | s(fun('type.semanticPrimitives.eq_result',X21),X19) = s(fun('type.semanticPrimitives.eq_result',X21),X20) | s(bool,'T_0') = s(bool,sK6(sK7(X19,X20,X21,'type.semanticPrimitives.eq_result'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',sK7(X19,X20,X21,'type.semanticPrimitives.eq_result'))) )),
  inference(superposition,[],[f63,f95])).
fof(f472,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK7(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f63,f74])).
fof(f470,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK7(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f63,f74])).
fof(f466,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,sK6(X16)) | s(bool,'T_0') = s(bool,X17) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16)) )),
  inference(subsumption_resolution,[],[f447,f47])).
fof(f447,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | s(bool,X17) != s(bool,sK6(X16)) | s(bool,'T_0') = s(bool,X17) | s('type.semanticPrimitives.eq_result',X16) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X16)) )),
  inference(superposition,[],[f76,f95])).
fof(f457,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f438])).
fof(f438,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f51,f95])).
fof(f435,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,happ(s(fun(bool,X1),X2),s(bool,'F_0'))) = s(X1,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result',X0),s(fun(bool,X1),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f62,f95])).
fof(f430,plain,(
  ( ! [X4,X2,X3,X1] : (s(X2,happ(s(fun(bool,X2),X3),s(bool,X1))) = s(X2,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))),s(fun(bool,X2),X3),s(X2,X4))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X1)))) )),
  inference(superposition,[],[f62,f95])).
fof(f424,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result',X1) | s(bool,'T_0') = s(bool,sK6(X1)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X1) | s(bool,'T_0') = s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f95,f95])).
fof(f463,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',happ(s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s(bool,X5))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result',happ(s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s(bool,X5))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))),s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s('type.semanticPrimitives.eq_result',X7))))) )),
  inference(forward_demodulation,[],[f417,f62])).
fof(f417,plain,(
  ( ! [X6,X7,X5] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result',happ(s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s(bool,X5))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))),s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s('type.semanticPrimitives.eq_result',X7)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X5))),s(fun(bool,'type.semanticPrimitives.eq_result'),X6),s('type.semanticPrimitives.eq_result',X7)))) )),
  inference(superposition,[],[f95,f62])).
fof(f462,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X4) | s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4))))) )),
  inference(forward_demodulation,[],[f416,f61])).
fof(f416,plain,(
  ( ! [X4,X3] : (s('type.semanticPrimitives.eq_result',X4) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,'type.semanticPrimitives.eq_result'),X3),s('type.semanticPrimitives.eq_result',X4)))) )),
  inference(superposition,[],[f95,f61])).
fof(f415,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0))))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6(X0)))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f95,f51])).
fof(f361,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X16))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X16))),s(fun(bool,bool),X17),s(bool,X18)))) )),
  inference(superposition,[],[f62,f50])).
fof(f359,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(bool,bool),X9),s(bool,X8))) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X8))),s(fun(bool,bool),X9),s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f62,f74])).
fof(f357,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X3) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X0))),s(fun(bool,bool),X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f62,f74])).
fof(f211,plain,(
  ( ! [X21,X20] : (s('type.semanticPrimitives.eq_result',X20) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X20) | s(bool,X21) = s(bool,sK6(X20)) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f51,f74])).
fof(f330,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f205,f47])).
fof(f205,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f76,f74])).
fof(f194,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f47,f74])).
fof(f177,plain,(
  ( ! [X21,X22] : (s('type.semanticPrimitives.eq_result',X21) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X22))) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,sK6(X21))) )),
  inference(superposition,[],[f51,f74])).
fof(f154,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X14),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f74,f61])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f74,f74])).
fof(f134,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X20),s(bool,X21)))) )),
  inference(superposition,[],[f74,f61])).
fof(f103,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.semanticPrimitives.eq_result_CASE_3'(s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0'),s(fun(bool,bool),X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f61,f50])).
fof(f95,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X0)) | s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') = s('type.semanticPrimitives.eq_result',X0)) )),
  inference(superposition,[],[f51,f50])).
fof(f76,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f56,f50])).
fof(f65,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f49,f52])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f57,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f52,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f53,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f41])).
fof(f49,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb',aHL_FALSITY)).
fof(f47,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.evalProps.evalProps_4|48|_prove_ltb',aHL_TRUTH)).
fof(f690,plain,(
  $false | ($spl4 | $spl8)),
  inference(global_subsumption,[],[f46,f47,f49,f50,f51,f53,f52,f57,f56,f61,f62,f63,f64,f65,f74,f76,f95,f103,f134,f144,f154,f177,f194,f330,f211,f343,f357,f359,f361,f415,f462,f463,f424,f430,f435,f457,f466,f451,f470,f472,f474,f475,f476,f478,f480,f481,f509,f505,f521,f522,f523,f525,f527,f529,f558,f564,f565,f566,f567,f573,f574,f581,f582,f583,f584,f590,f591,f605,f638,f679,f677,f668])).
fof(f688,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f687,f47])).
fof(f687,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f686,f638])).
fof(f686,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X30,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X30,X31),X33))),s(X30,X34))))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f656,f638])).
fof(f656,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(X31,bool),X32),s(X31,happ(s(fun(X30,X31),X33),s(X30,X34))))) | ~p(s(bool,happ(s(fun(X30,bool),'_dst_x0x1_2'(s(fun(X31,bool),X32),s(fun(X30,X31),X33))),s(X30,X34))))) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f583])).
fof(f685,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f684,f47])).
fof(f684,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f683,f638])).
fof(f683,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34))))))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f654,f638])).
fof(f654,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,happ(s(fun(X32,bool),'_dst_x0x1_2'(s(fun(X30,bool),X31),s(fun(X32,X30),X33))),s(X32,X34))) | ~p(s(bool,happ(s(fun(X30,bool),X31),s(X30,happ(s(fun(X32,X30),X33),s(X32,X34))))))) ) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f566])).
fof(f682,plain,(
  $false | ($spl3 | $spl4 | $spl8)),
  inference(global_subsumption,[],[f647,f46,f47,f49,f50,f51,f53,f52,f57,f56,f61,f62,f63,f64,f65,f74,f76,f95,f103,f134,f144,f154,f177,f194,f330,f211,f343,f357,f359,f361,f415,f462,f463,f424,f430,f435,f457,f466,f451,f470,f472,f474,f475,f476,f478,f480,f481,f509,f505,f521,f522,f523,f525,f527,f529,f558,f564,f565,f566,f567,f573,f574,f581,f582,f583,f584,f590,f591,f605,f638,f679,f677])).
fof(f647,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) | ($spl3 | $spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f469])).
fof(f681,plain,(
  $false | ($spl4 | $spl8)),
  inference(global_subsumption,[],[f46,f47,f49,f50,f51,f53,f52,f57,f56,f61,f62,f63,f64,f65,f74,f76,f95,f103,f134,f144,f154,f177,f194,f330,f211,f343,f357,f359,f361,f415,f462,f463,f424,f430,f435,f457,f466,f451,f470,f472,f474,f475,f476,f478,f480,f481,f509,f505,f521,f522,f523,f525,f527,f529,f558,f564,f565,f566,f567,f573,f574,f581,f582,f583,f584,f590,f591,f605,f638,f679,f677])).
fof(f680,plain,(
  $false | ($spl4 | $spl8)),
  inference(global_subsumption,[],[f46,f47,f49,f50,f51,f53,f52,f57,f56,f61,f62,f63,f64,f65,f74,f76,f95,f103,f134,f144,f154,f177,f194,f330,f211,f343,f357,f359,f361,f415,f462,f463,f424,f430,f435,f457,f466,f451,f470,f472,f474,f475,f476,f478,f480,f481,f509,f505,f521,f522,f523,f525,f527,f529,f558,f564,f565,f566,f567,f573,f574,f581,f582,f583,f584,f590,f591,f605,f638,f679])).
fof(f678,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f639,f47])).
fof(f639,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f638,f49])).
fof(f509,plain,(
  s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | $spl8),
  inference(cnf_transformation,[],[f509_D])).
fof(f509_D,plain,(
  s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f510,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f501,plain,(
  ( ! [X3] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))))) ) | $spl3),
  inference(superposition,[],[f482,f74])).
fof(f482,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))))) | $spl3),
  inference(unit_resulting_resolution,[],[f469,f51])).
fof(f500,plain,(
  ( ! [X2] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,X2))) | s(bool,'T_0') = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f482,f74])).
fof(f505,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f505_D])).
fof(f505_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f507,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | $spl6),
  inference(cnf_transformation,[],[f507_D])).
fof(f507_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f499,plain,(
  ( ! [X1] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,X1) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f482,f74])).
fof(f498,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'T_0'))) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,sK6('const.semanticPrimitives.Eq_val_1'(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f482,f74])).
fof(f467,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X0))) ) | $spl0),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X0))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f469,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | $spl3),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f451,plain,(
  ( ! [X0] : (s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_type_error_0') != s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.eq_result',X0) = s('type.semanticPrimitives.eq_result','const.semanticPrimitives.Eq_val_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X0))) )),
  inference(equality_factoring,[],[f95])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 40
% Passive clauses: 196
% Generated clauses: 1268
% Final active clauses: 30
% Final passive clauses: 110
% Input formulas: 15
% Initial clauses: 19
% 
% Duplicate literals: 94
% Trivial inequalities: 13
% Fw subsumption resolutions: 36
% Fw demodulations: 86
% Bw demodulations: 47
% Condensations: 98
% Global subsumptions: 5
% 
% Simple tautologies: 38
% Equational tautologies: 29
% Forward subsumptions: 651
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 43
% Forward superposition: 394
% Backward superposition: 304
% Self superposition: 47
% Equality factoring: 30
% Equality resolution: 3
% 
% Split clauses: 11
% Split components: 27
% SAT solver clauses: 235
% SAT solver unit clauses: 13
% SAT solver binary clauses: 68
% 
% Sat splits: 11
% Sat splitting refutations: 10
% 
% Memory used [KB]: 1918
% Time elapsed: 0.049 s
% ------------------------------
% ------------------------------
% Success in time 1.361 s
