Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_relationProps
% (16519)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 2826
% Generated clauses: 10108
% Final active clauses: 95
% Final passive clauses: 2584
% Input formulas: 30
% Initial clauses: 64
% 
% Duplicate literals: 507
% Trivial inequalities: 149
% Fw subsumption resolutions: 1603
% Fw demodulations: 268
% Bw demodulations: 1
% 
% Simple tautologies: 266
% Equational tautologies: 75
% Forward subsumptions: 3641
% 
% Binary resolution: 482
% Factoring: 2
% Forward superposition: 3682
% Backward superposition: 3045
% Self superposition: 195
% Equality factoring: 170
% Equality resolution: 27
% 
% Split clauses: 328
% Split components: 679
% SAT solver clauses: 919
% SAT solver unit clauses: 34
% SAT solver binary clauses: 368
% 
% Sat splits: 328
% 
% TWLsolver clauses: 919
% TWLsolver calls for satisfiability: 158
% 
% Memory used [KB]: 5628
% Time elapsed: 0.619 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_relationProps
% (16556)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 612
% Generated clauses: 11490
% Final active clauses: 108
% Final passive clauses: 448
% Input formulas: 30
% Initial clauses: 64
% 
% Duplicate literals: 360
% Trivial inequalities: 109
% Fw subsumption resolutions: 989
% Fw demodulations: 844
% Bw demodulations: 45
% Condensations: 1136
% Global subsumptions: 9
% 
% Simple tautologies: 152
% Equational tautologies: 36
% Forward subsumptions: 6064
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 116
% Factoring: 1
% Forward superposition: 3868
% Backward superposition: 3626
% Self superposition: 188
% Equality factoring: 81
% Equality resolution: 43
% 
% Split clauses: 67
% Split components: 137
% SAT solver clauses: 1211
% SAT solver unit clauses: 81
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 18
% 
% Sat splits: 67
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1175
% TWLsolver calls for satisfiability: 403
% 
% Memory used [KB]: 10746
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_relationProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_relationProps
% SZS output start Proof for thm.clos_relationProps
fof(f6106,plain,(
  $false),
  inference(subsumption_resolution,[],[f6044,f177])).
fof(f177,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f116,f113])).
fof(f113,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.bool.TRUTH')).
fof(f116,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.bool.NOT_CLAUSES|split|1')).
fof(f6044,plain,(
  p(s(bool,'F_0'))),
  inference(backward_demodulation,[],[f6002,f5797])).
fof(f5797,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.closSem.v'),sK3),s('type.list.list'('type.closSem.v'),sK8(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))))),
  inference(subsumption_resolution,[],[f5771,f113])).
fof(f5771,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(sK0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(sK0),sK2))))),s('type.list.list'('type.closSem.v'),sK3),s('type.list.list'('type.closSem.v'),sK8(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))))),
  inference(superposition,[],[f157,f237])).
fof(f237,plain,(
  s(bool,'T_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))),
  inference(unit_resulting_resolution,[],[f177,f112,f217])).
fof(f217,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f214,f177])).
fof(f214,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f167,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr',aHL_BOOL_CASES)).
fof(f167,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.bool.IMP_ANTISYM_AX')).
fof(f112,plain,(
  s(bool,'F_0') != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  s(bool,'F_0') != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f73])).
fof(f73,plain,(
  ? [X0,X1,X2,X3,X4] : s(bool,'F_0') != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X1)))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ~! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X1)))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X3))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s(bool,'F_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(X6),X8))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),X10))),s('type.closSem.state'(X6),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s(bool,'F_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X9))),s('type.closSem.state'(X6),X8))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X6)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),X10))),s('type.closSem.state'(X6),X7)))))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.clos_relationProps.res_rel_Rval_Rerr')).
fof(f157,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),sK8(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK8(X3,X2,X1,X0)))),s('type.closSem.state'(X0),sK9(X3,X2,X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),sK8(X3,X2,X1,X0))))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),sK9(X3,X2,X1,X0))))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),sK9(X3,X2,X1,X0)))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X6,X7] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X6))),s('type.closSem.state'(X0),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X6)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X7)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X7))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(rectify,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X16] : ! [X14,X17,X18] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X17))),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14)))) <=> ? [X19,X20] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X19))),s('type.closSem.state'(X16),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X16),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))))),s('type.list.list'('type.closSem.v'),X17),s('type.list.list'('type.closSem.v'),X19)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))),s('type.closSem.state'(X16),X18),s('type.closSem.state'(X16),X20)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X20)))))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.clos_relation.res_rel_rw|split|0')).
fof(f6002,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6001,f587])).
fof(f587,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.semanticPrimitives.result'(X0,X1),X2),'const.pair.,_2'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))),s(X2,X4))) != s('type.pair.prod'('type.semanticPrimitives.result'(X0,X1),X2),'const.pair.,_2'(s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X5))),s(X2,X6)))) )),
  inference(unit_resulting_resolution,[],[f172,f173])).
fof(f173,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X23] : ! [X14,X15,X24,X25] : (s('type.pair.prod'(X6,X23),'const.pair.,_2'(s(X6,X14),s(X23,X15))) = s('type.pair.prod'(X6,X23),'const.pair.,_2'(s(X6,X24),s(X23,X25))) <=> (s(X6,X14) = s(X6,X24) & s(X23,X15) = s(X23,X25)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.pair.CLOSED_PAIR_EQ')).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) != s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2,X3] : ~s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rval_1'(s(X0,X3))) = s('type.semanticPrimitives.result'(X0,X1),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X1),X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X23] : ! [X27,X24] : ~s('type.semanticPrimitives.result'(X6,X23),'const.semanticPrimitives.Rval_1'(s(X6,X24))) = s('type.semanticPrimitives.result'(X6,X23),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X23),X27)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.semanticPrimitives.result_distinct')).
fof(f6001,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK8(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))),s('type.closSem.state'(sK0),sK9(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f5978,f127])).
fof(f127,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.bool.EQ_CLAUSES|split|2')).
fof(f5978,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK8(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))),s('type.closSem.state'(sK0),sK9(sK2,sK3,'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)),sK0)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f156,f1434])).
fof(f1434,plain,(
  ( ! [X68] : (s(bool,X68) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(global_subsumption,[],[f112,f113,f177,f120,f123,f122,f127,f126,f147,f150,f155,f154,f153,f160,f159,f158,f157,f156,f165,f164,f163,f162,f161,f166,f170,f167,f171,f172,f175,f174,f173,f178,f187,f189,f217,f238,f237,f265,f469,f470,f500,f502,f504,f539,f540,f563,f587,f586,f588,f590,f592,f595,f596,f597,f598,f672,f720,f719,f718,f732,f734,f736,f739,f740,f741,f742,f777,f776,f853,f855,f857,f858,f860,f862,f1311,f1303,f1297,f913,f914,f1288,f1287,f991,f994,f1382,f1037,f1040,f1238,f1229,f1415,f1097,f1098,f1099,f1103,f1104,f1224,f1218,f1433])).
fof(f1433,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') != s(bool,X67) | s(bool,X68) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f1152,f122])).
fof(f1152,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') != s(bool,X67) | ~p(s(bool,X67)) | s(bool,X68) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f126,f540])).
fof(f1218,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X66)) )),
  inference(duplicate_literal_removal,[],[f1151])).
fof(f1151,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X66)) )),
  inference(superposition,[],[f120,f540])).
fof(f1224,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X1) = s(bool,X2)) )),
  inference(condensation,[],[f1223])).
fof(f1223,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X55) | s(bool,X56) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X53) = s(bool,X56)) )),
  inference(duplicate_literal_removal,[],[f1116])).
fof(f1116,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X53) != s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,X56) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X53) = s(bool,X56)) )),
  inference(superposition,[],[f470,f540])).
fof(f1104,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,'F_0'))) != s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X19),s(bool,X16))) | s(bool,X20) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X16) = s(bool,X20)) )),
  inference(superposition,[],[f718,f540])).
fof(f1103,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,'F_0'),s(X12,X14))) != s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13))) | s(bool,X15) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X11) = s(bool,X15)) )),
  inference(superposition,[],[f586,f540])).
fof(f1099,plain,(
  ( ! [X167,X165,X166,X164] : (s(X166,f(s(bool,X167))) != s(X166,g(s(bool,X167))) | s(fun(bool,X166),X164) = s(fun(bool,X166),X165) | s(bool,X167) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,sK12(X164,X165,X166,bool))) )),
  inference(superposition,[],[f171,f540])).
fof(f1098,plain,(
  ( ! [X163,X162] : (s('type.semanticPrimitives.error_result'(bool),X162) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X163))) | s('type.semanticPrimitives.error_result'(bool),X162) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X162,bool)))) | s(bool,X163) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,sK5(X162,bool))) )),
  inference(superposition,[],[f147,f540])).
fof(f1097,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,X161) != s(bool,f(s(X158,sK12(X159,X160,bool,X158)))) | s(fun(X158,bool),X159) = s(fun(X158,bool),X160) | s(bool,X161) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,g(s(X158,sK12(X159,X160,bool,X158))))) )),
  inference(superposition,[],[f171,f540])).
fof(f1415,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,X75) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(subsumption_resolution,[],[f1078,f178])).
fof(f1078,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,X75) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f178,f540])).
fof(f1229,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(condensation,[],[f1228])).
fof(f1228,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(duplicate_literal_removal,[],[f1067])).
fof(f1067,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,X46) != s(bool,X47) | s(bool,'T_0') = s(bool,X45) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f470,f540])).
fof(f1238,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(duplicate_literal_removal,[],[f1051])).
fof(f1051,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f120,f540])).
fof(f1040,plain,(
  ( ! [X123,X121,X118,X122,X120,X119] : (s('type.pair.prod'(X120,bool),'const.pair.,_2'(s(X120,X121),s(bool,X122))) != s('type.pair.prod'(X120,bool),'const.pair.,_2'(s(X120,X123),s(bool,X119))) | s(bool,X118) = s(bool,X122) | s(bool,X118) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f174,f540])).
fof(f1037,plain,(
  ( ! [X103,X101,X105,X102,X100,X104] : (s('type.pair.prod'(bool,X102),'const.pair.,_2'(s(bool,X101),s(X102,X103))) != s('type.pair.prod'(bool,X102),'const.pair.,_2'(s(bool,X104),s(X102,X105))) | s(bool,X100) = s(bool,X104) | s(bool,X100) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f173,f540])).
fof(f1382,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,X8) | s(bool,X8) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1006,f123])).
fof(f1006,plain,(
  ( ! [X8,X9] : (p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8) | s(bool,X8) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f127,f540])).
fof(f994,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X90))) != s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X91),s(bool,X87))) | s(bool,X90) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X87) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f174,f540])).
fof(f991,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X69),s(X70,X71))) != s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X72),s(X70,X73))) | s(bool,X72) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X69) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f173,f540])).
fof(f1287,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X24)) )),
  inference(duplicate_literal_removal,[],[f929])).
fof(f929,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f540,f120])).
fof(f1288,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(duplicate_literal_removal,[],[f927])).
fof(f927,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X14) = s(bool,X16) | s(bool,X14) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f540,f540])).
fof(f914,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f540,f187])).
fof(f913,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f540,f540])).
fof(f1297,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(duplicate_literal_removal,[],[f906])).
fof(f906,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f540,f540])).
fof(f1303,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f898])).
fof(f898,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f540,f120])).
fof(f1311,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X19) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X17) = s(bool,X19)) )),
  inference(duplicate_literal_removal,[],[f889])).
fof(f889,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X17) = s(bool,X19)) )),
  inference(superposition,[],[f540,f540])).
fof(f862,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK12(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK12(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f171,f120])).
fof(f860,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK12(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK12(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f171,f187])).
fof(f858,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK12(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f171,f187])).
fof(f857,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK12(X16,X17,X18,bool))) )),
  inference(superposition,[],[f171,f120])).
fof(f855,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK12(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f171,f187])).
fof(f853,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK12(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f171,f187])).
fof(f776,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,'type.pair.prod'(X1,bool)),'const.pair.,_2'(s(X0,X2),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,'T_0'))))) != s('type.pair.prod'(X0,'type.pair.prod'(X1,bool)),'const.pair.,_2'(s(X0,X4),s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X5),s(bool,'F_0')))))) )),
  inference(unit_resulting_resolution,[],[f718,f174])).
fof(f777,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.pair.prod'(X0,bool),X1),'const.pair.,_2'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X2),s(bool,'T_0'))),s(X1,X3))) != s('type.pair.prod'('type.pair.prod'(X0,bool),X1),'const.pair.,_2'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X4),s(bool,'F_0'))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f718,f173])).
fof(f742,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s('type.pair.prod'(X58,'type.semanticPrimitives.error_result'(X56)),'const.pair.,_2'(s(X58,X59),s('type.semanticPrimitives.error_result'(X56),X57))) != s('type.pair.prod'(X58,'type.semanticPrimitives.error_result'(X56)),'const.pair.,_2'(s(X58,X60),s('type.semanticPrimitives.error_result'(X56),X61))) | s('type.semanticPrimitives.error_result'(X56),X61) = s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,sK5(X57,X56)))) | s('type.semanticPrimitives.error_result'(X56),X57) = s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X57,X56))))) )),
  inference(superposition,[],[f174,f147])).
fof(f741,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s('type.pair.prod'(X52,'type.semanticPrimitives.error_result'(X50)),'const.pair.,_2'(s(X52,X53),s('type.semanticPrimitives.error_result'(X50),X51))) != s('type.pair.prod'(X52,'type.semanticPrimitives.error_result'(X50)),'const.pair.,_2'(s(X52,X54),s('type.semanticPrimitives.error_result'(X50),X55))) | s('type.semanticPrimitives.error_result'(X50),X55) = s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X51,X50)))) | s('type.semanticPrimitives.error_result'(X50),X51) = s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rraise_1'(s(X50,sK5(X51,X50))))) )),
  inference(superposition,[],[f174,f147])).
fof(f740,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.pair.prod'(X46,bool),'const.pair.,_2'(s(X46,X47),s(bool,'T_0'))) != s('type.pair.prod'(X46,bool),'const.pair.,_2'(s(X46,X48),s(bool,X49))) | s(bool,X49) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f174,f237])).
fof(f739,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X41))) != s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X44),s(bool,X45))) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f174,f120])).
fof(f736,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X26),s(bool,'F_0'))) != s('type.pair.prod'(X25,bool),'const.pair.,_2'(s(X25,X27),s(bool,X28))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f174,f120])).
fof(f734,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13))) != s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X16),s(bool,X17))) | s(bool,X12) = s(bool,X17) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f174,f187])).
fof(f732,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0'))) != s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X3),s(bool,X4))) | s(bool,X0) = s(bool,X4) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f174,f187])).
fof(f718,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))) != s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X2),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f178,f174])).
fof(f719,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,'type.semanticPrimitives.result'(X1,X2)),'const.pair.,_2'(s(X0,X3),s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rval_1'(s(X1,X4))))) != s('type.pair.prod'(X0,'type.semanticPrimitives.result'(X1,X2)),'const.pair.,_2'(s(X0,X5),s('type.semanticPrimitives.result'(X1,X2),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'(X2),X6)))))) )),
  inference(unit_resulting_resolution,[],[f172,f174])).
fof(f720,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,'type.pair.prod'(bool,X1)),'const.pair.,_2'(s(X0,X2),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X3))))) != s('type.pair.prod'(X0,'type.pair.prod'(bool,X1)),'const.pair.,_2'(s(X0,X4),s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X5)))))) )),
  inference(unit_resulting_resolution,[],[f586,f174])).
fof(f672,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'('type.pair.prod'(bool,X0),X1),'const.pair.,_2'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X2))),s(X1,X3))) != s('type.pair.prod'('type.pair.prod'(bool,X0),X1),'const.pair.,_2'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f586,f173])).
fof(f598,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (s('type.pair.prod'('type.semanticPrimitives.error_result'(X56),X58),'const.pair.,_2'(s('type.semanticPrimitives.error_result'(X56),X57),s(X58,X59))) != s('type.pair.prod'('type.semanticPrimitives.error_result'(X56),X58),'const.pair.,_2'(s('type.semanticPrimitives.error_result'(X56),X60),s(X58,X61))) | s('type.semanticPrimitives.error_result'(X56),X60) = s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rraise_1'(s(X56,sK5(X57,X56)))) | s('type.semanticPrimitives.error_result'(X56),X57) = s('type.semanticPrimitives.error_result'(X56),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X57,X56))))) )),
  inference(superposition,[],[f173,f147])).
fof(f597,plain,(
  ( ! [X54,X52,X50,X55,X53,X51] : (s('type.pair.prod'('type.semanticPrimitives.error_result'(X50),X52),'const.pair.,_2'(s('type.semanticPrimitives.error_result'(X50),X51),s(X52,X53))) != s('type.pair.prod'('type.semanticPrimitives.error_result'(X50),X52),'const.pair.,_2'(s('type.semanticPrimitives.error_result'(X50),X54),s(X52,X55))) | s('type.semanticPrimitives.error_result'(X50),X54) = s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X51,X50)))) | s('type.semanticPrimitives.error_result'(X50),X51) = s('type.semanticPrimitives.error_result'(X50),'const.semanticPrimitives.Rraise_1'(s(X50,sK5(X51,X50))))) )),
  inference(superposition,[],[f173,f147])).
fof(f596,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,'T_0'),s(X46,X47))) != s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,X48),s(X46,X49))) | s(bool,X48) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1)))))) )),
  inference(superposition,[],[f173,f237])).
fof(f595,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,X41),s(X42,X43))) != s('type.pair.prod'(bool,X42),'const.pair.,_2'(s(bool,X44),s(X42,X45))) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f173,f120])).
fof(f592,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s('type.pair.prod'(bool,X25),'const.pair.,_2'(s(bool,'F_0'),s(X25,X26))) != s('type.pair.prod'(bool,X25),'const.pair.,_2'(s(bool,X27),s(X25,X28))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f173,f120])).
fof(f590,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15))) != s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X16),s(X14,X17))) | s(bool,X12) = s(bool,X16) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f173,f187])).
fof(f588,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))) != s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X3),s(X1,X4))) | s(bool,X0) = s(bool,X3) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f173,f187])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1))) != s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X2)))) )),
  inference(unit_resulting_resolution,[],[f178,f173])).
fof(f563,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f562])).
fof(f562,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(duplicate_literal_removal,[],[f556])).
fof(f556,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f470,f120])).
fof(f540,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f526,f113])).
fof(f526,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f469,f237])).
fof(f539,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f524,f113])).
fof(f524,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f189,f237])).
fof(f504,plain,(
  ( ! [X8] : (s('type.semanticPrimitives.error_result'(bool),X8) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'F_0'))) | s('type.semanticPrimitives.error_result'(bool),X8) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X8,bool)))) | s(bool,'T_0') = s(bool,sK5(X8,bool))) )),
  inference(superposition,[],[f147,f120])).
fof(f502,plain,(
  ( ! [X4,X5] : (s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,X5))) | s('type.semanticPrimitives.error_result'(bool),X4) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X4,bool)))) | s(bool,'T_0') = s(bool,sK5(X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f147,f187])).
fof(f500,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rraise_1'(s(bool,'T_0'))) | s('type.semanticPrimitives.error_result'(bool),X0) = s('type.semanticPrimitives.error_result'(bool),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X0,bool)))) | s(bool,X1) = s(bool,sK5(X0,bool)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f147,f187])).
fof(f470,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f334,f113])).
fof(f334,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f189,f187])).
fof(f469,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f333,f113])).
fof(f333,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f170,f187])).
fof(f265,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f187,f187])).
fof(f238,plain,(
  p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK3))),s('type.closSem.state'(sK0),sK2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(sK0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),sK4))),s('type.closSem.state'(sK0),sK1))))))),
  inference(unit_resulting_resolution,[],[f177,f112,f167])).
fof(f189,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f126,f120])).
fof(f187,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f120,f120])).
fof(f178,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f177,f122])).
fof(f174,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f111])).
fof(f175,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f171,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr',aHL_EXT)).
fof(f170,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f76])).
fof(f166,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),X1),s('type.num.num',X2),s('type.closSem.v',X3),s('type.closSem.v',X4))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),X1),s('type.num.num',X2))),s('type.closSem.v',X3))),s('type.closSem.v',X4)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),X1),s('type.num.num',X2),s('type.closSem.v',X3),s('type.closSem.v',X4))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),X1),s('type.num.num',X2))),s('type.closSem.v',X3))),s('type.closSem.v',X4)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),X1),s('type.num.num',X2),s('type.closSem.v',X3),s('type.closSem.v',X4))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),X1),s('type.num.num',X2))),s('type.closSem.v',X3))),s('type.closSem.v',X4)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X16] : ! [X28,X29,X30,X31] : s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X16),X28),s('type.num.num',X29),s('type.closSem.v',X30),s('type.closSem.v',X31))) = s(bool,happ(s(fun('type.closSem.v',bool),happ(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X16),X28),s('type.num.num',X29))),s('type.closSem.v',X30))),s('type.closSem.v',X31)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','ARITY_2_4_const.clos_relation.val_rel')).
fof(f161,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',sK10(X3,X2,X1,X0)))))),s('type.closSem.state'(X0),sK11(X3,X2,X1,X0)))) | ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',sK10(X3,X2,X1,X0)))))),s('type.closSem.state'(X0),sK11(X3,X2,X1,X0)))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',sK10(X3,X2,X1,X0))))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),sK11(X3,X2,X1,X0))))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),sK11(X3,X2,X1,X0)))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X6,X7] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X6))))),s('type.closSem.state'(X0),X7))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X6)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X7)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X7))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(rectify,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5))))) & (! [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X4,X5] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X16] : ! [X14,X21,X18] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X21))))),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14)))) <=> ? [X22,X20] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X22))))),s('type.closSem.state'(X16),X20))) & p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X16),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))),s('type.closSem.v',X21),s('type.closSem.v',X22)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))),s('type.closSem.state'(X16),X18),s('type.closSem.state'(X16),X20)))) & s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X20)))))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.clos_relation.res_rel_rw|split|1')).
fof(f162,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',sK10(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f163,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),sK11(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f164,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),sK11(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f165,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X2))))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | ~p(s(bool,'const.clos_relation.val_rel_4'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.v',X2),s('type.closSem.v',X4)))) | s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rraise_1'(s('type.closSem.v',X4))))),s('type.closSem.state'(X0),X5)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f158,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),sK9(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f159,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) = s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),sK9(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f160,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))) != s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X5))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))),s('type.closSem.state'(X0),X3),s('type.closSem.state'(X0),X5)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.closSem.v',fun('type.closSem.v',bool)),'const.clos_relation.val_rel_2'(s('type.bool.itself'(X0),'const.bool.the_value_0'),s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X3))))),s('type.list.list'('type.closSem.v'),X2),s('type.list.list'('type.closSem.v'),X4)))) | s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X4))),s('type.closSem.state'(X0),X5)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),sK7(X2,X1,X0)))) | ~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),sK7(X2,X1,X0)))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),sK7(X2,X1,X0))))))) & (! [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X4] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X4))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X4)))))) & (! [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(rectify,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ? [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3)))))) & (! [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3))))) | p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) <=> ? [X3] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X16] : ! [X14,X18] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14)))) <=> ? [X20] : (s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X16),X20))) & p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X16),X18))),s('type.closSem.state'(X16),X18),s('type.closSem.state'(X16),X20))))))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.clos_relation.res_rel_rw|split|2')).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),sK7(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f155,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | ~p(s(bool,'const.clos_relation.state_rel_3'(s('type.num.num','const.closSem.state_clock_1'(s('type.closSem.state'(X0),X2))),s('type.closSem.state'(X0),X2),s('type.closSem.state'(X0),X3)))) | s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) != s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtimeout_error_0'))))),s('type.closSem.state'(X0),X3)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f150,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtype_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : s(bool,'T_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtype_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'T_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtype_error_0'))))),s('type.closSem.state'(X0),X2))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X16] : ! [X14,X18] : s(bool,'T_0') = s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rerr_1'(s('type.semanticPrimitives.error_result'('type.closSem.v'),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort','const.semanticPrimitives.Rtype_error_0'))))),s('type.closSem.state'(X16),X18))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X16)),X14)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.clos_relation.res_rel_rw|split|3')).
fof(f147,plain,(
  ( ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5(X1,X0))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,sK5(X1,X0)))) | s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',sK6(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1] : (? [X3] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rraise_1'(s(X0,X3))) | ? [X2] : s('type.semanticPrimitives.error_result'(X0),X1) = s('type.semanticPrimitives.error_result'(X0),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X2))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X26] : (? [X24] : s('type.semanticPrimitives.error_result'(X6),X26) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rraise_1'(s(X6,X24))) | ? [X24] : s('type.semanticPrimitives.error_result'(X6),X26) = s('type.semanticPrimitives.error_result'(X6),'const.semanticPrimitives.Rabort_1'(s('type.semanticPrimitives.abort',X24))))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.semanticPrimitives.error_result_nchotomy')).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f83])).
fof(f122,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.clos_relationProps.res_rel_Rval_Rerr','thm.bool.EQ_CLAUSES|split|1')).
fof(f123,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f156,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.clos_relation.res_rel_2'(s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),X2))),s('type.closSem.state'(X0),X3))),s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1)))) | s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),X1) = s('type.pair.prod'('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'type.closSem.state'(X0)),'const.pair.,_2'(s('type.semanticPrimitives.result'('type.list.list'('type.closSem.v'),'type.closSem.v'),'const.semanticPrimitives.Rval_1'(s('type.list.list'('type.closSem.v'),sK8(X3,X2,X1,X0)))),s('type.closSem.state'(X0),sK9(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f105])).
% SZS output end Proof for thm.clos_relationProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 62
% Passive clauses: 363
% Generated clauses: 5868
% Final active clauses: 46
% Final passive clauses: 164
% Input formulas: 30
% Initial clauses: 64
% 
% Duplicate literals: 562
% Trivial inequalities: 79
% Fw subsumption resolutions: 423
% Fw demodulations: 228
% Bw demodulations: 123
% Condensations: 1076
% Global subsumptions: 6
% 
% Simple tautologies: 129
% Equational tautologies: 86
% Forward subsumptions: 2762
% Fw demodulations to eq. taut.: 84
% Bw demodulations to eq. taut.: 17
% 
% Unit resulting resolution: 235
% Factoring: 1
% Forward superposition: 1327
% Backward superposition: 1472
% Self superposition: 120
% Equality factoring: 116
% Equality resolution: 14
% 
% Split clauses: 35
% Split components: 88
% SAT solver clauses: 444
% SAT solver unit clauses: 32
% SAT solver binary clauses: 94
% 
% Sat splits: 35
% Sat splitting refutations: 3
% 
% Memory used [KB]: 4221
% Time elapsed: 0.420 s
% ------------------------------
% ------------------------------
% Success in time 1.77 s
