Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (83110)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 153
% Passive clauses: 2959
% Generated clauses: 16845
% Final active clauses: 132
% Final passive clauses: 2573
% Input formulas: 68
% Initial clauses: 101
% 
% Duplicate literals: 512
% Trivial inequalities: 274
% Fw subsumption resolutions: 1325
% Fw demodulations: 2287
% Bw demodulations: 18
% 
% Simple tautologies: 235
% Equational tautologies: 43
% Forward subsumptions: 8407
% 
% Binary resolution: 379
% Factoring: 1
% Forward superposition: 5633
% Backward superposition: 5986
% Self superposition: 156
% Equality factoring: 70
% Equality resolution: 29
% 
% Split clauses: 206
% Split components: 414
% SAT solver clauses: 560
% SAT solver unit clauses: 10
% SAT solver binary clauses: 118
% 
% Sat splits: 206
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 559
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 13432
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (83135)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 166
% Passive clauses: 1464
% Generated clauses: 16256
% Final active clauses: 123
% Final passive clauses: 770
% Input formulas: 68
% Initial clauses: 101
% 
% Duplicate literals: 449
% Trivial inequalities: 169
% Fw subsumption resolutions: 787
% Fw demodulations: 1753
% Bw demodulations: 458
% Condensations: 967
% Global subsumptions: 8
% 
% Simple tautologies: 213
% Equational tautologies: 45
% Forward subsumptions: 7967
% Bw demodulations to eq. taut.: 99
% 
% Binary resolution: 333
% Factoring: 1
% Forward superposition: 5408
% Backward superposition: 4432
% Self superposition: 151
% Equality factoring: 82
% Equality resolution: 25
% 
% Split clauses: 88
% Split components: 176
% SAT solver clauses: 1897
% SAT solver unit clauses: 151
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 16
% 
% Sat splits: 88
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 1865
% TWLsolver calls for satisfiability: 400
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f88815,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f228,f88801])).
fof(f88801,plain,(
  ( ! [X577] : (~p(s(bool,X577))) )),
  inference(forward_demodulation,[],[f88800,f315])).
fof(f315,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X17] : ! [X26,X35,X36,X37,X38] : s(X17,X37) = s(X17,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X26),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s(X17,X36),s(X17,X37),s(X17,X38)))),
  file('thm.wordSem.result_distinct|split|5','thm.wordSem.result_case_def|split|3')).
fof(f88800,plain,(
  ( ! [X575,X576,X578,X574,X577] : (~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X574),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s(bool,X576),s(bool,X577),s(bool,X578))))) )),
  inference(subsumption_resolution,[],[f88236,f88139])).
fof(f88139,plain,(
  ( ! [X6,X10,X8] : (s(X6,X10) = s(X6,happ(s(fun('type.wordSem.word_loc'(sK0),X6),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X6)),X8),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1)))) )),
  inference(superposition,[],[f17248,f315])).
fof(f17248,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(sK0),X0),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X2),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),X0)),X1),s(X0,X3),s(X0,X4),s(X0,X5)))) )),
  inference(superposition,[],[f317,f227])).
fof(f227,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK0),sK2),s('type.wordSem.word_loc'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f183])).
fof(f183,plain,(
  ? [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : ~s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ~s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7)))),
  file('thm.wordSem.result_distinct|split|5','thm.wordSem.result_distinct|split|5')).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.wordSem.word_loc'(X1),X0),happ(s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s('type.wordSem.word_loc'(X1),X2))),s('type.wordSem.word_loc'(X1),X3))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X1),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X1),X2),s('type.wordSem.word_loc'(X1),X3))),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X4),s(fun('type.wordSem.word_loc'(X1),fun('type.wordSem.word_loc'(X1),X0)),X5),s(X0,X6),s(X0,X7),s(X0,X8)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X17,X6] : ! [X8,X7,X26,X35,X36,X37,X38] : s(X17,happ(s(fun('type.wordSem.word_loc'(X6),X17),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s('type.wordSem.word_loc'(X6),X8))),s('type.wordSem.word_loc'(X6),X7))) = s(X17,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X7))),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X26),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X17)),X35),s(X17,X36),s(X17,X37),s(X17,X38)))),
  file('thm.wordSem.result_distinct|split|5','thm.wordSem.result_case_def|split|1')).
fof(f88236,plain,(
  ( ! [X575,X576,X578,X574,X577] : (s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.word_loc'(sK0),bool),happ(s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s('type.wordSem.word_loc'(sK0),sK2))),s('type.wordSem.word_loc'(sK0),sK1))) | ~p(s(bool,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X574),s(fun('type.wordSem.word_loc'(sK0),fun('type.wordSem.word_loc'(sK0),bool)),X575),s(bool,X576),s(bool,X577),s(bool,X578))))) )),
  inference(superposition,[],[f242,f17248])).
fof(f242,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.wordSem.result_distinct|split|5','thm.bool.EQ_CLAUSES|split|2')).
fof(f228,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.wordSem.result_distinct|split|5',aHL_TRUTH)).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 333
% Passive clauses: 6679
% Generated clauses: 93225
% Final active clauses: 181
% Final passive clauses: 1718
% Input formulas: 68
% Initial clauses: 101
% Discarded non-redundant clauses: 722
% 
% Duplicate literals: 1641
% Trivial inequalities: 499
% Fw subsumption resolutions: 5814
% Fw demodulations: 20446
% Bw demodulations: 3890
% Condensations: 4410
% Global subsumptions: 157
% 
% Simple tautologies: 556
% Equational tautologies: 494
% Forward subsumptions: 49331
% Fw demodulations to eq. taut.: 2025
% Bw demodulations to eq. taut.: 168
% 
% Unit resulting resolution: 10603
% Factoring: 1
% Forward superposition: 21160
% Backward superposition: 19327
% Self superposition: 399
% Equality factoring: 183
% Equality resolution: 53
% 
% Split clauses: 46
% Split components: 92
% SAT solver clauses: 7791
% SAT solver unit clauses: 1666
% SAT solver binary clauses: 3140
% 
% Sat splits: 46
% Sat splitting refutations: 861
% 
% Memory used [KB]: 62813
% Time elapsed: 5.152 s
% ------------------------------
% ------------------------------
% Success in time 6.49 s
