Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.exh_to_patProof
% (139178)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 966
% Passive clauses: 5707
% Generated clauses: 18883
% Final active clauses: 952
% Final passive clauses: 4399
% Input formulas: 133
% Initial clauses: 464
% 
% Duplicate literals: 236
% Trivial inequalities: 48
% Fw subsumption resolutions: 1333
% Fw demodulations: 1357
% Bw demodulations: 6
% 
% Simple tautologies: 257
% Equational tautologies: 57
% Forward subsumptions: 9640
% 
% Binary resolution: 6723
% Factoring: 1175
% Forward superposition: 3030
% Backward superposition: 4332
% Self superposition: 17
% Equality factoring: 2
% Equality resolution: 32
% 
% Split clauses: 271
% Split components: 544
% SAT solver clauses: 542
% SAT solver unit clauses: 8
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 271
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 541
% TWLsolver calls for satisfiability: 466
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.exh_to_patProof
% (139201)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 674
% Passive clauses: 3250
% Generated clauses: 15240
% Final active clauses: 466
% Final passive clauses: 2011
% Input formulas: 133
% Initial clauses: 464
% 
% Duplicate literals: 92
% Trivial inequalities: 17
% Fw subsumption resolutions: 513
% Fw demodulations: 1980
% Bw demodulations: 637
% Condensations: 53
% Global subsumptions: 22
% 
% Simple tautologies: 248
% Equational tautologies: 39
% Forward subsumptions: 5594
% Bw demodulations to eq. taut.: 15
% 
% Binary resolution: 2479
% Factoring: 11
% Forward superposition: 3365
% Backward superposition: 2243
% Self superposition: 36
% Equality factoring: 2
% Equality resolution: 29
% 
% Split clauses: 101
% Split components: 204
% SAT solver clauses: 3673
% SAT solver unit clauses: 491
% SAT solver binary clauses: 187
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 23
% 
% Sat splits: 101
% Sat splitting refutations: 29
% 
% TWLsolver clauses: 3579
% TWLsolver calls for satisfiability: 7433
% 
% Memory used [KB]: 16502
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.exh_to_patProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.exh_to_patProof
% SZS output start Proof for thm.exh_to_patProof
fof(f12053,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f657,f1148,f1146,f656,f1147,f1145,f2219,f5057,f5067,f5066,f6035,f6038,f6040,f6042,f6044,f6048,f6050,f6054,f6056,f6068,f6070,f6082,f6084,f6095,f6106,f6125,f6154,f6156,f6165,f6170,f6198,f6207,f6218,f6221,f6226,f6228,f6247,f6252,f6306,f6308,f6317,f6322,f6350,f6359,f6370,f6373,f6378,f6380,f6399,f6404,f6445,f6447,f6449,f6458,f6460,f6462,f6478,f6480,f6488,f6490,f6506,f6508,f6516,f6518,f6534,f6536,f6544,f6546,f6559,f6561,f7290,f7291,f7292,f7293,f6586,f6588,f6590,f6594,f6597,f6601,f6603,f6605,f6609,f6612,f6620,f6626,f6634,f6636,f6638,f6645,f6647,f6654,f6656,f6663,f6665,f6672,f6674,f6682,f6699,f6716,f6759,f6761,f6772,f6774,f6776,f6778,f6798,f6800,f6802,f6811,f6813,f6815,f6817,f6819,f6838,f6840,f6850,f6859,f6861,f6877,f6879,f6887,f6889,f6905,f6907,f6915,f6917,f6933,f6935,f6943,f6945,f6959,f6976,f7005,f7007,f7020,f7022,f7035,f7037,f7050,f7070,f7072,f7083,f7085,f7087,f7101,f7103,f7112,f7119,f7126,f7139,f7142,f7144,f7146,f7148,f7150,f7152,f7154,f7156,f7158,f7175,f7177,f7179,f7181,f7192,f7194,f7204,f7206,f7211,f7213,f7217,f7219,f7221,f7223,f7227,f7229,f7231,f7233,f7237,f7239,f7241,f7243,f7247,f7249,f7251,f7253,f7255,f7259,f7261,f7263,f7265,f7267,f7269,f7271,f7273,f7305,f7306,f7307,f7308,f7309,f7310,f7311,f7312,f7313,f7314,f7315,f7316,f7317,f7318,f7319,f7320,f7321,f7322,f7323,f7324,f7325,f7326,f7327,f7328,f7329,f7330,f7331,f7332,f7333,f7334,f7335,f7336,f7337,f7338,f7339,f7340,f7341,f7342,f7343,f7344,f7345,f7346,f7347,f7348,f7349,f7350,f7351,f7352,f7353,f7354,f7355,f7356,f7357,f7358,f7359,f7360,f7361,f7362,f7363,f7364,f7365,f7366,f7367,f7368,f7369,f7370,f7371,f7372,f7373,f7374,f7375,f7376,f7377,f7378,f7379,f7380,f7381,f7382,f7383,f7384,f7385,f7386,f7387,f7388,f7389,f7390,f7391,f7392,f7393,f7394,f7395,f7396,f7397,f7398,f7399,f7400,f7401,f7402,f7403,f7404,f7405,f7406,f7407,f7408,f7409,f7410,f7411,f7412,f7413,f7414,f7415,f7416,f7417,f7418,f7419,f7420,f7421,f7422,f7423,f7424,f7425,f7426,f7427,f7428,f7429,f7430,f7431,f7432,f7433,f7434,f7435,f7436,f7437,f7438,f7439,f7440,f7441,f7442,f7443,f7444,f7445,f7446,f7447,f7448,f7449,f7450,f7451,f7452,f7453,f7454,f7455,f7456,f7457,f7458,f7459,f7460,f7461,f7462,f7463,f7464,f7465,f7466,f7467,f7468,f7469,f7470,f7471,f7472,f7473,f7474,f7475,f7476,f7477,f7478,f7479,f7482,f7483,f7484,f7485,f7486,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7495,f7496,f7497,f7498,f7499,f7500,f7501,f7502,f7503,f7504,f7505,f7506,f7507,f7508,f7509,f7510,f7511,f7512,f7513,f7514,f7515,f7516,f7517,f7518,f7519,f7520,f7521,f7522,f7523,f7524,f7525,f7526,f7527,f7528,f7529,f7530,f7531,f7532,f7533,f7534,f7535,f7536,f7537,f7538,f7539,f7540,f7541,f7542,f7543,f7544,f7545,f7546,f7547,f7548,f7549,f7550,f7551,f7552,f7553,f7554,f7555,f7556,f7557,f7558,f7559,f7560,f7561,f7562,f7563,f7564,f7565,f7566,f7567,f7568,f7569,f7570,f7571,f7572,f7573,f7574,f7575,f7576,f7577,f7578,f7579,f7580,f7581,f7582,f7583,f7584,f7585,f7586,f7587,f7588,f7589,f7933,f7935,f8713,f8715,f9354,f9355,f9403,f9404,f9463,f9464,f9516,f9517,f6036,f12052])).
fof(f12052,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12051,f1148])).
fof(f12051,plain,(
  s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) | $spl0),
  inference(backward_demodulation,[],[f11960,f11759])).
fof(f11759,plain,(
  s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182('const.patSem.Loc_1'(s('type.num.num',sK29)),sK28)))) | $spl0),
  inference(unit_resulting_resolution,[],[f2430,f2220,f3887,f3857,f8342,f939,f886])).
fof(f886,plain,(
  ( ! [X0,X1] : (~sP15(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182(X1,X0)))) | sP13(X0,X1) | sP12(X0,X1) | sP11(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK181(X1,X0)))) | sP14(X0,X1)) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0,X1] : ((~sP15(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK181(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK181(X1,X0))))) | sP11(X0,X1) | sP12(X0,X1) | sP13(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182(X1,X0))))) | sP14(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP11(X0,X1) & ~sP12(X0,X1) & ~sP13(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP14(X0,X1)) | sP15(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK181,sK182])],[f560])).
fof(f560,plain,(
  ! [X0,X1] : ((~sP15(X0,X1) | ? [X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X5)))) | sP11(X0,X1) | sP12(X0,X1) | sP13(X0,X1) | ? [X4] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP14(X0,X1)) & ((! [X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X3)))) & ~sP11(X0,X1) & ~sP12(X0,X1) & ~sP13(X0,X1) & ! [X2] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2)))) & ~sP14(X0,X1)) | sP15(X0,X1)))),
  inference(rectify,[],[f559])).
fof(f559,plain,(
  ! [X1,X0] : ((~sP15(X1,X0) | ? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP11(X1,X0) | sP12(X1,X0) | sP13(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP14(X1,X0)) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP11(X1,X0) & ~sP12(X1,X0) & ~sP13(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP14(X1,X0)) | sP15(X1,X0)))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X1,X0] : ((~sP15(X1,X0) | (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP11(X1,X0) | sP12(X1,X0) | sP13(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP14(X1,X0))) & ((! [X17] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) & ~sP11(X1,X0) & ~sP12(X1,X0) & ~sP13(X1,X0) & ! [X4] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) & ~sP14(X1,X0)) | sP15(X1,X0)))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X1,X0] : (sP15(X1,X0) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | sP11(X1,X0) | sP12(X1,X0) | sP13(X1,X0) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | sP14(X1,X0)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f939,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X57,X56] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X57))) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X56)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.v_distinct|split|3')).
fof(f8342,plain,(
  ( ! [X0,X1] : (~sP13(X0,'const.patSem.Loc_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f1105,f900])).
fof(f900,plain,(
  ( ! [X0,X1] : (~sP13(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK185(X1,X0)),s('type.list.list'('type.patLang.exp'),sK187(X1,X0)),s('type.num.num',sK189(X1,X0))))) )),
  inference(cnf_transformation,[],[f567])).
fof(f567,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK185(X1,X0)),s('type.list.list'('type.patLang.exp'),sK187(X1,X0)),s('type.num.num',sK189(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),sK186(X1,X0)),s('type.list.list'('type.patLang.exp'),sK188(X1,X0)),s('type.num.num',sK189(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK187(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK185(X1,X0)))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK188(X1,X0)))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK186(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),sK187(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK185(X1,X0)),s('type.list.list'('type.patSem.v'),sK186(X1,X0)))))))),s('type.list.list'('type.patLang.exp'),sK187(X1,X0)),s('type.list.list'('type.patLang.exp'),sK188(X1,X0))))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP13(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK185,sK186,sK187,sK188,sK189])],[f566])).
fof(f566,plain,(
  ! [X0,X1] : ((~sP13(X0,X1) | ? [X7,X8,X9,X10,X11] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patLang.exp'),X9),s('type.num.num',X11))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X11))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X7),s('type.list.list'('type.patSem.v'),X8))))))),s('type.list.list'('type.patLang.exp'),X9),s('type.list.list'('type.patLang.exp'),X10)))))) & (! [X2,X3,X4,X5,X6] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X4),s('type.num.num',X6))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patLang.exp'),X5),s('type.num.num',X6))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X5))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X4))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.list.list'('type.patLang.exp'),X4),s('type.list.list'('type.patLang.exp'),X5))))) | sP13(X0,X1)))),
  inference(rectify,[],[f565])).
fof(f565,plain,(
  ! [X1,X0] : ((~sP13(X1,X0) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8)))))) & (! [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | sP13(X1,X0)))),
  inference(nnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X1,X0] : (sP13(X1,X0) <=> ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f1105,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2,X3] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) != s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2,X3] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X3))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patLang.exp'),X1),s('type.num.num',X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X58,X28,X27,X56] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X56))) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X27),s('type.list.list'('type.patLang.exp'),X28),s('type.num.num',X58)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.v_distinct|split|12')).
fof(f3857,plain,(
  ( ! [X0,X1] : (~sP12(X0,'const.patSem.Loc_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f1064,f904])).
fof(f904,plain,(
  ( ! [X0,X1] : (~sP12(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK190(X1,X0)),s('type.patLang.exp',sK192(X1,X0))))) )),
  inference(cnf_transformation,[],[f570])).
fof(f570,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK190(X1,X0)),s('type.patLang.exp',sK192(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),sK191(X1,X0)),s('type.patLang.exp',sK193(X1,X0)))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK190(X1,X0)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),sK191(X1,X0)))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK190(X1,X0)),s('type.list.list'('type.patSem.v'),sK191(X1,X0)))))))),s('type.patLang.exp',sK192(X1,X0)))),s('type.patLang.exp',sK193(X1,X0))))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP12(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK190,sK191,sK192,sK193])],[f569])).
fof(f569,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ? [X6,X7,X8,X9] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X6),s('type.patLang.exp',X8))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X7),s('type.patLang.exp',X9))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7))))))),s('type.patLang.exp',X8))),s('type.patLang.exp',X9)))))) & (! [X2,X3,X4,X5] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X2),s('type.patLang.exp',X4))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X3),s('type.patLang.exp',X5))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X2))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X3))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),s('type.patLang.exp',X4))),s('type.patLang.exp',X5))))) | sP12(X0,X1)))),
  inference(rectify,[],[f568])).
fof(f568,plain,(
  ! [X1,X0] : ((~sP12(X1,X0) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13)))))) & (! [X10,X11,X12,X13] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) | ~p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | sP12(X1,X0)))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X1,X0] : (sP12(X1,X0) <=> ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1064,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X1),s('type.patLang.exp',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X28,X27,X56] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X56))) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X27),s('type.patLang.exp',X28)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.v_distinct|split|10')).
fof(f3887,plain,(
  ( ! [X0,X1] : (~sP11(X0,'const.patSem.Loc_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f1062,f908])).
fof(f908,plain,(
  ( ! [X0,X1] : (~sP11(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK194(X1,X0)),s('type.list.list'('type.patSem.v'),sK195(X1,X0))))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK194(X1,X0)),s('type.list.list'('type.patSem.v'),sK195(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',sK194(X1,X0)),s('type.list.list'('type.patSem.v'),sK196(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK195(X1,X0)),s('type.list.list'('type.patSem.v'),sK196(X1,X0))))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP11(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK194,sK195,sK196])],[f572])).
fof(f572,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ? [X5,X6,X7] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X6))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X5),s('type.list.list'('type.patSem.v'),X7))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patSem.v'),X7)))))) & (! [X2,X3,X4] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X3))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X2),s('type.list.list'('type.patSem.v'),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X3),s('type.list.list'('type.patSem.v'),X4))))) | sP11(X0,X1)))),
  inference(rectify,[],[f571])).
fof(f571,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16)))))) & (! [X14,X15,X16] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1062,plain,(
  ( ! [X2,X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) != s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X2))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X1),s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X28,X27,X56] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X56))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X27),s('type.list.list'('type.patSem.v'),X28)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.v_distinct|split|7')).
fof(f2220,plain,(
  sP15(sK28,'const.patSem.Loc_1'(s('type.num.num',sK29))) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1151])).
fof(f1151,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP15(X1,X0)) )),
  inference(forward_demodulation,[],[f912,f837])).
fof(f837,plain,(
  ( ! [X0,X1] : (s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X8,X9] : s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X8),s('type.patSem.v',X9))) = s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X8))),s('type.patSem.v',X9)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','ARITY_0_2_const.exh_to_patProof.v_rel')).
fof(f912,plain,(
  ( ! [X0,X1] : (sP15(X1,X0) | ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f574])).
fof(f574,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | sP15(X1,X0)) & (~sP15(X1,X0) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> sP15(X1,X0))),
  inference(definition_folding,[],[f232,f419,f418,f417,f416,f415])).
fof(f418,plain,(
  ! [X1,X0] : (sP14(X1,X0) <=> ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f232,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) <=> (? [X17] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X17)))) | ? [X14,X15,X16] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X15))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X14),s('type.list.list'('type.patSem.v'),X16))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X15),s('type.list.list'('type.patSem.v'),X16))))) | ? [X10,X11,X12,X13] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X10),s('type.patLang.exp',X12))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X11),s('type.patLang.exp',X13))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X10))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X11))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X10),s('type.list.list'('type.patSem.v'),X11))))))),s('type.patLang.exp',X12))),s('type.patLang.exp',X13))))) | ? [X5,X6,X7,X8,X9] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patLang.exp'),X7),s('type.num.num',X9))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X6),s('type.list.list'('type.patLang.exp'),X8),s('type.num.num',X9))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X5))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X8))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X6))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X7))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),s('type.list.list'('type.patLang.exp'),X7),s('type.list.list'('type.patLang.exp'),X8))))) | ? [X4] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X4)))) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X27,X28] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X27))),s('type.patSem.v',X28)))) <=> (? [X7] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X18,X19,X20] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20))))) | ? [X21,X22,X23,X24] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X22),s('type.patLang.exp',X24))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.patLang.exp',X23))),s('type.patLang.exp',X24))))) | ? [X21,X22,X25,X26,X13] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X13))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X22),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X26))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.list.list'('type.patLang.exp'),X25),s('type.list.list'('type.patLang.exp'),X26))))) | ? [X13] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13)))) | ? [X19,X20] : (s('type.patSem.v',X27) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X28) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20)))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_cases')).
fof(f2430,plain,(
  ( ! [X0,X1] : (~sP14(X0,'const.patSem.Loc_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f940,f896])).
fof(f896,plain,(
  ( ! [X0,X1] : (~sP14(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK183(X1,X0))))) )),
  inference(cnf_transformation,[],[f564])).
fof(f564,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK183(X1,X0)))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),sK184(X1,X0)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),sK183(X1,X0)),s('type.list.list'('type.patSem.v'),sK184(X1,X0))))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP14(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK183,sK184])],[f563])).
fof(f563,plain,(
  ! [X0,X1] : ((~sP14(X0,X1) | ? [X4,X5] : (s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X4))) & s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X4),s('type.list.list'('type.patSem.v'),X5)))))) & (! [X2,X3] : (s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP14(X0,X1)))),
  inference(rectify,[],[f562])).
fof(f562,plain,(
  ! [X1,X0] : ((~sP14(X1,X0) | ? [X2,X3] : (s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) & s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3)))))) & (! [X2,X3] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X2))) | s('type.patSem.v',X1) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X3))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.list.list'('type.patSem.v'),X2),s('type.list.list'('type.patSem.v'),X3))))) | sP14(X1,X0)))),
  inference(nnf_transformation,[],[f418])).
fof(f940,plain,(
  ( ! [X0,X1] : (s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) != s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X0)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X57,X56] : ~s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X56))) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X57)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.v_distinct|split|14')).
fof(f11960,plain,(
  s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182('const.patSem.Loc_1'(s('type.num.num',sK29)),sK28)))) | $spl0),
  inference(unit_resulting_resolution,[],[f2430,f2220,f3887,f3857,f8342,f939,f888])).
fof(f888,plain,(
  ( ! [X0,X1] : (~sP15(X0,X1) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK182(X1,X0)))) | sP13(X0,X1) | sP12(X0,X1) | sP11(X0,X1) | s('type.patSem.v',X1) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',sK181(X1,X0)))) | sP14(X0,X1)) )),
  inference(cnf_transformation,[],[f561])).
fof(f6036,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f6036_D])).
fof(f6036_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9517,plain,(
  ( ! [X5] : (p(s(bool,sK145(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f9517_D])).
fof(f9517_D,plain,(
  ( ! [X5] : (p(s(bool,sK145(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f9516,plain,(
  ( ! [X4,X5] : (p(s(bool,sK145(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9504,f781])).
fof(f781,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X122] : ! [X124] : s(X122,X124) = s(X122,happ(s(fun(X122,X122),'_i_0_0'),s(X122,X124)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','_i0_ax')).
fof(f9504,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK145(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f806,f781])).
fof(f806,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK145(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f513])).
fof(f513,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK144(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK145(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK144,sK145])],[f512])).
fof(f512,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f511])).
fof(f511,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X37] : ! [X39,X38] : ((? [X34] : p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34)))) | p(s(bool,X38))) <=> ? [X34] : (p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34)))) | p(s(bool,X38))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f9464,plain,(
  ( ! [X5] : (p(s(bool,sK144(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) | $spl20),
  inference(cnf_transformation,[],[f9464_D])).
fof(f9464_D,plain,(
  ( ! [X5] : (p(s(bool,sK144(X5,'_i_0_0',bool))) | p(s(bool,X5))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9463,plain,(
  ( ! [X4,X5] : (p(s(bool,sK144(X5,'_i_0_0',bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9451,f781])).
fof(f9451,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK144(X5,'_i_0_0',bool))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f804,f781])).
fof(f804,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK144(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f513])).
fof(f9404,plain,(
  ( ! [X5] : (p(s(bool,sK143('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f9404_D])).
fof(f9404_D,plain,(
  ( ! [X5] : (p(s(bool,sK143('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f9403,plain,(
  ( ! [X4,X5] : (p(s(bool,sK143('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9391,f781])).
fof(f9391,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK143('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f803,f781])).
fof(f803,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK143(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f509])).
fof(f509,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK142(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK143(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK142,sK143])],[f508])).
fof(f508,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f507])).
fof(f507,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f506])).
fof(f506,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X37] : ! [X39,X38] : ((p(s(bool,X39)) | ? [X34] : p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X34))))) <=> ? [X34] : (p(s(bool,X39)) | p(s(bool,happ(s(fun(X37,bool),X38),s(X37,X34))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f9355,plain,(
  ( ! [X5] : (p(s(bool,sK142('_i_0_0',X5,bool))) | p(s(bool,X5))) ) | $spl16),
  inference(cnf_transformation,[],[f9355_D])).
fof(f9355_D,plain,(
  ( ! [X5] : (p(s(bool,sK142('_i_0_0',X5,bool))) | p(s(bool,X5))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f9354,plain,(
  ( ! [X4,X5] : (p(s(bool,sK142('_i_0_0',X5,bool))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f9342,f781])).
fof(f9342,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK142('_i_0_0',X5,bool)))))) )),
  inference(superposition,[],[f801,f781])).
fof(f801,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK142(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f509])).
fof(f8715,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | $spl14),
  inference(cnf_transformation,[],[f8715_D])).
fof(f8715_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8713,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f8702,f781])).
fof(f8702,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'_i_0_0'))))))) )),
  inference(superposition,[],[f793,f781])).
fof(f793,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X37] : ! [X39,X34] : (p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34)))) => p(s(bool,happ(s(fun(X37,bool),X39),s(X37,'const.min.@_1'(s(fun(X37,bool),X39)))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.SELECT_AX')).
fof(f7935,plain,(
  p(s(bool,sK138('_i_0_0',bool))) | $spl12),
  inference(cnf_transformation,[],[f7935_D])).
fof(f7935_D,plain,(
  p(s(bool,sK138('_i_0_0',bool))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f7933,plain,(
  ( ! [X3] : (p(s(bool,sK138('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f7922,f781])).
fof(f7922,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK138('_i_0_0',bool)))))) )),
  inference(superposition,[],[f787,f781])).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK138(X1,X0)))))) )),
  inference(cnf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK137(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK138(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK137,sK138])],[f496])).
fof(f496,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f495])).
fof(f495,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X37] : ! [X39] : (~? [X34] : p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34)))) <=> ! [X34] : ~p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.NOT_EXISTS_THM')).
fof(f7589,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6049,f1149])).
fof(f1149,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f661,f658])).
fof(f658,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.TRUTH')).
fof(f661,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f60])).
fof(f60,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f6049,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1196,f6036,f838])).
fof(f838,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f369])).
fof(f369,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X30,X31] : ((p(s(bool,X30)) => p(s(bool,X31))) => ((p(s(bool,X31)) => p(s(bool,X30))) => s(bool,X30) = s(bool,X31)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f1196,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1149,f726])).
fof(f726,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X29] : (s(bool,'T_0') = s(bool,X29) <=> p(s(bool,X29)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f7588,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6055,f1149])).
fof(f6055,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1196,f6036,f838])).
fof(f7587,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6057,f1149])).
fof(f6057,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1234,f6036,f929])).
fof(f929,plain,(
  ( ! [X0,X1] : (~sP17(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1] : ((~sP17(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP17(X0,X1)))),
  inference(rectify,[],[f582])).
fof(f582,plain,(
  ! [X1,X0] : ((~sP17(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP17(X1,X0)))),
  inference(flattening,[],[f581])).
fof(f581,plain,(
  ! [X1,X0] : ((~sP17(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP17(X1,X0)))),
  inference(nnf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X1,X0] : (sP17(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f1234,plain,(
  sP17('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1149,f658,f931])).
fof(f931,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP17(X0,X1)) )),
  inference(cnf_transformation,[],[f583])).
fof(f7586,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6058,f6036])).
fof(f6058,plain,(
  ( ! [X0] : (p(s(bool,sK139(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f5078,f6036,f929])).
fof(f5078,plain,(
  ( ! [X0] : (sP17('T_0',sK139(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f658,f5067,f931])).
fof(f7585,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6060,f6036])).
fof(f6060,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2235,f6036,f929])).
fof(f2235,plain,(
  sP17('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f934])).
fof(f934,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP17(X0,X1)) )),
  inference(cnf_transformation,[],[f583])).
fof(f7584,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6061,f6036])).
fof(f6061,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1265,f6036,f929])).
fof(f1265,plain,(
  ( ! [X0] : (sP17('F_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1149,f1246])).
fof(f1246,plain,(
  ( ! [X0,X1] : (sP17(X1,'const.bool.BOUNDED_1'(s(bool,X0))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1244,f658])).
fof(f1244,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP17(X1,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(superposition,[],[f934,f665])).
fof(f665,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X63] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X63)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.BOUNDED_THM')).
fof(f7583,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6062,f6036])).
fof(f6062,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1465,f6036,f929])).
fof(f1465,plain,(
  ( ! [X0] : (sP17('F_0','const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f1149,f666,f934])).
fof(f666,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X37] : p(s(bool,'const.ind_type.ISO_2'(s(fun(X37,X37),'_i_0_0'),s(fun(X37,X37),'_i_0_0'))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.ind_type.ISO_REFL')).
fof(f7582,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6063,f1149])).
fof(f6063,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2234,f6036,f929])).
fof(f2234,plain,(
  sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f931])).
fof(f7581,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6064,f1149])).
fof(f6064,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6036,f929])).
fof(f1249,plain,(
  ( ! [X0] : (sP17('const.bool.BOUNDED_1'(s(bool,X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1149,f1237])).
fof(f1237,plain,(
  ( ! [X0,X1] : (sP17('const.bool.BOUNDED_1'(s(bool,X0)),X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1235,f658])).
fof(f1235,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP17('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f931,f665])).
fof(f7580,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6065,f1149])).
fof(f6065,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1464,f6036,f929])).
fof(f1464,plain,(
  ( ! [X0] : (sP17('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1149,f666,f931])).
fof(f7579,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6073,f1149])).
fof(f6073,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1243,f6036,f929])).
fof(f1243,plain,(
  sP17('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1149,f658,f934])).
fof(f7578,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6074,f1149])).
fof(f6074,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2235,f6036,f929])).
fof(f7577,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6075,f1149])).
fof(f6075,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1265,f6036,f929])).
fof(f7576,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6076,f1149])).
fof(f6076,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1465,f6036,f929])).
fof(f7575,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6077,f6036])).
fof(f6077,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2234,f6036,f929])).
fof(f7574,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6078,f6036])).
fof(f6078,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1249,f6036,f929])).
fof(f7573,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6079,f6036])).
fof(f6079,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1464,f6036,f929])).
fof(f7572,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6080,f6036])).
fof(f6080,plain,(
  ( ! [X0] : (p(s(bool,sK139(X0,'_i_0_0',bool)))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f5079,f6036,f929])).
fof(f5079,plain,(
  ( ! [X0] : (sP17(sK139(X0,'_i_0_0',bool),'T_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f658,f5067,f934])).
fof(f7571,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6124,f6302])).
fof(f6302,plain,(
  ( ! [X0] : (sP21('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1164,f1169,f6036,f994])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (~sP19(X1,X2,X0) | p(s(bool,X0)) | ~sP20(X0,X2,X1) | sP21(X0,X1,X2) | ~sP18(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0,X1,X2] : ((~sP21(X0,X1,X2) | (sP18(X2,X1,X0) & sP19(X1,X2,X0) & sP20(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP18(X2,X1,X0) | ~sP19(X1,X2,X0) | ~sP20(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP21(X0,X1,X2)))),
  inference(rectify,[],[f611])).
fof(f611,plain,(
  ! [X2,X1,X0] : ((~sP21(X2,X1,X0) | (sP18(X0,X1,X2) & sP19(X1,X0,X2) & sP20(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP18(X0,X1,X2) | ~sP19(X1,X0,X2) | ~sP20(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP21(X2,X1,X0)))),
  inference(flattening,[],[f610])).
fof(f610,plain,(
  ! [X2,X1,X0] : ((~sP21(X2,X1,X0) | (sP18(X0,X1,X2) & sP19(X1,X0,X2) & sP20(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP18(X0,X1,X2) | ~sP19(X1,X0,X2) | ~sP20(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP21(X2,X1,X0)))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X2,X1,X0] : (sP21(X2,X1,X0) <=> (sP18(X0,X1,X2) & sP19(X1,X0,X2) & sP20(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP21])])).
fof(f1169,plain,(
  ( ! [X0,X1] : (sP19(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1000])).
fof(f1000,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f618])).
fof(f618,plain,(
  ! [X0,X1,X2] : ((~sP19(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP19(X0,X1,X2)))),
  inference(rectify,[],[f617])).
fof(f617,plain,(
  ! [X1,X0,X2] : ((~sP19(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP19(X1,X0,X2)))),
  inference(flattening,[],[f616])).
fof(f616,plain,(
  ! [X1,X0,X2] : ((~sP19(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP19(X1,X0,X2)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X1,X0,X2] : (sP19(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP19])])).
fof(f1164,plain,(
  ( ! [X0,X1] : (sP20(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f996])).
fof(f996,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1,X2] : ((~sP20(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP20(X0,X1,X2)))),
  inference(rectify,[],[f614])).
fof(f614,plain,(
  ! [X2,X0,X1] : ((~sP20(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP20(X2,X0,X1)))),
  inference(flattening,[],[f613])).
fof(f613,plain,(
  ! [X2,X0,X1] : ((~sP20(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP20(X2,X0,X1)))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X2,X0,X1] : (sP20(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP20])])).
fof(f1174,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1004])).
fof(f1004,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(flattening,[],[f619])).
fof(f619,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2] : (sP18(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f6124,plain,(
  ( ! [X0] : (~sP21('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f658,f6036,f991])).
fof(f991,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f612])).
fof(f7570,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6126,f6461])).
fof(f6461,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1002])).
fof(f1002,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f618])).
fof(f6126,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f6036,f993])).
fof(f993,plain,(
  ( ! [X2,X0,X1] : (~sP19(X1,X2,X0) | p(s(bool,X1)) | ~sP20(X0,X2,X1) | sP21(X0,X1,X2) | ~sP18(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f612])).
fof(f2049,plain,(
  ~sP21('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f658,f1196,f1009])).
fof(f1009,plain,(
  ( ! [X2,X0,X1] : (~sP21(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP21(X2,X1,X0)) & (~sP21(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP21(X2,X1,X0))),
  inference(definition_folding,[],[f288,f428,f427,f426,f425])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X134,X138,X139] : ((p(s(bool,X139)) <=> s(bool,X134) = s(bool,X138)) <=> ((p(s(bool,X139)) | p(s(bool,X138)) | p(s(bool,X134))) & (p(s(bool,X139)) | ~p(s(bool,X134)) | ~p(s(bool,X138))) & (p(s(bool,X138)) | ~p(s(bool,X134)) | ~p(s(bool,X139))) & (p(s(bool,X134)) | ~p(s(bool,X138)) | ~p(s(bool,X139)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sat.sat_2|49|_prove_ltb')).
fof(f1175,plain,(
  ( ! [X0,X1] : (sP18(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f658,f1005])).
fof(f1005,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f620])).
fof(f7569,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6127,f6461])).
fof(f6127,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f6036,f993])).
fof(f5080,plain,(
  ( ! [X0] : (~sP21('T_0','T_0',sK139(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f658,f658,f5067,f991])).
fof(f7568,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6128,f6461])).
fof(f6128,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f1164,f6036,f993])).
fof(f7567,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6132,f6461])).
fof(f6132,plain,(
  ~sP19('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f6036,f993])).
fof(f2249,plain,(
  ~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1196,f1145,f1009])).
fof(f7566,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6134,f6461])).
fof(f6134,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f6036,f993])).
fof(f1165,plain,(
  ( ! [X0,X1] : (sP20(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f997])).
fof(f997,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f615])).
fof(f7565,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6135,f6461])).
fof(f6135,plain,(
  ~sP19('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f6036,f993])).
fof(f7564,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6136,f6461])).
fof(f6136,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f6036,f993])).
fof(f5096,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,sK139(X1,'_i_0_0',bool),X2)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f5067,f997])).
fof(f7563,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6137,f6461])).
fof(f6137,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f5096,f6036,f993])).
fof(f7562,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6144,f6448])).
fof(f6448,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f998])).
fof(f998,plain,(
  ( ! [X2,X0,X1] : (sP20(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f615])).
fof(f6144,plain,(
  ~sP20('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1169,f6036,f993])).
fof(f7561,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6145,f6448])).
fof(f6145,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1169,f6036,f993])).
fof(f7560,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6146,f6448])).
fof(f6146,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f1169,f6036,f993])).
fof(f7559,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6171,f6448])).
fof(f6171,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f2237,f6036,f993])).
fof(f2237,plain,(
  ( ! [X0,X1] : (sP19(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1000])).
fof(f7558,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6211,f6448])).
fof(f6211,plain,(
  ~sP20('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1170,f6036,f993])).
fof(f1170,plain,(
  ( ! [X0,X1] : (sP19(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f1001])).
fof(f1001,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f618])).
fof(f7557,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6212,f6448])).
fof(f6212,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1170,f6036,f993])).
fof(f7556,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6240,f6448])).
fof(f6240,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5101,f6036,f993])).
fof(f5101,plain,(
  ( ! [X2,X0,X1] : (sP19(X0,sK139(X1,'_i_0_0',bool),X2)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f5067,f1001])).
fof(f7555,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6241,f6448])).
fof(f6241,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f5101,f6036,f993])).
fof(f7554,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6278,f6461])).
fof(f6278,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f6036,f994])).
fof(f7553,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6279,f6461])).
fof(f6279,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f6036,f994])).
fof(f7552,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6280,f6461])).
fof(f6280,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f1164,f6036,f994])).
fof(f7551,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6284,f6461])).
fof(f6284,plain,(
  ~sP19('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f6036,f994])).
fof(f7550,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6286,f6461])).
fof(f6286,plain,(
  ~sP19('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f6036,f994])).
fof(f7549,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6287,f6461])).
fof(f6287,plain,(
  ~sP19('T_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f6036,f994])).
fof(f7548,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6288,f6461])).
fof(f6288,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f6036,f994])).
fof(f7547,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6289,f6461])).
fof(f6289,plain,(
  ( ! [X0] : (~sP19('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f5096,f6036,f994])).
fof(f7546,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6296,f6448])).
fof(f6296,plain,(
  ~sP20('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1169,f6036,f994])).
fof(f7545,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6297,f6448])).
fof(f6297,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1169,f6036,f994])).
fof(f7544,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6298,f6448])).
fof(f6298,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f1169,f6036,f994])).
fof(f7543,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6323,f6448])).
fof(f6323,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f2237,f6036,f994])).
fof(f7542,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6348,f6602])).
fof(f6602,plain,(
  ( ! [X0] : (~sP21(X0,'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2226,f6036,f1010])).
fof(f1010,plain,(
  ( ! [X2,X0,X1] : (~sP21(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f621])).
fof(f2226,plain,(
  s(bool,'T_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f841])).
fof(f841,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f370])).
fof(f6348,plain,(
  sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f1164,f2237,f6036,f994])).
fof(f2240,plain,(
  ( ! [X0,X1] : (sP18('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1006])).
fof(f1006,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f620])).
fof(f7541,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6363,f6448])).
fof(f6363,plain,(
  ~sP20('T_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1170,f6036,f994])).
fof(f7540,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6364,f6448])).
fof(f6364,plain,(
  ~sP20('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1170,f6036,f994])).
fof(f7539,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6392,f6448])).
fof(f6392,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5101,f6036,f994])).
fof(f7538,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6393,f6448])).
fof(f6393,plain,(
  ( ! [X0] : (~sP20('T_0',sK139(X0,'_i_0_0',bool),'T_0')) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f5080,f5101,f6036,f994])).
fof(f7537,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6401,f5080])).
fof(f6401,plain,(
  ( ! [X0] : (sP21('T_0','T_0',sK139(X0,'_i_0_0',bool))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1174,f1164,f5101,f6036,f994])).
fof(f7536,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6444,f6448])).
fof(f6444,plain,(
  ( ! [X0] : (~sP20('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f658,f6036,f995])).
fof(f995,plain,(
  ( ! [X2,X0,X1] : (~sP20(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f615])).
fof(f7535,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6457,f6461])).
fof(f6457,plain,(
  ( ! [X0] : (~sP19('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f658,f6036,f999])).
fof(f999,plain,(
  ( ! [X2,X0,X1] : (~sP19(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f618])).
fof(f7534,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6464,f6036])).
fof(f6464,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f1149,f6036,f1003])).
fof(f1003,plain,(
  ( ! [X2,X0,X1] : (~sP18(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f620])).
fof(f2239,plain,(
  ( ! [X0,X1] : (sP18(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1005])).
fof(f7533,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6465,f6036])).
fof(f6465,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f1149,f6036,f1003])).
fof(f1176,plain,(
  ( ! [X0,X1] : (sP18('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f658,f1006])).
fof(f7532,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6466,f6036])).
fof(f6466,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f1149,f6036,f1003])).
fof(f7531,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6467,f6036])).
fof(f6467,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f1003])).
fof(f7530,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6468,f6036])).
fof(f6468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f2238,plain,(
  ( ! [X0,X1] : (sP18(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1004])).
fof(f7529,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6470,f6036])).
fof(f6470,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7528,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6471,f6036])).
fof(f6471,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f6036,f6036,f1003])).
fof(f7527,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6472,f6036])).
fof(f6472,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7526,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6474,f6036])).
fof(f6474,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f1149,f6036,f1003])).
fof(f7525,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6475,f6036])).
fof(f6475,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f1149,f6036,f1003])).
fof(f7524,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6476,f6036])).
fof(f6476,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f1149,f6036,f1003])).
fof(f7523,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6482,f6036])).
fof(f6482,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f7522,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6483,f6036])).
fof(f6483,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f6036,f6036,f1003])).
fof(f7521,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6484,f6036])).
fof(f6484,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7520,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6485,f6036])).
fof(f6485,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f6036,f6036,f1003])).
fof(f7519,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6486,f6036])).
fof(f6486,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7518,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6492,f6036])).
fof(f6492,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f1149,f6036,f1003])).
fof(f7517,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6493,f6036])).
fof(f6493,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f1149,f6036,f1003])).
fof(f7516,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6494,f6036])).
fof(f6494,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f1149,f6036,f1003])).
fof(f7515,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6496,f6036])).
fof(f6496,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f7514,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6497,f6036])).
fof(f6497,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f6036,f6036,f1003])).
fof(f7513,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6498,f6036])).
fof(f6498,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7512,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6499,f6036])).
fof(f6499,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f6036,f6036,f1003])).
fof(f7511,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6500,f6036])).
fof(f6500,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7510,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6501,f6036])).
fof(f6501,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f1149,f6036,f1003])).
fof(f7509,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6502,f6036])).
fof(f6502,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f1149,f6036,f1003])).
fof(f7508,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6504,f6036])).
fof(f6504,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f1149,f6036,f1003])).
fof(f7507,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6509,f6036])).
fof(f6509,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f1003])).
fof(f7506,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6510,f6036])).
fof(f6510,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f7505,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6511,f6036])).
fof(f6511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f6036,f6036,f1003])).
fof(f7504,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6512,f6036])).
fof(f6512,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7503,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6514,f6036])).
fof(f6514,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7502,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6519,f6036])).
fof(f6519,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1149,f6036,f1003])).
fof(f7501,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6520,f6036])).
fof(f6520,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f1149,f6036,f1003])).
fof(f7500,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6522,f6036])).
fof(f6522,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f1149,f6036,f1003])).
fof(f7499,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6523,f6036])).
fof(f6523,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f1003])).
fof(f7498,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6524,f6036])).
fof(f6524,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f7497,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6526,f6036])).
fof(f6526,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7496,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6527,f6036])).
fof(f6527,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f6036,f6036,f1003])).
fof(f7495,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6528,f6036])).
fof(f6528,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7494,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6529,f6036])).
fof(f6529,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1149,f6036,f1003])).
fof(f7493,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6530,f6036])).
fof(f6530,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f1149,f6036,f1003])).
fof(f7492,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6532,f6036])).
fof(f6532,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f1149,f6036,f1003])).
fof(f7491,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6537,f6036])).
fof(f6537,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f1003])).
fof(f7490,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6538,f6036])).
fof(f6538,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2238,f6036,f6036,f1003])).
fof(f7489,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6539,f6036])).
fof(f6539,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f6036,f6036,f1003])).
fof(f7488,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6540,f6036])).
fof(f6540,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2239,f6036,f6036,f1003])).
fof(f7487,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6542,f6036])).
fof(f6542,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2240,f6036,f6036,f1003])).
fof(f7486,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6547,f1196])).
fof(f6547,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f2045,f6036,f1007])).
fof(f1007,plain,(
  ( ! [X2,X0,X1] : (sP21(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f621])).
fof(f2045,plain,(
  ~sP21('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f658,f1196,f1009])).
fof(f7485,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6548,f1196])).
fof(f6548,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f2049,f6036,f1007])).
fof(f7484,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6549,f5261])).
fof(f5261,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK139(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f658,f5067,f1131])).
fof(f1131,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1130])).
fof(f1130,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1078])).
fof(f1078,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f646])).
fof(f646,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f645])).
fof(f645,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f385])).
fof(f385,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X34,X135,X33,X136] : ((s(bool,X34) = s(bool,X135) & (p(s(bool,X135)) => s(bool,X33) = s(bool,X136))) => ((p(s(bool,X34)) => p(s(bool,X33))) <=> (p(s(bool,X135)) => p(s(bool,X136)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.IMP_CONG')).
fof(f6549,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK139(X0,'_i_0_0',bool))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f5080,f6036,f1007])).
fof(f7483,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6556,f1196])).
fof(f6556,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2249,f6036,f1007])).
fof(f7482,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6557,f1196])).
fof(f6557,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2250,f6036,f1007])).
fof(f2250,plain,(
  ~sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1196,f1145,f1009])).
fof(f7479,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6571,f781])).
fof(f6571,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f2028,f6036,f1010])).
fof(f2028,plain,(
  ( ! [X0] : (sP21('T_0',happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f658,f781,f1008])).
fof(f1008,plain,(
  ( ! [X2,X0,X1] : (sP21(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f621])).
fof(f7478,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f6573,f5072])).
fof(f5072,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK139(X0,'_i_0_0',bool))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1149,f5067,f838])).
fof(f6573,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,sK139(X0,'_i_0_0',bool))) ) | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f5087,f6036,f1010])).
fof(f5087,plain,(
  ( ! [X0] : (sP21('T_0',sK139(X0,'_i_0_0',bool),'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1174,f1165,f1170,f5067,f993])).
fof(f7477,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6574,f781])).
fof(f6574,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2246,f6036,f1010])).
fof(f2246,plain,(
  ( ! [X0] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f781,f1145,f1008])).
fof(f7476,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6575,f942])).
fof(f942,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2)))))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X37,X36] : ! [X34] : s(X36,X34) = s(X36,'const.sum.OUTR_1'(s('type.sum.sum'(X37,X36),'const.sum.INR_1'(s(X36,X34)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sum.OUTR')).
fof(f6575,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2247,f6036,f1010])).
fof(f2247,plain,(
  ( ! [X0,X1] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f942,f1145,f1008])).
fof(f7475,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6576,f943])).
fof(f943,plain,(
  ( ! [X2,X0,X1] : (s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : ! [X2] : s(X1,X2) = s(X1,'const.sum.OUTL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2)))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X36,X37] : ! [X34] : s(X37,X34) = s(X37,'const.sum.OUTL_1'(s('type.sum.sum'(X37,X36),'const.sum.INL_1'(s(X37,X34)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sum.OUTL')).
fof(f6576,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2248,f6036,f1010])).
fof(f2248,plain,(
  ( ! [X0,X1] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f943,f1145,f1008])).
fof(f7474,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6581,f781])).
fof(f6581,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2242,f6036,f1010])).
fof(f2242,plain,(
  ( ! [X0] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f781,f1145,f1008])).
fof(f7473,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6583,f942])).
fof(f6583,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTR_1'(s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2243,f6036,f1010])).
fof(f2243,plain,(
  ( ! [X0,X1] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'const.sum.OUTR_1'(s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,X1)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f942,f1145,f1008])).
fof(f7472,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6584,f943])).
fof(f6584,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.sum.OUTL_1'(s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,X0)))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2244,f6036,f1010])).
fof(f2244,plain,(
  ( ! [X0,X1] : (sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'const.sum.OUTL_1'(s('type.sum.sum'(bool,X0),'const.sum.INL_1'(s(bool,X1)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f943,f1145,f1008])).
fof(f7471,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6619,f7170])).
fof(f7170,plain,(
  ( ! [X0,X1] : (sP23('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1177,f6036,f1134])).
fof(f1134,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | p(s(bool,X0)) | ~sP22(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1017])).
fof(f1017,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP22(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f624])).
fof(f624,plain,(
  ! [X0,X1,X2] : ((~sP23(X0,X1,X2) | (sP22(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP22(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP23(X0,X1,X2)))),
  inference(rectify,[],[f623])).
fof(f623,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | (sP22(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP22(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP23(X2,X0,X1)))),
  inference(flattening,[],[f622])).
fof(f622,plain,(
  ! [X2,X0,X1] : ((~sP23(X2,X0,X1) | (sP22(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP22(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP23(X2,X0,X1)))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X2,X0,X1] : (sP23(X2,X0,X1) <=> (sP22(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP23])])).
fof(f1177,plain,(
  ( ! [X0,X1] : (sP22(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1019])).
fof(f1019,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP22(X0,X1,X2)))),
  inference(flattening,[],[f625])).
fof(f625,plain,(
  ! [X0,X1,X2] : ((~sP22(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP22(X0,X1,X2)))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0,X1,X2] : (sP22(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP22])])).
fof(f6619,plain,(
  ( ! [X0,X1] : (~sP23('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036,f1012])).
fof(f1012,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f624])).
fof(f7470,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6625,f7170])).
fof(f6625,plain,(
  ( ! [X0,X1] : (~sP23('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036,f1013])).
fof(f1013,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f624])).
fof(f7469,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6633,f6637])).
fof(f6637,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1021])).
fof(f1021,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f626])).
fof(f6633,plain,(
  ( ! [X0] : (~sP22('T_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f658,f6036,f1018])).
fof(f1018,plain,(
  ( ! [X2,X0,X1] : (~sP22(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f626])).
fof(f7468,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6641,f6036])).
fof(f6641,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2251,f6036,f1022])).
fof(f1022,plain,(
  ( ! [X2,X0,X1] : (sP23(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f628])).
fof(f628,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP23(X2,X0,X1)) & (~sP23(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f627])).
fof(f627,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP23(X2,X0,X1)) & (~sP23(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP23(X2,X0,X1))),
  inference(definition_folding,[],[f290,f431,f430])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X134,X138,X139] : ((p(s(bool,X139)) <=> (p(s(bool,X138)) & p(s(bool,X134)))) <=> ((p(s(bool,X139)) | ~p(s(bool,X138)) | ~p(s(bool,X134))) & (p(s(bool,X138)) | ~p(s(bool,X139))) & (p(s(bool,X134)) | ~p(s(bool,X139)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sat.sat_23_prove_ltb')).
fof(f2251,plain,(
  ( ! [X0] : (~sP23('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f1012])).
fof(f7467,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6642,f6036])).
fof(f6642,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2252,f6036,f1022])).
fof(f2252,plain,(
  ( ! [X0] : (~sP23('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f1013])).
fof(f7466,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6643,f1149])).
fof(f6643,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2504,f6036,f1022])).
fof(f2504,plain,(
  ~sP23('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f2499,f1011])).
fof(f1011,plain,(
  ( ! [X2,X0,X1] : (~sP23(X0,X1,X2) | sP22(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f624])).
fof(f2499,plain,(
  ~sP22('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f658,f1149,f658,f1018])).
fof(f7465,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6648,f1149])).
fof(f6648,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f1022])).
fof(f1274,plain,(
  ( ! [X0] : (~sP23('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1149,f1012])).
fof(f7464,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6649,f6036])).
fof(f6649,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1275,f6036,f1022])).
fof(f1275,plain,(
  ( ! [X0] : (~sP23('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f1149,f1013])).
fof(f7463,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6650,f1149])).
fof(f6650,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2251,f6036,f1022])).
fof(f7462,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6651,f6036])).
fof(f6651,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2252,f6036,f1022])).
fof(f7461,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6659,f6036])).
fof(f6659,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2251,f6036,f1023])).
fof(f1023,plain,(
  ( ! [X2,X0,X1] : (sP23(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f628])).
fof(f7460,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6660,f6036])).
fof(f6660,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2252,f6036,f1023])).
fof(f7459,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6661,f1149])).
fof(f6661,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2504,f6036,f1023])).
fof(f7458,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6666,f6036])).
fof(f6666,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f1023])).
fof(f7457,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6667,f1149])).
fof(f6667,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1275,f6036,f1023])).
fof(f7456,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6668,f6036])).
fof(f6668,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2251,f6036,f1023])).
fof(f7455,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6669,f1149])).
fof(f6669,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2252,f6036,f1023])).
fof(f7454,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6698,f6944])).
fof(f6944,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f6036,f6036,f1039])).
fof(f1039,plain,(
  ( ! [X2,X0,X1] : (sP25(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP25(X2,X0,X1)) & (~sP25(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f635])).
fof(f635,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP25(X2,X0,X1)) & (~sP25(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP25(X2,X0,X1))),
  inference(definition_folding,[],[f292,f434,f433])).
fof(f433,plain,(
  ! [X2,X0,X1] : (sP24(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP24])])).
fof(f434,plain,(
  ! [X2,X0,X1] : (sP25(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP24(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP25])])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X134,X138,X139] : ((p(s(bool,X139)) <=> (p(s(bool,X138)) | p(s(bool,X134)))) <=> ((p(s(bool,X139)) | ~p(s(bool,X138))) & (p(s(bool,X139)) | ~p(s(bool,X134))) & (p(s(bool,X138)) | p(s(bool,X134)) | ~p(s(bool,X139)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sat.sat_25_prove_ltb')).
fof(f6698,plain,(
  ( ! [X0,X1] : (~sP25(X0,X1,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036,f1028])).
fof(f1028,plain,(
  ( ! [X2,X0,X1] : (~sP25(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0,X1,X2] : ((~sP25(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP24(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP24(X0,X1,X2) | sP25(X0,X1,X2)))),
  inference(rectify,[],[f630])).
fof(f630,plain,(
  ! [X2,X0,X1] : ((~sP25(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP24(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP24(X2,X0,X1) | sP25(X2,X0,X1)))),
  inference(flattening,[],[f629])).
fof(f629,plain,(
  ! [X2,X0,X1] : ((~sP25(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP24(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP24(X2,X0,X1)) | sP25(X2,X0,X1)))),
  inference(nnf_transformation,[],[f434])).
fof(f7453,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6715,f6944])).
fof(f6715,plain,(
  ( ! [X0,X1] : (~sP25(X0,'T_0',X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036,f1029])).
fof(f1029,plain,(
  ( ! [X2,X0,X1] : (~sP25(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f631])).
fof(f7452,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6745,f6036])).
fof(f6745,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f2254,f6036,f1034])).
fof(f1034,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2) | ~sP24(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f631])).
fof(f2254,plain,(
  ( ! [X0] : (~sP25('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f1028])).
fof(f1184,plain,(
  ( ! [X0,X1] : (sP24('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f1038])).
fof(f1038,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0,X1,X2] : ((~sP24(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP24(X0,X1,X2)))),
  inference(rectify,[],[f633])).
fof(f633,plain,(
  ! [X2,X0,X1] : ((~sP24(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP24(X2,X0,X1)))),
  inference(flattening,[],[f632])).
fof(f632,plain,(
  ! [X2,X0,X1] : ((~sP24(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP24(X2,X0,X1)))),
  inference(nnf_transformation,[],[f433])).
fof(f7451,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6746,f6036])).
fof(f6746,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1184,f1289,f6036,f1034])).
fof(f1289,plain,(
  ( ! [X0] : (~sP25('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f1149,f1029])).
fof(f7450,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6747,f6036])).
fof(f6747,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1183,f1289,f6036,f1034])).
fof(f1183,plain,(
  ( ! [X0,X1] : (sP24(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f658,f1037])).
fof(f1037,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f634])).
fof(f7449,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6749,f6036])).
fof(f6749,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f1289,f6036,f1034])).
fof(f2262,plain,(
  ( ! [X0,X1] : (sP24(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1036])).
fof(f1036,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f634])).
fof(f7448,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6750,f6036])).
fof(f6750,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f2258,f6036,f1034])).
fof(f2258,plain,(
  ( ! [X0] : (~sP25('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f1029])).
fof(f7447,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6751,f6036])).
fof(f6751,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f2258,f6036,f1034])).
fof(f2263,plain,(
  ( ! [X0,X1] : (sP24(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1037])).
fof(f7446,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6753,f6036])).
fof(f6753,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f2258,f6036,f1034])).
fof(f7445,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6758,f6860])).
fof(f6860,plain,(
  ( ! [X2,X0,X1] : (sP24(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1038])).
fof(f6758,plain,(
  ~sP24('F_0','T_0','F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1149,f6036,f1034])).
fof(f7444,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6760,f6860])).
fof(f6760,plain,(
  ~sP24('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1149,f6036,f1034])).
fof(f7443,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6771,f6860])).
fof(f6771,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1034])).
fof(f1285,plain,(
  ( ! [X0] : (~sP25('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1149,f1028])).
fof(f7442,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6773,f6860])).
fof(f6773,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1034])).
fof(f7441,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6775,f6860])).
fof(f6775,plain,(
  ( ! [X0] : (~sP24('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1034])).
fof(f7440,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6777,f6860])).
fof(f6777,plain,(
  ( ! [X0] : (~sP24('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1034])).
fof(f7439,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6781,f6036])).
fof(f6781,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1182,f6036,f1034])).
fof(f1182,plain,(
  ( ! [X0,X1] : (sP24(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f1036])).
fof(f7438,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6783,f6036])).
fof(f6783,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1182,f6036,f1034])).
fof(f7437,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6784,f6036])).
fof(f6784,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f2262,f6036,f1034])).
fof(f7436,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6786,f6036])).
fof(f6786,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f2262,f6036,f1034])).
fof(f7435,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6788,f6036])).
fof(f6788,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f2263,f6036,f1034])).
fof(f7434,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6789,f6036])).
fof(f6789,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1184,f6036,f1034])).
fof(f7433,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6790,f6036])).
fof(f6790,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1184,f6036,f1034])).
fof(f7432,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6792,f6036])).
fof(f6792,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1184,f6036,f1034])).
fof(f7431,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6793,f6860])).
fof(f6793,plain,(
  ~sP24('F_0','F_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f6036,f1034])).
fof(f7430,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6794,f6860])).
fof(f6794,plain,(
  ~sP24('F_0','F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1149,f6036,f1034])).
fof(f7429,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6804,f6860])).
fof(f6804,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1034])).
fof(f7428,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6805,f6860])).
fof(f6805,plain,(
  ( ! [X0] : (~sP24('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1034])).
fof(f7427,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6806,f6860])).
fof(f6806,plain,(
  ( ! [X0] : (~sP24('F_0','T_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1034])).
fof(f7426,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6807,f6860])).
fof(f6807,plain,(
  ( ! [X0] : (~sP24('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1034])).
fof(f7425,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6834,f6036])).
fof(f6834,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f658,f6036,f1035])).
fof(f1035,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP24(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f634])).
fof(f7424,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6835,f6036])).
fof(f6835,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1183,f658,f6036,f1035])).
fof(f7423,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6836,f6036])).
fof(f6836,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f658,f6036,f1035])).
fof(f7422,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6837,f6860])).
fof(f6837,plain,(
  ( ! [X0] : (~sP24('T_0',X0,'F_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f658,f6036,f1035])).
fof(f7421,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6839,f6860])).
fof(f6839,plain,(
  ( ! [X0,X1] : (~sP24('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f658,f6036,f1035])).
fof(f7420,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6841,f6036])).
fof(f6841,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1182,f658,f6036,f1035])).
fof(f7419,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6842,f6036])).
fof(f6842,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f658,f6036,f1035])).
fof(f7418,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6844,f6036])).
fof(f6844,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f658,f6036,f1035])).
fof(f7417,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6849,f6860])).
fof(f6849,plain,(
  ( ! [X0] : (~sP24('T_0','F_0',X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f1149,f6036,f1035])).
fof(f7416,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6858,f6860])).
fof(f6858,plain,(
  ( ! [X0,X1] : (~sP24('T_0',X0,X1)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036,f6036,f1035])).
fof(f7415,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6863,f6036])).
fof(f6863,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1149,f6036,f1039])).
fof(f7414,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6864,f6036])).
fof(f6864,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1149,f6036,f1039])).
fof(f7413,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6865,f6036])).
fof(f6865,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1149,f6036,f1039])).
fof(f7412,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6867,f6036])).
fof(f6867,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7411,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6868,f6036])).
fof(f6868,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1039])).
fof(f7410,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6869,f6036])).
fof(f6869,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7409,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6870,f1149])).
fof(f6870,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f6036,f6036,f1039])).
fof(f2586,plain,(
  ~sP25('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f2573,f1030])).
fof(f1030,plain,(
  ( ! [X2,X0,X1] : (~sP25(X0,X1,X2) | sP24(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f631])).
fof(f2573,plain,(
  ~sP24('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1149,f1149,f658,f1035])).
fof(f7408,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6871,f1149])).
fof(f6871,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f4417,plain,(
  ~sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f4407,f1030])).
fof(f4407,plain,(
  ~sP24('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1149,f2832])).
fof(f2832,plain,(
  ( ! [X8,X7] : (~sP24('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X7,X8) | p(s(bool,X8)) | p(s(bool,X7))) ) | $spl0),
  inference(subsumption_resolution,[],[f2777,f658])).
fof(f2777,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP24('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X7,X8)) ) | $spl0),
  inference(superposition,[],[f1035,f2226])).
fof(f7407,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6872,f1149])).
fof(f6872,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1149,f6036,f1039])).
fof(f7406,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6873,f1149])).
fof(f6873,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1149,f6036,f1039])).
fof(f7405,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6875,f6036])).
fof(f6875,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f1149,f6036,f1039])).
fof(f7404,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6880,f1149])).
fof(f6880,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1039])).
fof(f7403,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6881,f1149])).
fof(f6881,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7402,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6882,f1149])).
fof(f6882,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1039])).
fof(f7401,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6883,f1149])).
fof(f6883,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7400,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6885,f6036])).
fof(f6885,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f7399,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6890,f6036])).
fof(f6890,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f6036,f1039])).
fof(f7398,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6891,f6036])).
fof(f6891,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1149,f6036,f1039])).
fof(f7397,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6893,f6036])).
fof(f6893,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1149,f6036,f1039])).
fof(f7396,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6894,f6036])).
fof(f6894,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1039])).
fof(f7395,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6895,f6036])).
fof(f6895,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7394,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6897,f6036])).
fof(f6897,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7393,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6898,f1149])).
fof(f6898,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f6036,f6036,f1039])).
fof(f7392,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6899,f1149])).
fof(f6899,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f7391,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6900,f1149])).
fof(f6900,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f6036,f1039])).
fof(f7390,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6901,f1149])).
fof(f6901,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1149,f6036,f1039])).
fof(f7389,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6903,f6036])).
fof(f6903,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f1149,f6036,f1039])).
fof(f7388,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6908,f1149])).
fof(f6908,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1039])).
fof(f7387,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6909,f1149])).
fof(f6909,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7386,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6910,f1149])).
fof(f6910,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1039])).
fof(f7385,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6911,f1149])).
fof(f6911,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7384,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6913,f6036])).
fof(f6913,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f7383,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6919,f6036])).
fof(f6919,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f1149,f6036,f1039])).
fof(f7382,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6920,f1149])).
fof(f6920,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f1149,f6036,f1039])).
fof(f7381,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6921,f1149])).
fof(f6921,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f1149,f6036,f1039])).
fof(f7380,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6922,f6036])).
fof(f6922,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f1039])).
fof(f7379,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6923,f6036])).
fof(f6923,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7378,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6925,f6036])).
fof(f6925,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7377,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6926,f1149])).
fof(f6926,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f6036,f6036,f1039])).
fof(f7376,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6927,f1149])).
fof(f6927,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f7375,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6929,f6036])).
fof(f6929,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f1149,f6036,f1039])).
fof(f7374,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6930,f1149])).
fof(f6930,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f1149,f6036,f1039])).
fof(f7373,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6931,f1149])).
fof(f6931,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f1149,f6036,f1039])).
fof(f7372,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6937,f6036])).
fof(f6937,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f6036,f6036,f1039])).
fof(f7371,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6938,f6036])).
fof(f6938,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f6036,f1039])).
fof(f7370,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6939,f6036])).
fof(f6939,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f6036,f6036,f1039])).
fof(f7369,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6940,f1149])).
fof(f6940,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f6036,f6036,f1039])).
fof(f7368,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6941,f1149])).
fof(f6941,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4417,f6036,f6036,f1039])).
fof(f7367,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6947,f6036])).
fof(f6947,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2264,f658,f6036,f1042])).
fof(f1042,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP25(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f636])).
fof(f2264,plain,(
  ( ! [X0] : (sP25('T_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f1040])).
fof(f1040,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP25(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f636])).
fof(f7366,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f6948,f6036])).
fof(f6948,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1298,f658,f6036,f1042])).
fof(f1298,plain,(
  ( ! [X0] : (sP25('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f658,f1041])).
fof(f1041,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP25(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f636])).
fof(f7365,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f6949,f6036])).
fof(f6949,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2269,f658,f6036,f1042])).
fof(f2269,plain,(
  ( ! [X0] : (sP25('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f1041])).
fof(f7364,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7000,f6036])).
fof(f7000,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1292,f658,f6036,f1042])).
fof(f1292,plain,(
  ( ! [X0] : (sP25('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f658,f1040])).
fof(f7363,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7001,f6036])).
fof(f7001,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2264,f658,f6036,f1042])).
fof(f7362,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7003,f6036])).
fof(f7003,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2269,f658,f6036,f1042])).
fof(f7361,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7008,f1149])).
fof(f7008,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1323,f6036,f1045])).
fof(f1045,plain,(
  ( ! [X2,X0,X1] : (~sP27(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0,X1,X2] : ((~sP27(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP26(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP26(X0,X1,X2) | sP27(X0,X1,X2)))),
  inference(rectify,[],[f638])).
fof(f638,plain,(
  ! [X2,X0,X1] : ((~sP27(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP26(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP26(X2,X0,X1) | sP27(X2,X0,X1)))),
  inference(flattening,[],[f637])).
fof(f637,plain,(
  ! [X2,X0,X1] : ((~sP27(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP26(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP26(X2,X0,X1)) | sP27(X2,X0,X1)))),
  inference(nnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X2,X0,X1] : (sP27(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP26(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP27])])).
fof(f1323,plain,(
  ( ! [X0] : (sP27('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1149,f658,f1057])).
fof(f1057,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP27(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP27(X2,X0,X1)) & (~sP27(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f643])).
fof(f643,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP27(X2,X0,X1)) & (~sP27(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP27(X2,X0,X1))),
  inference(definition_folding,[],[f384,f437,f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : (sP26(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP26])])).
fof(f384,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X134,X138,X139] : ((p(s(bool,X139)) <=> (p(s(bool,X138)) => p(s(bool,X134)))) <=> ((p(s(bool,X139)) | p(s(bool,X138))) & (p(s(bool,X139)) | ~p(s(bool,X134))) & (~p(s(bool,X138)) | p(s(bool,X134)) | ~p(s(bool,X139)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.sat.sat_27_prove_ltb')).
fof(f7360,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7009,f6036])).
fof(f7009,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1339,f6036,f1045])).
fof(f1339,plain,(
  ( ! [X0] : (sP27('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f658,f1058])).
fof(f1058,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP27(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f644])).
fof(f7359,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7010,f6036])).
fof(f7010,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2279,f6036,f1045])).
fof(f2279,plain,(
  ( ! [X0] : (sP27('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f1058])).
fof(f7358,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7012,f1149])).
fof(f7012,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2288,f6036,f1045])).
fof(f2288,plain,(
  ( ! [X0] : (sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1186,f1145,f1132])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP26(X0,X1,X2) | sP27(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1048])).
fof(f1048,plain,(
  ( ! [X2,X0,X1] : (sP27(X0,X1,X2) | ~sP26(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f639])).
fof(f1186,plain,(
  ( ! [X0,X1] : (sP26(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1149,f1053])).
fof(f1053,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X0,X1,X2] : ((~sP26(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP26(X0,X1,X2)))),
  inference(rectify,[],[f641])).
fof(f641,plain,(
  ! [X2,X0,X1] : ((~sP26(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP26(X2,X0,X1)))),
  inference(flattening,[],[f640])).
fof(f640,plain,(
  ! [X2,X0,X1] : ((~sP26(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP26(X2,X0,X1)))),
  inference(nnf_transformation,[],[f436])).
fof(f7357,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7013,f6036])).
fof(f7013,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f6036,f1045])).
fof(f2289,plain,(
  ( ! [X0] : (sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1188,f1145,f1132])).
fof(f1188,plain,(
  ( ! [X0,X1] : (sP26(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f658,f1054])).
fof(f1054,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f642])).
fof(f7356,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7014,f6036])).
fof(f7014,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2283,f6036,f1045])).
fof(f2283,plain,(
  ( ! [X0] : (sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1145,f1058])).
fof(f7355,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7015,f1149])).
fof(f7015,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1406,f6036,f1045])).
fof(f1406,plain,(
  ( ! [X0,X1] : (sP27('const.bool.BOUNDED_1'(s(bool,X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1149,f1326])).
fof(f1326,plain,(
  ( ! [X2,X0,X1] : (sP27('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f1324,f658])).
fof(f1324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP27('const.bool.BOUNDED_1'(s(bool,X0)),X2,X1)) )),
  inference(superposition,[],[f1057,f665])).
fof(f7354,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7016,f6036])).
fof(f7016,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1416,f6036,f1045])).
fof(f1416,plain,(
  ( ! [X0,X1] : (sP27('const.bool.BOUNDED_1'(s(bool,X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f658,f1342])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP27('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(subsumption_resolution,[],[f1340,f658])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | sP27('const.bool.BOUNDED_1'(s(bool,X0)),X1,X2)) )),
  inference(superposition,[],[f1058,f665])).
fof(f7353,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7017,f6036])).
fof(f7017,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2294,f6036,f1045])).
fof(f2294,plain,(
  ( ! [X0,X1] : (sP27('const.bool.BOUNDED_1'(s(bool,X0)),'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1342])).
fof(f7352,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7018,f6036])).
fof(f7018,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6036,f1045])).
fof(f1419,plain,(
  ( ! [X2,X0,X1] : (sP27('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(subsumption_resolution,[],[f1417,f658])).
fof(f1417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP27('const.bool.BOUNDED_1'(s(bool,X1)),'const.bool.BOUNDED_1'(s(bool,X0)),X2)) )),
  inference(superposition,[],[f1342,f665])).
fof(f7351,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7026,f1149])).
fof(f7026,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2655,f6036,f1045])).
fof(f2655,plain,(
  sP27('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1189,f658,f1149,f1051])).
fof(f1051,plain,(
  ( ! [X2,X0,X1] : (sP27(X0,X1,X2) | ~sP26(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f639])).
fof(f1189,plain,(
  ( ! [X0,X1] : (sP26('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f1055])).
fof(f1055,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f642])).
fof(f7350,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7027,f6036])).
fof(f7027,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2288,f6036,f1045])).
fof(f7349,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7028,f6036])).
fof(f7028,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2289,f6036,f1045])).
fof(f7348,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7029,f6036])).
fof(f7029,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2283,f6036,f1045])).
fof(f7347,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7030,f6036])).
fof(f7030,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1406,f6036,f1045])).
fof(f7346,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7031,f6036])).
fof(f7031,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1416,f6036,f1045])).
fof(f7345,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7032,f6036])).
fof(f7032,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2294,f6036,f1045])).
fof(f7344,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7033,f6036])).
fof(f7033,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1419,f6036,f1045])).
fof(f7343,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7069,f7086])).
fof(f7086,plain,(
  ( ! [X2,X0,X1] : (sP26(X0,X1,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1055])).
fof(f7069,plain,(
  ~sP26('F_0','T_0','T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f1313,f658,f6036,f1051])).
fof(f1313,plain,(
  ( ! [X0] : (~sP27('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f1149,f1046])).
fof(f1046,plain,(
  ( ! [X2,X0,X1] : (~sP27(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f639])).
fof(f7342,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7071,f7086])).
fof(f7071,plain,(
  ~sP26('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f658,f6036,f1051])).
fof(f2274,plain,(
  ( ! [X0] : (~sP27('F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1149,f1145,f1046])).
fof(f7341,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7082,f7086])).
fof(f7082,plain,(
  ( ! [X0] : (~sP26('T_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f658,f658,f6036,f1052])).
fof(f1052,plain,(
  ( ! [X2,X0,X1] : (~sP26(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f642])).
fof(f7340,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7096,f1149])).
fof(f7096,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f1313,f658,f6036,f1056])).
fof(f1056,plain,(
  ( ! [X2,X0,X1] : (sP27(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f644])).
fof(f7339,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7097,f1149])).
fof(f7097,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f658,f6036,f1056])).
fof(f7338,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7099,f6036])).
fof(f7099,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4512,f658,f6036,f1056])).
fof(f4512,plain,(
  ~sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1149,f2838])).
fof(f2838,plain,(
  ( ! [X19,X20] : (~sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X20,X19) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl0),
  inference(subsumption_resolution,[],[f2783,f658])).
fof(f2783,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X20,X19)) ) | $spl0),
  inference(superposition,[],[f1059,f2226])).
fof(f1059,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP27(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f644])).
fof(f7337,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7105,f6036])).
fof(f7105,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2274,f658,f6036,f1056])).
fof(f7336,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7106,f1149])).
fof(f7106,plain,(
  p(s(bool,'F_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f2672,f658,f6036,f1056])).
fof(f2672,plain,(
  ~sP27('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f2666,f1047])).
fof(f1047,plain,(
  ( ! [X2,X0,X1] : (~sP27(X0,X1,X2) | sP26(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f639])).
fof(f2666,plain,(
  ~sP26('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f658,f658,f1149,f1052])).
fof(f7335,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7107,f1149])).
fof(f7107,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4512,f658,f6036,f1056])).
fof(f7334,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7174,f1177])).
fof(f7174,plain,(
  ( ! [X0] : (~sP22(X0,'F_0','T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f1134])).
fof(f7333,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7176,f1177])).
fof(f7176,plain,(
  ( ! [X0] : (~sP22('F_0',X0,'T_0')) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1275,f6036,f1134])).
fof(f7332,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7178,f2253])).
fof(f2253,plain,(
  ( ! [X0,X1] : (sP22(X0,X1,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1019])).
fof(f7178,plain,(
  ( ! [X0] : (~sP22(X0,'F_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2251,f6036,f1134])).
fof(f7331,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7180,f2253])).
fof(f7180,plain,(
  ( ! [X0] : (~sP22('F_0',X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2252,f6036,f1134])).
fof(f7330,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7184,f6036])).
fof(f7184,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2228,f6036,f1152])).
fof(f1152,plain,(
  ( ! [X0,X1] : (~sP16(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f914,f1149])).
fof(f914,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP16(X0,X1)) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ((~sP16(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP16(X0,X1)))),
  inference(flattening,[],[f575])).
fof(f575,plain,(
  ! [X0,X1] : ((~sP16(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP16(X0,X1)))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1] : (sP16(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f2228,plain,(
  ( ! [X0] : (sP16(X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f915])).
fof(f915,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP16(X0,X1)) )),
  inference(cnf_transformation,[],[f576])).
fof(f7329,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7185,f6036])).
fof(f7185,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1195,f6036,f1152])).
fof(f1195,plain,(
  ( ! [X2,X3] : (sP16(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(subsumption_resolution,[],[f1193,f658])).
fof(f1193,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP16(X3,'const.bool.BOUNDED_1'(s(bool,X2)))) )),
  inference(superposition,[],[f915,f665])).
fof(f7328,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7186,f6036])).
fof(f7186,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1458,f6036,f1152])).
fof(f1458,plain,(
  ( ! [X0,X1] : (sP16(X0,'const.ind_type.ISO_2'(s(fun(X1,X1),'_i_0_0'),s(fun(X1,X1),'_i_0_0')))) )),
  inference(unit_resulting_resolution,[],[f666,f915])).
fof(f7327,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7187,f6036])).
fof(f7187,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1163,f6036,f1152])).
fof(f1163,plain,(
  ( ! [X0] : (sP16('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f658,f916])).
fof(f916,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP16(X0,X1)) )),
  inference(cnf_transformation,[],[f576])).
fof(f7326,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7188,f6036])).
fof(f7188,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2229,f6036,f1152])).
fof(f2229,plain,(
  ( ! [X0] : (sP16('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f916])).
fof(f7325,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7189,f6036])).
fof(f7189,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1194,f6036,f1152])).
fof(f1194,plain,(
  ( ! [X0,X1] : (sP16('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(subsumption_resolution,[],[f1192,f658])).
fof(f1192,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP16('const.bool.BOUNDED_1'(s(bool,X0)),X1)) )),
  inference(superposition,[],[f916,f665])).
fof(f7324,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7190,f6036])).
fof(f7190,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1459,f6036,f1152])).
fof(f1459,plain,(
  ( ! [X0,X1] : (sP16('const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f666,f916])).
fof(f7323,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7195,f6036])).
fof(f7195,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1162,f6036,f1152])).
fof(f1162,plain,(
  ( ! [X0] : (sP16(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f915])).
fof(f7322,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7196,f6036])).
fof(f7196,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2228,f6036,f1152])).
fof(f7321,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7197,f6036])).
fof(f7197,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1195,f6036,f1152])).
fof(f7320,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7198,f6036])).
fof(f7198,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1458,f6036,f1152])).
fof(f7319,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7200,f6036])).
fof(f7200,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2229,f6036,f1152])).
fof(f7318,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7201,f6036])).
fof(f7201,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.BOUNDED_1'(s(bool,X0))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1194,f6036,f1152])).
fof(f7317,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f7202,f6036])).
fof(f7202,plain,(
  ( ! [X0] : (p(s(bool,'const.ind_type.ISO_2'(s(fun(X0,X0),'_i_0_0'),s(fun(X0,X0),'_i_0_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1459,f6036,f1152])).
fof(f7316,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7218,f6036])).
fof(f7218,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f6036,f2832])).
fof(f7315,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7220,f6036])).
fof(f7220,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1183,f6036,f2832])).
fof(f7314,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7222,f6036])).
fof(f7222,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f6036,f2832])).
fof(f7313,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7226,f6036])).
fof(f7226,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1182,f6036,f2832])).
fof(f7312,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7228,f6036])).
fof(f7228,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2262,f6036,f2832])).
fof(f7311,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7232,f6036])).
fof(f7232,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2263,f6036,f2832])).
fof(f7310,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7238,f6036])).
fof(f7238,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2268,f6036,f2835])).
fof(f2835,plain,(
  ( ! [X14,X13] : (~sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X14,X13) | p(s(bool,X14)) | p(s(bool,X13))) ) | $spl0),
  inference(subsumption_resolution,[],[f2780,f658])).
fof(f2780,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | p(s(bool,X13)) | p(s(bool,X14)) | ~sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X14,X13)) ) | $spl0),
  inference(superposition,[],[f1042,f2226])).
fof(f2268,plain,(
  ( ! [X0] : (sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1145,f1040])).
fof(f7309,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7240,f6036])).
fof(f7240,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2291,f6036,f2835])).
fof(f2291,plain,(
  ( ! [X0] : (sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1183,f1145,f1133])).
fof(f1133,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP24(X0,X1,X2) | sP25(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1031])).
fof(f1031,plain,(
  ( ! [X2,X0,X1] : (sP25(X0,X1,X2) | ~sP24(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f631])).
fof(f7308,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7242,f6036])).
fof(f7242,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f6036,f2835])).
fof(f2273,plain,(
  ( ! [X0] : (sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f1145,f1041])).
fof(f7307,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7246,f6036])).
fof(f7246,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2290,f6036,f2835])).
fof(f2290,plain,(
  ( ! [X0] : (sP25('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X0,'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1182,f1145,f1133])).
fof(f7306,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7248,f6036])).
fof(f7248,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2268,f6036,f2835])).
fof(f7305,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f7252,f6036])).
fof(f7252,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2273,f6036,f2835])).
fof(f7273,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f669,f6036,f883])).
fof(f883,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))))))),
  inference(nnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X40,X41] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X40),s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X41)))))) <=> s('type.patSem.v',X40) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X41))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_lit|split|3')).
fof(f669,plain,(
  ( ! [X0] : (s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X0))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X41] : s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X41))) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num','const.bool.COND_3'(s(bool,X41),s('type.num.num','const.conLang.true_tag_0'),s('type.num.num','const.conLang.false_tag_0'))),s('type.list.list'('type.patSem.v'),'const.list.NIL_0')))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.patSem.Boolv_def')).
fof(f7271,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f781,f6036,f881])).
fof(f881,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))))) )),
  inference(cnf_transformation,[],[f555])).
fof(f555,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))))))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X40,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X40),s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))))) <=> s('type.patSem.v',X40) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_lit|split|1')).
fof(f7269,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f669,f6036,f879])).
fof(f879,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6,X41] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X41))),s('type.patSem.v',X6)))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Boolv_1'(s(bool,X41))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_lit|split|2')).
fof(f7267,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f781,f6036,f877])).
fof(f877,plain,(
  ( ! [X0,X1] : (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))) )),
  inference(cnf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1)))) & (s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0))))))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X1))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))),s('type.patSem.v',X6)))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_lit|split|0')).
fof(f7265,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2221,f6036,f871])).
fof(f871,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP10(X0,X1)) & (~sP10(X0,X1) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> sP10(X0,X1))),
  inference(definition_folding,[],[f371,f413,f412,f411,f410,f409])).
fof(f409,plain,(
  ! [X2,X4,X3] : (sP6(X2,X4,X3) <=> ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f410,plain,(
  ! [X2,X4,X3] : (sP7(X2,X4,X3) <=> ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f411,plain,(
  ! [X2,X4,X3] : (sP8(X2,X4,X3) <=> ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f412,plain,(
  ! [X2,X4,X3] : (sP9(X2,X4,X3) <=> ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f413,plain,(
  ! [X0,X1] : (sP10(X0,X1) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | sP6(X2,X4,X3) | sP7(X2,X4,X3) | sP8(X2,X4,X3) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | sP9(X2,X4,X3)) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f371,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (? [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) & ~p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) | p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(ennf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) <=> ! [X2] : (! [X3,X4] : ((? [X20] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X20)))) | ? [X17,X18,X19] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X18))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X17),s('type.list.list'('type.patSem.v'),X19))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X18),s('type.list.list'('type.patSem.v'),X19))))) | ? [X13,X14,X15,X16] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X13),s('type.patLang.exp',X15))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X14),s('type.patLang.exp',X16))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X13))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X14))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X13),s('type.list.list'('type.patSem.v'),X14))))))),s('type.patLang.exp',X15))),s('type.patLang.exp',X16))))) | ? [X8,X9,X10,X11,X12] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patLang.exp'),X10),s('type.num.num',X12))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X9),s('type.list.list'('type.patLang.exp'),X11),s('type.num.num',X12))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X8))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X11))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X9))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X10))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X8),s('type.list.list'('type.patSem.v'),X9))))))),s('type.list.list'('type.patLang.exp'),X10),s('type.list.list'('type.patLang.exp'),X11))))) | ? [X7] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7)))) | ? [X5,X6] : (s('type.patSem.v',X3) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X5))) & s('type.patSem.v',X4) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X6))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.list.list'('type.patSem.v'),X5),s('type.list.list'('type.patSem.v'),X6)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X3))),s('type.patSem.v',X4))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X2),s('type.patSem.v',X0))),s('type.patSem.v',X1))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X5,X14] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X5),s('type.patSem.v',X14)))) <=> ! [X15] : (! [X16,X17] : ((? [X7] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Litv_1'(s('type.ast.lit',X7)))) | ? [X18,X19,X20] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Conv_2'(s('type.num.num',X18),s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20))))) | ? [X21,X22,X23,X24] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X21),s('type.patLang.exp',X23))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Closure_2'(s('type.list.list'('type.patSem.v'),X22),s('type.patLang.exp',X24))) & p(s(bool,happ(s(fun('type.patLang.exp',bool),happ(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bind_1'(s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.patLang.exp',X23))),s('type.patLang.exp',X24))))) | ? [X21,X22,X25,X26,X13] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patLang.exp'),X25),s('type.num.num',X13))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Recclosure_3'(s('type.list.list'('type.patSem.v'),X22),s('type.list.list'('type.patLang.exp'),X26),s('type.num.num',X13))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patLang.exp',fun('type.patLang.exp',bool)),'const.exh_to_patProof.exp_rel_3'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X21))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X26))))),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patSem.v'),X22))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.bindn_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.patLang.exp'),X25))))),s(fun('type.num.num',fun('type.num.num',bool)),'const.exh_to_patProof.env_rel_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X21),s('type.list.list'('type.patSem.v'),X22))))))),s('type.list.list'('type.patLang.exp'),X25),s('type.list.list'('type.patLang.exp'),X26))))) | ? [X13] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13)))) | ? [X19,X20] : (s('type.patSem.v',X16) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X19))) & s('type.patSem.v',X17) = s('type.patSem.v','const.patSem.Vectorv_1'(s('type.list.list'('type.patSem.v'),X20))) & p(s(bool,'const.list.LIST_REL_3'(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.list.list'('type.patSem.v'),X19),s('type.list.list'('type.patSem.v'),X20)))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.patSem.v',X16))),s('type.patSem.v',X17))))) => p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),X15),s('type.patSem.v',X5))),s('type.patSem.v',X14))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_def')).
fof(f2221,plain,(
  sP10('const.patSem.Loc_1'(s('type.num.num',sK29)),sK28) | $spl0),
  inference(unit_resulting_resolution,[],[f1145,f870])).
fof(f870,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1)))) | sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f550])).
fof(f7263,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2220,f6036,f1150])).
fof(f1150,plain,(
  ( ! [X0,X1] : (~sP15(X1,X0) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',X0),s('type.patSem.v',X1))))) )),
  inference(forward_demodulation,[],[f913,f837])).
fof(f913,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.patSem.v',bool),happ(s(fun('type.patSem.v',fun('type.patSem.v',bool)),'const.exh_to_patProof.v_rel_0'),s('type.patSem.v',X0))),s('type.patSem.v',X1)))) | ~sP15(X1,X0)) )),
  inference(cnf_transformation,[],[f574])).
fof(f7261,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f658,f2283,f6036,f2838])).
fof(f7259,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f658,f2289,f6036,f2838])).
fof(f7255,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f4512,f6036,f2836])).
fof(f2836,plain,(
  ( ! [X15,X16] : (sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X16,X15) | p(s(bool,X15))) ) | $spl0),
  inference(subsumption_resolution,[],[f2781,f658])).
fof(f2781,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | sP27('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X16,X15)) ) | $spl0),
  inference(superposition,[],[f1057,f2226])).
fof(f7253,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2273,f6036,f2835])).
fof(f7251,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2291,f6036,f2835])).
fof(f7249,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2268,f6036,f2835])).
fof(f7247,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2290,f6036,f2835])).
fof(f7243,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2273,f6036,f2835])).
fof(f7241,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2291,f6036,f2835])).
fof(f7239,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2268,f6036,f2835])).
fof(f7237,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2290,f6036,f2835])).
fof(f7233,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2263,f6036,f2832])).
fof(f7231,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f1183,f6036,f2832])).
fof(f7229,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2262,f6036,f2832])).
fof(f7227,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f1182,f6036,f2832])).
fof(f7223,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f2263,f6036,f2832])).
fof(f7221,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1149,f1183,f6036,f2832])).
fof(f7219,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f2262,f6036,f2832])).
fof(f7217,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f6036,f1182,f6036,f2832])).
fof(f7213,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2232,f6036,f2831])).
fof(f2831,plain,(
  ( ! [X6] : (sP17(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | p(s(bool,X6))) ) | $spl0),
  inference(subsumption_resolution,[],[f2776,f658])).
fof(f2776,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP17(X6,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) ) | $spl0),
  inference(superposition,[],[f934,f2226])).
fof(f2232,plain,(
  ~sP17('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f930])).
fof(f930,plain,(
  ( ! [X0,X1] : (~sP17(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f7211,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2230,f6036,f2830])).
fof(f2830,plain,(
  ( ! [X5] : (sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X5) | p(s(bool,X5))) ) | $spl0),
  inference(subsumption_resolution,[],[f2775,f658])).
fof(f2775,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),X5)) ) | $spl0),
  inference(superposition,[],[f931,f2226])).
fof(f2230,plain,(
  ~sP17('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f658,f1145,f930])).
fof(f7206,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1162,f6036,f6036,f1152])).
fof(f7204,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1162,f1149,f6036,f1152])).
fof(f7194,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1162,f6036,f6036,f1152])).
fof(f7192,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1163,f1149,f6036,f1152])).
fof(f7181,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1180,f2252,f6036,f1134])).
fof(f1180,plain,(
  ( ! [X0,X1] : (sP22('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f1021])).
fof(f7179,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1178,f2251,f6036,f1134])).
fof(f1178,plain,(
  ( ! [X0,X1] : (sP22(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1149,f1020])).
fof(f1020,plain,(
  ( ! [X2,X0,X1] : (sP22(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f626])).
fof(f7177,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1180,f1275,f6036,f1134])).
fof(f7175,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1178,f1274,f6036,f1134])).
fof(f7158,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f947,f6036,f1131])).
fof(f947,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X37,X36] : ! [X34,X33] : s(X36,X33) = s(X36,'const.pair.SND_1'(s('type.pair.prod'(X37,X36),'const.pair.,_2'(s(X37,X34),s(X36,X33)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.pair.SND')).
fof(f7156,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f943,f6036,f1131])).
fof(f7154,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f946,f6036,f1131])).
fof(f946,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X36,X37] : ! [X34,X33] : s(X37,X34) = s(X37,'const.pair.FST_1'(s('type.pair.prod'(X37,X36),'const.pair.,_2'(s(X37,X34),s(X36,X33)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.pair.FST')).
fof(f7152,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f942,f6036,f1131])).
fof(f7150,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f792,f6036,f1131])).
fof(f792,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X37] : ! [X56,X101] : s(X37,X56) = s(X37,'const.ind_type.FCONS_3'(s(X37,X56),s(fun('type.num.num',X37),X101),s('type.num.num','const.num.|48|_0')))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.ind_type.FCONS|split|0')).
fof(f7148,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f966,f6036,f1131])).
fof(f966,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK206(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),sK206(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),sK206(X3,X2,X1,X0)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),sK206(X3,X2,X1,X0)),s('type.list.list'(X1),X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK206])],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ! [X2,X3] : ? [X4] : (s(X0,X2) = s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.NIL_0'))) & ! [X5,X6] : s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X5),s('type.list.list'(X1),X6))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun('type.list.list'(X1),fun(X0,X0)),happ(s(fun(X1,fun('type.list.list'(X1),fun(X0,X0))),X3),s(X1,X5))),s('type.list.list'(X1),X6))),s(X0,happ(s(fun('type.list.list'(X1),X0),X4),s('type.list.list'(X1),X6))))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X36,X37] : ! [X128,X102] : ? [X129] : (s(X36,X128) = s(X36,happ(s(fun('type.list.list'(X37),X36),X129),s('type.list.list'(X37),'const.list.NIL_0'))) & ! [X27,X28] : s(X36,happ(s(fun('type.list.list'(X37),X36),X129),s('type.list.list'(X37),'const.list.CONS_2'(s(X37,X27),s('type.list.list'(X37),X28))))) = s(X36,happ(s(fun(X36,X36),happ(s(fun('type.list.list'(X37),fun(X36,X36)),happ(s(fun(X37,fun('type.list.list'(X37),fun(X36,X36))),X102),s(X37,X27))),s('type.list.list'(X37),X28))),s(X36,happ(s(fun('type.list.list'(X37),X36),X129),s('type.list.list'(X37),X28))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.list.list_Axiom')).
fof(f7146,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f781,f6036,f1131])).
fof(f7144,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f658,f2226,f6036,f1131])).
fof(f7142,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f665,f6036,f1131])).
fof(f7139,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1339,f658,f658,f6036,f1059])).
fof(f7126,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2672,f658,f6036,f1057])).
fof(f7119,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1313,f658,f6036,f6036,f1056])).
fof(f7112,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2672,f658,f1149,f6036,f1056])).
fof(f7103,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1313,f6036,f658,f6036,f1056])).
fof(f7101,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1313,f1149,f658,f6036,f1056])).
fof(f7087,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2666,f6036,f1055])).
fof(f7085,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2666,f6036,f1053])).
fof(f7083,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1188,f658,f658,f6036,f1052])).
fof(f7072,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1189,f2274,f658,f6036,f1051])).
fof(f7070,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1189,f1313,f658,f6036,f1051])).
fof(f7050,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1339,f658,f6036,f1046])).
fof(f7037,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1323,f6036,f6036,f1045])).
fof(f7035,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2655,f1149,f6036,f1045])).
fof(f7022,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1323,f6036,f6036,f1045])).
fof(f7020,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1323,f1149,f6036,f1045])).
fof(f7007,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1292,f6036,f658,f6036,f1042])).
fof(f7005,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1292,f1149,f658,f6036,f1042])).
fof(f6976,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1292,f658,f6036,f6036,f1042])).
fof(f6959,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1298,f658,f1149,f6036,f1042])).
fof(f6945,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f6036,f1039])).
fof(f6943,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1149,f6036,f6036,f1039])).
fof(f6935,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f1149,f6036,f1039])).
fof(f6933,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2586,f1149,f1149,f6036,f1039])).
fof(f6917,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f6036,f1039])).
fof(f6915,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f6036,f6036,f1039])).
fof(f6907,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f1149,f6036,f1039])).
fof(f6905,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f1149,f6036,f1039])).
fof(f6889,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f6036,f6036,f6036,f1039])).
fof(f6887,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1149,f6036,f6036,f1039])).
fof(f6879,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f6036,f1149,f6036,f1039])).
fof(f6877,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1149,f1149,f6036,f1039])).
fof(f6861,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2573,f6036,f1038])).
fof(f6859,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1182,f658,f6036,f6036,f1035])).
fof(f6850,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1182,f658,f1149,f6036,f1035])).
fof(f6840,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1182,f6036,f658,f6036,f1035])).
fof(f6838,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1183,f1149,f658,f6036,f1035])).
fof(f6819,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1184,f6036,f6036,f1034])).
fof(f6817,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2258,f2263,f6036,f6036,f1034])).
fof(f6815,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1289,f1183,f6036,f6036,f1034])).
fof(f6813,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f2262,f6036,f6036,f1034])).
fof(f6811,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1182,f6036,f6036,f1034])).
fof(f6802,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1184,f1149,f6036,f1034])).
fof(f6800,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2254,f2262,f1149,f6036,f1034])).
fof(f6798,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1285,f1182,f1149,f6036,f1034])).
fof(f6778,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f2258,f6036,f6036,f1034])).
fof(f6776,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1184,f1289,f6036,f6036,f1034])).
fof(f6774,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f2254,f6036,f6036,f1034])).
fof(f6772,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1184,f1285,f6036,f6036,f1034])).
fof(f6761,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1184,f2258,f1149,f6036,f1034])).
fof(f6759,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1184,f1289,f1149,f6036,f1034])).
fof(f6716,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1298,f658,f6036,f1029])).
fof(f6699,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1292,f658,f6036,f1028])).
fof(f6682,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2486,f658,f658,f6036,f1027])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (~sP23(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f628])).
fof(f2486,plain,(
  sP23('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1177,f658,f658,f1014])).
fof(f1014,plain,(
  ( ! [X2,X0,X1] : (sP23(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP22(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f624])).
fof(f6674,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f6036,f1023])).
fof(f6672,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1275,f1149,f6036,f1023])).
fof(f6665,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f6036,f1023])).
fof(f6663,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2504,f1149,f6036,f1023])).
fof(f6656,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f6036,f1022])).
fof(f6654,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f1149,f6036,f1022])).
fof(f6647,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1274,f6036,f6036,f1022])).
fof(f6645,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2504,f1149,f6036,f1022])).
fof(f6638,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2499,f6036,f1021])).
fof(f6636,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2499,f6036,f1020])).
fof(f6634,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1177,f658,f658,f6036,f1018])).
fof(f6626,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2486,f658,f6036,f1013])).
fof(f6620,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2486,f658,f6036,f1012])).
fof(f6612,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2248,f943,f6036,f1010])).
fof(f6609,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2247,f942,f6036,f1010])).
fof(f6605,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2246,f781,f6036,f1010])).
fof(f6603,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2760,f2226,f6036,f1010])).
fof(f2760,plain,(
  sP21('T_0','T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))) | $spl0),
  inference(unit_resulting_resolution,[],[f658,f2226,f1008])).
fof(f6601,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2031,f665,f6036,f1010])).
fof(f2031,plain,(
  ( ! [X0] : (sP21('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f658,f665,f1008])).
fof(f6597,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2244,f943,f6036,f1010])).
fof(f6594,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2243,f942,f6036,f1010])).
fof(f6590,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2028,f781,f6036,f1010])).
fof(f6588,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f2758,f2226,f6036,f1010])).
fof(f2758,plain,(
  sP21('T_0','const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f658,f2226,f1008])).
fof(f6586,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2026,f665,f6036,f1010])).
fof(f2026,plain,(
  ( ! [X0] : (sP21('T_0','const.bool.BOUNDED_1'(s(bool,X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f658,f665,f1008])).
fof(f7293,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f6579])).
fof(f6579,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3803,f6036,f1010])).
fof(f3803,plain,(
  sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2238,f1165,f1170,f1149,f993])).
fof(f7292,plain,(
  $false | ($spl0 | $spl8)),
  inference(trivial_inequality_removal,[],[f6578])).
fof(f6578,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f3769,f6036,f1010])).
fof(f3769,plain,(
  sP21('const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f1176,f1164,f2237,f658,f992])).
fof(f992,plain,(
  ( ! [X2,X0,X1] : (~sP19(X1,X2,X0) | ~p(s(bool,X2)) | ~sP20(X0,X2,X1) | sP21(X0,X1,X2) | ~sP18(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f612])).
fof(f7291,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6569])).
fof(f6569,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3802,f6036,f1010])).
fof(f3802,plain,(
  sP21('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1174,f1165,f1170,f1149,f993])).
fof(f7290,plain,(
  $false | $spl8),
  inference(trivial_inequality_removal,[],[f6568])).
fof(f6568,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl8),
  inference(unit_resulting_resolution,[],[f3764,f6036,f1010])).
fof(f3764,plain,(
  sP21('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1176,f1164,f1169,f658,f992])).
fof(f6561,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2049,f1196,f6036,f1007])).
fof(f6559,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2045,f1196,f6036,f1007])).
fof(f6546,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f6036,f1003])).
fof(f6544,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1149,f6036,f6036,f1003])).
fof(f6536,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f1149,f6036,f1003])).
fof(f6534,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1149,f1149,f6036,f1003])).
fof(f6518,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f6036,f1003])).
fof(f6516,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f1149,f6036,f6036,f1003])).
fof(f6508,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f6036,f1149,f6036,f1003])).
fof(f6506,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f1149,f1149,f6036,f1003])).
fof(f6490,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f6036,f6036,f1003])).
fof(f6488,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f1149,f6036,f6036,f1003])).
fof(f6480,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1174,f6036,f1149,f6036,f1003])).
fof(f6478,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1176,f1149,f1149,f6036,f1003])).
fof(f6462,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2461,f6036,f1002])).
fof(f2461,plain,(
  ~sP19('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f658,f1149,f658,f999])).
fof(f6460,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2461,f6036,f1001])).
fof(f6458,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1169,f658,f658,f6036,f999])).
fof(f6449,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2451,f6036,f998])).
fof(f2451,plain,(
  ~sP20('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f658,f1149,f658,f995])).
fof(f6447,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2451,f6036,f997])).
fof(f6445,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2687,f658,f658,f6036,f995])).
fof(f2687,plain,(
  ( ! [X0] : (sP20('T_0','T_0','const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f2026,f990])).
fof(f990,plain,(
  ( ! [X2,X0,X1] : (~sP21(X0,X1,X2) | sP20(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f612])).
fof(f6404,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f5101,f6036,f994])).
fof(f6399,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f5101,f6036,f994])).
fof(f6380,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f1170,f6036,f994])).
fof(f6378,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f1170,f6036,f994])).
fof(f6373,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f1170,f6036,f994])).
fof(f6370,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f1170,f6036,f994])).
fof(f6359,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f2237,f6036,f994])).
fof(f6350,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f2237,f6036,f994])).
fof(f6322,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f1169,f6036,f994])).
fof(f6317,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f1169,f6036,f994])).
fof(f6308,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f1169,f6036,f994])).
fof(f6306,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f1169,f6036,f994])).
fof(f6252,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f5101,f6036,f993])).
fof(f6247,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f5101,f6036,f993])).
fof(f6228,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f1170,f6036,f993])).
fof(f6226,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f1170,f6036,f993])).
fof(f6221,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f1170,f6036,f993])).
fof(f6218,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f1170,f6036,f993])).
fof(f6207,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1165,f2237,f6036,f993])).
fof(f6198,plain,(
  $false | ($spl0 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f2249,f1164,f2237,f6036,f993])).
fof(f6170,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f5096,f1169,f6036,f993])).
fof(f6165,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1165,f1169,f6036,f993])).
fof(f6156,plain,(
  $false | ($spl6 | $spl8)),
  inference(unit_resulting_resolution,[],[f1175,f5080,f1164,f1169,f6036,f993])).
fof(f6154,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1175,f2049,f1164,f1169,f6036,f993])).
fof(f6125,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f2031,f658,f658,f6036,f991])).
fof(f6106,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1232,f658,f6036,f934])).
fof(f1232,plain,(
  ~sP17('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f658,f658,f930])).
fof(f6095,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1232,f658,f6036,f931])).
fof(f6084,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1234,f6036,f6036,f929])).
fof(f6082,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1243,f1149,f6036,f929])).
fof(f6070,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1234,f6036,f6036,f929])).
fof(f6068,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1234,f1149,f6036,f929])).
fof(f6056,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1196,f6036,f838])).
fof(f6054,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1196,f6036,f838])).
fof(f6050,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1149,f1196,f6036,f838])).
fof(f6048,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f6036,f1196,f6036,f838])).
fof(f6044,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1196,f6036,f731])).
fof(f731,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X29] : (s(bool,'F_0') = s(bool,X29) <=> ~p(s(bool,X29)))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f6042,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f665,f6036,f726])).
fof(f6040,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f658,f6036])).
fof(f6038,plain,(
  p(s(bool,sK137('_i_0_0',bool))) | $spl10),
  inference(cnf_transformation,[],[f6038_D])).
fof(f6038_D,plain,(
  p(s(bool,sK137('_i_0_0',bool))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f6035,plain,(
  ( ! [X3] : (p(s(bool,sK137('_i_0_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f6024,f781])).
fof(f6024,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK137('_i_0_0',bool)))))) )),
  inference(superposition,[],[f786,f781])).
fof(f786,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK137(X1,X0)))))) )),
  inference(cnf_transformation,[],[f497])).
fof(f5066,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl4),
  inference(cnf_transformation,[],[f5066_D])).
fof(f5066_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5067,plain,(
  ( ! [X3] : (~p(s(bool,sK139(X3,'_i_0_0',bool)))) ) | $spl6),
  inference(cnf_transformation,[],[f5067_D])).
fof(f5067_D,plain,(
  ( ! [X3] : (~p(s(bool,sK139(X3,'_i_0_0',bool)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5057,plain,(
  ( ! [X3] : (~p(s(bool,sK139(X3,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f795,f781])).
fof(f795,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK139(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK139(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK139(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK139])],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X37] : ! [X39,X29] : (! [X34] : (s(X37,X29) = s(X37,X34) => p(s(bool,happ(s(fun(X37,bool),X39),s(X37,X34))))) => p(s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.bool.JRH_INDUCT_UTIL')).
fof(f2219,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2218,f1149])).
fof(f2218,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2198,f1215])).
fof(f1215,plain,(
  s(bool,'F_0') = s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK28),s('type.patSem.v',sK28))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1161,f731])).
fof(f1161,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK28),s('type.patSem.v',sK28)))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1146,f1147])).
fof(f2198,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v',sK28),s('type.patSem.v',sK28)))) | $spl2),
  inference(superposition,[],[f667,f1147])).
fof(f667,plain,(
  ( ! [X0] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13] : p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))),s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X13))))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_rules|split|4')).
fof(f1145,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | $spl0),
  inference(cnf_transformation,[],[f1145_D])).
fof(f1145_D,plain,(
  p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1147,plain,(
  s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) | $spl2),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f656,plain,(
  s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) | p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))))),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | s('type.patSem.v',sK28) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | s('type.patSem.v',sK28) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29])],[f439])).
fof(f439,plain,(
  ? [X0,X1] : ((p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))) & (~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) | s('type.patSem.v',X0) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f338])).
fof(f338,plain,(
  ? [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) <~> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))),s('type.patSem.v',X0)))) <=> s('type.patSem.v',X0) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))),s('type.patSem.v',X6)))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))),s('type.patSem.v',X6)))) <=> s('type.patSem.v',X6) = s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',X7))))),
  file('thm.exh_to_patProof.v_rel_loc|split|0','thm.exh_to_patProof.v_rel_loc|split|0')).
fof(f1146,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) | $spl1),
  inference(cnf_transformation,[],[f1146_D])).
fof(f1146_D,plain,(
  ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1148,plain,(
  s('type.patSem.v',sK28) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) | $spl3),
  inference(cnf_transformation,[],[f1148_D])).
fof(f1148_D,plain,(
  s('type.patSem.v',sK28) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f657,plain,(
  s('type.patSem.v',sK28) != s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))) | ~p(s(bool,'const.exh_to_patProof.v_rel_2'(s('type.patSem.v','const.patSem.Loc_1'(s('type.num.num',sK29))),s('type.patSem.v',sK28))))),
  inference(cnf_transformation,[],[f440])).
% SZS output end Proof for thm.exh_to_patProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 632
% Passive clauses: 2291
% Generated clauses: 11411
% Final active clauses: 584
% Final passive clauses: 1530
% Input formulas: 133
% Initial clauses: 464
% 
% Duplicate literals: 8
% Trivial inequalities: 8
% Fw subsumption resolutions: 332
% Fw demodulations: 1830
% Bw demodulations: 37
% Condensations: 19
% 
% Simple tautologies: 88
% Equational tautologies: 37
% Forward subsumptions: 5600
% Fw demodulations to eq. taut.: 720
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 4944
% Factoring: 1
% Forward superposition: 2288
% Backward superposition: 1406
% Self superposition: 16
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 10
% Split components: 20
% SAT solver clauses: 2781
% SAT solver unit clauses: 970
% SAT solver binary clauses: 1286
% 
% Sat splits: 10
% Sat splitting refutations: 468
% 
% Memory used [KB]: 11257
% Time elapsed: 0.592 s
% ------------------------------
% ------------------------------
% Success in time 1.914 s
