Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.evalProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.evalProps
% SZS output start Proof for thm.evalProps
fof(f3209,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3208,f580,f292,f444,f293,f446,f442,f3156,f294,f447,f445,f443])).
fof(f443,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl1),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f445,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | $spl3),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f447,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',sK0) | $spl5),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',sK0) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f294,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',sK0) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),sK1) | s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f263])).
fof(f263,plain,(
  ? [X0,X1] : ((s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X1) | s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',X0)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ? [X0,X1] : ((s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X0))) & (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) | (s('type.list.list'('type.ast.t'),'const.list.NIL_0') != s('type.list.list'('type.ast.t'),X1) | s('type.ast.tctor','const.ast.TC_word8_0') != s('type.ast.tctor',X0))))),
  inference(nnf_transformation,[],[f257])).
fof(f257,plain,(
  ? [X0,X1] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) <~> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X0)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ~! [X0,X1] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X1),s('type.ast.tctor',X0))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X1) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X7),s('type.ast.tctor',X6))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X7) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X7),s('type.ast.tctor',X6))) <=> (s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),X7) & s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',X6)))),
  file('thm.evalProps.Tword_simp|split|18','thm.evalProps.Tword_simp|split|18')).
fof(f3156,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f3155,f445])).
fof(f3155,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f3149])).
fof(f3149,plain,(
  s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t','const.ast.Tword8_0') | s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | ($spl0 | $spl4)),
  inference(superposition,[],[f2017,f579])).
fof(f579,plain,(
  s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0')))),
  inference(forward_demodulation,[],[f573,f300])).
fof(f300,plain,(
  s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0')))),
  inference(cnf_transformation,[],[f100])).
fof(f100,axiom,(
  s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0')))),
  file('thm.evalProps.Tword_simp|split|18','thm.ast.Tword8_def')).
fof(f573,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0')))),
  inference(superposition,[],[f375,f297])).
fof(f297,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W8_0')))),
  inference(cnf_transformation,[],[f97])).
fof(f97,axiom,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size','const.ast.W8_0')))),
  file('thm.evalProps.Tword_simp|split|18','thm.ast.TC_word_def|split|0')).
fof(f375,plain,(
  ( ! [X0] : (s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X0))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X16] : s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size',X16))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word_1'(s('type.ast.word_size',X16)))))),
  file('thm.evalProps.Tword_simp|split|18','thm.ast.Tword_def')).
fof(f2017,plain,(
  ( ! [X12,X11] : (s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X11),s('type.ast.tctor',X12))) | s('type.list.list'('type.ast.t'),sK1) = s('type.list.list'('type.ast.t'),X11)) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f437,f583])).
fof(f583,plain,(
  s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f582,f300])).
fof(f582,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f442,f446])).
fof(f437,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & (s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2,X3] : ((s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3))) & ((s('type.list.list'('type.ast.t'),X0) != s('type.list.list'('type.ast.t'),X2) | s('type.ast.tctor',X1) != s('type.ast.tctor',X3)) | s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3)))))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) <=> (s('type.list.list'('type.ast.t'),X0) = s('type.list.list'('type.ast.t'),X2) & s('type.ast.tctor',X1) = s('type.ast.tctor',X3)))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X12,X13,X14,X15] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X12),s('type.ast.tctor',X13))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X14),s('type.ast.tctor',X15))) <=> (s('type.list.list'('type.ast.t'),X12) = s('type.list.list'('type.ast.t'),X14) & s('type.ast.tctor',X13) = s('type.ast.tctor',X15)))),
  file('thm.evalProps.Tword_simp|split|18','thm.ast.t_|49||49||split|2')).
fof(f442,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl0),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f446,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0) | $spl4),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f293,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f444,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | $spl2),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f292,plain,(
  s('type.list.list'('type.ast.t'),'const.list.NIL_0') = s('type.list.list'('type.ast.t'),sK1) | s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0)))),
  inference(cnf_transformation,[],[f264])).
fof(f580,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f579,f524])).
fof(f524,plain,(
  s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f523,f300])).
fof(f523,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),'const.list.NIL_0'),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f522,f444])).
fof(f522,plain,(
  s('type.ast.t','const.ast.Tword_1'(s('type.ast.word_size','const.ast.W8_0'))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor','const.ast.TC_word8_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f443,f446])).
fof(f3208,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f3206,f447])).
fof(f3206,plain,(
  s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0) | $spl0),
  inference(trivial_inequality_removal,[],[f3200])).
fof(f3200,plain,(
  s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t','const.ast.Tword8_0') | s('type.ast.tctor','const.ast.TC_word8_0') = s('type.ast.tctor',sK0) | $spl0),
  inference(superposition,[],[f3182,f579])).
fof(f3182,plain,(
  ( ! [X6,X7] : (s('type.ast.t','const.ast.Tword8_0') != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X6),s('type.ast.tctor',X7))) | s('type.ast.tctor',sK0) = s('type.ast.tctor',X7)) ) | $spl0),
  inference(superposition,[],[f438,f3177])).
fof(f3177,plain,(
  s('type.ast.t','const.ast.Tword8_0') = s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),sK1),s('type.ast.tctor',sK0))) | $spl0),
  inference(forward_demodulation,[],[f442,f300])).
fof(f438,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X0),s('type.ast.tctor',X1))) != s('type.ast.t','const.ast.Tapp_2'(s('type.list.list'('type.ast.t'),X2),s('type.ast.tctor',X3))) | s('type.ast.tctor',X1) = s('type.ast.tctor',X3)) )),
  inference(cnf_transformation,[],[f291])).
% SZS output end Proof for thm.evalProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 205
% Passive clauses: 559
% Generated clauses: 2795
% Final active clauses: 172
% Final passive clauses: 212
% Input formulas: 120
% Initial clauses: 148
% 
% Duplicate literals: 223
% Trivial inequalities: 57
% Fw subsumption resolutions: 290
% Fw demodulations: 101
% Bw demodulations: 7
% 
% Simple tautologies: 137
% Equational tautologies: 51
% Forward subsumptions: 1218
% 
% Binary resolution: 35
% Factoring: 1
% Forward superposition: 938
% Backward superposition: 712
% Self superposition: 82
% Equality factoring: 85
% Equality resolution: 13
% 
% Split clauses: 109
% Split components: 219
% SAT solver clauses: 197
% SAT solver unit clauses: 40
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 109
% Sat splitting refutations: 27
% 
% TWLsolver clauses: 194
% TWLsolver calls for satisfiability: 54
% 
% Memory used [KB]: 1791
% Time elapsed: 0.061 s
% ------------------------------
% ------------------------------
% Success in time 0.065 s
