Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.benchmark
% (129418)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 161
% Passive clauses: 3428
% Generated clauses: 18125
% Final active clauses: 142
% Final passive clauses: 3082
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 550
% Trivial inequalities: 314
% Fw subsumption resolutions: 1498
% Fw demodulations: 2568
% Bw demodulations: 19
% 
% Simple tautologies: 262
% Equational tautologies: 46
% Forward subsumptions: 8968
% 
% Binary resolution: 513
% Factoring: 1
% Forward superposition: 6241
% Backward superposition: 6038
% Self superposition: 149
% Equality factoring: 70
% Equality resolution: 32
% 
% Split clauses: 252
% Split components: 506
% SAT solver clauses: 676
% SAT solver unit clauses: 9
% SAT solver binary clauses: 151
% 
% Sat splits: 252
% 
% TWLsolver clauses: 676
% TWLsolver calls for satisfiability: 185
% 
% Memory used [KB]: 13816
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.benchmark
% (129438)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1340
% Generated clauses: 14834
% Final active clauses: 117
% Final passive clauses: 979
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 475
% Trivial inequalities: 170
% Fw subsumption resolutions: 851
% Fw demodulations: 1540
% Bw demodulations: 218
% Condensations: 1118
% Global subsumptions: 8
% 
% Simple tautologies: 192
% Equational tautologies: 46
% Forward subsumptions: 6818
% Bw demodulations to eq. taut.: 23
% 
% Binary resolution: 292
% Factoring: 1
% Forward superposition: 4716
% Backward superposition: 3640
% Self superposition: 146
% Equality factoring: 78
% Equality resolution: 23
% 
% Split clauses: 108
% Split components: 216
% SAT solver clauses: 2288
% SAT solver unit clauses: 150
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 16
% 
% Sat splits: 108
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2256
% TWLsolver calls for satisfiability: 419
% 
% Memory used [KB]: 16375
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.benchmark
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.benchmark
% SZS output start Proof for thm.benchmark
fof(f49165,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1268,f1269,f1261,f1273,f1807,f1810,f1808,f3367,f3373,f3861,f3844,f3864,f3999,f4000,f4003,f4005,f4006,f4008,f4011,f4013,f4015,f4016,f4017,f4019,f4021,f4023,f4025,f4027,f4029,f4031,f4033,f4036,f4038,f4040,f4042,f4043,f4044,f4056,f4058,f4061,f4071,f4073,f4075,f4077,f4079,f4081,f4083,f4085,f4086,f4087,f4089,f4092,f4094,f4096,f4098,f4101,f4104,f4105,f4106,f4108,f4110,f4112,f4114,f4116,f4117,f7710,f7712,f7695,f7721,f7714,f8529,f8601,f8675,f8677,f8678,f8679,f8733,f8736,f8734,f9139,f9140,f9149,f9154,f9159,f9160,f9173,f9175,f9178,f9181,f9183,f9203,f9209,f9211,f9213,f9215,f9217,f9219,f9221,f9223,f9224,f9225,f9226,f9227,f9228,f9230,f9232,f11262,f11264,f11266,f11270,f11272,f11276,f11278,f11319,f11323,f11736,f11737,f11740,f11744,f11748,f11749,f11750,f11751,f11753,f11755,f11757,f11758,f11759,f11760,f11761,f11762,f11764,f11766,f11767,f11768,f11769,f11770,f11771,f11772,f11773,f11774,f19732,f19813,f15028,f14125,f4538,f4555,f3269,f3374,f3286,f1271,f3862,f20787,f21986,f20868,f37047,f37048,f37049,f37050,f37118,f37051,f37115,f37119,f37053,f37054,f37122,f37055,f37056,f37123,f37124,f37125,f37126,f37117,f37121,f37127,f37114,f37112,f37139,f40337,f44421,f44423,f44422,f44560,f44562,f44564,f44568,f44570,f44572,f44574,f44576,f44578,f44580,f44582,f44584,f44586,f44588,f44590,f44592,f44594,f44598,f44600,f44602,f44604,f44606,f44608,f44610,f44612,f44614,f44616,f44618,f44620,f44622,f44624,f44629,f44720,f44722,f44724,f44726,f44728,f44730,f44732,f44748,f44755,f44761,f44765,f44767,f44867,f44868,f44869,f44870,f44871,f44872,f44873,f44874,f44875,f44876,f44877,f44878,f44879,f44880,f44881,f44882,f44883,f44884,f44885,f44886,f44887,f44888,f44889,f44890,f44891,f44892,f44893,f44894,f44895,f44896,f44897,f44898,f44899,f44900,f44901,f44902,f45347,f40435,f8600,f45674,f45676,f45675,f1809,f48688,f48690,f48691,f48693,f48694,f48696,f48697,f48699,f48701,f48700,f48702,f48703,f48705,f48706,f48709,f48710,f48712,f48713,f48715,f48716,f48718,f48719,f48721,f48724,f48725,f48726,f48727,f48728,f48729,f48689,f48731,f48732,f48734,f48735,f48736,f48698,f48737,f48742,f48743,f48746,f48747,f48692,f48751,f48752,f48755,f48756,f48758,f48760,f48759,f48762,f48763,f48695,f48770,f48707,f48722,f48781,f48783,f48784,f48785,f48786,f48787,f48788,f48789,f48790,f48791,f48782,f48792,f48794,f48795,f48796,f48797,f48798,f48799,f48800,f48801,f48802,f48793,f48811,f48812,f48813,f48816,f48818,f48820,f48824,f48826,f48828,f48830,f48832,f48834,f48836,f48838,f48840,f48842,f48844,f48846,f48850,f48852,f48854,f48856,f48858,f48860,f48862,f48864,f48866,f48868,f48870,f48872,f48877,f48970,f48972,f48974,f48976,f48978,f48980,f48994,f49001,f49007,f49011,f49013,f49015,f49118,f49119,f49120,f49121,f49122,f49123,f49124,f49125,f49126,f49127,f49128,f49129,f49130,f49131,f49132,f49133,f49134,f49135,f49136,f49137,f49138,f49139,f49140,f49141,f49142,f49143,f49144,f49145,f49146,f49147,f49148,f49163,f49164])).
fof(f49164,plain,(
  $false | ($spl62 | $spl84 | $spl104)),
  inference(subsumption_resolution,[],[f48752,f49162])).
fof(f49162,plain,(
  ( ! [X386] : (p(s(bool,X386))) ) | ($spl62 | $spl84)),
  inference(subsumption_resolution,[],[f49161,f231])).
fof(f231,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.benchmark.digit_distinct|split|2',aHL_TRUTH)).
fof(f49161,plain,(
  ( ! [X386] : (~p(s(bool,'T_0')) | p(s(bool,X386))) ) | ($spl62 | $spl84)),
  inference(forward_demodulation,[],[f48719,f48693])).
fof(f49163,plain,(
  $false | ($spl62 | $spl84 | $spl106)),
  inference(subsumption_resolution,[],[f48756,f49162])).
fof(f49148,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48819,f49000])).
fof(f49000,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X43,X45,X46] : (p(s(bool,happ(s(fun(X43,bool),happ(s(fun(X43,fun(X43,bool)),'=_0'),s(X43,X45))),s(X43,X46)))) <=> s(X43,X45) = s(X43,X46))),
  file('thm.benchmark.digit_distinct|split|2',arity_0_HL_EQ)).
fof(f48814,plain,(
  ( ! [X386] : (p(s(bool,X386))) ) | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48719,f48707])).
fof(f48819,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f245])).
fof(f245,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X24] : (s(bool,'F_0') = s(bool,X24) <=> ~p(s(bool,X24)))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f49147,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48823,f231])).
fof(f48823,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f48814,f298])).
fof(f298,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f344,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f337,f241])).
fof(f241,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X24] : (s(bool,'T_0') = s(bool,X24) <=> p(s(bool,X24)))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f337,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f234,f231])).
fof(f234,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f45])).
fof(f45,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.NOT_CLAUSES|split|2')).
fof(f49146,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48825,f231])).
fof(f48825,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12126,f48814,f298])).
fof(f12126,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl21),
  inference(unit_resulting_resolution,[],[f7713,f241])).
fof(f7713,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl21),
  inference(cnf_transformation,[],[f7713_D])).
fof(f7713_D,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f49145,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48827,f231])).
fof(f48827,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f31171,f48814,f298])).
fof(f31171,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f31051,f337])).
fof(f31051,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(superposition,[],[f241,f29559])).
fof(f29559,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f315,f28362,f1700])).
fof(f1700,plain,(
  ( ! [X43,X46] : (s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1699,f231])).
fof(f1699,plain,(
  ( ! [X43,X46] : (~p(s(bool,'T_0')) | s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(forward_demodulation,[],[f1613,f266])).
fof(f266,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X25] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X6,X25)))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.DATATYPE_TAG_THM')).
fof(f1613,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X44,X45)))) | s(bool,X43) != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f362,f1239])).
fof(f1239,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1226,f231])).
fof(f1226,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,'const.bool.DATATYPE_1'(s(X24,X25))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f1212,f266])).
fof(f1212,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f1211,f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X14] : ! [X16] : s(X14,X16) = s(X14,happ(s(fun(X14,X14),'_i_0_0'),s(X14,X16)))),
  file('thm.benchmark.digit_distinct|split|2','_i0_ax')).
fof(f1211,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f1182,f231])).
fof(f1182,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f298,f392])).
fof(f392,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f267,f238])).
fof(f238,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.benchmark.digit_distinct|split|2',aHL_BOOL_CASES)).
fof(f362,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f245,f238])).
fof(f28362,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f28263,f7712])).
fof(f28263,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK6('_i_0_0',bool)))) | ($spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f11921,f28234])).
fof(f28234,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f28108,f337])).
fof(f28108,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f241,f19873])).
fof(f19873,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl26)),
  inference(subsumption_resolution,[],[f18856,f7712])).
fof(f18856,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,sK6('_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X27))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f12804,f11920])).
fof(f11920,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f11874,f344])).
fof(f11874,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f7712,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f238,f238])).
fof(f12804,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,X33) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f12767,f7712])).
fof(f12767,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK6('_i_0_0',bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl26),
  inference(superposition,[],[f8601,f267])).
fof(f11921,plain,(
  s(bool,'T_0') != s(bool,sK6('_i_0_0',bool)) | $spl18),
  inference(subsumption_resolution,[],[f11880,f337])).
fof(f11880,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK6('_i_0_0',bool)) | $spl18),
  inference(superposition,[],[f241,f7712])).
fof(f315,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,happ(s(fun(X1,X0),X3),s(X1,X2))) = s(X0,happ(s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,X2))),s(fun(X1,X0),X3)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X12,X14] : ! [X16,X17] : s(X12,happ(s(fun(X14,X12),X17),s(X14,X16))) = s(X12,happ(s(fun(fun(X14,X12),X12),'_dstsel_0_x0_1'(s(X14,X16))),s(fun(X14,X12),X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dstsel_0_x0_ax')).
fof(f49144,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48829,f231])).
fof(f48829,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28362,f48814,f298])).
fof(f49143,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48831,f231])).
fof(f48831,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21339,f48814,f298])).
fof(f21339,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1273,f241])).
fof(f49142,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48833,f231])).
fof(f48833,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11921,f48814,f298])).
fof(f49141,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48837,f48814])).
fof(f48837,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12126,f48814,f298])).
fof(f49140,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48839,f48814])).
fof(f48839,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f31171,f48814,f298])).
fof(f49139,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48841,f48814])).
fof(f48841,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28362,f48814,f298])).
fof(f49138,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48843,f48814])).
fof(f48843,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21339,f48814,f298])).
fof(f49137,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48845,f48814])).
fof(f48845,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11921,f48814,f298])).
fof(f49136,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48849,f231])).
fof(f48849,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f48814,f298])).
fof(f49135,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48851,f231])).
fof(f48851,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12126,f48814,f298])).
fof(f49134,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48853,f231])).
fof(f48853,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f31171,f48814,f298])).
fof(f49133,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48855,f231])).
fof(f48855,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28362,f48814,f298])).
fof(f49132,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48857,f231])).
fof(f48857,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21339,f48814,f298])).
fof(f49131,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48859,f231])).
fof(f48859,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11921,f48814,f298])).
fof(f49130,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48863,f48814])).
fof(f48863,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12126,f48814,f298])).
fof(f49129,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48865,f48814])).
fof(f48865,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f31171,f48814,f298])).
fof(f49128,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48867,f48814])).
fof(f48867,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28362,f48814,f298])).
fof(f49127,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48869,f48814])).
fof(f48869,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21339,f48814,f298])).
fof(f49126,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48871,f48814])).
fof(f48871,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11921,f48814,f298])).
fof(f49125,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48875,f48814])).
fof(f48875,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0'))))) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814,f306])).
fof(f306,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X45,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X45))),s(bool,X46)))) <=> (p(s(bool,X45)) => p(s(bool,X46))))),
  file('thm.benchmark.digit_distinct|split|2','==>_HL_0')).
fof(f49124,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48969,f49000])).
fof(f48969,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f48814,f362])).
fof(f49123,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48971,f49000])).
fof(f48971,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) ) | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12126,f48814,f362])).
fof(f49122,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48973,f49000])).
fof(f48973,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) ) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f31171,f48814,f362])).
fof(f49121,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48975,f49000])).
fof(f48975,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f28362,f48814,f362])).
fof(f49120,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48977,f49000])).
fof(f48977,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK4(X1,'_i_0_0',bool))) ) | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21339,f48814,f362])).
fof(f49119,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f48979,f49000])).
fof(f48979,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK6('_i_0_0',bool))) ) | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f11921,f48814,f362])).
fof(f49118,plain,(
  $false | ($spl76 | $spl84)),
  inference(subsumption_resolution,[],[f49010,f48814])).
fof(f49010,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814,f306])).
fof(f49015,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814,f8244])).
fof(f8244,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'F_0')))) | p(s(bool,X0))) )),
  inference(condensation,[],[f8243])).
fof(f8243,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f8122,f241])).
fof(f8122,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f303,f238])).
fof(f303,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X45,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X45))),s(bool,X46)))) <=> (p(s(bool,X45)) | p(s(bool,X46))))),
  file('thm.benchmark.digit_distinct|split|2','|bslash||slash|_HL_0')).
fof(f49013,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814,f8178])).
fof(f8178,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X0)))) | p(s(bool,X0))) )),
  inference(condensation,[],[f8177])).
fof(f8177,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f8080,f241])).
fof(f8080,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f303,f238])).
fof(f49011,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f337,f48814,f306])).
fof(f49007,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f337,f48814,f303])).
fof(f49001,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f48814,f320])).
fof(f48994,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f344,f48814,f1212])).
fof(f48980,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f7712,f11921,f48814,f362])).
fof(f48978,plain,(
  $false | ($spl4 | $spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f21370,f21339,f48814,f362])).
fof(f21370,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | ($spl4 | $spl18)),
  inference(forward_demodulation,[],[f21353,f7712])).
fof(f21353,plain,(
  ( ! [X0] : (s(bool,sK6('_i_0_0',bool)) = s(bool,sK4(X0,'_i_0_0',bool))) ) | ($spl4 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1273,f418])).
fof(f418,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f412,f337])).
fof(f412,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f295,f238])).
fof(f295,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f48976,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f267,f28362,f48814,f362])).
fof(f48974,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f29559,f31171,f48814,f362])).
fof(f48972,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12143,f12126,f48814,f362])).
fof(f12143,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl21)),
  inference(forward_demodulation,[],[f12133,f7712])).
fof(f12133,plain,(
  s(bool,sK6('_i_0_0',bool)) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f11921,f7713,f418])).
fof(f48970,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12143,f344,f48814,f362])).
fof(f48877,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814,f48814,f306])).
fof(f48872,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f11921,f48814,f298])).
fof(f48870,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f21339,f48814,f298])).
fof(f48868,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f28362,f48814,f298])).
fof(f48866,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f31171,f48814,f298])).
fof(f48864,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f12126,f48814,f298])).
fof(f48862,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f344,f48814,f298])).
fof(f48860,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f11921,f48814,f298])).
fof(f48858,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f21339,f48814,f298])).
fof(f48856,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f28362,f48814,f298])).
fof(f48854,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f31171,f48814,f298])).
fof(f48852,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f12126,f48814,f298])).
fof(f48850,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f344,f48814,f298])).
fof(f48846,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f11921,f48814,f298])).
fof(f48844,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f21339,f48814,f298])).
fof(f48842,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f28362,f48814,f298])).
fof(f48840,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f31171,f48814,f298])).
fof(f48838,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f12126,f48814,f298])).
fof(f48836,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f48814,f344,f48814,f298])).
fof(f48834,plain,(
  $false | ($spl18 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f11921,f48814,f298])).
fof(f48832,plain,(
  $false | ($spl4 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f21339,f48814,f298])).
fof(f48830,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f28362,f48814,f298])).
fof(f48828,plain,(
  $false | ($spl18 | $spl26 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f31171,f48814,f298])).
fof(f48826,plain,(
  $false | ($spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f12126,f48814,f298])).
fof(f48824,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f231,f344,f48814,f298])).
fof(f48820,plain,(
  $false | ($spl18 | $spl21 | $spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f12143,f48814,f245])).
fof(f48818,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f344,f48814,f242])).
fof(f242,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f192])).
fof(f48816,plain,(
  $false | ($spl76 | $spl84)),
  inference(unit_resulting_resolution,[],[f337,f48814])).
fof(f48813,plain,(
  $false | ($spl56 | $spl60)),
  inference(subsumption_resolution,[],[f48689,f48692])).
fof(f48812,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f48810,f337])).
fof(f48810,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(backward_demodulation,[],[f48698,f48393])).
fof(f48393,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.benchmark.times'(sK0),X0),X2),s(X0,happ(s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))),s('type.benchmark.times'(sK0),sK3))))) )),
  inference(forward_demodulation,[],[f48392,f30353])).
fof(f30353,plain,(
  ( ! [X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (s(X16,happ(s(fun(X20,X16),'_dst_x0x1x0_3'(s(fun(X18,fun(X17,X16)),X19),s(fun(X20,X18),X21),s(X17,happ(s(fun(X23,X17),X24),s(X23,X25))))),s(X20,X22))) = s(X16,happ(s(fun(X23,X16),'_dst_x0x0x1_3'(s(fun(X18,fun(X17,X16)),X19),s(X18,happ(s(fun(X20,X18),X21),s(X20,X22))),s(fun(X23,X17),X24))),s(X23,X25)))) )),
  inference(superposition,[],[f327,f326])).
fof(f326,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X2,X0)),X4),s(X1,X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,X5))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X12,X13,X39,X14] : ! [X15,X16,X40,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x0x1_3'(s(fun(X13,fun(X39,X12)),X15),s(X13,X16),s(fun(X14,X39),X40))),s(X14,X17))) = s(X12,happ(s(fun(X39,X12),happ(s(fun(X13,fun(X39,X12)),X15),s(X13,X16))),s(X39,happ(s(fun(X14,X39),X40),s(X14,X17)))))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x0x0x1_ax')).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X2,X0),'_dst_x0x1x0_3'(s(fun(X1,fun(X3,X0)),X4),s(fun(X2,X1),X5),s(X3,X6))),s(X2,X7))) = s(X0,happ(s(fun(X3,X0),happ(s(fun(X1,fun(X3,X0)),X4),s(X1,happ(s(fun(X2,X1),X5),s(X2,X7))))),s(X3,X6)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X12,X13,X14,X39] : ! [X15,X16,X40,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x1x0_3'(s(fun(X13,fun(X39,X12)),X15),s(fun(X14,X13),X16),s(X39,X40))),s(X14,X17))) = s(X12,happ(s(fun(X39,X12),happ(s(fun(X13,fun(X39,X12)),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))),s(X39,X40)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x0x1x0_ax')).
fof(f48392,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.benchmark.times'(sK0),X0),X2),s('type.benchmark.times'(sK0),sK3))),s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48391,f318])).
fof(f318,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.benchmark.times'(X1),X0),X4),s('type.benchmark.times'(X1),X2))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.One_1'(s('type.benchmark.times'(X1),X2))),s(X0,X3),s(fun('type.benchmark.times'(X1),X0),X4),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X5)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun('type.benchmark.times'(X1),X0),X4),s('type.benchmark.times'(X1),X2))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.One_1'(s('type.benchmark.times'(X1),X2))),s(X0,X3),s(fun('type.benchmark.times'(X1),X0),X4),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X5)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun('type.benchmark.times'(X1),X0),X4),s('type.benchmark.times'(X1),X2))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.One_1'(s('type.benchmark.times'(X1),X2))),s(X0,X3),s(fun('type.benchmark.times'(X1),X0),X4),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X5)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X26,X6] : ! [X9,X35,X21,X36] : s(X26,happ(s(fun('type.benchmark.times'(X6),X26),X21),s('type.benchmark.times'(X6),X9))) = s(X26,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9))),s(X26,X35),s(fun('type.benchmark.times'(X6),X26),X21),s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),X26)),X36)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_case_def|split|1')).
fof(f48391,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(X0,X1),s(fun('type.benchmark.times'(sK0),X0),X2),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3))),s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48390,f28873])).
fof(f28873,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X13,X8),'_dst_x0x0x1_3'(s(fun(X10,fun(X9,X8)),X11),s(X10,X12),s(fun(X13,X9),X14))),s(X13,X15))) = s(X8,happ(s(fun(X10,X8),'_dst_x1x0_2'(s(fun(X10,fun(X9,X8)),X11),s(X9,happ(s(fun(X13,X9),X14),s(X13,X15))))),s(X10,X12)))) )),
  inference(superposition,[],[f326,f324])).
fof(f324,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X12,X14,X13] : ! [X15,X16,X17] : s(X12,happ(s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),X15),s(X14,X17))),s(X13,X16))) = s(X12,happ(s(fun(X14,X12),'_dst_x1x0_2'(s(fun(X14,fun(X13,X12)),X15),s(X13,X16))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x1x0_ax')).
fof(f48390,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))),s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(X0,X1),s(fun('type.benchmark.times'(sK0),X0),X2),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3))))))) )),
  inference(forward_demodulation,[],[f47512,f324])).
fof(f47512,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(X0,X1),s(fun('type.benchmark.times'(sK0),X0),X2),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3))))),s(X0,happ(s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X3),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))))) )),
  inference(unit_resulting_resolution,[],[f32462,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f32462,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.benchmark.times'(sK0),X0),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X1),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(X0,X2),s(fun('type.benchmark.times'(sK0),X0),X3),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),X0)),X1)))) )),
  inference(superposition,[],[f319,f230])).
fof(f230,plain,(
  s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) = s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),sK2),s('type.benchmark.times'(sK0),sK1)))),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) = s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),sK2),s('type.benchmark.times'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f182])).
fof(f182,plain,(
  ? [X0,X1,X2,X3] : s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X3))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3] : s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X3))) != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X3))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9))) = s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9))) = s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_distinct|split|2')).
fof(f319,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.benchmark.times'(X1),X0),happ(s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6),s('type.benchmark.times'(X1),X2))),s('type.benchmark.times'(X1),X3))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.Two_2'(s('type.benchmark.times'(X1),X2),s('type.benchmark.times'(X1),X3))),s(X0,X4),s(fun('type.benchmark.times'(X1),X0),X5),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6)))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.benchmark.times'(X1),X0),happ(s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6),s('type.benchmark.times'(X1),X2))),s('type.benchmark.times'(X1),X3))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.Two_2'(s('type.benchmark.times'(X1),X2),s('type.benchmark.times'(X1),X3))),s(X0,X4),s(fun('type.benchmark.times'(X1),X0),X5),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.benchmark.times'(X1),X0),happ(s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6),s('type.benchmark.times'(X1),X2))),s('type.benchmark.times'(X1),X3))) = s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.Two_2'(s('type.benchmark.times'(X1),X2),s('type.benchmark.times'(X1),X3))),s(X0,X4),s(fun('type.benchmark.times'(X1),X0),X5),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X6)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X26,X6] : ! [X8,X7,X35,X21,X36] : s(X26,happ(s(fun('type.benchmark.times'(X6),X26),happ(s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),X26)),X36),s('type.benchmark.times'(X6),X8))),s('type.benchmark.times'(X6),X7))) = s(X26,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7))),s(X26,X35),s(fun('type.benchmark.times'(X6),X26),X21),s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),X26)),X36)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_case_def|split|2')).
fof(f48811,plain,(
  $false | ($spl68 | $spl76)),
  inference(subsumption_resolution,[],[f48809,f337])).
fof(f48809,plain,(
  p(s(bool,'F_0')) | ($spl68 | $spl76)),
  inference(backward_demodulation,[],[f48698,f48707])).
fof(f48793,plain,(
  ( ! [X1089,X1090] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1090),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1089),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) ) | $spl126),
  inference(cnf_transformation,[],[f48793_D])).
fof(f48793_D,plain,(
  ( ! [X1089,X1090] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1090),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1089),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f48802,plain,(
  ( ! [X1130,X1134,X1133] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1133,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1133,X1134),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1130),s('type.benchmark.times'(sK0),sK3)))))))) ) | $spl136),
  inference(cnf_transformation,[],[f48802_D])).
fof(f48802_D,plain,(
  ( ! [X1130,X1134,X1133] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1133,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1133,X1134),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1130),s('type.benchmark.times'(sK0),sK3)))))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f48801,plain,(
  ( ! [X1130,X1134,X1132,X1131,X1133] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1133,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1133,X1134),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1130),s('type.benchmark.times'(sK0),sK3))))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1132),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1131),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f48066,f318])).
fof(f48066,plain,(
  ( ! [X1130,X1134,X1132,X1131,X1129,X1133] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1132),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1131),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1133,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1133,X1134),s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1129),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1130),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1131)))))))) )),
  inference(superposition,[],[f20500,f32462])).
fof(f20500,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X27),s('type.benchmark.times'(sK0),X26))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X24,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X24,X25),s('type.benchmark.times'(sK0),X26)))))) )),
  inference(superposition,[],[f20244,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X26] : ! [X25,X31] : s(X26,X31) = s(X26,'const.pair.SND_1'(s('type.pair.prod'(X6,X26),'const.pair.,_2'(s(X6,X25),s(X26,X31)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.pair.SND')).
fof(f20244,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X0),s('type.benchmark.times'(sK0),X1))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),X1)) )),
  inference(superposition,[],[f292,f230])).
fof(f292,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4))) & (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4) | s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4))) & ((s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4)) | s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) <=> (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) <=> (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X8,X7,X32,X33] : (s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7))) = s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X32),s('type.benchmark.times'(X6),X33))) <=> (s('type.benchmark.times'(X6),X8) = s('type.benchmark.times'(X6),X32) & s('type.benchmark.times'(X6),X7) = s('type.benchmark.times'(X6),X33)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_|49||49||split|1')).
fof(f48800,plain,(
  ( ! [X1118,X1122,X1121] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1121),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1118),s('type.benchmark.times'(sK0),sK3))),s(X1121,X1122)))))) ) | $spl134),
  inference(cnf_transformation,[],[f48800_D])).
fof(f48800_D,plain,(
  ( ! [X1118,X1122,X1121] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1121),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1118),s('type.benchmark.times'(sK0),sK3))),s(X1121,X1122)))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f48799,plain,(
  ( ! [X1118,X1122,X1120,X1119,X1121] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1121),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1118),s('type.benchmark.times'(sK0),sK3))),s(X1121,X1122))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1120),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1119),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f48064,f318])).
fof(f48064,plain,(
  ( ! [X1118,X1122,X1120,X1119,X1117,X1121] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1120),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1119),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1121),'const.pair.,_2'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1117),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1118),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1119))),s(X1121,X1122)))))) )),
  inference(superposition,[],[f20499,f32462])).
fof(f20499,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X23),s('type.benchmark.times'(sK0),X21))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X20),'const.pair.,_2'(s('type.benchmark.times'(sK0),X21),s(X20,X22)))))) )),
  inference(superposition,[],[f20244,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X26,X6] : ! [X25,X31] : s(X6,X25) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X26),'const.pair.,_2'(s(X6,X25),s(X26,X31)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.pair.FST')).
fof(f48798,plain,(
  ( ! [X1108,X1111] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1108),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1111),s('type.num.num','const.num.|48|_0')))) ) | $spl132),
  inference(cnf_transformation,[],[f48798_D])).
fof(f48798_D,plain,(
  ( ! [X1108,X1111] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1108),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1111),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f48797,plain,(
  ( ! [X1110,X1108,X1111,X1109] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1108),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1111),s('type.num.num','const.num.|48|_0'))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1110),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1109),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f48062,f318])).
fof(f48062,plain,(
  ( ! [X1110,X1108,X1107,X1111,X1109] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1110),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1109),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1107),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1108),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1109))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1111),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f20497,f32462])).
fof(f20497,plain,(
  ( ! [X14,X15,X13] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X15),s('type.benchmark.times'(sK0),X13))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),X13),s(fun('type.num.num','type.benchmark.times'(sK0)),X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f20244,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X9,X21] : s(X6,X9) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X21),s('type.num.num','const.num.|48|_0')))),
  file('thm.benchmark.digit_distinct|split|2','thm.ind_type.FCONS|split|0')).
fof(f48796,plain,(
  ( ! [X1099,X1096,X1100] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun(X1099,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1096),s('type.benchmark.times'(sK0),sK3))))),s(X1099,X1100)))) ) | $spl130),
  inference(cnf_transformation,[],[f48796_D])).
fof(f48796_D,plain,(
  ( ! [X1099,X1096,X1100] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun(X1099,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1096),s('type.benchmark.times'(sK0),sK3))))),s(X1099,X1100)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f48795,plain,(
  ( ! [X1099,X1097,X1098,X1096,X1100] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun(X1099,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1096),s('type.benchmark.times'(sK0),sK3))))),s(X1099,X1100))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1098),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1097),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f48060,f318])).
fof(f48060,plain,(
  ( ! [X1099,X1097,X1095,X1098,X1096,X1100] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1098),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1097),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun(X1099,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1095),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1096),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1097))))),s(X1099,X1100)))) )),
  inference(superposition,[],[f20494,f32462])).
fof(f20494,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X5),s('type.benchmark.times'(sK0),X3))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun(X2,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f20244,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X14,X12] : ! [X15,X17] : s(X12,X15) = s(X12,happ(s(fun(X14,X12),'_dst_x0_1'(s(X12,X15))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x0_ax')).
fof(f48794,plain,(
  ( ! [X1088] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1088),s('type.benchmark.times'(sK0),sK3)))))) ) | $spl128),
  inference(cnf_transformation,[],[f48794_D])).
fof(f48794_D,plain,(
  ( ! [X1088] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1088),s('type.benchmark.times'(sK0),sK3)))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f48792,plain,(
  ( ! [X1089,X1090,X1088] : (s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1088),s('type.benchmark.times'(sK0),sK3))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1090),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1089),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f48058,f318])).
fof(f48058,plain,(
  ( ! [X1087,X1089,X1090,X1088] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1090),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1089),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1087),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1088),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1089)))))) )),
  inference(superposition,[],[f20493,f32462])).
fof(f20493,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X1),s('type.benchmark.times'(sK0),X0))) | s('type.benchmark.times'(sK0),sK1) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),X0)))) )),
  inference(superposition,[],[f20244,f267])).
fof(f48782,plain,(
  ( ! [X1035,X1036] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1035),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1036)))) ) | $spl114),
  inference(cnf_transformation,[],[f48782_D])).
fof(f48782_D,plain,(
  ( ! [X1035,X1036] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1035),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1036)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f48791,plain,(
  ( ! [X1079,X1080,X1076] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1079,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1079,X1080),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1076),s('type.benchmark.times'(sK0),sK3)))))))) ) | $spl124),
  inference(cnf_transformation,[],[f48791_D])).
fof(f48791_D,plain,(
  ( ! [X1079,X1080,X1076] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1079,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1079,X1080),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1076),s('type.benchmark.times'(sK0),sK3)))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f48790,plain,(
  ( ! [X1079,X1077,X1080,X1078,X1076] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1079,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1079,X1080),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1076),s('type.benchmark.times'(sK0),sK3))))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1077),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1078)))) )),
  inference(forward_demodulation,[],[f48056,f318])).
fof(f48056,plain,(
  ( ! [X1075,X1079,X1077,X1080,X1078,X1076] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1077),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1078))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X1079,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X1079,X1080),s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1075),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1076),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1077)))))))) )),
  inference(superposition,[],[f20402,f32462])).
fof(f20402,plain,(
  ( ! [X26,X24,X27,X25] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X26),s('type.benchmark.times'(sK0),X27))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.SND_1'(s('type.pair.prod'(X24,'type.benchmark.times'(sK0)),'const.pair.,_2'(s(X24,X25),s('type.benchmark.times'(sK0),X26)))))) )),
  inference(superposition,[],[f19960,f313])).
fof(f19960,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X0),s('type.benchmark.times'(sK0),X1))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),X0)) )),
  inference(superposition,[],[f291,f230])).
fof(f291,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3)) )),
  inference(cnf_transformation,[],[f215])).
fof(f48789,plain,(
  ( ! [X1064,X1068,X1067] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1067),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1064),s('type.benchmark.times'(sK0),sK3))),s(X1067,X1068)))))) ) | $spl122),
  inference(cnf_transformation,[],[f48789_D])).
fof(f48789_D,plain,(
  ( ! [X1064,X1068,X1067] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1067),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1064),s('type.benchmark.times'(sK0),sK3))),s(X1067,X1068)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f48788,plain,(
  ( ! [X1066,X1064,X1068,X1067,X1065] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1067),'const.pair.,_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1064),s('type.benchmark.times'(sK0),sK3))),s(X1067,X1068))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1065),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1066)))) )),
  inference(forward_demodulation,[],[f48054,f318])).
fof(f48054,plain,(
  ( ! [X1066,X1064,X1068,X1067,X1065,X1063] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1065),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1066))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X1067),'const.pair.,_2'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1063),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1064),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1065))),s(X1067,X1068)))))) )),
  inference(superposition,[],[f20401,f32462])).
fof(f20401,plain,(
  ( ! [X23,X21,X22,X20] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X21),s('type.benchmark.times'(sK0),X23))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(sK0),X20),'const.pair.,_2'(s('type.benchmark.times'(sK0),X21),s(X20,X22)))))) )),
  inference(superposition,[],[f19960,f311])).
fof(f48787,plain,(
  ( ! [X1054,X1057] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1054),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1057),s('type.num.num','const.num.|48|_0')))) ) | $spl120),
  inference(cnf_transformation,[],[f48787_D])).
fof(f48787_D,plain,(
  ( ! [X1054,X1057] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1054),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1057),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f48786,plain,(
  ( ! [X1054,X1056,X1055,X1057] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1054),s('type.benchmark.times'(sK0),sK3))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1057),s('type.num.num','const.num.|48|_0'))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1055),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1056)))) )),
  inference(forward_demodulation,[],[f48052,f318])).
fof(f48052,plain,(
  ( ! [X1054,X1056,X1055,X1053,X1057] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1055),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1056))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1053),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1054),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1055))),s(fun('type.num.num','type.benchmark.times'(sK0)),X1057),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f20399,f32462])).
fof(f20399,plain,(
  ( ! [X14,X15,X13] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X13),s('type.benchmark.times'(sK0),X15))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),'const.ind_type.FCONS_3'(s('type.benchmark.times'(sK0),X13),s(fun('type.num.num','type.benchmark.times'(sK0)),X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f19960,f276])).
fof(f48785,plain,(
  ( ! [X1042,X1046,X1045] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun(X1045,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1042),s('type.benchmark.times'(sK0),sK3))))),s(X1045,X1046)))) ) | $spl118),
  inference(cnf_transformation,[],[f48785_D])).
fof(f48785_D,plain,(
  ( ! [X1042,X1046,X1045] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun(X1045,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1042),s('type.benchmark.times'(sK0),sK3))))),s(X1045,X1046)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f48784,plain,(
  ( ! [X1042,X1046,X1044,X1043,X1045] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun(X1045,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1042),s('type.benchmark.times'(sK0),sK3))))),s(X1045,X1046))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1043),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1044)))) )),
  inference(forward_demodulation,[],[f48050,f318])).
fof(f48050,plain,(
  ( ! [X1042,X1046,X1044,X1043,X1041,X1045] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1043),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1044))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun(X1045,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1041),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1042),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1043))))),s(X1045,X1046)))) )),
  inference(superposition,[],[f20396,f32462])).
fof(f20396,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X3),s('type.benchmark.times'(sK0),X5))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun(X2,'type.benchmark.times'(sK0)),'_dst_x0_1'(s('type.benchmark.times'(sK0),X3))),s(X2,X4)))) )),
  inference(superposition,[],[f19960,f312])).
fof(f48783,plain,(
  ( ! [X1034] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1034),s('type.benchmark.times'(sK0),sK3)))))) ) | $spl116),
  inference(cnf_transformation,[],[f48783_D])).
fof(f48783_D,plain,(
  ( ! [X1034] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1034),s('type.benchmark.times'(sK0),sK3)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f48781,plain,(
  ( ! [X1035,X1034,X1036] : (s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1034),s('type.benchmark.times'(sK0),sK3))))) | s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1035),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1036)))) )),
  inference(forward_demodulation,[],[f48048,f318])).
fof(f48048,plain,(
  ( ! [X1035,X1033,X1034,X1036] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1035),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))),s('type.benchmark.times'(sK0),X1036))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s('type.benchmark.times'(sK0),X1033),s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),X1034),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0))),X1035)))))) )),
  inference(superposition,[],[f20395,f32462])).
fof(f20395,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))) != s('type.benchmark.digit'(sK0),'const.benchmark.Two_2'(s('type.benchmark.times'(sK0),X0),s('type.benchmark.times'(sK0),X1))) | s('type.benchmark.times'(sK0),sK2) = s('type.benchmark.times'(sK0),happ(s(fun('type.benchmark.times'(sK0),'type.benchmark.times'(sK0)),'_i_0_0'),s('type.benchmark.times'(sK0),X0)))) )),
  inference(superposition,[],[f19960,f267])).
fof(f48722,plain,(
  ( ! [X389,X390] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389),s('type.benchmark.times'(sK0),sK2))),s(bool,X390))),s('type.benchmark.times'(sK0),sK1))))) ) | $spl86),
  inference(cnf_transformation,[],[f48722_D])).
fof(f48722_D,plain,(
  ( ! [X389,X390] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389),s('type.benchmark.times'(sK0),sK2))),s(bool,X390))),s('type.benchmark.times'(sK0),sK1))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f48707,plain,(
  ( ! [X368] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X368),s('type.benchmark.times'(sK0),sK3))))) ) | $spl76),
  inference(cnf_transformation,[],[f48707_D])).
fof(f48707_D,plain,(
  ( ! [X368] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X368),s('type.benchmark.times'(sK0),sK3))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f48770,plain,(
  ( ! [X580,X582,X581] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X580),s('type.benchmark.times'(sK0),sK3)))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X581),s('type.benchmark.times'(sK0),sK2))),s(bool,X582))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48769,f318])).
fof(f48769,plain,(
  ( ! [X580,X582,X581,X579] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X581),s('type.benchmark.times'(sK0),sK2))),s(bool,X582))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X579),s(fun('type.benchmark.times'(sK0),bool),X580),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X581))))) )),
  inference(forward_demodulation,[],[f47986,f30352])).
fof(f30352,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(X8,happ(s(fun(X12,X8),'_dst_x0x1x0_3'(s(fun(X10,fun(X9,X8)),X11),s(fun(X12,X10),X13),s(X9,X15))),s(X12,X14))) = s(X8,happ(s(fun(X10,X8),'_dst_x1x0_2'(s(fun(X10,fun(X9,X8)),X11),s(X9,X15))),s(X10,happ(s(fun(X12,X10),X13),s(X12,X14)))))) )),
  inference(superposition,[],[f327,f324])).
fof(f47986,plain,(
  ( ! [X580,X582,X581,X579] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X582))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X581),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X579),s(fun('type.benchmark.times'(sK0),bool),X580),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X581))))) )),
  inference(superposition,[],[f26083,f32462])).
fof(f26083,plain,(
  ( ! [X220,X219] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X220))),s(bool,X219)))) | p(s(bool,X219))) )),
  inference(superposition,[],[f307,f324])).
fof(f307,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f48695,plain,(
  ( ! [X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X351),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) | $spl64),
  inference(cnf_transformation,[],[f48695_D])).
fof(f48695_D,plain,(
  ( ! [X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X351),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f48763,plain,(
  ( ! [X557] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X557),s('type.benchmark.times'(sK0),sK3)))))) ) | $spl112),
  inference(cnf_transformation,[],[f48763_D])).
fof(f48763_D,plain,(
  ( ! [X557] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X557),s('type.benchmark.times'(sK0),sK3)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f48762,plain,(
  ( ! [X557,X558] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X557),s('type.benchmark.times'(sK0),sK3))))) | s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X558),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) | ($spl18 | $spl26)),
  inference(forward_demodulation,[],[f47979,f318])).
fof(f47979,plain,(
  ( ! [X557,X556,X558] : (s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X558),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X556),s(fun('type.benchmark.times'(sK0),bool),X557),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X558)))))) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f12804,f32462])).
fof(f48759,plain,(
  ( ! [X552] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X552),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) | $spl108),
  inference(cnf_transformation,[],[f48759_D])).
fof(f48759_D,plain,(
  ( ! [X552] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X552),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f48760,plain,(
  ( ! [X551] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X551),s('type.benchmark.times'(sK0),sK3)))))) ) | $spl110),
  inference(cnf_transformation,[],[f48760_D])).
fof(f48760_D,plain,(
  ( ! [X551] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X551),s('type.benchmark.times'(sK0),sK3)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f48758,plain,(
  ( ! [X551,X552] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X551),s('type.benchmark.times'(sK0),sK3))))) | s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X552),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) ) | $spl16),
  inference(forward_demodulation,[],[f47977,f318])).
fof(f47977,plain,(
  ( ! [X551,X552,X550] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X552),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X550),s(fun('type.benchmark.times'(sK0),bool),X551),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X552)))))) ) | $spl16),
  inference(superposition,[],[f10428,f32462])).
fof(f10428,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl16),
  inference(forward_demodulation,[],[f10385,f7949])).
fof(f7949,plain,(
  s(bool,'T_0') = s(bool,sK7('_i_0_0',bool)) | $spl16),
  inference(unit_resulting_resolution,[],[f344,f7926,f1212])).
fof(f7926,plain,(
  p(s(bool,sK7('_i_0_0',bool))) | $spl16),
  inference(forward_demodulation,[],[f7925,f267])).
fof(f7925,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('_i_0_0',bool))))) | $spl16),
  inference(subsumption_resolution,[],[f7903,f231])).
fof(f7903,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK7('_i_0_0',bool))))) | $spl16),
  inference(superposition,[],[f301,f4200])).
fof(f4200,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))) | $spl16),
  inference(unit_resulting_resolution,[],[f344,f3864,f1212])).
fof(f301,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0)))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X43,X44] : (p(s(bool,happ(s(fun(fun(X43,bool),bool),'?_1'),s(fun(X43,bool),X44)))) <=> ? [X45] : p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X45)))))),
  file('thm.benchmark.digit_distinct|split|2','?_HL_QUANT')).
fof(f10385,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X33)))) ) | $spl16),
  inference(superposition,[],[f8069,f267])).
fof(f8069,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f8011,f231])).
fof(f8011,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,sK7('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X10)) ) | $spl16),
  inference(superposition,[],[f362,f7949])).
fof(f48756,plain,(
  ( ! [X546] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s(bool,'F_0'))),s('type.benchmark.times'(sK0),sK1))))) ) | $spl106),
  inference(cnf_transformation,[],[f48756_D])).
fof(f48756_D,plain,(
  ( ! [X546] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s(bool,'F_0'))),s('type.benchmark.times'(sK0),sK1))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f48755,plain,(
  ( ! [X546,X545] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X545),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s(bool,'F_0'))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48754,f318])).
fof(f48754,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s(bool,'F_0'))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X544),s(fun('type.benchmark.times'(sK0),bool),X545),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546))))) )),
  inference(forward_demodulation,[],[f48753,f30352])).
fof(f48753,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x1x0_2'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X544),s(fun('type.benchmark.times'(sK0),bool),X545),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546))))) )),
  inference(forward_demodulation,[],[f47975,f324])).
fof(f47975,plain,(
  ( ! [X546,X545,X544] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))),s(bool,'F_0')))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X544),s(fun('type.benchmark.times'(sK0),bool),X545),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X546))))) )),
  inference(superposition,[],[f8244,f32462])).
fof(f48752,plain,(
  ( ! [X543] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1))))) ) | $spl104),
  inference(cnf_transformation,[],[f48752_D])).
fof(f48752_D,plain,(
  ( ! [X543] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f48751,plain,(
  ( ! [X542,X543] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X542),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48750,f318])).
fof(f48750,plain,(
  ( ! [X542,X541,X543] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X541),s(fun('type.benchmark.times'(sK0),bool),X542),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543))))) )),
  inference(forward_demodulation,[],[f47974,f326])).
fof(f47974,plain,(
  ( ! [X542,X541,X543] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X541),s(fun('type.benchmark.times'(sK0),bool),X542),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X543))))) )),
  inference(superposition,[],[f8178,f32462])).
fof(f48692,plain,(
  ( ! [X348] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X348),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) ) | $spl60),
  inference(cnf_transformation,[],[f48692_D])).
fof(f48692_D,plain,(
  ( ! [X348] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X348),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f48747,plain,(
  ( ! [X532,X534] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X532),s('type.benchmark.times'(sK0),sK3))))))))) ) | $spl102),
  inference(cnf_transformation,[],[f48747_D])).
fof(f48747_D,plain,(
  ( ! [X532,X534] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X532),s('type.benchmark.times'(sK0),sK3))))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f48746,plain,(
  ( ! [X532,X534,X533] : (p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X532),s('type.benchmark.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X533),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47971,f318])).
fof(f47971,plain,(
  ( ! [X532,X534,X533,X531] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X533),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X534,bool),bool),'!_1'),s(fun(X534,bool),'_dst_x0_1'(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X531),s(fun('type.benchmark.times'(sK0),bool),X532),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X533))))))))) )),
  inference(superposition,[],[f7696,f32462])).
fof(f7696,plain,(
  ( ! [X37,X38] : (~p(s(bool,X38)) | p(s(bool,happ(s(fun(fun(X37,bool),bool),'!_1'),s(fun(X37,bool),'_dst_x0_1'(s(bool,X38))))))) )),
  inference(superposition,[],[f300,f312])).
fof(f300,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X1,X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) <=> ! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X43,X44] : (p(s(bool,happ(s(fun(fun(X43,bool),bool),'!_1'),s(fun(X43,bool),X44)))) <=> ! [X45] : p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X45)))))),
  file('thm.benchmark.digit_distinct|split|2','!_HL_QUANT')).
fof(f48743,plain,(
  ( ! [X516,X518] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X516),s('type.benchmark.times'(sK0),sK3))))))))) ) | $spl100),
  inference(cnf_transformation,[],[f48743_D])).
fof(f48743_D,plain,(
  ( ! [X516,X518] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X516),s('type.benchmark.times'(sK0),sK3))))))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f48742,plain,(
  ( ! [X516,X518,X517] : (p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X516),s('type.benchmark.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X517),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47966,f318])).
fof(f47966,plain,(
  ( ! [X516,X518,X517,X515] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X517),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(fun(X518,bool),bool),'?_1'),s(fun(X518,bool),'_dst_x0_1'(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X515),s(fun('type.benchmark.times'(sK0),bool),X516),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X517))))))))) )),
  inference(superposition,[],[f3845,f32462])).
fof(f3845,plain,(
  ( ! [X50,X49] : (~p(s(bool,X50)) | p(s(bool,happ(s(fun(fun(X49,bool),bool),'?_1'),s(fun(X49,bool),'_dst_x0_1'(s(bool,X50))))))) )),
  inference(superposition,[],[f302,f312])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f48737,plain,(
  ( ! [X466,X465] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X465),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))))))) ) | $spl98),
  inference(cnf_transformation,[],[f48737_D])).
fof(f48737_D,plain,(
  ( ! [X466,X465] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X465),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f48698,plain,(
  ( ! [X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X353),s('type.benchmark.times'(sK0),sK3)))) ) | $spl68),
  inference(cnf_transformation,[],[f48698_D])).
fof(f48698_D,plain,(
  ( ! [X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X353),s('type.benchmark.times'(sK0),sK3)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f48736,plain,(
  ( ! [X466,X465,X464] : (s(bool,'F_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X464),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X465),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))))))) )),
  inference(forward_demodulation,[],[f47953,f318])).
fof(f47953,plain,(
  ( ! [X466,X463,X465,X464] : (p(s(bool,'|dollar|exists_1'(s(fun(X466,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X465),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))))) | s(bool,'F_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X463),s(fun('type.benchmark.times'(sK0),bool),X464),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X465)))) )),
  inference(superposition,[],[f1427,f32462])).
fof(f1427,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,X22)))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(forward_demodulation,[],[f1426,f267])).
fof(f1426,plain,(
  ( ! [X23,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1416,f231])).
fof(f1416,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X23,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22)))))))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1262,f392])).
fof(f1262,plain,(
  ( ! [X21,X22] : (~p(s(bool,X22)) | p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),'_dst_x0_1'(s(bool,X22))))))) )),
  inference(superposition,[],[f280,f312])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK4(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X23,X24] : (! [X25] : (s(X6,X24) = s(X6,X25) => p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X25))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X23)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.bool.JRH_INDUCT_UTIL')).
fof(f48735,plain,(
  ( ! [X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X460),s('type.benchmark.times'(sK0),sK3))))))))) ) | $spl96),
  inference(cnf_transformation,[],[f48735_D])).
fof(f48735_D,plain,(
  ( ! [X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X460),s('type.benchmark.times'(sK0),sK3))))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f48734,plain,(
  ( ! [X461,X460,X462] : (p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X460),s('type.benchmark.times'(sK0),sK3)))))))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X461),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47952,f318])).
fof(f47952,plain,(
  ( ! [X461,X459,X460,X462] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X461),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X462,bool),'_dst_x0_1'(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X459),s(fun('type.benchmark.times'(sK0),bool),X460),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X461))))))))) )),
  inference(superposition,[],[f1262,f32462])).
fof(f48732,plain,(
  ( ! [X453,X451] : (s(bool,X453) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X451),s('type.benchmark.times'(sK0),sK3))) | s(bool,'F_0') = s(bool,X453)) ) | $spl94),
  inference(cnf_transformation,[],[f48732_D])).
fof(f48732_D,plain,(
  ( ! [X453,X451] : (s(bool,X453) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X451),s('type.benchmark.times'(sK0),sK3))) | s(bool,'F_0') = s(bool,X453)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f48731,plain,(
  ( ! [X453,X451,X452] : (s(bool,X453) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X451),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X452),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'F_0') = s(bool,X453)) )),
  inference(forward_demodulation,[],[f47950,f318])).
fof(f47950,plain,(
  ( ! [X453,X451,X452,X450] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X452),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,X453) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X450),s(fun('type.benchmark.times'(sK0),bool),X451),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X452))) | s(bool,'F_0') = s(bool,X453)) )),
  inference(superposition,[],[f1212,f32462])).
fof(f48689,plain,(
  ( ! [X345] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X345),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) ) | $spl56),
  inference(cnf_transformation,[],[f48689_D])).
fof(f48689_D,plain,(
  ( ! [X345] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X345),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f48729,plain,(
  ( ! [X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X423),s('type.benchmark.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X425)) ) | $spl92),
  inference(cnf_transformation,[],[f48729_D])).
fof(f48729_D,plain,(
  ( ! [X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X423),s('type.benchmark.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X425)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f48728,plain,(
  ( ! [X424,X423,X425] : (s(bool,X425) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X423),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X424),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,X425)) )),
  inference(forward_demodulation,[],[f47944,f318])).
fof(f47944,plain,(
  ( ! [X422,X424,X423,X425] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X424),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,X425) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X422),s(fun('type.benchmark.times'(sK0),bool),X423),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X424))) | s(bool,'T_0') = s(bool,X425)) )),
  inference(superposition,[],[f418,f32462])).
fof(f48727,plain,(
  ( ! [X416,X418] : (s(bool,X418) != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X416),s('type.benchmark.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X418)) ) | $spl90),
  inference(cnf_transformation,[],[f48727_D])).
fof(f48727_D,plain,(
  ( ! [X416,X418] : (s(bool,X418) != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X416),s('type.benchmark.times'(sK0),sK3))) | s(bool,'T_0') = s(bool,X418)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f48726,plain,(
  ( ! [X416,X418,X417] : (s(bool,X418) != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X416),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X417),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,X418)) )),
  inference(forward_demodulation,[],[f47942,f318])).
fof(f47942,plain,(
  ( ! [X416,X418,X415,X417] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X417),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,X418) != s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X415),s(fun('type.benchmark.times'(sK0),bool),X416),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X417))) | s(bool,'T_0') = s(bool,X418)) )),
  inference(superposition,[],[f362,f32462])).
fof(f48725,plain,(
  ( ! [X396,X398] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.benchmark.times'(sK0),bool),X396))),s('type.benchmark.times'(sK0),sK3))))) ) | $spl88),
  inference(cnf_transformation,[],[f48725_D])).
fof(f48725_D,plain,(
  ( ! [X396,X398] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.benchmark.times'(sK0),bool),X396))),s('type.benchmark.times'(sK0),sK3))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f48724,plain,(
  ( ! [X397,X396,X398] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398),s(fun('type.benchmark.times'(sK0),bool),X396))),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X397),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48723,f326])).
fof(f48723,plain,(
  ( ! [X397,X396,X398] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X396),s('type.benchmark.times'(sK0),sK3)))))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X397),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47937,f318])).
fof(f47937,plain,(
  ( ! [X397,X395,X396,X398] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X397),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X398))),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X395),s(fun('type.benchmark.times'(sK0),bool),X396),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X397))))))) )),
  inference(superposition,[],[f308,f32462])).
fof(f308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f48721,plain,(
  ( ! [X389,X388,X390] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X388),s('type.benchmark.times'(sK0),sK3)))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389),s('type.benchmark.times'(sK0),sK2))),s(bool,X390))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48720,f318])).
fof(f48720,plain,(
  ( ! [X389,X387,X388,X390] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389),s('type.benchmark.times'(sK0),sK2))),s(bool,X390))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X387),s(fun('type.benchmark.times'(sK0),bool),X388),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389))))) )),
  inference(forward_demodulation,[],[f47935,f327])).
fof(f47935,plain,(
  ( ! [X389,X387,X388,X390] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))),s(bool,X390)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X387),s(fun('type.benchmark.times'(sK0),bool),X388),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X389))))) )),
  inference(superposition,[],[f307,f32462])).
fof(f48719,plain,(
  ( ! [X384,X386] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),X384),s(bool,X386))),s('type.benchmark.times'(sK0),sK3)))) | p(s(bool,X386))) ) | $spl84),
  inference(cnf_transformation,[],[f48719_D])).
fof(f48719_D,plain,(
  ( ! [X384,X386] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),X384),s(bool,X386))),s('type.benchmark.times'(sK0),sK3)))) | p(s(bool,X386))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f48718,plain,(
  ( ! [X385,X384,X386] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.benchmark.times'(sK0),bool),X384),s(bool,X386))),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X385),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X386))) )),
  inference(forward_demodulation,[],[f48717,f327])).
fof(f48717,plain,(
  ( ! [X385,X384,X386] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X384),s('type.benchmark.times'(sK0),sK3))))),s(bool,X386)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X385),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X386))) )),
  inference(forward_demodulation,[],[f47934,f318])).
fof(f47934,plain,(
  ( ! [X383,X385,X384,X386] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X385),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X386)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X383),s(fun('type.benchmark.times'(sK0),bool),X384),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X385))))),s(bool,X386))))) )),
  inference(superposition,[],[f306,f32462])).
fof(f48716,plain,(
  ( ! [X380,X382] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.benchmark.times'(sK0),bool),X380))),s('type.benchmark.times'(sK0),sK3))))) ) | $spl82),
  inference(cnf_transformation,[],[f48716_D])).
fof(f48716_D,plain,(
  ( ! [X380,X382] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.benchmark.times'(sK0),bool),X380))),s('type.benchmark.times'(sK0),sK3))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f48715,plain,(
  ( ! [X380,X382,X381] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382),s(fun('type.benchmark.times'(sK0),bool),X380))),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X381),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48714,f326])).
fof(f48714,plain,(
  ( ! [X380,X382,X381] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X380),s('type.benchmark.times'(sK0),sK3)))))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X381),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47933,f318])).
fof(f47933,plain,(
  ( ! [X379,X380,X382,X381] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X381),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X382))),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X379),s(fun('type.benchmark.times'(sK0),bool),X380),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X381))))))) )),
  inference(superposition,[],[f305,f32462])).
fof(f305,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f48713,plain,(
  ( ! [X376,X378] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),X376),s(bool,X378))),s('type.benchmark.times'(sK0),sK3))))) ) | $spl80),
  inference(cnf_transformation,[],[f48713_D])).
fof(f48713_D,plain,(
  ( ! [X376,X378] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),X376),s(bool,X378))),s('type.benchmark.times'(sK0),sK3))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f48712,plain,(
  ( ! [X377,X376,X378] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),X376),s(bool,X378))),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X377),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f48711,f327])).
fof(f48711,plain,(
  ( ! [X377,X376,X378] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X376),s('type.benchmark.times'(sK0),sK3))))),s(bool,X378)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X377),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47932,f318])).
fof(f47932,plain,(
  ( ! [X377,X376,X378,X375] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X377),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X375),s(fun('type.benchmark.times'(sK0),bool),X376),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X377))))),s(bool,X378))))) )),
  inference(superposition,[],[f304,f32462])).
fof(f304,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f48710,plain,(
  ( ! [X374,X373] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X374))) ) | $spl78),
  inference(cnf_transformation,[],[f48710_D])).
fof(f48710_D,plain,(
  ( ! [X374,X373] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X374))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f48709,plain,(
  ( ! [X372,X374,X373] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X372),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X374))) )),
  inference(forward_demodulation,[],[f48708,f318])).
fof(f48708,plain,(
  ( ! [X372,X374,X373,X371] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373),s('type.benchmark.times'(sK0),sK2))))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X374)) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X371),s(fun('type.benchmark.times'(sK0),bool),X372),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373))))) )),
  inference(forward_demodulation,[],[f47931,f326])).
fof(f47931,plain,(
  ( ! [X372,X374,X373,X371] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X374))),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))))) | p(s(bool,X374)) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X371),s(fun('type.benchmark.times'(sK0),bool),X372),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X373))))) )),
  inference(superposition,[],[f303,f32462])).
fof(f48706,plain,(
  ( ! [X370,X369] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369),s('type.benchmark.times'(sK0),sK2))),s(bool,X370))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X370))) ) | $spl74),
  inference(cnf_transformation,[],[f48706_D])).
fof(f48706_D,plain,(
  ( ! [X370,X369] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369),s('type.benchmark.times'(sK0),sK2))),s(bool,X370))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X370))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f48705,plain,(
  ( ! [X368,X370,X369] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X368),s('type.benchmark.times'(sK0),sK3)))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369),s('type.benchmark.times'(sK0),sK2))),s(bool,X370))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X370))) )),
  inference(forward_demodulation,[],[f48704,f318])).
fof(f48704,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369),s('type.benchmark.times'(sK0),sK2))),s(bool,X370))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X367),s(fun('type.benchmark.times'(sK0),bool),X368),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369)))) | p(s(bool,X370))) )),
  inference(forward_demodulation,[],[f47930,f327])).
fof(f47930,plain,(
  ( ! [X368,X370,X367,X369] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))),s(bool,X370)))) | p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X367),s(fun('type.benchmark.times'(sK0),bool),X368),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X369)))) | p(s(bool,X370))) )),
  inference(superposition,[],[f303,f32462])).
fof(f48703,plain,(
  ( ! [X364,X366] : (s(bool,X366) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X364),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,X366))) ) | $spl72),
  inference(cnf_transformation,[],[f48703_D])).
fof(f48703_D,plain,(
  ( ! [X364,X366] : (s(bool,X366) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X364),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,X366))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f48702,plain,(
  ( ! [X364,X366,X365] : (s(bool,X366) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X364),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X365),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | ~p(s(bool,X366))) )),
  inference(forward_demodulation,[],[f47929,f318])).
fof(f47929,plain,(
  ( ! [X364,X366,X365,X363] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X365),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | ~p(s(bool,X366)) | s(bool,X366) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X363),s(fun('type.benchmark.times'(sK0),bool),X364),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X365)))) )),
  inference(superposition,[],[f298,f32462])).
fof(f48700,plain,(
  ( ! [X356,X358] : (s(bool,X358) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X356),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,X358))) ) | $spl70),
  inference(cnf_transformation,[],[f48700_D])).
fof(f48700_D,plain,(
  ( ! [X356,X358] : (s(bool,X358) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X356),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,X358))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f48701,plain,(
  ( ! [X360,X362,X361] : (s(bool,X362) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X360),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X361),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X362))) )),
  inference(forward_demodulation,[],[f47928,f318])).
fof(f47928,plain,(
  ( ! [X360,X362,X359,X361] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X361),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X362)) | s(bool,X362) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X359),s(fun('type.benchmark.times'(sK0),bool),X360),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X361)))) )),
  inference(superposition,[],[f295,f32462])).
fof(f48699,plain,(
  ( ! [X356,X358,X357] : (s(bool,X358) = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X356),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X357),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X358))) )),
  inference(forward_demodulation,[],[f47927,f318])).
fof(f47927,plain,(
  ( ! [X356,X358,X357,X355] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X357),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | p(s(bool,X358)) | s(bool,X358) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X355),s(fun('type.benchmark.times'(sK0),bool),X356),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X357)))) )),
  inference(superposition,[],[f295,f32462])).
fof(f48697,plain,(
  ( ! [X354,X353] : (s(bool,'F_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X353),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X354),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47926,f318])).
fof(f47926,plain,(
  ( ! [X352,X354,X353] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X354),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'F_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X352),s(fun('type.benchmark.times'(sK0),bool),X353),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X354)))) )),
  inference(superposition,[],[f246,f32462])).
fof(f246,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f194])).
fof(f48696,plain,(
  ( ! [X350] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X350),s('type.benchmark.times'(sK0),sK3))))) ) | $spl66),
  inference(cnf_transformation,[],[f48696_D])).
fof(f48696_D,plain,(
  ( ! [X350] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X350),s('type.benchmark.times'(sK0),sK3))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f48694,plain,(
  ( ! [X350,X351] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X350),s('type.benchmark.times'(sK0),sK3)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X351),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f47925,f318])).
fof(f47925,plain,(
  ( ! [X350,X349,X351] : (s(bool,'F_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X351),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))) | ~p(s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X349),s(fun('type.benchmark.times'(sK0),bool),X350),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X351))))) )),
  inference(superposition,[],[f245,f32462])).
fof(f48693,plain,(
  ( ! [X347] : (s(bool,'T_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X347),s('type.benchmark.times'(sK0),sK3)))) ) | $spl62),
  inference(cnf_transformation,[],[f48693_D])).
fof(f48693_D,plain,(
  ( ! [X347] : (s(bool,'T_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X347),s('type.benchmark.times'(sK0),sK3)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f48691,plain,(
  ( ! [X347,X348] : (s(bool,'T_0') = s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X347),s('type.benchmark.times'(sK0),sK3))) | ~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X348),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47924,f318])).
fof(f47924,plain,(
  ( ! [X346,X347,X348] : (~p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X348),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X346),s(fun('type.benchmark.times'(sK0),bool),X347),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X348)))) )),
  inference(superposition,[],[f242,f32462])).
fof(f48690,plain,(
  ( ! [X344] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X344),s('type.benchmark.times'(sK0),sK3)))) ) | $spl58),
  inference(cnf_transformation,[],[f48690_D])).
fof(f48690_D,plain,(
  ( ! [X344] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X344),s('type.benchmark.times'(sK0),sK3)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f48688,plain,(
  ( ! [X344,X345] : (s(bool,'T_0') != s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),X344),s('type.benchmark.times'(sK0),sK3))) | p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X345),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f47923,f318])).
fof(f47923,plain,(
  ( ! [X344,X343,X345] : (p(s(bool,happ(s(fun('type.benchmark.times'(sK0),bool),happ(s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X345),s('type.benchmark.times'(sK0),sK2))),s('type.benchmark.times'(sK0),sK1)))) | s(bool,'T_0') != s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(sK0),'const.benchmark.One_1'(s('type.benchmark.times'(sK0),sK3))),s(bool,X343),s(fun('type.benchmark.times'(sK0),bool),X344),s(fun('type.benchmark.times'(sK0),fun('type.benchmark.times'(sK0),bool)),X345)))) )),
  inference(superposition,[],[f241,f32462])).
fof(f1809,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl8),
  inference(cnf_transformation,[],[f1809_D])).
fof(f1809_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f45675,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) ) | $spl54),
  inference(cnf_transformation,[],[f45675_D])).
fof(f45675_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f45676,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,'F_0')))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f45549,f241])).
fof(f45549,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,'F_0')))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f8244,f359])).
fof(f45674,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f45548,f241])).
fof(f45548,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,'F_0')))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f8244,f359])).
fof(f8600,plain,(
  ( ! [X60] : (s(bool,'T_0') = s(bool,X60)) ) | $spl24),
  inference(cnf_transformation,[],[f8600_D])).
fof(f8600_D,plain,(
  ( ! [X60] : (s(bool,'T_0') = s(bool,X60)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f40435,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X65)) ) | $spl50),
  inference(cnf_transformation,[],[f40435_D])).
fof(f40435_D,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X65)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f45347,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76)) ) | $spl44),
  inference(duplicate_literal_removal,[],[f45047])).
fof(f45047,plain,(
  ( ! [X76,X75] : (s(bool,X76) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) ) | $spl44),
  inference(superposition,[],[f1700,f37846])).
fof(f37846,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(trivial_inequality_removal,[],[f37836])).
fof(f37836,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(superposition,[],[f37753,f359])).
fof(f37753,plain,(
  s(bool,'T_0') != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(subsumption_resolution,[],[f37645,f337])).
fof(f37645,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(superposition,[],[f241,f37121])).
fof(f44902,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44563,f44754])).
fof(f44754,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1809,f320])).
fof(f44563,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f1809,f245])).
fof(f44901,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44567,f231])).
fof(f44567,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f344,f1809,f298])).
fof(f44900,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f44569,f231])).
fof(f44569,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f1809,f298])).
fof(f44899,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44571,f231])).
fof(f44571,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f31171,f1809,f298])).
fof(f44898,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44573,f231])).
fof(f44573,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28362,f1809,f298])).
fof(f44897,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f44575,f231])).
fof(f44575,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f21339,f1809,f298])).
fof(f44896,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f44577,f231])).
fof(f44577,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1809,f298])).
fof(f44895,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44579,f231])).
fof(f44579,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37753,f1809,f298])).
fof(f44894,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f44583,f1809])).
fof(f44583,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f1809,f298])).
fof(f44893,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44585,f1809])).
fof(f44585,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f31171,f1809,f298])).
fof(f44892,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44587,f1809])).
fof(f44587,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28362,f1809,f298])).
fof(f44891,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f44589,f1809])).
fof(f44589,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f21339,f1809,f298])).
fof(f44890,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f44591,f1809])).
fof(f44591,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1809,f298])).
fof(f44889,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44593,f1809])).
fof(f44593,plain,(
  ~p(s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37753,f1809,f298])).
fof(f44888,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44597,f231])).
fof(f44597,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(unit_resulting_resolution,[],[f344,f1809,f298])).
fof(f44887,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f44599,f231])).
fof(f44599,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f1809,f298])).
fof(f44886,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44601,f231])).
fof(f44601,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f31171,f1809,f298])).
fof(f44885,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44603,f231])).
fof(f44603,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28362,f1809,f298])).
fof(f44884,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f44605,f231])).
fof(f44605,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f21339,f1809,f298])).
fof(f44883,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f44607,f231])).
fof(f44607,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1809,f298])).
fof(f44882,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44609,f231])).
fof(f44609,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37753,f1809,f298])).
fof(f44881,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f44613,f1809])).
fof(f44613,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f1809,f298])).
fof(f44880,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44615,f1809])).
fof(f44615,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f31171,f1809,f298])).
fof(f44879,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44617,f1809])).
fof(f44617,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28362,f1809,f298])).
fof(f44878,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f44619,f1809])).
fof(f44619,plain,(
  ( ! [X0] : (~p(s(bool,sK4(X0,'_i_0_0',bool)))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f21339,f1809,f298])).
fof(f44877,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f44621,f1809])).
fof(f44621,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1809,f298])).
fof(f44876,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44623,f1809])).
fof(f44623,plain,(
  ~p(s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37753,f1809,f298])).
fof(f44875,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44627,f1809])).
fof(f44627,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'F_0'))))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f337,f1809,f306])).
fof(f44874,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44719,f44754])).
fof(f44719,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl8),
  inference(unit_resulting_resolution,[],[f344,f1809,f362])).
fof(f44873,plain,(
  $false | ($spl8 | $spl21)),
  inference(subsumption_resolution,[],[f44721,f44754])).
fof(f44721,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) ) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f1809,f362])).
fof(f44872,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44723,f44754])).
fof(f44723,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(fun(bool,bool),bool),'_dstsel_0_x0_1'(s(bool,'F_0'))),s(fun(bool,bool),'_i_0_0')))) ) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f31171,f1809,f362])).
fof(f44871,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f44725,f44754])).
fof(f44725,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f28362,f1809,f362])).
fof(f44870,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f44727,f44754])).
fof(f44727,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,sK4(X1,'_i_0_0',bool))) ) | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f21339,f1809,f362])).
fof(f44869,plain,(
  $false | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f44729,f44754])).
fof(f44729,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK6('_i_0_0',bool))) ) | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f11921,f1809,f362])).
fof(f44868,plain,(
  $false | ($spl8 | $spl44)),
  inference(subsumption_resolution,[],[f44731,f44754])).
fof(f44731,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) ) | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37753,f1809,f362])).
fof(f44867,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f44764,f1809])).
fof(f44764,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(unit_resulting_resolution,[],[f337,f1809,f306])).
fof(f44767,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f337,f1809,f8178])).
fof(f44765,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f231,f337,f1809,f306])).
fof(f44761,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f337,f337,f1809,f303])).
fof(f44755,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f344,f1809,f320])).
fof(f44748,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f344,f344,f1809,f1212])).
fof(f44732,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f37121,f37753,f1809,f362])).
fof(f44730,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f7712,f11921,f1809,f362])).
fof(f44728,plain,(
  $false | ($spl4 | $spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f21370,f21339,f1809,f362])).
fof(f44726,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f267,f28362,f1809,f362])).
fof(f44724,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f29559,f31171,f1809,f362])).
fof(f44722,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12143,f12126,f1809,f362])).
fof(f44720,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12143,f344,f1809,f362])).
fof(f44629,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f337,f1809,f1809,f306])).
fof(f44624,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f1809,f37753,f1809,f298])).
fof(f44622,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f1809,f11921,f1809,f298])).
fof(f44620,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1809,f21339,f1809,f298])).
fof(f44618,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1809,f28362,f1809,f298])).
fof(f44616,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1809,f31171,f1809,f298])).
fof(f44614,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1809,f12126,f1809,f298])).
fof(f44612,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1809,f344,f1809,f298])).
fof(f44610,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f231,f37753,f1809,f298])).
fof(f44608,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f231,f11921,f1809,f298])).
fof(f44606,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f231,f21339,f1809,f298])).
fof(f44604,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f231,f28362,f1809,f298])).
fof(f44602,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f231,f31171,f1809,f298])).
fof(f44600,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f231,f12126,f1809,f298])).
fof(f44598,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f231,f344,f1809,f298])).
fof(f44594,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f1809,f37753,f1809,f298])).
fof(f44592,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f1809,f11921,f1809,f298])).
fof(f44590,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f1809,f21339,f1809,f298])).
fof(f44588,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1809,f28362,f1809,f298])).
fof(f44586,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f1809,f31171,f1809,f298])).
fof(f44584,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1809,f12126,f1809,f298])).
fof(f44582,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f1809,f344,f1809,f298])).
fof(f44580,plain,(
  $false | ($spl8 | $spl44)),
  inference(unit_resulting_resolution,[],[f231,f37753,f1809,f298])).
fof(f44578,plain,(
  $false | ($spl8 | $spl18)),
  inference(unit_resulting_resolution,[],[f231,f11921,f1809,f298])).
fof(f44576,plain,(
  $false | ($spl4 | $spl8)),
  inference(unit_resulting_resolution,[],[f231,f21339,f1809,f298])).
fof(f44574,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f231,f28362,f1809,f298])).
fof(f44572,plain,(
  $false | ($spl8 | $spl18 | $spl26)),
  inference(unit_resulting_resolution,[],[f231,f31171,f1809,f298])).
fof(f44570,plain,(
  $false | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f231,f12126,f1809,f298])).
fof(f44568,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f231,f344,f1809,f298])).
fof(f44564,plain,(
  $false | ($spl8 | $spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f12143,f1809,f245])).
fof(f44562,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f344,f1809,f242])).
fof(f44560,plain,(
  $false | $spl8),
  inference(unit_resulting_resolution,[],[f337,f1809])).
fof(f44422,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl52),
  inference(cnf_transformation,[],[f44422_D])).
fof(f44422_D,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f44423,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X6)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f44308,f241])).
fof(f44308,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X6)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f8178,f359])).
fof(f44421,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f44307,f241])).
fof(f44307,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,X4)))) | p(s(bool,X3)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f8178,f359])).
fof(f40337,plain,(
  ( ! [X64,X65] : (s(bool,X65) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X65)) ) | $spl38),
  inference(duplicate_literal_removal,[],[f40066])).
fof(f40066,plain,(
  ( ! [X64,X65] : (s(bool,X65) != s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X64) | s(bool,'F_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X64)) ) | $spl38),
  inference(superposition,[],[f1700,f37115])).
fof(f37139,plain,(
  $false | ($spl34 | $spl44)),
  inference(subsumption_resolution,[],[f37138,f344])).
fof(f37138,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl34 | $spl44)),
  inference(backward_demodulation,[],[f37112,f37121])).
fof(f37112,plain,(
  s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl34),
  inference(cnf_transformation,[],[f37112_D])).
fof(f37112_D,plain,(
  s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f37114,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | $spl36),
  inference(cnf_transformation,[],[f37114_D])).
fof(f37114_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f37127,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | ($spl4 | $spl18)),
  inference(forward_demodulation,[],[f37060,f21370])).
fof(f37060,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,sK4(X12,'_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) ) | $spl4),
  inference(superposition,[],[f34338,f21395])).
fof(f21395,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,sK4(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(subsumption_resolution,[],[f21355,f231])).
fof(f21355,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,sK4(X1,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(superposition,[],[f1273,f359])).
fof(f34338,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(X0,bool),sK7('?_1',fun(X0,bool))),s(X0,sK7(sK7('?_1',fun(X0,bool)),X0)))))) )),
  inference(unit_resulting_resolution,[],[f33830,f301])).
fof(f33830,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),sK7('?_1',fun(X0,bool))))))) )),
  inference(unit_resulting_resolution,[],[f33140,f301])).
fof(f33140,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(fun(X0,bool),bool),bool),'?_1'),s(fun(fun(X0,bool),bool),'?_1'))))) )),
  inference(unit_resulting_resolution,[],[f32025,f302])).
fof(f32025,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f231,f3845])).
fof(f37121,plain,(
  s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl44),
  inference(cnf_transformation,[],[f37121_D])).
fof(f37121_D,plain,(
  s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f37117,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | $spl40),
  inference(cnf_transformation,[],[f37117_D])).
fof(f37117_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f37126,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl2),
  inference(forward_demodulation,[],[f37059,f19894])).
fof(f19894,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl2),
  inference(unit_resulting_resolution,[],[f344,f1271,f1212])).
fof(f37059,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl2),
  inference(superposition,[],[f34338,f19938])).
fof(f19938,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl2),
  inference(subsumption_resolution,[],[f19903,f337])).
fof(f19903,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f1271,f1593])).
fof(f1593,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f1239,f1239])).
fof(f37125,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl18),
  inference(forward_demodulation,[],[f37058,f7712])).
fof(f37058,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,sK6('_i_0_0',bool))))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl18),
  inference(superposition,[],[f34338,f11920])).
fof(f37124,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl16),
  inference(forward_demodulation,[],[f37057,f7949])).
fof(f37057,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,sK7('_i_0_0',bool))))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | $spl16),
  inference(superposition,[],[f34338,f7982])).
fof(f7982,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,sK7('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f7956,f337])).
fof(f7956,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,sK7('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X6)) ) | $spl16),
  inference(superposition,[],[f7926,f1593])).
fof(f37123,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9)) ) | $spl48),
  inference(cnf_transformation,[],[f37123_D])).
fof(f37123_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f37056,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X10)))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f34338,f1593])).
fof(f37055,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X9)))) | s(bool,'F_0') = s(bool,X9) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f34338,f1593])).
fof(f37122,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) ) | $spl46),
  inference(cnf_transformation,[],[f37122_D])).
fof(f37122_D,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f37054,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X8) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f34338,f1593])).
fof(f37053,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,X7) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f34338,f1593])).
fof(f37119,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))),
  inference(forward_demodulation,[],[f37052,f266])).
fof(f37052,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'const.bool.DATATYPE_1'(s(X4,X5)))))) | s(bool,'F_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f34338,f1239])).
fof(f37115,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f37115_D])).
fof(f37115_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f37051,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X3) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f34338,f359])).
fof(f37118,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f37118_D])).
fof(f37118_D,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f37050,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f34338,f359])).
fof(f37049,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))) )),
  inference(superposition,[],[f34338,f359])).
fof(f37048,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f34338,f359])).
fof(f37047,plain,(
  p(s(bool,happ(s(fun(bool,bool),sK7('?_1',fun(bool,bool))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK7(sK7('?_1',fun(bool,bool)),bool))),
  inference(superposition,[],[f34338,f238])).
fof(f20868,plain,(
  ( ! [X99,X100] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X100)) ) | $spl32),
  inference(cnf_transformation,[],[f20868_D])).
fof(f20868_D,plain,(
  ( ! [X99,X100] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f21986,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f21805])).
fof(f21805,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X98)) ) | $spl4),
  inference(superposition,[],[f1700,f21395])).
fof(f20787,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f20610])).
fof(f20610,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,sK4(X99,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X98)) ) | $spl0),
  inference(superposition,[],[f1700,f20089])).
fof(f20089,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,sK4(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f19989,f344])).
fof(f19989,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,sK4(X2,'_i_0_0',bool)) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1269,f359])).
fof(f3862,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f3862_D])).
fof(f3862_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1271,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | $spl2),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3286,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3257,f231])).
fof(f3257,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f3209,f392])).
fof(f3209,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f3095,f241])).
fof(f3095,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f280,f693])).
fof(f693,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f692,f337])).
fof(f692,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f279,f238])).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f3374,plain,(
  ( ! [X33] : (~p(s(bool,X33))) ) | $spl12),
  inference(cnf_transformation,[],[f3374_D])).
fof(f3374_D,plain,(
  ( ! [X33] : (~p(s(bool,X33))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3269,plain,(
  ( ! [X44] : (~p(s(bool,X44)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3209,f267])).
fof(f4555,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4524,f231])).
fof(f4524,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f3209,f392])).
fof(f4538,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3209,f267])).
fof(f14125,plain,(
  ( ! [X49] : (~p(s(bool,X49)) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f3209,f267])).
fof(f15028,plain,(
  ( ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f14937,f231])).
fof(f14937,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f3209,f392])).
fof(f19813,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X67)) ) | $spl30),
  inference(cnf_transformation,[],[f19813_D])).
fof(f19813_D,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X67)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f19732,plain,(
  ( ! [X66,X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67)) ) | ($spl2 | $spl18 | $spl21)),
  inference(duplicate_literal_removal,[],[f19547])).
fof(f19547,plain,(
  ( ! [X66,X67] : (s(bool,X67) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) ) | ($spl2 | $spl18 | $spl21)),
  inference(superposition,[],[f1700,f19182])).
fof(f19182,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl18 | $spl21)),
  inference(trivial_inequality_removal,[],[f19166])).
fof(f19166,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl18 | $spl21)),
  inference(superposition,[],[f19042,f359])).
fof(f19042,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | ($spl2 | $spl18 | $spl21)),
  inference(forward_demodulation,[],[f18954,f12143])).
fof(f18954,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))) | ($spl2 | $spl21)),
  inference(unit_resulting_resolution,[],[f12126,f18750])).
fof(f18750,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(subsumption_resolution,[],[f18641,f337])).
fof(f18641,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl2),
  inference(superposition,[],[f241,f2040])).
fof(f2040,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'T_0') = s(bool,X24)) ) | $spl2),
  inference(forward_demodulation,[],[f1922,f1286])).
fof(f1286,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | $spl2),
  inference(unit_resulting_resolution,[],[f1271,f242])).
fof(f1922,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) ) | $spl2),
  inference(superposition,[],[f1347,f267])).
fof(f1347,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f1324,f231])).
fof(f1324,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X10)) ) | $spl2),
  inference(superposition,[],[f1212,f1286])).
fof(f11774,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763,f11765,f11318,f11274,f11268,f11265])).
fof(f11265,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11259,f245])).
fof(f11268,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11259,f11259,f298])).
fof(f11274,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f11259,f11259,f298])).
fof(f11318,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f362])).
fof(f11765,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl20),
  inference(subsumption_resolution,[],[f11680,f231])).
fof(f11680,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f7032])).
fof(f11763,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl20),
  inference(subsumption_resolution,[],[f11674,f231])).
fof(f11674,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f7004])).
fof(f11756,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl20),
  inference(subsumption_resolution,[],[f11625,f231])).
fof(f11625,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f6971])).
fof(f11754,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl20),
  inference(subsumption_resolution,[],[f11593,f231])).
fof(f11593,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f6928])).
fof(f11752,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl20),
  inference(subsumption_resolution,[],[f11576,f231])).
fof(f11576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f320])).
fof(f11747,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(condensation,[],[f11746])).
fof(f11746,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl20),
  inference(condensation,[],[f11745])).
fof(f11745,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(subsumption_resolution,[],[f11504,f11259])).
fof(f11504,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'T_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f6957])).
fof(f11743,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl20),
  inference(condensation,[],[f11742])).
fof(f11742,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl20),
  inference(condensation,[],[f11741])).
fof(f11741,plain,(
  ( ! [X61,X62,X63] : (s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(subsumption_resolution,[],[f11499,f11259])).
fof(f11499,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f6914])).
fof(f11322,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f1212])).
fof(f11259,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl20),
  inference(forward_demodulation,[],[f11258,f267])).
fof(f11258,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11237,f231])).
fof(f11237,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))))) ) | $spl20),
  inference(superposition,[],[f299,f9619])).
fof(f11222,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl20),
  inference(subsumption_resolution,[],[f11169,f231])).
fof(f11169,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl20),
  inference(superposition,[],[f362,f9619])).
fof(f11221,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7),s(fun(fun(bool,bool),bool),'!_1'))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11220,f326])).
fof(f11220,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11166,f231])).
fof(f11166,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X7))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(superposition,[],[f308,f9619])).
fof(f11218,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4),s(fun(fun(bool,bool),bool),'!_1'))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11217,f326])).
fof(f11217,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11163,f231])).
fof(f11163,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))) ) | $spl20),
  inference(superposition,[],[f305,f9619])).
fof(f11216,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(fun(bool,bool),bool),'!_1'),s(bool,X3))),s(fun(bool,bool),'_i_0_0'))))) ) | $spl20),
  inference(forward_demodulation,[],[f11215,f327])).
fof(f11215,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X3))))) ) | $spl20),
  inference(subsumption_resolution,[],[f11162,f231])).
fof(f11162,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X3))))) ) | $spl20),
  inference(superposition,[],[f304,f9619])).
fof(f9653,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl20),
  inference(subsumption_resolution,[],[f9634,f337])).
fof(f9634,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | s(bool,'F_0') = s(bool,X4)) ) | $spl20),
  inference(superposition,[],[f7714,f1593])).
fof(f9619,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f231,f7714,f298])).
fof(f9621,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f7714,f362])).
fof(f9622,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))) | $spl20),
  inference(unit_resulting_resolution,[],[f267,f7714,f362])).
fof(f9623,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),s(X0,X1)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f312,f7714,f362])).
fof(f9624,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(fun('type.num.num',bool),X0),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f276,f7714,f362])).
fof(f9625,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(X0,X1)))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f311,f7714,f362])).
fof(f9626,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f313,f7714,f362])).
fof(f9627,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))),s(fun('type.benchmark.times'(X0),bool),X1),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),bool)),X2)))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f317,f7714,f362])).
fof(f9628,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f7714,f1262])).
fof(f450,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f276,f238])).
fof(f492,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f311,f238])).
fof(f538,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f312,f238])).
fof(f588,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f313,f238])).
fof(f1207,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20)))) )),
  inference(forward_demodulation,[],[f1167,f267])).
fof(f1167,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X19))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20)))) )),
  inference(superposition,[],[f392,f312])).
fof(f1208,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1168,f267])).
fof(f1168,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X21))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f392,f276])).
fof(f1209,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X24),s(X23,X25)))))) )),
  inference(forward_demodulation,[],[f1169,f267])).
fof(f1169,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X24),s(X23,X25)))))) )),
  inference(superposition,[],[f392,f311])).
fof(f1210,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(forward_demodulation,[],[f1170,f267])).
fof(f1170,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X28))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(superposition,[],[f392,f313])).
fof(f1213,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f1185,f231])).
fof(f1185,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | s(bool,X18) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(superposition,[],[f362,f392])).
fof(f1238,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1225,f231])).
fof(f1225,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X23) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X22))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f1212,f392])).
fof(f1358,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),X16))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X15),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X15),X16)))) )),
  inference(superposition,[],[f309,f238])).
fof(f1383,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X80),X81))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X80),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X80),X81)))) )),
  inference(superposition,[],[f392,f309])).
fof(f1548,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f1520,f231])).
fof(f1520,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | s(bool,X28) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f1212,f1450])).
fof(f1564,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,'const.bool.DATATYPE_1'(s(X51,X52))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X50)))) )),
  inference(superposition,[],[f1239,f267])).
fof(f1565,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X54) = s(bool,'const.bool.DATATYPE_1'(s(X56,X57))) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),'_dst_x0_1'(s(bool,X54))),s(X53,X55)))) )),
  inference(superposition,[],[f1239,f312])).
fof(f1567,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X62) = s(bool,'const.bool.DATATYPE_1'(s(X64,X65))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X62),s(fun('type.num.num',bool),X63),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1239,f276])).
fof(f1568,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X67) = s(bool,'const.bool.DATATYPE_1'(s(X69,X70))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X67),s(X66,X68)))))) )),
  inference(superposition,[],[f1239,f311])).
fof(f1569,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,X73) = s(bool,'const.bool.DATATYPE_1'(s(X74,X75))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X71,bool),'const.pair.,_2'(s(X71,X72),s(bool,X73)))))) )),
  inference(superposition,[],[f1239,f313])).
fof(f1687,plain,(
  ( ! [X59,X58] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) )),
  inference(forward_demodulation,[],[f1566,f266])).
fof(f1566,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'const.bool.DATATYPE_1'(s(X60,X61))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),X59))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X58),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X58),X59)))) )),
  inference(superposition,[],[f1239,f309])).
fof(f1811,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18))))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1787,f337])).
fof(f1787,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X17,bool),'_dst_x0_1'(s(bool,X18)))))) )),
  inference(superposition,[],[f1427,f238])).
fof(f1893,plain,(
  ( ! [X0,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1892,f231])).
fof(f1892,plain,(
  ( ! [X0,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f1872,f266])).
fof(f1872,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X1,X2)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X3)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f304,f1239])).
fof(f1896,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X27))),s(bool,X26)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(forward_demodulation,[],[f1895,f327])).
fof(f1895,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26))))),s(bool,X27)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f1881,f231])).
fof(f1881,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26))))),s(bool,X27)))) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f304,f392])).
fof(f2120,plain,(
  ( ! [X59,X58] : (s(bool,X58) = s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)))) )),
  inference(superposition,[],[f1593,f267])).
fof(f2121,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63) | s(bool,'F_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62)))) )),
  inference(superposition,[],[f1593,f312])).
fof(f2122,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65)))) )),
  inference(superposition,[],[f1593,f309])).
fof(f2123,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) = s(bool,X69) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X67),s(fun('type.num.num',bool),X68),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1593,f276])).
fof(f2124,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,X71) = s(bool,X73) | s(bool,'F_0') = s(bool,X73) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72)))))) )),
  inference(superposition,[],[f1593,f311])).
fof(f2125,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76)))))) )),
  inference(superposition,[],[f1593,f313])).
fof(f2181,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49))),s(X48,X50))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f1593,f312])).
fof(f2182,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53))) | s(bool,X54) = s(bool,happ(s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X52),X53))) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f1593,f309])).
fof(f2183,plain,(
  ( ! [X57,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X57) = s(bool,'const.ind_type.FCONS_3'(s(bool,X55),s(fun('type.num.num',bool),X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X57)) )),
  inference(superposition,[],[f1593,f276])).
fof(f2184,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X59),s(X58,X60))))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f1593,f311])).
fof(f2185,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X62,bool),'const.pair.,_2'(s(X62,X63),s(bool,X64))))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f1593,f313])).
fof(f2339,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f245,f1593])).
fof(f2722,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2375,f337])).
fof(f2375,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f295,f1593])).
fof(f2733,plain,(
  ( ! [X76,X74,X77,X75] : (s(X74,X76) = s(X74,sK4(X76,X75,X74)) | s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(X74,bool),X75))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f2399,f337])).
fof(f2399,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,'F_0')) | s(X74,X76) = s(X74,sK4(X76,X75,X74)) | s(bool,X77) = s(bool,'|dollar|exists_1'(s(fun(X74,bool),X75))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f279,f1593])).
fof(f2734,plain,(
  ( ! [X80,X81,X82] : (s(bool,'F_0') = s(bool,X81) | s(bool,X82) = s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2401,f337])).
fof(f2401,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X81) | s(bool,X82) = s(bool,'|dollar|exists_1'(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))))) | s(bool,'F_0') = s(bool,X82)) )),
  inference(superposition,[],[f1427,f1593])).
fof(f2780,plain,(
  ( ! [X14,X17] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f2779,f231])).
fof(f2779,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f2755,f266])).
fof(f2755,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f305,f1239])).
fof(f2784,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47),s(fun(bool,bool),'_i_0_0'))),s(bool,X46)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f2783,f326])).
fof(f2783,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2766,f231])).
fof(f2766,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f305,f392])).
fof(f2837,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58)))) )),
  inference(superposition,[],[f1700,f267])).
fof(f2838,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) != s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62)))) )),
  inference(superposition,[],[f1700,f312])).
fof(f2839,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'T_0') = s(bool,X66) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65)))) )),
  inference(superposition,[],[f1700,f309])).
fof(f2840,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X67),s(fun('type.num.num',bool),X68),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1700,f276])).
fof(f2841,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72)))))) )),
  inference(superposition,[],[f1700,f311])).
fof(f2842,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76)))))) )),
  inference(superposition,[],[f1700,f313])).
fof(f2865,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X59) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X58))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f1700,f267])).
fof(f2866,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X61) != s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),'_dst_x0_1'(s(bool,X61))),s(X60,X62))) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f1700,f312])).
fof(f2867,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),X65))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X64),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X64),X65))) | s(bool,'F_0') = s(bool,X66)) )),
  inference(superposition,[],[f1700,f309])).
fof(f2868,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X67),s(fun('type.num.num',bool),X68),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f1700,f276])).
fof(f2869,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,X71) != s(bool,X73) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72))))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f1700,f311])).
fof(f2870,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76))))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f1700,f313])).
fof(f2936,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,X61) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f2340,f267])).
fof(f2937,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | s(bool,X65) = s(bool,happ(s(fun(X62,bool),'_dst_x0_1'(s(bool,X63))),s(X62,X64))) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f2340,f312])).
fof(f2938,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),X67))) | s(bool,X68) = s(bool,happ(s(fun('type.pair.prod'(bool,X66),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X66),X67))) | s(bool,'F_0') = s(bool,X68)) )),
  inference(superposition,[],[f2340,f309])).
fof(f2939,plain,(
  ( ! [X70,X71,X69] : (s(bool,'T_0') != s(bool,X69) | s(bool,X71) = s(bool,'const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f2340,f276])).
fof(f2940,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,X73) | s(bool,X75) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),'const.pair.,_2'(s(bool,X73),s(X72,X74))))) | s(bool,'F_0') = s(bool,X75)) )),
  inference(superposition,[],[f2340,f311])).
fof(f2941,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X79) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X76,bool),'const.pair.,_2'(s(X76,X77),s(bool,X78))))) | s(bool,'F_0') = s(bool,X79)) )),
  inference(superposition,[],[f2340,f313])).
fof(f3050,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X38))),s(bool,X39)))) | s(bool,X38) = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f2982,f337])).
fof(f2982,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X38))),s(bool,X39)))) | p(s(bool,'F_0')) | s(bool,X38) = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f307,f1593])).
fof(f3066,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X20))),s(bool,X21))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f3032,f337])).
fof(f3032,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'F_0')) | p(s(bool,X20)) | s(bool,X22) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X20))),s(bool,X21))) | s(bool,'F_0') = s(bool,X22)) )),
  inference(superposition,[],[f307,f1593])).
fof(f3067,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,sK4(X3,X4,bool))) )),
  inference(superposition,[],[f693,f238])).
fof(f3072,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X17) = s(bool,'const.bool.DATATYPE_1'(s(X19,X20))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X18))) | s(bool,'F_0') = s(bool,sK4(X17,X18,bool))) )),
  inference(superposition,[],[f693,f1239])).
fof(f3074,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X24))) | s(bool,X25) = s(bool,sK4(X23,X24,bool)) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f693,f1593])).
fof(f3076,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,sK4(X29,X30,bool))) )),
  inference(superposition,[],[f693,f1593])).
fof(f3110,plain,(
  ( ! [X57,X58] : (s(bool,'F_0') != s(bool,X57) | ~p(s(bool,sK4(X57,X58,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f245,f693])).
fof(f3111,plain,(
  ( ! [X59,X60] : (p(s(bool,X59)) | s(bool,'F_0') = s(bool,sK4(X59,X60,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f246,f693])).
fof(f3124,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X97))) | s(bool,'F_0') = s(bool,sK4(X97,X98,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X98)))) )),
  inference(superposition,[],[f392,f693])).
fof(f3127,plain,(
  ( ! [X107,X108,X106] : (~p(s(bool,X106)) | s(bool,X108) = s(bool,sK4(X106,X107,bool)) | s(bool,'F_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f1212,f693])).
fof(f3132,plain,(
  ( ! [X121,X122,X120] : (p(s(bool,'|dollar|exists_1'(s(fun(X122,bool),'_dst_x0_1'(s(bool,X120)))))) | s(bool,'F_0') = s(bool,sK4(X120,X121,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X121)))) )),
  inference(superposition,[],[f1427,f693])).
fof(f3137,plain,(
  ( ! [X136,X135,X137] : (s(bool,X135) != s(bool,X137) | s(bool,'T_0') = s(bool,sK4(X135,X136,bool)) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X136)))) )),
  inference(superposition,[],[f1700,f693])).
fof(f3138,plain,(
  ( ! [X140,X138,X139] : (s(bool,X138) != s(bool,X140) | s(bool,'T_0') = s(bool,X140) | s(bool,'F_0') = s(bool,sK4(X138,X139,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X139)))) )),
  inference(superposition,[],[f1700,f693])).
fof(f3139,plain,(
  ( ! [X142,X143,X141] : (s(bool,'T_0') != s(bool,X141) | s(bool,X143) = s(bool,sK4(X141,X142,bool)) | s(bool,'F_0') = s(bool,X143) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X142)))) )),
  inference(superposition,[],[f2340,f693])).
fof(f3290,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(subsumption_resolution,[],[f3289,f231])).
fof(f3289,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(forward_demodulation,[],[f3263,f266])).
fof(f3263,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X23,X24)))) | p(s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(superposition,[],[f3209,f1239])).
fof(f3327,plain,(
  ( ! [X14,X17] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f3326,f231])).
fof(f3326,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f3300,f266])).
fof(f3300,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X17))),s(bool,X14)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f308,f1239])).
fof(f3331,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47),s(fun(bool,bool),'_i_0_0'))),s(bool,X46)))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f3330,f326])).
fof(f3330,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f3311,f231])).
fof(f3311,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X47))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X46)))))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f308,f392])).
fof(f3454,plain,(
  ( ! [X23,X19,X20] : (p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(subsumption_resolution,[],[f3453,f231])).
fof(f3453,plain,(
  ( ! [X23,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(forward_demodulation,[],[f3440,f266])).
fof(f3440,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X21,X22)))) | p(s(bool,happ(s(fun(X19,bool),X20),s(X19,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X19,bool),bool),'!_1'),s(fun(X19,bool),X20)))) )),
  inference(superposition,[],[f299,f1239])).
fof(f3867,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(subsumption_resolution,[],[f3866,f231])).
fof(f3866,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(forward_demodulation,[],[f3837,f266])).
fof(f3837,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X23,X24)))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X21)))) | s(bool,'F_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,X22)))) )),
  inference(superposition,[],[f302,f1239])).
fof(f5481,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f5423,f337])).
fof(f5423,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X13))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f1871,f1593])).
fof(f5716,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f5648,f337])).
fof(f5648,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f2749,f1593])).
fof(f5859,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f5781,f337])).
fof(f5781,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f5578,f1593])).
fof(f6118,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f284,f238])).
fof(f6125,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X32),s('type.num.num',X33))) | s(bool,X34) = s(bool,'const.ind_type.FCONS_3'(s(bool,X31),s(fun('type.num.num',bool),X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f284,f1593])).
fof(f6127,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X42) = s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) )),
  inference(superposition,[],[f284,f1593])).
fof(f6162,plain,(
  ( ! [X138,X136,X137] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,happ(s(fun('type.num.num',bool),X137),s('type.num.num',X138))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X136),s(fun('type.num.num',bool),X137),s('type.num.num','const.num.SUC_1'(s('type.num.num',X138)))))) )),
  inference(superposition,[],[f392,f284])).
fof(f6175,plain,(
  ( ! [X189,X187,X190,X188] : (s(bool,X190) != s(bool,happ(s(fun('type.num.num',bool),X188),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X187),s(fun('type.num.num',bool),X188),s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))))) | s(bool,'F_0') = s(bool,X190)) )),
  inference(superposition,[],[f1700,f284])).
fof(f6176,plain,(
  ( ! [X191,X194,X192,X193] : (s(bool,X194) != s(bool,happ(s(fun('type.num.num',bool),X192),s('type.num.num',X193))) | s(bool,'T_0') = s(bool,X194) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X191),s(fun('type.num.num',bool),X192),s('type.num.num','const.num.SUC_1'(s('type.num.num',X193)))))) )),
  inference(superposition,[],[f1700,f284])).
fof(f6177,plain,(
  ( ! [X198,X196,X197,X195] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X196),s('type.num.num',X197))) | s(bool,X198) = s(bool,'const.ind_type.FCONS_3'(s(bool,X195),s(fun('type.num.num',bool),X196),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197))))) | s(bool,'F_0') = s(bool,X198)) )),
  inference(superposition,[],[f2340,f284])).
fof(f6206,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(forward_demodulation,[],[f6123,f266])).
fof(f6123,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'const.bool.DATATYPE_1'(s(X26,X27))) = s(bool,happ(s(fun('type.num.num',bool),X24),s('type.num.num',X25))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X23),s(fun('type.num.num',bool),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f284,f1239])).
fof(f6350,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f344,f2968,f1212])).
fof(f6435,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f6343,f326])).
fof(f6343,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f2968,f308])).
fof(f6436,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f6342,f326])).
fof(f6342,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f2968,f305])).
fof(f6437,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))),s(bool,X0))))) )),
  inference(forward_demodulation,[],[f6341,f327])).
fof(f6341,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f2968,f304])).
fof(f6540,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X5),'const.benchmark.Zero_0'),s(bool,X6),s(fun('type.benchmark.times'(X5),bool),X7),s(fun('type.benchmark.times'(X5),fun('type.benchmark.times'(X5),bool)),X8)))) )),
  inference(superposition,[],[f317,f238])).
fof(f6545,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,X30) = s(bool,'const.bool.DATATYPE_1'(s(X33,X34))) | s(bool,'F_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X29),'const.benchmark.Zero_0'),s(bool,X30),s(fun('type.benchmark.times'(X29),bool),X31),s(fun('type.benchmark.times'(X29),fun('type.benchmark.times'(X29),bool)),X32)))) )),
  inference(superposition,[],[f317,f1239])).
fof(f6547,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'F_0') = s(bool,X40) | s(bool,X43) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X39),'const.benchmark.Zero_0'),s(bool,X40),s(fun('type.benchmark.times'(X39),bool),X41),s(fun('type.benchmark.times'(X39),fun('type.benchmark.times'(X39),bool)),X42))) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f317,f1593])).
fof(f6549,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) = s(bool,X53) | s(bool,'F_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X49),'const.benchmark.Zero_0'),s(bool,X50),s(fun('type.benchmark.times'(X49),bool),X51),s(fun('type.benchmark.times'(X49),fun('type.benchmark.times'(X49),bool)),X52)))) )),
  inference(superposition,[],[f317,f1593])).
fof(f6598,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,X239) != s(bool,X242) | s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X238),'const.benchmark.Zero_0'),s(bool,X239),s(fun('type.benchmark.times'(X238),bool),X240),s(fun('type.benchmark.times'(X238),fun('type.benchmark.times'(X238),bool)),X241))) | s(bool,'F_0') = s(bool,X242)) )),
  inference(superposition,[],[f1700,f317])).
fof(f6599,plain,(
  ( ! [X243,X246,X244,X247,X245] : (s(bool,X244) != s(bool,X247) | s(bool,'T_0') = s(bool,X247) | s(bool,'F_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X243),'const.benchmark.Zero_0'),s(bool,X244),s(fun('type.benchmark.times'(X243),bool),X245),s(fun('type.benchmark.times'(X243),fun('type.benchmark.times'(X243),bool)),X246)))) )),
  inference(superposition,[],[f1700,f317])).
fof(f6600,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'T_0') != s(bool,X249) | s(bool,X252) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X248),'const.benchmark.Zero_0'),s(bool,X249),s(fun('type.benchmark.times'(X248),bool),X250),s(fun('type.benchmark.times'(X248),fun('type.benchmark.times'(X248),bool)),X251))) | s(bool,'F_0') = s(bool,X252)) )),
  inference(superposition,[],[f2340,f317])).
fof(f6889,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f6818,f337])).
fof(f6818,plain,(
  ( ! [X14,X13] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X13))),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f3294,f1593])).
fof(f6895,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f344,f320])).
fof(f6907,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X42)))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f320,f238])).
fof(f6914,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f320,f1593])).
fof(f6916,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,X69)))) | s(bool,X67) = s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f320,f1593])).
fof(f6922,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X84))),s(bool,X85)))) | s(bool,'F_0') = s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f320,f1593])).
fof(f6924,plain,(
  ( ! [X90,X91] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X90))),s(bool,X91)))) | s(bool,'F_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f320,f238])).
fof(f6950,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'F_0')))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f320,f238])).
fof(f6957,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'F_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f320,f1593])).
fof(f7008,plain,(
  ( ! [X58,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X58)))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f6912,f266])).
fof(f6912,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.DATATYPE_1'(s(X56,X57))))),s(bool,X58)))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f320,f1239])).
fof(f7036,plain,(
  ( ! [X58,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'T_0')))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f6955,f266])).
fof(f6955,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'const.bool.DATATYPE_1'(s(X56,X57)))))) | s(bool,X55) = s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f320,f1239])).
fof(f7057,plain,(
  ( ! [X21,X19,X20] : (s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(subsumption_resolution,[],[f7056,f231])).
fof(f7056,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(forward_demodulation,[],[f6989,f266])).
fof(f6989,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X22,X23)))) | s(X19,X20) = s(X19,X21) | s(bool,'F_0') = s(bool,happ(s(fun(X19,bool),happ(s(fun(X19,fun(X19,bool)),'=_0'),s(X19,X20))),s(X19,X21)))) )),
  inference(superposition,[],[f320,f1239])).
fof(f7210,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f7126,f337])).
fof(f7126,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X15))))) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f6744,f1593])).
fof(f7220,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f6892,f246])).
fof(f7221,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f6892,f241])).
fof(f7260,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X0))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f7214,f327])).
fof(f7214,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f6892,f307])).
fof(f7265,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f7243,f231])).
fof(f7243,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f6892,f359])).
fof(f7311,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X42)))) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f321,f238])).
fof(f7318,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'F_0'))),s(bool,X62)))) | s(bool,X61) != s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7320,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X68))),s(bool,X69)))) | s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7326,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X84))),s(bool,X85)))) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7328,plain,(
  ( ! [X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X90))),s(bool,X91)))) | s(bool,'F_0') != s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f321,f238])).
fof(f7354,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X42))),s(bool,'F_0')))) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f321,f238])).
fof(f7361,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'F_0')))) | s(bool,X61) != s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7363,plain,(
  ( ! [X68,X69,X67] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X69))),s(bool,X68)))) | s(bool,X67) != s(bool,X69) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7369,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X85))),s(bool,X84)))) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7371,plain,(
  ( ! [X90,X91] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X91))),s(bool,X90)))) | s(bool,'F_0') != s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f321,f238])).
fof(f7484,plain,(
  ( ! [X58,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X58)))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f7316,f266])).
fof(f7316,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'const.bool.DATATYPE_1'(s(X56,X57))))),s(bool,X58)))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f321,f1239])).
fof(f7512,plain,(
  ( ! [X58,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'T_0')))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(forward_demodulation,[],[f7359,f266])).
fof(f7359,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X58))),s(bool,'const.bool.DATATYPE_1'(s(X56,X57)))))) | s(bool,X55) != s(bool,X58) | s(bool,'F_0') = s(bool,X55)) )),
  inference(superposition,[],[f321,f1239])).
fof(f7538,plain,(
  ( ! [X30,X28,X29,X27] : (s(X27,X28) != s(X27,X29) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),happ(s(fun(X27,fun(X27,bool)),'=_0'),s(X27,X28))),s(X27,X29))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f7395,f337])).
fof(f7395,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'F_0')) | s(X27,X28) != s(X27,X29) | s(bool,X30) = s(bool,happ(s(fun(X27,bool),happ(s(fun(X27,fun(X27,bool)),'=_0'),s(X27,X28))),s(X27,X29))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f321,f1593])).
fof(f7598,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7572,f337])).
fof(f7572,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f7403,f1593])).
fof(f7656,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f7627,f337])).
fof(f7627,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f7545,f1593])).
fof(f7668,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK6(X0,bool))) )),
  inference(superposition,[],[f300,f238])).
fof(f7675,plain,(
  ( ! [X14,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X13)))) | s(bool,X14) = s(bool,sK6(X13,bool)) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f300,f1593])).
fof(f7677,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X17)))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,sK6(X17,bool))) )),
  inference(superposition,[],[f300,f1593])).
fof(f7706,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X9)))) | s(bool,'F_0') = s(bool,sK6(X9,bool))) )),
  inference(forward_demodulation,[],[f7673,f266])).
fof(f7673,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'const.bool.DATATYPE_1'(s(X10,X11)))))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X9)))) | s(bool,'F_0') = s(bool,sK6(X9,bool))) )),
  inference(superposition,[],[f300,f1239])).
fof(f7717,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK6(X15,X14))))) )),
  inference(subsumption_resolution,[],[f7716,f231])).
fof(f7716,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK6(X15,X14))))) )),
  inference(forward_demodulation,[],[f7687,f266])).
fof(f7687,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,happ(s(fun(fun(X14,bool),bool),'!_1'),s(fun(X14,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(X14,bool),X15),s(X14,sK6(X15,X14))))) )),
  inference(superposition,[],[f300,f1239])).
fof(f7929,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK7(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(subsumption_resolution,[],[f7928,f231])).
fof(f7928,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK7(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(forward_demodulation,[],[f7909,f266])).
fof(f7909,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK7(X15,X14))))) | s(bool,'F_0') = s(bool,happ(s(fun(fun(X14,bool),bool),'?_1'),s(fun(X14,bool),X15)))) )),
  inference(superposition,[],[f301,f1239])).
fof(f8079,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'F_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f337,f337,f303])).
fof(f8187,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X43))),s(bool,X44)))) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8095,f337])).
fof(f8095,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X43))),s(bool,X44)))) | p(s(bool,'F_0')) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f303,f1593])).
fof(f8253,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X44))),s(bool,X43)))) | p(s(bool,X44)) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f8137,f337])).
fof(f8137,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X44))),s(bool,X43)))) | p(s(bool,X44)) | p(s(bool,'F_0')) | s(bool,X43) = s(bool,X45) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f303,f1593])).
fof(f8311,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(subsumption_resolution,[],[f8310,f231])).
fof(f8310,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(forward_demodulation,[],[f8169,f266])).
fof(f8169,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X14))),s(bool,X15)))) )),
  inference(superposition,[],[f303,f1239])).
fof(f8683,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f231,f337,f306])).
fof(f8727,plain,(
  ( ! [X14,X17] : (p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f8726,f231])).
fof(f8726,plain,(
  ( ! [X14,X17] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(forward_demodulation,[],[f8689,f266])).
fof(f8689,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X15,X16)))) | p(s(bool,X17)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X14))),s(bool,X17)))) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f306,f1239])).
fof(f8738,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X52))),s(bool,X51)))) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X51)) )),
  inference(forward_demodulation,[],[f8737,f327])).
fof(f8737,plain,(
  ( ! [X52,X51] : (p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X51))))),s(bool,X52)))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f8702,f231])).
fof(f8702,plain,(
  ( ! [X52,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X51))))),s(bool,X52)))) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f306,f392])).
fof(f299,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f3223,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | p(s(bool,'|dollar|exists_1'(s(fun(X26,bool),sK4(X27,X28,fun(X26,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f3209,f693])).
fof(f3288,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3259,f231])).
fof(f3259,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f3209,f359])).
fof(f3452,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f3436,f231])).
fof(f3436,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f299,f359])).
fof(f3794,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X29)))) | p(s(bool,happ(s(fun(fun(X26,bool),bool),'?_1'),s(fun(X26,bool),sK4(X27,X28,fun(X26,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f302,f693])).
fof(f3865,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3833,f231])).
fof(f3833,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X4,bool),bool),'?_1'),s(fun(X4,bool),X5)))) | s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f302,f359])).
fof(f7715,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK6(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7683,f231])).
fof(f7683,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X3)))) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK6(X3,X2)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f300,f359])).
fof(f7698,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f7659,f267])).
fof(f7659,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,happ(s(fun(fun(X0,bool),bool),'!_1'),s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f300,f267])).
fof(f7699,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK6(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(forward_demodulation,[],[f7660,f312])).
fof(f7660,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK6(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | p(s(bool,happ(s(fun(fun(X2,bool),bool),'!_1'),s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))))))) )),
  inference(superposition,[],[f300,f312])).
fof(f7701,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(forward_demodulation,[],[f7662,f276])).
fof(f7662,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK6('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | p(s(bool,happ(s(fun(fun(X9,bool),bool),'!_1'),s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f300,f276])).
fof(f7703,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),X18)))) | ~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(forward_demodulation,[],[f7664,f311])).
fof(f7664,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK6('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | p(s(bool,happ(s(fun(fun(X16,bool),bool),'!_1'),s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))))))) )),
  inference(superposition,[],[f300,f311])).
fof(f7704,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),X23)))) | ~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK6('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(forward_demodulation,[],[f7665,f313])).
fof(f7665,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK6('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | p(s(bool,happ(s(fun(fun(X20,bool),bool),'!_1'),s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))))))) )),
  inference(superposition,[],[f300,f313])).
fof(f7705,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),X26)))) | ~p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK6('const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28)),X24)))))) )),
  inference(forward_demodulation,[],[f7666,f317])).
fof(f7666,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK6('const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28)),X24))))) | p(s(bool,happ(s(fun(fun(X24,bool),bool),'!_1'),s(fun(X24,bool),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28))))))) )),
  inference(superposition,[],[f300,f317])).
fof(f7667,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(X29,bool),X30),s(X29,sK6(sK4(X30,X31,fun(X29,bool)),X29))))) | p(s(bool,happ(s(fun(fun(X29,bool),bool),'!_1'),s(fun(X29,bool),sK4(X30,X31,fun(X29,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X29,bool),bool),X31)))) )),
  inference(superposition,[],[f300,f693])).
fof(f7927,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f7905,f231])).
fof(f7905,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X2,bool),X3),s(X2,sK7(X3,X2))))) | s(bool,X4) = s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f301,f359])).
fof(f7917,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | ~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f7894,f267])).
fof(f7894,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(fun(X0,bool),bool),'?_1'),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))),s(X0,sK7(happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(superposition,[],[f301,f267])).
fof(f7918,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),X4),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2))))) | ~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4))))) )),
  inference(forward_demodulation,[],[f7895,f312])).
fof(f7895,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(fun(X2,bool),bool),'?_1'),s(fun(X2,bool),X4)))) | p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5))),s(X2,sK7(happ(s(fun(X3,fun(X2,bool)),'_dst_x0_1'(s(fun(X2,bool),X4))),s(X3,X5)),X2)))))) )),
  inference(superposition,[],[f301,f312])).
fof(f7920,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9))))) | ~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10))))) )),
  inference(forward_demodulation,[],[f7897,f276])).
fof(f7897,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun(fun(X9,bool),bool),'?_1'),s(fun(X9,bool),X10)))) | p(s(bool,happ(s(fun(X9,bool),'const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0'))),s(X9,sK7('const.ind_type.FCONS_3'(s(fun(X9,bool),X10),s(fun('type.num.num',fun(X9,bool)),X11),s('type.num.num','const.num.|48|_0')),X9)))))) )),
  inference(superposition,[],[f301,f276])).
fof(f7922,plain,(
  ( ! [X19,X17,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),X18),s(X16,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))))) | ~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X18))))) )),
  inference(forward_demodulation,[],[f7899,f311])).
fof(f7899,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,happ(s(fun(fun(X16,bool),bool),'?_1'),s(fun(X16,bool),X18)))) | p(s(bool,happ(s(fun(X16,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19))))),s(X16,sK7('const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16)))))) )),
  inference(superposition,[],[f301,f311])).
fof(f7923,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),X23),s(X20,sK7('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20))))) | ~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23))))) )),
  inference(forward_demodulation,[],[f7900,f313])).
fof(f7900,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(fun(X20,bool),bool),'?_1'),s(fun(X20,bool),X23)))) | p(s(bool,happ(s(fun(X20,bool),'const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23))))),s(X20,sK7('const.pair.SND_1'(s('type.pair.prod'(X21,fun(X20,bool)),'const.pair.,_2'(s(X21,X22),s(fun(X20,bool),X23)))),X20)))))) )),
  inference(superposition,[],[f301,f313])).
fof(f7924,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,happ(s(fun(X24,bool),X26),s(X24,sK7('const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28)),X24))))) | ~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X26))))) )),
  inference(forward_demodulation,[],[f7901,f317])).
fof(f7901,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,happ(s(fun(fun(X24,bool),bool),'?_1'),s(fun(X24,bool),X26)))) | p(s(bool,happ(s(fun(X24,bool),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28))),s(X24,sK7('const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X25),'const.benchmark.Zero_0'),s(fun(X24,bool),X26),s(fun('type.benchmark.times'(X25),fun(X24,bool)),X27),s(fun('type.benchmark.times'(X25),fun('type.benchmark.times'(X25),fun(X24,bool))),X28)),X24)))))) )),
  inference(superposition,[],[f301,f317])).
fof(f6906,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X38))),s(X37,X40)))) | s(X37,X40) = s(X37,sK4(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f320,f693])).
fof(f6949,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X40))),s(X37,X38)))) | s(X37,X40) = s(X37,sK4(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f320,f693])).
fof(f7310,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X38))),s(X37,X40)))) | s(X37,X40) != s(X37,sK4(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f321,f693])).
fof(f7353,plain,(
  ( ! [X39,X37,X38,X40] : (p(s(bool,happ(s(fun(X37,bool),happ(s(fun(X37,fun(X37,bool)),'=_0'),s(X37,X40))),s(X37,X38)))) | s(X37,X40) != s(X37,sK4(X38,X39,X37)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X39)))) )),
  inference(superposition,[],[f321,f693])).
fof(f7440,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7439,f324])).
fof(f7439,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f7438,f317])).
fof(f7438,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.Zero_0'),s(X0,X2),s(fun('type.benchmark.times'(X1),X0),X3),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X4))))))) )),
  inference(forward_demodulation,[],[f7289,f324])).
fof(f7289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X1),'const.benchmark.Zero_0'),s(X0,X2),s(fun('type.benchmark.times'(X1),X0),X3),s(fun('type.benchmark.times'(X1),fun('type.benchmark.times'(X1),X0)),X4))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f317,f321])).
fof(f7702,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | ~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK6('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f7663,f284])).
fof(f7663,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK6('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | p(s(bool,happ(s(fun(fun(X12,bool),bool),'!_1'),s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))))) )),
  inference(superposition,[],[f300,f284])).
fof(f7921,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12))))) | ~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15))))))) )),
  inference(forward_demodulation,[],[f7898,f284])).
fof(f7898,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun(fun(X12,bool),bool),'?_1'),s(fun(X12,bool),happ(s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num',X15)))))) | p(s(bool,happ(s(fun(X12,bool),'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s(X12,sK7('const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15)))),X12)))))) )),
  inference(superposition,[],[f301,f284])).
fof(f7700,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | ~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK6(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(forward_demodulation,[],[f7661,f309])).
fof(f7661,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK6(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | p(s(bool,happ(s(fun(fun(X6,bool),bool),'!_1'),s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(superposition,[],[f300,f309])).
fof(f7919,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6))))) | ~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8))))))) )),
  inference(forward_demodulation,[],[f7896,f309])).
fof(f7896,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(fun(X6,bool),bool),'?_1'),s(fun(X6,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X6,bool),X7),X8)))))) | p(s(bool,happ(s(fun(X6,bool),happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8))),s(X6,sK7(happ(s(fun('type.pair.prod'(fun(X6,bool),X7),fun(X6,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X6,bool),X7),X8)),X6)))))) )),
  inference(superposition,[],[f301,f309])).
fof(f7474,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1),s(fun(X0,X0),'_i_0_0'))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f7266,f326])).
fof(f7266,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f267,f321])).
fof(f3434,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(fun(X26,bool),bool),'!_1'),s(fun(X26,bool),X27)))) | p(s(bool,happ(s(fun(X26,bool),sK4(X27,X28,fun(X26,bool))),s(X26,X29)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f299,f693])).
fof(f7902,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,happ(s(fun(fun(X29,bool),bool),'?_1'),s(fun(X29,bool),X30)))) | p(s(bool,happ(s(fun(X29,bool),sK4(X30,X31,fun(X29,bool))),s(X29,sK7(sK4(X30,X31,fun(X29,bool)),X29))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X29,bool),bool),X31)))) )),
  inference(superposition,[],[f301,f693])).
fof(f7669,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK6(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f300,f359])).
fof(f7670,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK6(X3,bool))) )),
  inference(superposition,[],[f300,f359])).
fof(f8742,plain,(
  ( ! [X55] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X55)))) | p(s(bool,X55))) )),
  inference(forward_demodulation,[],[f8741,f1450])).
fof(f8741,plain,(
  ( ! [X54,X55] : (p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X54,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X55))))) )),
  inference(subsumption_resolution,[],[f8704,f231])).
fof(f8704,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X54,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(bool,X55))))) )),
  inference(superposition,[],[f306,f1450])).
fof(f7332,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X99)))) | s(bool,X99) != s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) )),
  inference(superposition,[],[f321,f693])).
fof(f7312,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X44)))) | s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f321,f359])).
fof(f6928,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X99)))) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) )),
  inference(superposition,[],[f320,f693])).
fof(f6908,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X44)))) | s(bool,X43) = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f320,f359])).
fof(f1894,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,X5)))) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1873,f231])).
fof(f1873,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X4))),s(bool,X5)))) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f304,f359])).
fof(f2778,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2751,f231])).
fof(f2751,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f305,f359])).
fof(f8121,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X110))),s(bool,X112)))) | p(s(bool,sK4(X110,X111,bool))) | p(s(bool,X112)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X111)))) )),
  inference(superposition,[],[f303,f693])).
fof(f8163,plain,(
  ( ! [X111,X112,X110] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X112))),s(bool,X110)))) | p(s(bool,X112)) | p(s(bool,sK4(X110,X111,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X111)))) )),
  inference(superposition,[],[f303,f693])).
fof(f8180,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,X0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8179])).
fof(f8179,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f8082,f241])).
fof(f8082,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | p(s(bool,X7)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f303,f359])).
fof(f8246,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8245])).
fof(f8245,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f8124,f241])).
fof(f8124,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X7))),s(bool,X6)))) | p(s(bool,X7)) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f303,f359])).
fof(f8309,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8165,f231])).
fof(f8165,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X2))),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f303,f359])).
fof(f3047,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f3046])).
fof(f3046,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2971,f241])).
fof(f2971,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X6))),s(bool,X7)))) | p(s(bool,X5)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f307,f359])).
fof(f3063,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(subsumption_resolution,[],[f3025,f337])).
fof(f3025,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) )),
  inference(superposition,[],[f307,f238])).
fof(f3117,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X76))),s(bool,X78)))) | p(s(bool,sK4(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) )),
  inference(superposition,[],[f307,f693])).
fof(f3325,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3296,f231])).
fof(f3296,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X3))),s(bool,X2)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f308,f359])).
fof(f8725,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f8685,f231])).
fof(f8685,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X2))),s(bool,X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f306,f359])).
fof(f6909,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f320,f359])).
fof(f6920,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,'T_0') = s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f320,f359])).
fof(f7313,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X47))),s(bool,X48)))) | s(bool,X46) != s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f321,f359])).
fof(f7324,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X78))),s(bool,X79)))) | s(bool,'T_0') != s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f321,f359])).
fof(f7356,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X48))),s(bool,X47)))) | s(bool,X46) != s(bool,X48) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f321,f359])).
fof(f7367,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X79))),s(bool,X78)))) | s(bool,'T_0') != s(bool,X79) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f321,f359])).
fof(f3115,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,X70)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,sK4(X70,X71,bool)))),s(bool,X72)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f304,f693])).
fof(f3211,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))))),s(bool,X41)))) | s(X39,X42) = s(X39,sK4(X42,X40,X39))) )),
  inference(subsumption_resolution,[],[f3165,f231])).
fof(f3165,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))))),s(bool,X41)))) | s(X39,X42) = s(X39,sK4(X42,X40,X39))) )),
  inference(superposition,[],[f304,f693])).
fof(f8743,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X58)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))))),s(bool,X58)))) | s(X56,X59) = s(X56,sK4(X59,X57,X56))) )),
  inference(subsumption_resolution,[],[f8705,f231])).
fof(f8705,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))))),s(bool,X58)))) | s(X56,X59) = s(X56,sK4(X59,X57,X56))) )),
  inference(superposition,[],[f306,f693])).
fof(f8724,plain,(
  ( ! [X111,X109,X110] : (~p(s(bool,X109)) | p(s(bool,X111)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,sK4(X109,X110,bool)))),s(bool,X111)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X110)))) )),
  inference(superposition,[],[f306,f693])).
fof(f7375,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X99))),s(bool,'T_0')))) | s(bool,X99) != s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) )),
  inference(superposition,[],[f321,f693])).
fof(f7355,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X44))),s(bool,'T_0')))) | s(bool,X43) != s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f321,f359])).
fof(f6971,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X99))),s(bool,'T_0')))) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) )),
  inference(superposition,[],[f320,f693])).
fof(f6951,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X44))),s(bool,'T_0')))) | s(bool,X43) = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f320,f359])).
fof(f3212,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X45))),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))))) | s(X43,X46) = s(X43,sK4(X46,X44,X43))) )),
  inference(subsumption_resolution,[],[f3166,f231])).
fof(f3166,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X45))),s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))))) | s(X43,X46) = s(X43,sK4(X46,X44,X43))) )),
  inference(superposition,[],[f305,f693])).
fof(f3338,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X55))),s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))))) | s(X53,X56) = s(X53,sK4(X56,X54,X53))) )),
  inference(subsumption_resolution,[],[f3315,f231])).
fof(f3315,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X55))),s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))))) | s(X53,X56) = s(X53,sK4(X56,X54,X53))) )),
  inference(superposition,[],[f308,f693])).
fof(f3116,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X75))),s(bool,sK4(X73,X74,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f305,f693])).
fof(f3324,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X84)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X86))),s(bool,sK4(X84,X85,bool))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X85)))) )),
  inference(superposition,[],[f308,f693])).
fof(f1905,plain,(
  ( ! [X39,X38] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X39))),s(bool,X38)))) | ~p(s(bool,X38))) )),
  inference(forward_demodulation,[],[f1886,f327])).
fof(f1886,plain,(
  ( ! [X39,X38] : (~p(s(bool,X38)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X38))))),s(bool,X39))))) )),
  inference(superposition,[],[f304,f267])).
fof(f5450,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))),s(bool,X0))))) )),
  inference(forward_demodulation,[],[f5360,f327])).
fof(f5360,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f1871,f304])).
fof(f5675,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,X1))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5585,f327])).
fof(f5585,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f2749,f304])).
fof(f6846,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6751,f327])).
fof(f6751,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(unit_resulting_resolution,[],[f3294,f304])).
fof(f8768,plain,(
  ( ! [X80,X81] : (~p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(bool,bool),'_i_0_0'),s(bool,X81))),s(bool,X80)))) | ~p(s(bool,X80)) | p(s(bool,X81))) )),
  inference(forward_demodulation,[],[f8716,f327])).
fof(f8716,plain,(
  ( ! [X80,X81] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X80))))),s(bool,X81))))) )),
  inference(superposition,[],[f306,f267])).
fof(f2795,plain,(
  ( ! [X61,X60] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X61),s(fun(bool,bool),'_i_0_0'))),s(bool,X60)))) | ~p(s(bool,X60))) )),
  inference(forward_demodulation,[],[f2772,f326])).
fof(f2772,plain,(
  ( ! [X61,X60] : (~p(s(bool,X60)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X61))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X60))))))) )),
  inference(superposition,[],[f305,f267])).
fof(f3343,plain,(
  ( ! [X64,X65] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X65),s(fun(bool,bool),'_i_0_0'))),s(bool,X64)))) | ~p(s(bool,X64))) )),
  inference(forward_demodulation,[],[f3318,f326])).
fof(f3318,plain,(
  ( ! [X64,X65] : (~p(s(bool,X64)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X65))),s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X64))))))) )),
  inference(superposition,[],[f308,f267])).
fof(f5448,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f5362,f326])).
fof(f5362,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f1871,f308])).
fof(f5449,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))),s(bool,X1))))) )),
  inference(forward_demodulation,[],[f5361,f326])).
fof(f5361,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))) )),
  inference(unit_resulting_resolution,[],[f1871,f305])).
fof(f5673,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5587,f326])).
fof(f5587,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f2749,f308])).
fof(f5674,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f5586,f326])).
fof(f5586,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f2749,f305])).
fof(f6844,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6753,f326])).
fof(f6753,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f308])).
fof(f6845,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f6752,f326])).
fof(f6752,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f305])).
fof(f6743,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f302])).
fof(f7544,plain,(
  p(s(bool,happ(s(fun(fun('type.num.num',bool),bool),'?_1'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f7403,f302])).
fof(f3846,plain,(
  ( ! [X52,X53] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),X53)))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X52),bool),bool),'?_1'),s(fun('type.pair.prod'(bool,X52),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f302,f309])).
fof(f7213,plain,(
  ~p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f6892,f299])).
fof(f7697,plain,(
  ( ! [X39] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X39),sK6('const.pair.FST_0','type.pair.prod'(bool,X39)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'(bool,X39),bool),bool),'!_1'),s(fun('type.pair.prod'(bool,X39),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f300,f309])).
fof(f6997,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))))))) )),
  inference(forward_demodulation,[],[f6897,f324])).
fof(f6897,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))))),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))))) )),
  inference(unit_resulting_resolution,[],[f283,f320])).
fof(f6998,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f6896,f324])).
fof(f6896,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))))) )),
  inference(unit_resulting_resolution,[],[f275,f320])).
fof(f6894,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))))))) )),
  inference(unit_resulting_resolution,[],[f283,f320])).
fof(f6893,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f275,f320])).
fof(f7406,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7405,f324])).
fof(f7405,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7404,f269])).
fof(f7404,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),'_dst_x1x0_2'(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1))))))) )),
  inference(forward_demodulation,[],[f7300,f324])).
fof(f7300,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f269,f321])).
fof(f7457,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.times'(X0),bool),'_dst_x0x0x1_3'(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))),s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'))),s('type.benchmark.times'(X0),X1))))) )),
  inference(forward_demodulation,[],[f7282,f326])).
fof(f7282,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.benchmark.digit'(X0),bool),happ(s(fun('type.benchmark.digit'(X0),fun('type.benchmark.digit'(X0),bool)),'=_0'),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))))),s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1))))))) )),
  inference(unit_resulting_resolution,[],[f269,f321])).
fof(f6119,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',X9))) | s(bool,X10) = s(bool,'const.ind_type.FCONS_3'(s(bool,X7),s(fun('type.num.num',bool),X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f284,f359])).
fof(f6120,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) = s(bool,happ(s(fun('type.num.num',bool),X12),s('type.num.num',X13))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X11),s(fun('type.num.num',bool),X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13)))))) )),
  inference(superposition,[],[f284,f359])).
fof(f6220,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,happ(s(fun('type.num.num',bool),X105),s('type.num.num',X106)))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X105),s(bool,X107))),s('type.num.num',X106))))) )),
  inference(forward_demodulation,[],[f6219,f284])).
fof(f6219,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.num.num',bool),X105),s(bool,X107))),s('type.num.num',X106)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X104),s(fun('type.num.num',bool),X105),s('type.num.num','const.num.SUC_1'(s('type.num.num',X106))))))) )),
  inference(forward_demodulation,[],[f6154,f327])).
fof(f6154,plain,(
  ( ! [X107,X105,X106,X104] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun('type.num.num',bool),X105),s('type.num.num',X106))))),s(bool,X107)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X104),s(fun('type.num.num',bool),X105),s('type.num.num','const.num.SUC_1'(s('type.num.num',X106))))))) )),
  inference(superposition,[],[f307,f284])).
fof(f6224,plain,(
  ( ! [X144,X145,X143] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',X145))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X143),s(fun('type.num.num',bool),X144),s('type.num.num','const.num.SUC_1'(s('type.num.num',X145)))))) )),
  inference(forward_demodulation,[],[f6164,f266])).
fof(f6164,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.bool.DATATYPE_1'(s(X146,X147))) != s(bool,happ(s(fun('type.num.num',bool),X144),s('type.num.num',X145))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X143),s(fun('type.num.num',bool),X144),s('type.num.num','const.num.SUC_1'(s('type.num.num',X145)))))) )),
  inference(superposition,[],[f643,f284])).
fof(f8239,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X94),s(bool,X96))),s('type.num.num',X95)))) | p(s(bool,X96))) )),
  inference(forward_demodulation,[],[f8238,f284])).
fof(f8238,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),X94),s(bool,X96))),s('type.num.num',X95)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95)))))) | p(s(bool,X96))) )),
  inference(forward_demodulation,[],[f8117,f327])).
fof(f8117,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))))),s(bool,X96)))) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95)))))) | p(s(bool,X96))) )),
  inference(superposition,[],[f303,f284])).
fof(f8305,plain,(
  ( ! [X94,X95,X96] : (p(s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X96),s(fun('type.num.num',bool),X94))),s('type.num.num',X95)))) | p(s(bool,X96))) )),
  inference(forward_demodulation,[],[f8304,f284])).
fof(f8304,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X96),s(fun('type.num.num',bool),X94))),s('type.num.num',X95)))) | p(s(bool,X96)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))))))) )),
  inference(forward_demodulation,[],[f8159,f326])).
fof(f8159,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X96))),s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95)))))) | p(s(bool,X96)) | p(s(bool,'const.ind_type.FCONS_3'(s(bool,X93),s(fun('type.num.num',bool),X94),s('type.num.num','const.num.SUC_1'(s('type.num.num',X95))))))) )),
  inference(superposition,[],[f303,f284])).
fof(f7467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f7466,f284])).
fof(f7466,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun('type.num.num',X0),X1),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))),s('type.num.num',X2))))) )),
  inference(forward_demodulation,[],[f7273,f327])).
fof(f7273,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X1),s('type.num.num',X2))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X3),s(fun('type.num.num',X0),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) )),
  inference(unit_resulting_resolution,[],[f284,f321])).
fof(f7480,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f7479,f284])).
fof(f7479,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7306,f327])).
fof(f7306,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f321,f284])).
fof(f7004,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f7003,f284])).
fof(f7003,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(fun('type.num.num',X16),X18),s(X16,X20))),s('type.num.num',X19)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f6902,f327])).
fof(f6902,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))))),s(X16,X20)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f320,f284])).
fof(f7590,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s(bool,X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7552,f327])).
fof(f7552,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f7403,f304])).
fof(f7508,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) != s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f7507,f284])).
fof(f7507,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f7349,f326])).
fof(f7349,plain,(
  ( ! [X19,X17,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) != s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f321,f284])).
fof(f7032,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19))) | ~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19))))) )),
  inference(forward_demodulation,[],[f7031,f284])).
fof(f7031,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20),s(fun('type.num.num',X16),X18))),s('type.num.num',X19)))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(forward_demodulation,[],[f6945,f326])).
fof(f6945,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X16,fun(X16,bool)),'=_0'),s(X16,X20))),s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))))) | s(X16,X20) = s(X16,'const.ind_type.FCONS_3'(s(X16,X17),s(fun('type.num.num',X16),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) )),
  inference(superposition,[],[f320,f284])).
fof(f7434,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f7433,f326])).
fof(f7433,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f7432,f324])).
fof(f7432,plain,(
  ( ! [X2,X0,X3] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(forward_demodulation,[],[f7431,f284])).
fof(f7431,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) )),
  inference(forward_demodulation,[],[f7291,f324])).
fof(f7291,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))),s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))))))) )),
  inference(unit_resulting_resolution,[],[f284,f321])).
fof(f7588,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7554,f326])).
fof(f7554,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f7403,f308])).
fof(f7589,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x0x0x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f7553,f326])).
fof(f7553,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f7403,f305])).
fof(f687,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f279,f267])).
fof(f688,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f279,f312])).
fof(f689,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | s(X8,X11) = s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) )),
  inference(superposition,[],[f279,f276])).
fof(f690,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | s(X12,X16) = s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) )),
  inference(superposition,[],[f279,f311])).
fof(f691,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | s(X17,X21) = s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) )),
  inference(superposition,[],[f279,f313])).
fof(f3078,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) )),
  inference(superposition,[],[f693,f267])).
fof(f3079,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5))) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) )),
  inference(superposition,[],[f693,f312])).
fof(f3081,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))) | s(X12,X15) = s(X12,sK4(X15,'const.ind_type.FCONS_3'(s(fun(X12,bool),X13),s(fun('type.num.num',fun(X12,bool)),X14),s('type.num.num','const.num.|48|_0')),X12))) )),
  inference(superposition,[],[f693,f276])).
fof(f3082,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X16,bool),X18))) | s(X16,X20) = s(X16,sK4(X20,'const.pair.FST_1'(s('type.pair.prod'(fun(X16,bool),X17),'const.pair.,_2'(s(fun(X16,bool),X18),s(X17,X19)))),X16))) )),
  inference(superposition,[],[f693,f311])).
fof(f3083,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X24))) | s(X21,X25) = s(X21,sK4(X25,'const.pair.SND_1'(s('type.pair.prod'(X22,fun(X21,bool)),'const.pair.,_2'(s(X22,X23),s(fun(X21,bool),X24)))),X21))) )),
  inference(superposition,[],[f693,f313])).
fof(f3084,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),X27))) | s(X26,X29) = s(X26,sK4(X29,sK4(X27,X28,fun(X26,bool)),X26)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X26,bool),bool),X28)))) )),
  inference(superposition,[],[f693,f693])).
fof(f3142,plain,(
  ( ! [X154,X152,X151,X153] : (p(s(bool,'|dollar|exists_1'(s(fun(X151,bool),X152)))) | s(X151,X154) = s(X151,sK4(X154,sK4(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) )),
  inference(superposition,[],[f279,f693])).
fof(f3210,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | s(bool,X37) = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(subsumption_resolution,[],[f3164,f231])).
fof(f3164,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | s(bool,X37) = s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))) | s(X35,X38) = s(X35,sK4(X38,X36,X35))) )),
  inference(superposition,[],[f298,f693])).
fof(f3214,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))))))) | s(X92,X95) = s(X92,sK4(X95,X93,X92))) )),
  inference(subsumption_resolution,[],[f3177,f231])).
fof(f3177,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))))))) | s(X92,X95) = s(X92,sK4(X95,X93,X92))) )),
  inference(superposition,[],[f1262,f693])).
fof(f6609,plain,(
  ( ! [X294,X296,X298,X295,X297,X299] : (p(s(bool,'|dollar|exists_1'(s(fun(X294,bool),X296)))) | s(X294,X299) = s(X294,sK4(X299,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X295),'const.benchmark.Zero_0'),s(fun(X294,bool),X296),s(fun('type.benchmark.times'(X295),fun(X294,bool)),X297),s(fun('type.benchmark.times'(X295),fun('type.benchmark.times'(X295),fun(X294,bool))),X298)),X294))) )),
  inference(superposition,[],[f279,f317])).
fof(f6610,plain,(
  ( ! [X300,X302,X304,X301,X303,X305] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X300,bool),X302))) | s(X300,X305) = s(X300,sK4(X305,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X301),'const.benchmark.Zero_0'),s(fun(X300,bool),X302),s(fun('type.benchmark.times'(X301),fun(X300,bool)),X303),s(fun('type.benchmark.times'(X301),fun('type.benchmark.times'(X301),fun(X300,bool))),X304)),X300))) )),
  inference(superposition,[],[f693,f317])).
fof(f6187,plain,(
  ( ! [X241,X239,X237,X240,X238] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X237,bool),happ(s(fun('type.num.num',fun(X237,bool)),X239),s('type.num.num',X240))))) | s(X237,X241) = s(X237,sK4(X241,'const.ind_type.FCONS_3'(s(fun(X237,bool),X238),s(fun('type.num.num',fun(X237,bool)),X239),s('type.num.num','const.num.SUC_1'(s('type.num.num',X240)))),X237))) )),
  inference(superposition,[],[f693,f284])).
fof(f6186,plain,(
  ( ! [X235,X233,X236,X234,X232] : (p(s(bool,'|dollar|exists_1'(s(fun(X232,bool),happ(s(fun('type.num.num',fun(X232,bool)),X234),s('type.num.num',X235)))))) | s(X232,X236) = s(X232,sK4(X236,'const.ind_type.FCONS_3'(s(fun(X232,bool),X233),s(fun('type.num.num',fun(X232,bool)),X234),s('type.num.num','const.num.SUC_1'(s('type.num.num',X235)))),X232))) )),
  inference(superposition,[],[f279,f284])).
fof(f1389,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'|dollar|exists_1'(s(fun(X99,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X99,bool),X100),X101)))))) | s(X99,X102) = s(X99,sK4(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) )),
  inference(superposition,[],[f279,f309])).
fof(f3080,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X8,bool),X9),X10))))) | s(X8,X11) = s(X8,sK4(X11,happ(s(fun('type.pair.prod'(fun(X8,bool),X9),fun(X8,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X8,bool),X9),X10)),X8))) )),
  inference(superposition,[],[f693,f309])).
fof(f1425,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1408,f231])).
fof(f1408,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1262,f359])).
fof(f3129,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | p(s(bool,'|dollar|exists_1'(s(fun(X115,bool),'_dst_x0_1'(s(bool,sK4(X113,X114,bool))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X114)))) )),
  inference(superposition,[],[f1262,f693])).
fof(f3108,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') != s(bool,sK4(X53,X54,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) )),
  inference(superposition,[],[f241,f693])).
fof(f3109,plain,(
  ( ! [X56,X55] : (~p(s(bool,X55)) | s(bool,'T_0') = s(bool,sK4(X55,X56,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f242,f693])).
fof(f3112,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X61)) | p(s(bool,X63)) | s(bool,X63) = s(bool,sK4(X61,X62,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X62)))) )),
  inference(superposition,[],[f295,f693])).
fof(f3114,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X69)) | s(bool,X69) = s(bool,sK4(X67,X68,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68)))) )),
  inference(superposition,[],[f298,f693])).
fof(f3123,plain,(
  ( ! [X94,X95,X96] : (~p(s(bool,X94)) | s(bool,X96) != s(bool,sK4(X94,X95,bool)) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X95)))) )),
  inference(superposition,[],[f362,f693])).
fof(f3125,plain,(
  ( ! [X101,X99,X100] : (p(s(bool,X99)) | s(bool,X101) = s(bool,sK4(X99,X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) )),
  inference(superposition,[],[f418,f693])).
fof(f3148,plain,(
  ( ! [X175,X173,X176,X174] : (s('type.benchmark.times'(X173),X174) != s('type.benchmark.times'(X173),X176) | s('type.benchmark.digit'(X173),'const.benchmark.One_1'(s('type.benchmark.times'(X173),X176))) = s('type.benchmark.digit'(X173),'const.benchmark.One_1'(s('type.benchmark.times'(X173),sK4(X174,X175,'type.benchmark.times'(X173))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.benchmark.times'(X173),bool),X175)))) )),
  inference(superposition,[],[f282,f693])).
fof(f1354,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),X4))) | s(bool,X5) = s(bool,happ(s(fun('type.pair.prod'(bool,X3),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X3),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f309,f359])).
fof(f1356,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X9),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X9),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f309,f359])).
fof(f1404,plain,(
  ( ! [X85,X86] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) )),
  inference(forward_demodulation,[],[f1385,f266])).
fof(f1385,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'const.bool.DATATYPE_1'(s(X87,X88))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),X86))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'(bool,X85),bool),'const.pair.FST_0'),s('type.pair.prod'(bool,X85),X86)))) )),
  inference(superposition,[],[f643,f309])).
fof(f339,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.CONSTR_3'(s('type.num.num','const.num.|48|_0'),s('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),happ(s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'_dst_x1x0_2'(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'))))))),s(fun('type.benchmark.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),sK5(X0)))))) )),
  inference(forward_demodulation,[],[f294,f285])).
fof(f294,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),happ(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),happ(s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s('type.benchmark.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'_dst_x1x0_2'(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'))))))),s(fun('type.benchmark.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),sK5(X0)))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),happ(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),happ(s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s('type.benchmark.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'_dst_x1x0_2'(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'))))))),s(fun('type.benchmark.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),sK5(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f129])).
fof(f129,plain,(
  ! [X0] : ? [X1] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),happ(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),happ(s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s('type.benchmark.times'(X0),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'_dst_x1x0_2'(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s('type.benchmark.times'(X0),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),bool)),'_dst_x00x11_2'(s(fun(fun('type.benchmark.times'(X0),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),'_i_0_0'))))))),s(fun('type.benchmark.digit'(X0),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X0),'type.benchmark.times'(X0)))),X1))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ? [X38] : p(s(bool,'const.bool.TYPE_DEFINITION_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),bool),bool),'!_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),bool),'_dst_x0x1_2'(s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),bool),'!_0'),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool)),'_dst_x00x01x11_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x1x0_2'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool)),'=_0'),s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),happ(s(fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),happ(s(fun('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)),happ(s(fun('type.benchmark.times'(X6),'type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),happ(s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),'type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'const.pair.,_0'),s('type.benchmark.times'(X6),'const.bool.ARB_0'))),s('type.benchmark.times'(X6),'const.bool.ARB_0'))))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'const.ind_type.BOTTOM_0'))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.benchmark.times'(X6),bool)),'_dst_x00x10x01_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_i_0_0'),s(fun('type.benchmark.times'(X6),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_dst_x0x1x0_3'(s(fun('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(fun('type.benchmark.times'(X6),'type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'_dst_x1x0_2'(s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),'type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'const.pair.,_0'),s('type.benchmark.times'(X6),'const.bool.ARB_0'))),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'const.ind_type.BOTTOM_0'))))))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),'_dst_x0x1_2'(s(fun(fun('type.benchmark.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.benchmark.times'(X6),bool)),'_dst_x00x11_2'(s(fun(fun('type.benchmark.times'(X6),bool),bool),'?_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),bool))),'_dst_x000x100x011_3'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool)),'=_0'),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_i_0_0'),s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))))),'_dst_x00x11x00_3'(s(fun('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)),fun(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))))),'const.ind_type.CONSTR_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),'type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'const.pair.,_0'),s(fun('type.num.num','type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_dst_x0_1'(s('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'const.ind_type.BOTTOM_0'))))))))))))))),s(fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool)),'_i_0_0'))))),s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),fun(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),bool),bool)),'_dstsel_1_x10_1'(s(fun('type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6))),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),'_i_0_0'))))))),s(fun('type.benchmark.digit'(X6),'type.ind_type.recspace'('type.pair.prod'('type.benchmark.times'(X6),'type.benchmark.times'(X6)))),X38))))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_TY_DEF')).
fof(f7617,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f344,f7545,f1212])).
fof(f7618,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f7545,f1262])).
fof(f7545,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f7403,f3209])).
fof(f7556,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X0),s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f284,f7403,f362])).
fof(f7563,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f344,f7403,f1212])).
fof(f7564,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f7403,f1262])).
fof(f7531,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(subsumption_resolution,[],[f7388,f337])).
fof(f7388,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) )),
  inference(superposition,[],[f321,f238])).
fof(f7437,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7436,f324])).
fof(f7436,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7435,f309])).
fof(f7435,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(forward_demodulation,[],[f7290,f324])).
fof(f7290,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2))))),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))))))) )),
  inference(unit_resulting_resolution,[],[f309,f321])).
fof(f7403,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f7402,f324])).
fof(f7402,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f7401,f268])).
fof(f7401,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),'_dst_x1x0_2'(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))))))) )),
  inference(forward_demodulation,[],[f7301,f324])).
fof(f7301,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'=_0'),s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'))))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f268,f321])).
fof(f7076,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f344,f6744,f1212])).
fof(f7077,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f6744,f1262])).
fof(f7055,plain,(
  ( ! [X6,X4,X5,X3] : (s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6985,f231])).
fof(f6985,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X4) = s(X3,X5) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),happ(s(fun(X3,fun(X3,bool)),'=_0'),s(X3,X4))),s(X3,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f320,f359])).
fof(f6892,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(unit_resulting_resolution,[],[f344,f320])).
fof(f6744,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f3209])).
fof(f6761,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f344,f3294,f1212])).
fof(f6762,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X1))),s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f3294,f1262])).
fof(f6740,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f6708,f337])).
fof(f6708,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f6334,f1593])).
fof(f6690,plain,(
  ( ! [X4] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X4)))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6334,f238])).
fof(f6688,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0)))))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f6334,f1593])).
fof(f6686,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0')))))),
  inference(unit_resulting_resolution,[],[f344,f6334,f1212])).
fof(f6687,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f6334,f1262])).
fof(f6615,plain,(
  ( ! [X333,X329,X331,X332,X334,X330] : (s('type.benchmark.digit'(X329),'const.benchmark.One_1'(s('type.benchmark.times'(X329),X331))) != s('type.benchmark.digit'(X329),'const.benchmark.One_1'(s('type.benchmark.times'(X329),X334))) | s('type.benchmark.times'(X329),X334) = s('type.benchmark.times'(X329),'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X330),'const.benchmark.Zero_0'),s('type.benchmark.times'(X329),X331),s(fun('type.benchmark.times'(X330),'type.benchmark.times'(X329)),X332),s(fun('type.benchmark.times'(X330),fun('type.benchmark.times'(X330),'type.benchmark.times'(X329))),X333)))) )),
  inference(superposition,[],[f281,f317])).
fof(f6587,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,X184) != s(bool,'const.bool.DATATYPE_1'(s(X187,X188))) | s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X183),'const.benchmark.Zero_0'),s(bool,X184),s(fun('type.benchmark.times'(X183),bool),X185),s(fun('type.benchmark.times'(X183),fun('type.benchmark.times'(X183),bool)),X186)))) )),
  inference(superposition,[],[f643,f317])).
fof(f6542,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X14),'const.benchmark.Zero_0'),s(bool,X15),s(fun('type.benchmark.times'(X14),bool),X16),s(fun('type.benchmark.times'(X14),fun('type.benchmark.times'(X14),bool)),X17)))) )),
  inference(superposition,[],[f317,f359])).
fof(f6541,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,'T_0') = s(bool,X10) | s(bool,X13) = s(bool,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X9),'const.benchmark.Zero_0'),s(bool,X10),s(fun('type.benchmark.times'(X9),bool),X11),s(fun('type.benchmark.times'(X9),fun('type.benchmark.times'(X9),bool)),X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f317,f359])).
fof(f6333,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f2968,f302])).
fof(f6334,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))))))),
  inference(unit_resulting_resolution,[],[f2968,f3209])).
fof(f6351,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f2968,f1262])).
fof(f6191,plain,(
  ( ! [X255,X257,X259,X256,X258] : (s('type.benchmark.digit'(X255),'const.benchmark.One_1'(s('type.benchmark.times'(X255),X259))) != s('type.benchmark.digit'(X255),'const.benchmark.One_1'(s('type.benchmark.times'(X255),happ(s(fun('type.num.num','type.benchmark.times'(X255)),X257),s('type.num.num',X258))))) | s('type.benchmark.times'(X255),X259) = s('type.benchmark.times'(X255),'const.ind_type.FCONS_3'(s('type.benchmark.times'(X255),X256),s(fun('type.num.num','type.benchmark.times'(X255)),X257),s('type.num.num','const.num.SUC_1'(s('type.num.num',X258)))))) )),
  inference(superposition,[],[f281,f284])).
fof(f5734,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f344,f5578,f1212])).
fof(f5735,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))))))))))) )),
  inference(unit_resulting_resolution,[],[f5578,f1262])).
fof(f5578,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))))))) )),
  inference(unit_resulting_resolution,[],[f2749,f3209])).
fof(f5594,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0')))) )),
  inference(unit_resulting_resolution,[],[f344,f2749,f1212])).
fof(f5595,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X1))),s(bool,'T_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f2749,f1262])).
fof(f5547,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f5517,f337])).
fof(f5517,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f5353,f1593])).
fof(f5501,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0)))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f5353,f359])).
fof(f5499,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0')))))),
  inference(unit_resulting_resolution,[],[f344,f5353,f1212])).
fof(f5500,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f5353,f1262])).
fof(f5353,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))))))),
  inference(unit_resulting_resolution,[],[f1871,f3209])).
fof(f5369,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f344,f1871,f1212])).
fof(f5370,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X1))))))))) )),
  inference(unit_resulting_resolution,[],[f1871,f1262])).
fof(f3294,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f231,f308])).
fof(f2749,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0'))))) )),
  inference(unit_resulting_resolution,[],[f231,f305])).
fof(f2968,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'F_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f337,f307])).
fof(f1871,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0))))) )),
  inference(unit_resulting_resolution,[],[f231,f304])).
fof(f3370,plain,(
  ( ! [X14,X15] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(subsumption_resolution,[],[f3369,f231])).
fof(f3369,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(forward_demodulation,[],[f3361,f266])).
fof(f3361,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X16,X17)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X14),bool),'const.pair.FST_0')))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),X15)))) )),
  inference(superposition,[],[f3271,f1239])).
fof(f3368,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3357,f231])).
fof(f3357,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X2),bool),'const.pair.FST_0')))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f3271,f359])).
fof(f3271,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),X49)))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X48),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f3209,f309])).
fof(f3213,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X69) != s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68))) | s(bool,'T_0') = s(bool,X69) | s(X67,X70) = s(X67,sK4(X70,X68,X67))) )),
  inference(subsumption_resolution,[],[f3171,f231])).
fof(f3171,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | s(bool,X69) != s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68))) | s(bool,'T_0') = s(bool,X69) | s(X67,X70) = s(X67,sK4(X70,X68,X67))) )),
  inference(superposition,[],[f362,f693])).
fof(f3146,plain,(
  ( ! [X167,X165,X166,X168] : (s('type.benchmark.digit'(X165),'const.benchmark.One_1'(s('type.benchmark.times'(X165),X166))) != s('type.benchmark.digit'(X165),'const.benchmark.One_1'(s('type.benchmark.times'(X165),X168))) | s('type.benchmark.times'(X165),X168) = s('type.benchmark.times'(X165),sK4(X166,X167,'type.benchmark.times'(X165))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.benchmark.times'(X165),bool),X167)))) )),
  inference(superposition,[],[f281,f693])).
fof(f3141,plain,(
  ( ! [X149,X147,X150,X148] : (~p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK4(X150,sK4(X148,X149,fun(X147,bool)),X147))))) | p(s(bool,'|dollar|exists_1'(s(fun(X147,bool),sK4(X148,X149,fun(X147,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X147,bool),bool),X149)))) )),
  inference(superposition,[],[f280,f693])).
fof(f3126,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,X102) != s(bool,'const.bool.DATATYPE_1'(s(X104,X105))) | s(bool,'T_0') = s(bool,sK4(X102,X103,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X103)))) )),
  inference(superposition,[],[f643,f693])).
fof(f3069,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK4(X8,X9,bool))) )),
  inference(superposition,[],[f693,f359])).
fof(f3068,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6))) | s(bool,X7) = s(bool,sK4(X5,X6,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f693,f359])).
fof(f2906,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f2340,f359])).
fof(f2891,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2856,f344])).
fof(f2856,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f1700,f1593])).
fof(f2425,plain,(
  ( ! [X158,X156,X157] : (~p(s(bool,happ(s(fun(bool,bool),X157),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157)))) | s(bool,X158) = s(bool,sK4(X156,X157,bool)) | s(bool,'F_0') = s(bool,X158)) )),
  inference(superposition,[],[f280,f1593])).
fof(f2340,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f344,f1593])).
fof(f2277,plain,(
  ( ! [X167,X169,X168] : (~p(s(bool,happ(s(fun(bool,bool),X168),s(bool,X169)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X168)))) | s(bool,'F_0') = s(bool,X169) | s(bool,'F_0') = s(bool,sK4(X167,X168,bool))) )),
  inference(superposition,[],[f280,f1593])).
fof(f2154,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f1593,f1593])).
fof(f1708,plain,(
  ( ! [X182,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,'F_0') = s(bool,sK4(X182,X183,bool))) )),
  inference(forward_demodulation,[],[f1642,f266])).
fof(f1642,plain,(
  ( ! [X185,X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X183),s(bool,'const.bool.DATATYPE_1'(s(X184,X185)))))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | s(bool,'F_0') = s(bool,sK4(X182,X183,bool))) )),
  inference(superposition,[],[f280,f1239])).
fof(f1705,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK4(X90,X89,X88))))) )),
  inference(subsumption_resolution,[],[f1704,f231])).
fof(f1704,plain,(
  ( ! [X90,X88,X89] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK4(X90,X89,X88))))) )),
  inference(forward_demodulation,[],[f1620,f266])).
fof(f1620,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'const.bool.DATATYPE_1'(s(X91,X92)))) | p(s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) | s(bool,'F_0') = s(bool,happ(s(fun(X88,bool),X89),s(X88,sK4(X90,X89,X88))))) )),
  inference(superposition,[],[f280,f1239])).
fof(f1547,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1516,f231])).
fof(f1516,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | s(bool,X20) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f362,f1450])).
fof(f1496,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1450,f359])).
fof(f1486,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s('type.benchmark.digit'(X21),'const.benchmark.One_1'(s('type.benchmark.times'(X21),X24))) != s('type.benchmark.digit'(X21),'const.benchmark.One_1'(s('type.benchmark.times'(X21),X25))) | s('type.benchmark.times'(X21),X25) = s('type.benchmark.times'(X21),'const.pair.SND_1'(s('type.pair.prod'(X22,'type.benchmark.times'(X21)),'const.pair.,_2'(s(X22,X23),s('type.benchmark.times'(X21),X24)))))) )),
  inference(superposition,[],[f281,f313])).
fof(f1485,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s('type.benchmark.digit'(X16),'const.benchmark.One_1'(s('type.benchmark.times'(X16),X18))) != s('type.benchmark.digit'(X16),'const.benchmark.One_1'(s('type.benchmark.times'(X16),X20))) | s('type.benchmark.times'(X16),X20) = s('type.benchmark.times'(X16),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(X16),X17),'const.pair.,_2'(s('type.benchmark.times'(X16),X18),s(X17,X19)))))) )),
  inference(superposition,[],[f281,f311])).
fof(f1484,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.benchmark.digit'(X12),'const.benchmark.One_1'(s('type.benchmark.times'(X12),X13))) != s('type.benchmark.digit'(X12),'const.benchmark.One_1'(s('type.benchmark.times'(X12),X15))) | s('type.benchmark.times'(X12),X15) = s('type.benchmark.times'(X12),'const.ind_type.FCONS_3'(s('type.benchmark.times'(X12),X13),s(fun('type.num.num','type.benchmark.times'(X12)),X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f281,f276])).
fof(f1483,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.benchmark.digit'(X8),'const.benchmark.One_1'(s('type.benchmark.times'(X8),X11))) != s('type.benchmark.digit'(X8),'const.benchmark.One_1'(s('type.benchmark.times'(X8),'const.pair.FST_1'(s('type.pair.prod'('type.benchmark.times'(X8),X9),X10))))) | s('type.benchmark.times'(X8),X11) = s('type.benchmark.times'(X8),happ(s(fun('type.pair.prod'('type.benchmark.times'(X8),X9),'type.benchmark.times'(X8)),'const.pair.FST_0'),s('type.pair.prod'('type.benchmark.times'(X8),X9),X10)))) )),
  inference(superposition,[],[f281,f309])).
fof(f1482,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s('type.benchmark.digit'(X3),'const.benchmark.One_1'(s('type.benchmark.times'(X3),X5))) != s('type.benchmark.digit'(X3),'const.benchmark.One_1'(s('type.benchmark.times'(X3),X7))) | s('type.benchmark.times'(X3),X7) = s('type.benchmark.times'(X3),happ(s(fun(X4,'type.benchmark.times'(X3)),'_dst_x0_1'(s('type.benchmark.times'(X3),X5))),s(X4,X6)))) )),
  inference(superposition,[],[f281,f312])).
fof(f1481,plain,(
  ( ! [X2,X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) != s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))) | s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.times'(X0)),'_i_0_0'),s('type.benchmark.times'(X0),X1)))) )),
  inference(superposition,[],[f281,f267])).
fof(f1450,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f344,f1407,f1212])).
fof(f1451,plain,(
  ( ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f1407,f1262])).
fof(f1407,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f231,f1262])).
fof(f1406,plain,(
  ( ! [X97,X95,X98,X96] : (p(s(bool,'|dollar|exists_1'(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97)))))) | ~p(s(bool,happ(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97))),s(X95,sK4(X98,happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97)),X95)))))) )),
  inference(forward_demodulation,[],[f1388,f309])).
fof(f1388,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X95,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X95,bool),X96),X97))),s(X95,sK4(X98,happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97)),X95))))) | p(s(bool,'|dollar|exists_1'(s(fun(X95,bool),happ(s(fun('type.pair.prod'(fun(X95,bool),X96),fun(X95,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X95,bool),X96),X97))))))) )),
  inference(superposition,[],[f280,f309])).
fof(f1369,plain,(
  ( ! [X45,X44] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),sK4(X45,'const.pair.FST_0','type.pair.prod'(bool,X44)))))) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X44),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f280,f309])).
fof(f1272,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK4(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1256,f231])).
fof(f1256,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK4(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f280,f359])).
fof(f1254,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f280,f238])).
fof(f1252,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f280,f359])).
fof(f1250,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f280,f359])).
fof(f1267,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),X20)))) | ~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17)))))) )),
  inference(forward_demodulation,[],[f1249,f313])).
fof(f1249,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,happ(s(fun(X17,bool),X20),s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))))) | p(s(bool,'|dollar|exists_1'(s(fun(X17,bool),'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20))))))))) )),
  inference(superposition,[],[f280,f313])).
fof(f1266,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X14)))) | ~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12)))))) )),
  inference(forward_demodulation,[],[f1248,f311])).
fof(f1248,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~p(s(bool,happ(s(fun(X12,bool),X14),s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))))) | p(s(bool,'|dollar|exists_1'(s(fun(X12,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15))))))))) )),
  inference(superposition,[],[f280,f311])).
fof(f1265,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X9)))) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8)))))) )),
  inference(forward_demodulation,[],[f1247,f276])).
fof(f1247,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))))) | p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f280,f276])).
fof(f1264,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),X5)))) | ~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3)))))) )),
  inference(forward_demodulation,[],[f1246,f312])).
fof(f1246,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X5),s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X3,bool),happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6))))))) )),
  inference(superposition,[],[f280,f312])).
fof(f1263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0)))))) )),
  inference(forward_demodulation,[],[f1245,f267])).
fof(f1245,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1))))))) )),
  inference(superposition,[],[f280,f267])).
fof(f1233,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1217,f231])).
fof(f1217,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1212,f359])).
fof(f1133,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f867,f231])).
fof(f867,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f362,f359])).
fof(f1132,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f866,f231])).
fof(f866,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f298,f359])).
fof(f849,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,'const.bool.DATATYPE_1'(s(bool,X28))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f266,f359])).
fof(f756,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,X31) | s(bool,X32) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X29,bool),'const.pair.,_2'(s(X29,X30),s(bool,X31))))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f359,f313])).
fof(f755,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),'const.pair.,_2'(s(bool,X26),s(X25,X27))))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f359,f311])).
fof(f754,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X24) = s(bool,'const.ind_type.FCONS_3'(s(bool,X22),s(fun('type.num.num',bool),X23),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f359,f276])).
fof(f753,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,happ(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))),s(X18,X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f359,f312])).
fof(f752,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f359,f267])).
fof(f742,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f359,f359])).
fof(f727,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X37),s(bool,X38)))))) )),
  inference(superposition,[],[f359,f313])).
fof(f726,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X33),s(X32,X34)))))) )),
  inference(superposition,[],[f359,f311])).
fof(f725,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X29),s(fun('type.num.num',bool),X30),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f359,f276])).
fof(f724,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26))),s(X25,X27)))) )),
  inference(superposition,[],[f359,f312])).
fof(f723,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X23)))) )),
  inference(superposition,[],[f359,f267])).
fof(f662,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,X29) != s(bool,'const.bool.DATATYPE_1'(s(X30,X31))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X29)))))) )),
  inference(superposition,[],[f643,f313])).
fof(f661,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,X23) != s(bool,'const.bool.DATATYPE_1'(s(X25,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X23),s(X22,X24)))))) )),
  inference(superposition,[],[f643,f311])).
fof(f660,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X18) != s(bool,'const.bool.DATATYPE_1'(s(X20,X21))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X18),s(fun('type.num.num',bool),X19),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f643,f276])).
fof(f659,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,X14) != s(bool,'const.bool.DATATYPE_1'(s(X16,X17))) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),'_dst_x0_1'(s(bool,X14))),s(X13,X15)))) )),
  inference(superposition,[],[f643,f312])).
fof(f658,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,'const.bool.DATATYPE_1'(s(X11,X12))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X10)))) )),
  inference(superposition,[],[f643,f267])).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f630,f231])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,'const.bool.DATATYPE_1'(s(X0,X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f362,f266])).
fof(f333,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ! [X6,X7,X8,X9] : s(fun(X3,fun(X4,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X5,X0)),X6),s(X1,happ(s(fun(X2,X1),X7),s(X2,X9))),s(fun(X3,fun(X4,X5)),X8))) = s(fun(X3,fun(X4,X0)),happ(s(fun(X2,fun(X3,fun(X4,X0))),'_dst_x000x100x011_3'(s(fun(X1,fun(X5,X0)),X6),s(fun(X2,X1),X7),s(fun(X3,fun(X4,X5)),X8))),s(X2,X9)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X12,X41,X14,X13,X39,X42] : ! [X15,X16,X40,X17] : s(fun(X13,fun(X39,X12)),'_dst_x00x00x11_3'(s(fun(X41,fun(X42,X12)),X15),s(X41,happ(s(fun(X14,X41),X16),s(X14,X17))),s(fun(X13,fun(X39,X42)),X40))) = s(fun(X13,fun(X39,X12)),happ(s(fun(X14,fun(X13,fun(X39,X12))),'_dst_x000x100x011_3'(s(fun(X41,fun(X42,X12)),X15),s(fun(X14,X41),X16),s(fun(X13,fun(X39,X42)),X40))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x000x100x011_ax')).
fof(f332,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x01x11_3'(s(fun(X1,fun(X3,X0)),X5),s(fun(X2,X1),X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X12,X39,X13,X41,X14] : ! [X15,X16,X40,X17] : s(fun(X13,X12),'_dst_x0x1x1_3'(s(fun(X39,fun(X41,X12)),X15),s(fun(X13,X39),X16),s(fun(X13,X41),happ(s(fun(X14,fun(X13,X41)),X40),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x01x11_3'(s(fun(X39,fun(X41,X12)),X15),s(fun(X13,X39),X16),s(fun(X14,fun(X13,X41)),X40))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x00x01x11_ax')).
fof(f331,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X1,X0),'_dst_x0x1x0_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X6),s(X3,X8))),s(X4,X7))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11x00_3'(s(fun(X2,fun(X4,X0)),X5),s(fun(X3,fun(X1,X2)),X6),s(X4,X7))),s(X3,X8)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12,X13,X39,X14,X41] : ! [X15,X16,X40,X17] : s(fun(X13,X12),'_dst_x0x1x0_3'(s(fun(X39,fun(X41,X12)),X15),s(fun(X13,X39),happ(s(fun(X14,fun(X13,X39)),X16),s(X14,X17))),s(X41,X40))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x11x00_3'(s(fun(X39,fun(X41,X12)),X15),s(fun(X14,fun(X13,X39)),X16),s(X41,X40))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x00x11x00_ax')).
fof(f330,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X2,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X2,X3),happ(s(fun(X4,fun(X2,X3)),X7),s(X4,X8))))) = s(fun(X2,X0),happ(s(fun(X4,fun(X2,X0)),'_dst_x00x00x11_3'(s(fun(X1,fun(X3,X0)),X5),s(X1,X6),s(fun(X4,fun(X2,X3)),X7))),s(X4,X8)))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X12,X39,X13,X41,X14] : ! [X15,X16,X40,X17] : s(fun(X13,X12),'_dst_x0x0x1_3'(s(fun(X39,fun(X41,X12)),X15),s(X39,X16),s(fun(X13,X41),happ(s(fun(X14,fun(X13,X41)),X40),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x00x11_3'(s(fun(X39,fun(X41,X12)),X15),s(X39,X16),s(fun(X14,fun(X13,X41)),X40))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x00x00x11_ax')).
fof(f329,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4] : ! [X5,X6,X7,X8] : s(fun(X3,X0),'_dst_x0x0x1_3'(s(fun(X1,fun(X4,X0)),X5),s(X1,happ(s(fun(X2,X1),X6),s(X2,X8))),s(fun(X3,X4),X7))) = s(fun(X3,X0),happ(s(fun(X2,fun(X3,X0)),'_dst_x00x10x01_3'(s(fun(X1,fun(X4,X0)),X5),s(fun(X2,X1),X6),s(fun(X3,X4),X7))),s(X2,X8)))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12,X39,X14,X13,X41] : ! [X15,X16,X40,X17] : s(fun(X13,X12),'_dst_x0x0x1_3'(s(fun(X39,fun(X41,X12)),X15),s(X39,happ(s(fun(X14,X39),X16),s(X14,X17))),s(fun(X13,X41),X40))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x10x01_3'(s(fun(X39,fun(X41,X12)),X15),s(fun(X14,X39),X16),s(fun(X13,X41),X40))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x00x10x01_ax')).
fof(f328,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X12,X13,X39,X14] : ! [X15,X16,X40,X17] : s(X12,happ(s(fun(X14,X12),'_dst_x0x1x1_3'(s(fun(X13,fun(X39,X12)),X15),s(fun(X14,X13),X16),s(fun(X14,X39),X40))),s(X14,X17))) = s(X12,happ(s(fun(X39,X12),happ(s(fun(X13,fun(X39,X12)),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))),s(X39,happ(s(fun(X14,X39),X40),s(X14,X17)))))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x0x1x1_ax')).
fof(f325,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X1,X0),'_dst_x0x1_2'(s(fun(X2,X0),X4),s(fun(X1,X2),happ(s(fun(X3,fun(X1,X2)),X5),s(X3,X6))))) = s(fun(X1,X0),happ(s(fun(X3,fun(X1,X0)),'_dst_x00x11_2'(s(fun(X2,X0),X4),s(fun(X3,fun(X1,X2)),X5))),s(X3,X6)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X12,X13,X39,X14] : ! [X15,X16,X17] : s(fun(X13,X12),'_dst_x0x1_2'(s(fun(X39,X12),X15),s(fun(X13,X39),happ(s(fun(X14,fun(X13,X39)),X16),s(X14,X17))))) = s(fun(X13,X12),happ(s(fun(X14,fun(X13,X12)),'_dst_x00x11_2'(s(fun(X39,X12),X15),s(fun(X14,fun(X13,X39)),X16))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x00x11_ax')).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(X0,fun('type.num.num',X0)),'_dst_x1x0_2'(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(fun('type.num.num',X0),X2))),s(X0,X1))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f324,f286])).
fof(f340,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(X1,fun(fun('type.num.num',X0),X0)),'_dst_x1x0_2'(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(X1,X5))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(backward_demodulation,[],[f324,f310])).
fof(f323,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X13,X14] : ! [X15,X16,X17] : s(X12,happ(s(fun(X13,X12),X15),s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))) = s(X12,happ(s(fun(X14,X12),'_dst_x0x1_2'(s(fun(X13,X12),X15),s(fun(X14,X13),X16))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dst_x0x1_ax')).
fof(f322,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(fun(fun(X1,X0),X0),'_dstsel_0_x0_1'(s(X1,happ(s(fun(X2,X1),X3),s(X2,X4))))) = s(fun(fun(X1,X0),X0),happ(s(fun(X2,fun(fun(X1,X0),X0)),'_dstsel_1_x10_1'(s(fun(X2,X1),X3))),s(X2,X4)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X12,X13,X14] : ! [X16,X17] : s(fun(fun(X13,X12),X12),'_dstsel_0_x0_1'(s(X13,happ(s(fun(X14,X13),X16),s(X14,X17))))) = s(fun(fun(X13,X12),X12),happ(s(fun(X14,fun(fun(X13,X12),X12)),'_dstsel_1_x10_1'(s(fun(X14,X13),X16))),s(X14,X17)))),
  file('thm.benchmark.digit_distinct|split|2','_dstsel_1_x10_ax')).
fof(f317,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'),s(X1,X2),s(fun('type.benchmark.times'(X0),X1),X3),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),X1)),X4)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3,X4] : s(X1,X2) = s(X1,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'),s(X1,X2),s(fun('type.benchmark.times'(X0),X1),X3),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),X1)),X4)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X1,X2) = s(X1,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X0),'const.benchmark.Zero_0'),s(X1,X2),s(fun('type.benchmark.times'(X0),X1),X3),s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),X1)),X4)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X26] : ! [X35,X21,X36] : s(X26,X35) = s(X26,'const.benchmark.digit_CASE_4'(s('type.benchmark.digit'(X6),'const.benchmark.Zero_0'),s(X26,X35),s(fun('type.benchmark.times'(X6),X26),X21),s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),X26)),X36)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_case_def|split|0')).
fof(f316,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK9(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK9(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.benchmark.digit_distinct|split|2',aHL_EXT)).
fof(f310,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK8(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X26,X6] : ! [X27] : ? [X21] : ! [X28,X29,X30] : s(X26,happ(s(fun('type.ind_type.recspace'(X6),X26),X21),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X28),s(X6,X29),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30))))) = s(X26,happ(s(fun(fun('type.num.num',X26),X26),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X26),X26)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X26),X26))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X26),X26)))),X27),s('type.num.num',X28))),s(X6,X29))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30))),s(fun('type.num.num',X26),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X26),X21),s(fun('type.num.num','type.ind_type.recspace'(X6)),X30)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.ind_type.CONSTR_REC')).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : ! [X2] : s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,happ(s(fun('type.pair.prod'(X0,X1),X0),'const.pair.FST_0'),s('type.pair.prod'(X0,X1),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X26] : ! [X18] : s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X26),X18))) = s(X6,happ(s(fun('type.pair.prod'(X6,X26),X6),'const.pair.FST_0'),s('type.pair.prod'(X6,X26),X18)))),
  file('thm.benchmark.digit_distinct|split|2','ARITY_0_1_const.pair.FST')).
fof(f293,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4) | s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))) )),
  inference(cnf_transformation,[],[f215])).
fof(f288,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4))) & (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4) | s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4))) & ((s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4)) | s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) <=> (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) <=> (s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X3) & s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X8,X7,X32,X33] : (s('type.benchmark.times'(X6),'const.benchmark.Twice_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7))) = s('type.benchmark.times'(X6),'const.benchmark.Twice_2'(s('type.benchmark.times'(X6),X32),s('type.benchmark.times'(X6),X33))) <=> (s('type.benchmark.times'(X6),X8) = s('type.benchmark.times'(X6),X32) & s('type.benchmark.times'(X6),X7) = s('type.benchmark.times'(X6),X33)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.times_|49||49||split|1')).
fof(f289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) != s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4))) | s('type.benchmark.times'(X0),X2) = s('type.benchmark.times'(X0),X4)) )),
  inference(cnf_transformation,[],[f213])).
fof(f290,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X3) | s('type.benchmark.times'(X0),X2) != s('type.benchmark.times'(X0),X4) | s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.times'(X0),'const.benchmark.Twice_2'(s('type.benchmark.times'(X0),X3),s('type.benchmark.times'(X0),X4)))) )),
  inference(cnf_transformation,[],[f213])).
fof(f287,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X3)))))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X3)))))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X3)))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X21,X8,X7] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.benchmark.times_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.times'(X6),X8))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.times'(X6),X7)))))))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_size_def|split|2')).
fof(f286,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X18,X19,X20] : s(X6,'const.ind_type.FCONS_3'(s(X6,X18),s(fun('type.num.num',X6),X19),s('type.num.num',X20))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X18))),s(fun('type.num.num',X6),X19))),s('type.num.num',X20)))),
  file('thm.benchmark.digit_distinct|split|2','ARITY_0_3_const.ind_type.FCONS')).
fof(f285,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X18,X19,X20] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X18),s(X6,X19),s(fun('type.num.num','type.ind_type.recspace'(X6)),X20))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X18))),s(X6,X19))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X20)))),
  file('thm.benchmark.digit_distinct|split|2','ARITY_1_3_const.ind_type.CONSTR')).
fof(f284,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9,X21,X22] : s(X6,happ(s(fun('type.num.num',X6),X21),s('type.num.num',X22))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X9),s(fun('type.num.num',X6),X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.ind_type.FCONS|split|1')).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') != s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') = s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X2),s('type.benchmark.times'(X0),X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X8] : ~s('type.benchmark.digit'(X6),'const.benchmark.Zero_0') = s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X8),s('type.benchmark.times'(X6),X7)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_distinct|split|1')).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) != s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))) | s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X2)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) != s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))) | s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X2)) & (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X2) | s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2)))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))) <=> s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X2))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : (s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))) <=> s('type.benchmark.times'(X0),X1) = s('type.benchmark.times'(X0),X2))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X9,X34] : (s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9))) = s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X34))) <=> s('type.benchmark.times'(X6),X9) = s('type.benchmark.times'(X6),X34))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_|49||49||split|0')).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (s('type.benchmark.times'(X0),X1) != s('type.benchmark.times'(X0),X2) | s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2)))) )),
  inference(cnf_transformation,[],[f211])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0))),'const.benchmark.Two_0'),s('type.benchmark.times'(X0),X1))),s('type.benchmark.times'(X0),X2)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0))),'const.benchmark.Two_0'),s('type.benchmark.times'(X0),X1))),s('type.benchmark.times'(X0),X2)))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s('type.benchmark.digit'(X0),'const.benchmark.Two_2'(s('type.benchmark.times'(X0),X1),s('type.benchmark.times'(X0),X2))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),happ(s(fun('type.benchmark.times'(X0),fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0))),'const.benchmark.Two_0'),s('type.benchmark.times'(X0),X1))),s('type.benchmark.times'(X0),X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X18,X19] : s('type.benchmark.digit'(X6),'const.benchmark.Two_2'(s('type.benchmark.times'(X6),X18),s('type.benchmark.times'(X6),X19))) = s('type.benchmark.digit'(X6),happ(s(fun('type.benchmark.times'(X6),'type.benchmark.digit'(X6)),happ(s(fun('type.benchmark.times'(X6),fun('type.benchmark.times'(X6),'type.benchmark.digit'(X6))),'const.benchmark.Two_0'),s('type.benchmark.times'(X6),X18))),s('type.benchmark.times'(X6),X19)))),
  file('thm.benchmark.digit_distinct|split|2','ARITY_0_2_const.benchmark.Two')).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.times'(X0),X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X21,X9] : s('type.num.num','const.benchmark.digit_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.benchmark.times_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.times'(X6),X9)))))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_size_def|split|1')).
fof(f275,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') != s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') != s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : ~s('type.benchmark.digit'(X0),'const.benchmark.Zero_0') = s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X9] : ~s('type.benchmark.digit'(X6),'const.benchmark.Zero_0') = s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X9)))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_distinct|split|0')).
fof(f269,plain,(
  ( ! [X0,X1] : (s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s('type.benchmark.digit'(X0),'const.benchmark.One_1'(s('type.benchmark.times'(X0),X1))) = s('type.benchmark.digit'(X0),happ(s(fun('type.benchmark.times'(X0),'type.benchmark.digit'(X0)),'const.benchmark.One_0'),s('type.benchmark.times'(X0),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X18] : s('type.benchmark.digit'(X6),'const.benchmark.One_1'(s('type.benchmark.times'(X6),X18))) = s('type.benchmark.digit'(X6),happ(s(fun('type.benchmark.times'(X6),'type.benchmark.digit'(X6)),'const.benchmark.One_0'),s('type.benchmark.times'(X6),X18)))),
  file('thm.benchmark.digit_distinct|split|2','ARITY_0_1_const.benchmark.One')).
fof(f268,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0')))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0')))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.benchmark.digit_size_2'(s(fun(X0,'type.num.num'),X1),s('type.benchmark.digit'(X0),'const.benchmark.Zero_0')))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X21] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.benchmark.digit_size_2'(s(fun(X6,'type.num.num'),X21),s('type.benchmark.digit'(X6),'const.benchmark.Zero_0')))),
  file('thm.benchmark.digit_distinct|split|2','thm.benchmark.digit_size_def|split|0')).
fof(f11773,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763,f11765,f11318,f11274,f11268])).
fof(f11772,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11269,f231])).
fof(f11269,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f298])).
fof(f11771,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11271,f11259])).
fof(f11271,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f298])).
fof(f11770,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763,f11765,f11318,f11274])).
fof(f11769,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11275,f231])).
fof(f11275,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f298])).
fof(f11768,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11277,f11259])).
fof(f11277,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f298])).
fof(f11767,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763,f11765,f11318])).
fof(f11766,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763,f11765])).
fof(f11764,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756,f11763])).
fof(f11762,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11665,f231])).
fof(f11665,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f6997])).
fof(f11761,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11664,f231])).
fof(f11664,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f6998])).
fof(f11760,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11663,f231])).
fof(f11663,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f6894])).
fof(f11759,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11662,f231])).
fof(f11662,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f6893])).
fof(f11758,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11658,f231])).
fof(f11658,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f7213])).
fof(f11757,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754,f11756])).
fof(f11755,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752,f11754])).
fof(f11753,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747,f11752])).
fof(f11751,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11544,f11259])).
fof(f11544,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11322,f8683])).
fof(f11750,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11538,f11259])).
fof(f11538,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11322,f8079])).
fof(f11749,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11511,f11322])).
fof(f11511,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl20),
  inference(backward_demodulation,[],[f11322,f7221])).
fof(f11748,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743,f11747])).
fof(f11744,plain,(
  $false | $spl20),
  inference(global_subsumption,[],[f230,f231,f238,f242,f241,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f418,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1233,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1593,f1700,f1705,f1708,f1807,f1810,f2154,f2277,f2340,f2425,f2891,f2906,f3068,f3069,f3126,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f1871,f2968,f2749,f3294,f5370,f5369,f5353,f5500,f5499,f5501,f5547,f5595,f5594,f5578,f5735,f5734,f6191,f6351,f6334,f6333,f6541,f6542,f6587,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6744,f6892,f7055,f7077,f7076,f7403,f7437,f7531,f7564,f7563,f7556,f7545,f7618,f7617,f7710,f7695,f8733,f8736,f339,f1404,f1356,f1354,f3148,f3125,f3123,f3114,f3112,f3109,f3108,f3129,f1425,f1262,f3080,f1389,f6186,f6187,f6610,f6609,f3214,f3210,f3142,f3084,f3083,f3082,f3081,f3079,f3078,f691,f690,f689,f688,f687,f279,f7589,f7588,f7434,f7032,f7508,f7590,f7004,f7480,f7467,f8305,f8239,f6224,f6220,f6120,f6119,f7457,f7406,f6893,f6894,f6998,f6997,f7714,f7696,f7697,f7213,f3845,f3846,f7544,f6743,f6845,f6844,f5674,f5673,f5449,f5448,f3343,f2795,f8768,f6846,f5675,f5450,f1905,f3324,f3116,f3338,f3212,f6951,f6971,f7355,f7375,f8724,f8743,f3211,f3115,f7367,f7356,f7324,f7313,f6920,f6909,f8725,f3325,f3117,f3063,f3047,f308,f307,f306,f8309,f8246,f8180,f8163,f8121,f2778,f1894,f305,f304,f303,f6908,f6928,f7312,f7332,f8742,f7670,f7669,f7902,f3434,f7474,f7919,f7700,f7921,f7702,f7440,f7353,f7310,f6949,f6906,f321,f320,f7924,f7923,f7922,f7920,f7918,f7917,f301,f7927,f7667,f7705,f7704,f7703,f7701,f7699,f7698,f300,f7715,f3865,f3794,f3452,f3288,f3223,f3209,f302,f299,f8738,f8727,f8683,f8311,f8253,f8244,f8187,f8178,f8079,f7929,f7717,f7706,f7677,f7675,f7668,f7656,f7598,f7538,f7512,f7484,f7371,f7369,f7363,f7361,f7354,f7328,f7326,f7320,f7318,f7311,f7265,f7260,f7221,f7220,f7210,f7057,f7036,f7008,f6957,f6950,f6924,f6922,f6916,f6914,f6907,f6895,f6889,f6600,f6599,f6598,f6549,f6547,f6545,f6540,f6437,f6436,f6435,f6350,f6206,f6177,f6176,f6175,f6162,f6127,f6125,f6118,f5859,f5716,f5481,f3867,f3454,f3331,f3327,f3290,f3139,f3138,f3137,f3132,f3127,f3124,f3111,f3110,f3076,f3074,f3072,f3067,f3066,f3050,f2941,f2940,f2939,f2938,f2937,f2936,f2870,f2869,f2868,f2867,f2866,f2865,f2842,f2841,f2840,f2839,f2838,f2837,f2784,f2780,f2734,f2733,f2722,f2339,f2185,f2184,f2183,f2182,f2181,f2125,f2124,f2123,f2122,f2121,f2120,f1896,f1893,f1811,f1687,f1569,f1568,f1567,f1565,f1564,f1548,f1427,f1383,f1358,f1239,f1238,f1213,f1212,f1210,f1209,f1208,f1207,f588,f538,f492,f450,f392,f337,f246,f245,f9628,f9627,f9626,f9625,f9624,f9623,f9622,f9621,f9619,f9653,f11216,f11218,f11221,f11222,f11259,f11322,f11743])).
fof(f11740,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11497,f11259])).
fof(f11497,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl20),
  inference(backward_demodulation,[],[f11322,f6895])).
fof(f11737,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11376,f231])).
fof(f11376,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f11322,f337])).
fof(f11736,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f11735,f11259])).
fof(f11735,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f11374,f11322])).
fof(f11374,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f11322,f245])).
fof(f11323,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f344,f344,f11259,f1212])).
fof(f11319,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f267,f344,f11259,f362])).
fof(f11278,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f11259,f344,f11259,f298])).
fof(f11276,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f231,f344,f11259,f298])).
fof(f11272,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f11259,f344,f11259,f298])).
fof(f11270,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f231,f344,f11259,f298])).
fof(f11266,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f267,f11259,f245])).
fof(f11264,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f344,f11259,f242])).
fof(f11262,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f337,f11259])).
fof(f9232,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216,f9218,f9220,f9229,f9231])).
fof(f9231,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl24),
  inference(subsumption_resolution,[],[f9080,f231])).
fof(f9080,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f7032])).
fof(f9229,plain,(
  ( ! [X19,X20,X18,X16] : (s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl24),
  inference(subsumption_resolution,[],[f9077,f231])).
fof(f9077,plain,(
  ( ! [X19,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,X20) = s(X16,happ(s(fun('type.num.num',X16),X18),s('type.num.num',X19)))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f7004])).
fof(f9220,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9109,f231])).
fof(f9109,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(condensation,[],[f9041])).
fof(f9041,plain,(
  ( ! [X80,X81] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f8768])).
fof(f9218,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl24),
  inference(subsumption_resolution,[],[f9030,f231])).
fof(f9030,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f6971])).
fof(f9216,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X58)) | s(X56,X59) = s(X56,sK4(X59,X57,X56))) ) | $spl24),
  inference(subsumption_resolution,[],[f9026,f231])).
fof(f9026,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | s(X56,X59) = s(X56,sK4(X59,X57,X56))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f8743])).
fof(f9214,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9124,f231])).
fof(f9124,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(condensation,[],[f9011])).
fof(f9011,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f306])).
fof(f9212,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9131,f231])).
fof(f9131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f9001])).
fof(f9001,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f303])).
fof(f9210,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl24),
  inference(subsumption_resolution,[],[f8999,f231])).
fof(f8999,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(X97,X100) = s(X97,sK4(X100,X98,X97))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f6928])).
fof(f9208,plain,(
  ( ! [X55] : (p(s(bool,X55))) ) | $spl24),
  inference(subsumption_resolution,[],[f8996,f231])).
fof(f8996,plain,(
  ( ! [X55] : (~p(s(bool,'T_0')) | p(s(bool,X55))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f8742])).
fof(f9202,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f8979,f231])).
fof(f8979,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X2)) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f320])).
fof(f9180,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9179,f231])).
fof(f9179,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(forward_demodulation,[],[f8948,f5594])).
fof(f8948,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,X0))),s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f8244])).
fof(f9177,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f9176,f231])).
fof(f9176,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | $spl24),
  inference(forward_demodulation,[],[f8946,f5369])).
fof(f8946,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f8178])).
fof(f9158,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f9157,f231])).
fof(f9157,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f9156])).
fof(f9156,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(condensation,[],[f9155])).
fof(f9155,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(forward_demodulation,[],[f8910,f8600])).
fof(f8910,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,X62))),s(bool,'T_0')))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f6957])).
fof(f9153,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f9152,f231])).
fof(f9152,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | $spl24),
  inference(condensation,[],[f9151])).
fof(f9151,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(condensation,[],[f9150])).
fof(f9150,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(forward_demodulation,[],[f8905,f8600])).
fof(f8905,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,X62)))) | s(bool,X61) = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f6914])).
fof(f9138,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f8776,f8600])).
fof(f8776,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f8600,f245])).
fof(f9230,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216,f9218,f9220,f9229])).
fof(f9228,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9067,f231])).
fof(f9067,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f6893])).
fof(f9227,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9066,f231])).
fof(f9066,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f6894])).
fof(f9226,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9065,f231])).
fof(f9065,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f6998])).
fof(f9225,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9064,f231])).
fof(f9064,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f6997])).
fof(f9224,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9060,f231])).
fof(f9060,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f7213])).
fof(f9223,plain,(
  $false | ($spl16 | $spl24)),
  inference(global_subsumption,[],[f9222,f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216,f9218,f9220])).
fof(f9222,plain,(
  ( ! [X53] : (p(s(bool,X53))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9051,f231])).
fof(f9051,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | p(s(bool,X53))) ) | ($spl16 | $spl24)),
  inference(backward_demodulation,[],[f8600,f8740])).
fof(f8740,plain,(
  ( ! [X53] : (~p(s(bool,happ(s(fun(fun(bool,bool),bool),'_dst_x0x1x0_3'(s(fun(bool,fun(bool,bool)),'==>_0'),s(fun(fun(bool,bool),bool),'?_1'),s(bool,X53))),s(fun(bool,bool),'_i_0_0')))) | p(s(bool,X53))) ) | $spl16),
  inference(forward_demodulation,[],[f8739,f327])).
fof(f8739,plain,(
  ( ! [X53] : (p(s(bool,X53)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X53))))) ) | $spl16),
  inference(subsumption_resolution,[],[f8703,f231])).
fof(f8703,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))),s(bool,X53))))) ) | $spl16),
  inference(superposition,[],[f306,f4200])).
fof(f9221,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216,f9218,f9220])).
fof(f9219,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216,f9218])).
fof(f9217,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214,f9216])).
fof(f9215,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212,f9214])).
fof(f9213,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210,f9212])).
fof(f9211,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208,f9210])).
fof(f9209,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202,f9208])).
fof(f9203,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180,f9202])).
fof(f9183,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9182,f231])).
fof(f9182,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8951,f6761])).
fof(f8951,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8600,f8683])).
fof(f9181,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177,f9180])).
fof(f9178,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158,f9177])).
fof(f9175,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9174,f231])).
fof(f9174,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8945,f5369])).
fof(f8945,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|bslash||slash|_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8600,f8079])).
fof(f9173,plain,(
  $false | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f8941,f8600])).
fof(f8941,plain,(
  s(bool,'T_0') != s(bool,sK6('_i_0_0',bool)) | ($spl19 | $spl24)),
  inference(backward_demodulation,[],[f8600,f7711])).
fof(f7711,plain,(
  s(bool,'F_0') != s(bool,sK6('_i_0_0',bool)) | $spl19),
  inference(cnf_transformation,[],[f7711_D])).
fof(f7711_D,plain,(
  s(bool,'F_0') != s(bool,sK6('_i_0_0',bool)) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f9160,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f8917,f8600])).
fof(f8917,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0'))) | $spl24),
  inference(backward_demodulation,[],[f8600,f7221])).
fof(f9159,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153,f9158])).
fof(f9154,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138,f9153])).
fof(f9149,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f9148,f231])).
fof(f9148,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(forward_demodulation,[],[f8903,f8600])).
fof(f8903,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,'T_0')))) | $spl24),
  inference(backward_demodulation,[],[f8600,f6895])).
fof(f9140,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f8778,f231])).
fof(f8778,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f8600,f337])).
fof(f9139,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f298,f295,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1369,f1406,f1407,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f2120,f2121,f2123,f2124,f2125,f2154,f2181,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2837,f2838,f2840,f2841,f2842,f2891,f2865,f2866,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f3067,f3068,f3069,f3072,f3074,f3076,f3124,f3126,f3137,f3138,f3139,f3141,f3146,f3213,f3271,f3368,f3370,f3367,f3861,f3844,f299,f302,f3209,f3223,f3288,f3290,f3452,f3454,f3794,f3865,f3867,f300,f301,f320,f321,f3434,f1871,f2968,f303,f304,f305,f1893,f1894,f2778,f2780,f306,f307,f308,f3047,f3050,f3063,f3066,f3117,f3325,f3327,f3115,f3211,f2749,f3294,f3212,f3338,f3116,f3324,f1905,f1896,f2795,f2784,f3343,f3331,f3845,f3846,f279,f687,f688,f689,f690,f691,f3078,f3079,f3081,f3082,f3083,f3084,f3142,f3210,f3214,f3080,f1389,f1262,f1425,f1427,f3132,f3129,f3108,f3109,f3110,f3111,f3112,f3114,f3123,f3125,f3127,f3148,f1354,f1356,f1358,f1383,f1404,f1687,f2122,f2182,f2839,f2867,f2938,f339,f5370,f5369,f5448,f5449,f5450,f5353,f5481,f5500,f5499,f5501,f5547,f5595,f5594,f5673,f5674,f5675,f5578,f5716,f5735,f5734,f5859,f6118,f6119,f6120,f6206,f6125,f6127,f6220,f6162,f6224,f6175,f6176,f6177,f6186,f6187,f6191,f6351,f6350,f6435,f6436,f6437,f6334,f6333,f6540,f6541,f6542,f6545,f6547,f6549,f6587,f6598,f6599,f6600,f6609,f6610,f6615,f6687,f6686,f6688,f6690,f6740,f6762,f6761,f6844,f6845,f6846,f6744,f6743,f6889,f6997,f6998,f6895,f6894,f6893,f6892,f7004,f6906,f6907,f6908,f6909,f7008,f6914,f6916,f6920,f6922,f6924,f6928,f7032,f6949,f6950,f6951,f7036,f6957,f6971,f7055,f7057,f7077,f7076,f7210,f7221,f7220,f7260,f7213,f7265,f7403,f7406,f7434,f7437,f7440,f7457,f7467,f7474,f7480,f7310,f7311,f7312,f7313,f7484,f7318,f7320,f7324,f7326,f7328,f7332,f7508,f7353,f7354,f7355,f7356,f7512,f7361,f7363,f7367,f7369,f7371,f7375,f7531,f7538,f7564,f7563,f7556,f7588,f7589,f7590,f7545,f7544,f7598,f7618,f7617,f7656,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7667,f7668,f7669,f7670,f7706,f7675,f7677,f7710,f7715,f7717,f7695,f7696,f7697,f7917,f7918,f7919,f7920,f7921,f7922,f7923,f7924,f7902,f7927,f7929,f8079,f8178,f8180,f8187,f8239,f8121,f8244,f8246,f8253,f8305,f8163,f8309,f8311,f8683,f8725,f8727,f8733,f8736,f8738,f8742,f8743,f8768,f8724,f8600,f9138])).
fof(f8734,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) | $spl28),
  inference(cnf_transformation,[],[f8734_D])).
fof(f8734_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8736,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8735])).
fof(f8735,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X41)))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f8698,f362])).
fof(f8698,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X41)))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f306,f359])).
fof(f8733,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X0)))) | p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f8732])).
fof(f8732,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f8697,f362])).
fof(f8697,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,'T_0'))),s(bool,X38)))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f306,f359])).
fof(f8679,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8655,f7711])).
fof(f8655,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f7721,f246])).
fof(f8678,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8656,f7711])).
fof(f8656,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f337,f7721,f295])).
fof(f8677,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8658,f7711])).
fof(f8658,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f337,f7721,f295])).
fof(f8675,plain,(
  $false | ($spl19 | $spl23)),
  inference(subsumption_resolution,[],[f8665,f7711])).
fof(f8665,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl23),
  inference(unit_resulting_resolution,[],[f344,f7721,f418])).
fof(f8601,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,sK6('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X61)) ) | $spl26),
  inference(cnf_transformation,[],[f8601_D])).
fof(f8601_D,plain,(
  ( ! [X61] : (s(bool,X61) != s(bool,sK6('_i_0_0',bool)) | s(bool,'F_0') = s(bool,X61)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8529,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X61)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f8398])).
fof(f8398,plain,(
  ( ! [X61,X60] : (s(bool,X61) != s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X60)) ) | $spl18),
  inference(superposition,[],[f1700,f7796])).
fof(f7796,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(subsumption_resolution,[],[f7728,f344])).
fof(f7728,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK6('_i_0_0',bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f7712,f359])).
fof(f7714,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | $spl20),
  inference(cnf_transformation,[],[f7714_D])).
fof(f7714_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7721,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | $spl23),
  inference(cnf_transformation,[],[f7721_D])).
fof(f7721_D,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f7695,plain,(
  ~p(s(bool,sK6('_i_0_0',bool))) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0'))))),
  inference(superposition,[],[f300,f267])).
fof(f7712,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) | $spl18),
  inference(cnf_transformation,[],[f7712_D])).
fof(f7712_D,plain,(
  s(bool,'F_0') = s(bool,sK6('_i_0_0',bool)) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f7710,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK6('_i_0_0',bool))),
  inference(subsumption_resolution,[],[f7681,f231])).
fof(f7681,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'!_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK6('_i_0_0',bool))),
  inference(superposition,[],[f300,f392])).
fof(f4117,plain,(
  $false | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f3864,f3374])).
fof(f4116,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103,f4107,f4109,f4111,f4113,f4115])).
fof(f4115,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(subsumption_resolution,[],[f3288,f3374])).
fof(f4113,plain,(
  ( ! [X6,X4,X3] : (s(bool,X6) = s(bool,happ(s(fun(fun(X3,bool),bool),'!_1'),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(subsumption_resolution,[],[f3452,f3374])).
fof(f4111,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl12),
  inference(subsumption_resolution,[],[f3865,f3374])).
fof(f4109,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK4(X3,X2,X1)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f1272,f3374])).
fof(f4107,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl12),
  inference(subsumption_resolution,[],[f321,f3374])).
fof(f4103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4102])).
fof(f4102,plain,(
  ( ! [X6,X4] : (s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(subsumption_resolution,[],[f1894,f3374])).
fof(f4100,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4099])).
fof(f4099,plain,(
  ( ! [X4,X2] : (s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f2778,f3374])).
fof(f4095,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(subsumption_resolution,[],[f3047,f3374])).
fof(f4093,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3063,f3374])).
fof(f4091,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4090])).
fof(f4090,plain,(
  ( ! [X4,X2] : (s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl12),
  inference(subsumption_resolution,[],[f3325,f3374])).
fof(f4088,plain,(
  ( ! [X39,X42,X40] : (s(X39,X42) = s(X39,sK4(X42,X40,X39))) ) | $spl12),
  inference(subsumption_resolution,[],[f3211,f3374])).
fof(f4084,plain,(
  ( ! [X43,X46,X44] : (s(X43,X46) = s(X43,sK4(X46,X44,X43))) ) | $spl12),
  inference(subsumption_resolution,[],[f3212,f3374])).
fof(f4082,plain,(
  ( ! [X54,X56,X53] : (s(X53,X56) = s(X53,sK4(X56,X54,X53))) ) | $spl12),
  inference(subsumption_resolution,[],[f3338,f3374])).
fof(f4080,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f279,f3374])).
fof(f4078,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f687,f3374])).
fof(f4076,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl12),
  inference(subsumption_resolution,[],[f688,f3374])).
fof(f4074,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X11) = s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) ) | $spl12),
  inference(subsumption_resolution,[],[f689,f3374])).
fof(f4072,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(X12,X16) = s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) ) | $spl12),
  inference(subsumption_resolution,[],[f690,f3374])).
fof(f4070,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(X17,X21) = s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl12),
  inference(subsumption_resolution,[],[f691,f3374])).
fof(f4060,plain,(
  ( ! [X92,X95,X93] : (s(X92,X95) = s(X92,sK4(X95,X93,X92))) ) | $spl12),
  inference(subsumption_resolution,[],[f3214,f3374])).
fof(f4057,plain,(
  ( ! [X101,X99,X102,X100] : (s(X99,X102) = s(X99,sK4(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl12),
  inference(subsumption_resolution,[],[f1389,f3374])).
fof(f4055,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(condensation,[],[f4054])).
fof(f4054,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl12),
  inference(subsumption_resolution,[],[f1425,f3374])).
fof(f4053,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) ) | $spl12),
  inference(subsumption_resolution,[],[f4052,f3374])).
fof(f4052,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3108,f4051])).
fof(f4051,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100))) | s(bool,X101) = s(bool,sK4(X99,X100,bool))) ) | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050])).
fof(f4050,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) = s(bool,sK4(X99,X100,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3125,f3374])).
fof(f4045,plain,(
  ( ! [X68] : (s(bool,'F_0') = s(bool,X68)) ) | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938])).
fof(f4114,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103,f4107,f4109,f4111,f4113])).
fof(f4112,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103,f4107,f4109,f4111])).
fof(f4110,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103,f4107,f4109])).
fof(f4108,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103,f4107])).
fof(f4106,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1871,f3374])).
fof(f4105,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2968,f3374])).
fof(f4104,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100,f4103])).
fof(f4101,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095,f4100])).
fof(f4098,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f4097,f3374])).
fof(f4097,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl12),
  inference(subsumption_resolution,[],[f307,f3374])).
fof(f4096,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093,f4095])).
fof(f4094,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091,f4093])).
fof(f4092,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088,f4091])).
fof(f4089,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084,f4088])).
fof(f4087,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2749,f3374])).
fof(f4086,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3294,f3374])).
fof(f4085,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082,f4084])).
fof(f4083,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080,f4082])).
fof(f4081,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078,f4080])).
fof(f4079,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076,f4078])).
fof(f4077,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074,f4076])).
fof(f4075,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072,f4074])).
fof(f4073,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070,f4072])).
fof(f4071,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060,f4070])).
fof(f4061,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057,f4060])).
fof(f4058,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055,f4057])).
fof(f4056,plain,(
  $false | $spl12),
  inference(global_subsumption,[],[f230,f231,f337,f238,f242,f241,f246,f245,f266,f267,f268,f269,f275,f276,f277,f278,f280,f279,f282,f281,f283,f284,f285,f286,f287,f290,f289,f288,f293,f292,f291,f339,f298,f295,f300,f299,f302,f301,f305,f304,f303,f308,f307,f306,f309,f310,f311,f312,f313,f315,f316,f317,f318,f319,f321,f320,f322,f323,f324,f340,f341,f325,f326,f327,f328,f329,f330,f331,f332,f333,f344,f359,f362,f392,f418,f450,f492,f538,f588,f643,f658,f659,f660,f661,f662,f687,f688,f689,f690,f691,f693,f723,f724,f725,f726,f727,f742,f752,f753,f754,f755,f756,f849,f1132,f1133,f1207,f1208,f1209,f1210,f1212,f1213,f1233,f1238,f1239,f1263,f1264,f1265,f1266,f1267,f1250,f1252,f1254,f1268,f1272,f1261,f1262,f1354,f1356,f1358,f1369,f1383,f1404,f1406,f1389,f1407,f1425,f1427,f1451,f1450,f1481,f1482,f1483,f1484,f1485,f1486,f1496,f1547,f1548,f1564,f1565,f1687,f1567,f1568,f1569,f1593,f1700,f1705,f1708,f1807,f1810,f1811,f1871,f1893,f1894,f1896,f1905,f2120,f2121,f2122,f2123,f2124,f2125,f2154,f2181,f2182,f2183,f2184,f2185,f2277,f2339,f2340,f2722,f2733,f2734,f2425,f2749,f2778,f2780,f2784,f2795,f2837,f2838,f2839,f2840,f2841,f2842,f2891,f2865,f2866,f2867,f2868,f2869,f2870,f2906,f2936,f2937,f2939,f2940,f2941,f2968,f3047,f3050,f3063,f3066,f3067,f3068,f3069,f3072,f3074,f3076,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3209,f3108,f3109,f3110,f3111,f3112,f3114,f3115,f3116,f3117,f3123,f3124,f3125,f3126,f3127,f3129,f3132,f3137,f3138,f3139,f3141,f3142,f3146,f3148,f3210,f3211,f3212,f3213,f3214,f3223,f3288,f3290,f3271,f3294,f3325,f3327,f3331,f3338,f3343,f3324,f3368,f3370,f3367,f3434,f3452,f3454,f3794,f3861,f3865,f3867,f3844,f3845,f3846,f3374,f2938,f4045,f4050,f4051,f4053,f4055])).
fof(f4044,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f339,f3374])).
fof(f4043,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3975,f337])).
fof(f3975,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3862,f339])).
fof(f4042,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4041,f3995])).
fof(f3995,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(subsumption_resolution,[],[f3890,f337])).
fof(f3890,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f321])).
fof(f4041,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3978,f337])).
fof(f3978,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3952])).
fof(f3952,plain,(
  ( ! [X2,X0] : (p(s(bool,'F_0')) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f1425])).
fof(f4040,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4039,f3995])).
fof(f4039,plain,(
  ( ! [X101,X99,X102,X100] : (s(X99,X102) = s(X99,sK4(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl14),
  inference(subsumption_resolution,[],[f3950,f337])).
fof(f3950,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,'F_0')) | s(X99,X102) = s(X99,sK4(X102,happ(s(fun('type.pair.prod'(fun(X99,bool),X100),fun(X99,bool)),'const.pair.FST_0'),s('type.pair.prod'(fun(X99,bool),X100),X101)),X99))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f1389])).
fof(f4038,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4037,f3995])).
fof(f4037,plain,(
  ( ! [X92,X95,X93] : (s(X92,X95) = s(X92,sK4(X95,X93,X92))) ) | $spl14),
  inference(subsumption_resolution,[],[f3948,f4001])).
fof(f4001,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3902,f337])).
fof(f3902,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f304])).
fof(f3948,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,'F_0')))))) | s(X92,X95) = s(X92,sK4(X95,X93,X92))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3214])).
fof(f4036,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4035,f3995])).
fof(f4035,plain,(
  ( ! [X154,X152,X151,X153] : (s(X151,X154) = s(X151,sK4(X154,sK4(X152,X153,fun(X151,bool)),X151))) ) | $spl14),
  inference(subsumption_resolution,[],[f4034,f3995])).
fof(f4034,plain,(
  ( ! [X154,X152,X151,X153] : (s(X151,X154) = s(X151,sK4(X154,sK4(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) ) | $spl14),
  inference(subsumption_resolution,[],[f3946,f337])).
fof(f3946,plain,(
  ( ! [X154,X152,X151,X153] : (p(s(bool,'F_0')) | s(X151,X154) = s(X151,sK4(X154,sK4(X152,X153,fun(X151,bool)),X151)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X151,bool),bool),X153)))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3142])).
fof(f4033,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4032,f3995])).
fof(f4032,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(X17,X21) = s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl14),
  inference(subsumption_resolution,[],[f3939,f337])).
fof(f3939,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'F_0')) | s(X17,X21) = s(X17,sK4(X21,'const.pair.SND_1'(s('type.pair.prod'(X18,fun(X17,bool)),'const.pair.,_2'(s(X18,X19),s(fun(X17,bool),X20)))),X17))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f691])).
fof(f4031,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4030,f3995])).
fof(f4030,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(X12,X16) = s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) ) | $spl14),
  inference(subsumption_resolution,[],[f3938,f337])).
fof(f3938,plain,(
  ( ! [X14,X12,X15,X13,X16] : (p(s(bool,'F_0')) | s(X12,X16) = s(X12,sK4(X16,'const.pair.FST_1'(s('type.pair.prod'(fun(X12,bool),X13),'const.pair.,_2'(s(fun(X12,bool),X14),s(X13,X15)))),X12))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f690])).
fof(f4029,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4028,f3995])).
fof(f4028,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X11) = s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) ) | $spl14),
  inference(subsumption_resolution,[],[f3937,f337])).
fof(f3937,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | s(X8,X11) = s(X8,sK4(X11,'const.ind_type.FCONS_3'(s(fun(X8,bool),X9),s(fun('type.num.num',fun(X8,bool)),X10),s('type.num.num','const.num.|48|_0')),X8))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f689])).
fof(f4027,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4026,f3995])).
fof(f4026,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl14),
  inference(subsumption_resolution,[],[f3936,f337])).
fof(f3936,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,'F_0')) | s(X3,X7) = s(X3,sK4(X7,happ(s(fun(X4,fun(X3,bool)),'_dst_x0_1'(s(fun(X3,bool),X5))),s(X4,X6)),X3))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f688])).
fof(f4025,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4024,f3995])).
fof(f4024,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3935,f337])).
fof(f3935,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,happ(s(fun(fun(X0,bool),fun(X0,bool)),'_i_0_0'),s(fun(X0,bool),X1)),X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f687])).
fof(f4023,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4022,f3995])).
fof(f4022,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK4(X2,X1,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f3934,f337])).
fof(f3934,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f279])).
fof(f4021,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4020,f3995])).
fof(f4020,plain,(
  ( ! [X54,X56,X53] : (s(X53,X56) = s(X53,sK4(X56,X54,X53))) ) | $spl14),
  inference(subsumption_resolution,[],[f3923,f337])).
fof(f3923,plain,(
  ( ! [X54,X56,X53] : (p(s(bool,'F_0')) | s(X53,X56) = s(X53,sK4(X56,X54,X53))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3338])).
fof(f4019,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4018,f3995])).
fof(f4018,plain,(
  ( ! [X43,X46,X44] : (s(X43,X46) = s(X43,sK4(X46,X44,X43))) ) | $spl14),
  inference(subsumption_resolution,[],[f3922,f337])).
fof(f3922,plain,(
  ( ! [X43,X46,X44] : (p(s(bool,'F_0')) | s(X43,X46) = s(X43,sK4(X46,X44,X43))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3212])).
fof(f4017,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3921,f337])).
fof(f3921,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3862,f3294])).
fof(f4016,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3920,f337])).
fof(f3920,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3862,f2749])).
fof(f4015,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4014,f3995])).
fof(f4014,plain,(
  ( ! [X39,X42,X40] : (s(X39,X42) = s(X39,sK4(X42,X40,X39))) ) | $spl14),
  inference(subsumption_resolution,[],[f3919,f337])).
fof(f3919,plain,(
  ( ! [X39,X42,X40] : (p(s(bool,'F_0')) | s(X39,X42) = s(X39,sK4(X42,X40,X39))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3211])).
fof(f4013,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4012,f3995])).
fof(f4012,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3979,f337])).
fof(f3979,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3916])).
fof(f3916,plain,(
  ( ! [X4,X2] : (p(s(bool,'F_0')) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3325])).
fof(f4011,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4010,f3995])).
fof(f4010,plain,(
  ( ! [X77] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(subsumption_resolution,[],[f4009,f4001])).
fof(f4009,plain,(
  ( ! [X76,X77] : (p(s(bool,sK4(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(subsumption_resolution,[],[f3915,f337])).
fof(f3915,plain,(
  ( ! [X76,X77] : (p(s(bool,'F_0')) | p(s(bool,sK4(X76,X77,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3117])).
fof(f4008,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4007,f3995])).
fof(f4007,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3911,f337])).
fof(f3911,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f3047])).
fof(f4006,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3982,f337])).
fof(f3982,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(condensation,[],[f3909])).
fof(f3909,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f307])).
fof(f4005,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4004,f3995])).
fof(f4004,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3984,f337])).
fof(f3984,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3906])).
fof(f3906,plain,(
  ( ! [X4,X2] : (p(s(bool,'F_0')) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f2778])).
fof(f4003,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4002,f3995])).
fof(f4002,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f3985,f337])).
fof(f3985,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'F_0'))) ) | $spl14),
  inference(condensation,[],[f3905])).
fof(f3905,plain,(
  ( ! [X6,X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl14),
  inference(backward_demodulation,[],[f3862,f1894])).
fof(f4000,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3900,f337])).
fof(f3900,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3862,f2968])).
fof(f3999,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3899,f337])).
fof(f3899,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(backward_demodulation,[],[f3862,f1871])).
fof(f3864,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | $spl16),
  inference(cnf_transformation,[],[f3864_D])).
fof(f3864_D,plain,(
  p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3844,plain,(
  ( ! [X48] : (~p(s(bool,X48)) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f302,f267])).
fof(f3861,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3831,f231])).
fof(f3831,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(fun(bool,bool),bool),'?_1'),s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f302,f392])).
fof(f3373,plain,(
  ( ! [X32] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) ) | $spl10),
  inference(cnf_transformation,[],[f3373_D])).
fof(f3373_D,plain,(
  ( ! [X32] : (p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3367,plain,(
  ( ! [X33,X32] : (~p(s(bool,X33)) | p(s(bool,'|dollar|exists_1'(s(fun('type.pair.prod'(bool,X32),bool),'const.pair.FST_0'))))) )),
  inference(superposition,[],[f3271,f311])).
fof(f1808,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) ) | $spl6),
  inference(cnf_transformation,[],[f1808_D])).
fof(f1808_D,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1810,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15)))))) )),
  inference(subsumption_resolution,[],[f1786,f241])).
fof(f1786,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15)))))) )),
  inference(superposition,[],[f1427,f359])).
fof(f1807,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12)))))) )),
  inference(subsumption_resolution,[],[f1785,f241])).
fof(f1785,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X11,bool),'_dst_x0_1'(s(bool,X12))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1427,f359])).
fof(f1273,plain,(
  ( ! [X20] : (~p(s(bool,sK4(X20,'_i_0_0',bool)))) ) | $spl4),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  ( ! [X20] : (~p(s(bool,sK4(X20,'_i_0_0',bool)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1261,plain,(
  ( ! [X20] : (~p(s(bool,sK4(X20,'_i_0_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0'))))) )),
  inference(superposition,[],[f280,f267])).
fof(f1269,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) | $spl0),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1268,plain,(
  ( ! [X0] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(subsumption_resolution,[],[f1255,f231])).
fof(f1255,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'_i_0_0')))) | s(bool,'F_0') = s(bool,sK4(X0,'_i_0_0',bool))) )),
  inference(superposition,[],[f280,f392])).
% SZS output end Proof for thm.benchmark
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 282
% Passive clauses: 3595
% Generated clauses: 49733
% Final active clauses: 158
% Final passive clauses: 2345
% Input formulas: 68
% Initial clauses: 104
% 
% Duplicate literals: 1143
% Trivial inequalities: 211
% Fw subsumption resolutions: 2162
% Fw demodulations: 10540
% Bw demodulations: 945
% Condensations: 2073
% Global subsumptions: 52
% 
% Simple tautologies: 266
% Equational tautologies: 294
% Forward subsumptions: 27408
% Fw demodulations to eq. taut.: 1636
% Bw demodulations to eq. taut.: 43
% 
% Unit resulting resolution: 5713
% Factoring: 1
% Forward superposition: 13979
% Backward superposition: 11349
% Self superposition: 273
% Equality factoring: 98
% Equality resolution: 29
% 
% Split clauses: 77
% Split components: 154
% SAT solver clauses: 4079
% SAT solver unit clauses: 1066
% SAT solver binary clauses: 1726
% 
% Sat splits: 77
% Sat splitting refutations: 276
% 
% Memory used [KB]: 50148
% Time elapsed: 2.355 s
% ------------------------------
% ------------------------------
% Success in time 3.7 s
