Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_numberProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_numberProof
% SZS output start Proof for thm.clos_numberProof
fof(f1262,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1261,f130,f139,f137,f1009,f77,f123,f148,f160,f1078,f1052])).
fof(f1052,plain,(
  s(bool,'F_0') = s(bool,sK1) | $spl32),
  inference(cnf_transformation,[],[f1052_D])).
fof(f1052_D,plain,(
  s(bool,'F_0') = s(bool,sK1) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1078,plain,(
  s(bool,'F_0') = s(bool,sK1) | $spl10),
  inference(backward_demodulation,[],[f139,f73])).
fof(f73,plain,(
  s(bool,sK0) = s(bool,sK1)),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  s(bool,sK0) = s(bool,sK1) & ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f43])).
fof(f43,plain,(
  ? [X0,X1] : (s(bool,X0) = s(bool,X1) & ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0)))))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : (s(bool,X0) = s(bool,X1) => p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s(bool,X6) = s(bool,X7) => p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X7))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s(bool,X6) = s(bool,X7) => p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X7))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X6)))))))),
  file('thm.clos_numberProof.clos_numberProof_35_prove_ltb','thm.clos_numberProof.clos_numberProof_35_prove_ltb')).
fof(f160,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3)) ) | $spl16),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'F_0') = s(bool,X3)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f148,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl1),
  inference(resolution,[],[f112,f123])).
fof(f112,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.clos_numberProof.clos_numberProof_35_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f123,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f77,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.clos_numberProof.clos_numberProof_35_prove_ltb',aHL_FALSITY)).
fof(f1009,plain,(
  $false | ($spl1 | $spl8 | $spl16)),
  inference(subsumption_resolution,[],[f997,f123])).
fof(f997,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl16)),
  inference(superposition,[],[f964,f194])).
fof(f194,plain,(
  s(bool,'F_0') = s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))))) | ($spl8 | $spl16)),
  inference(resolution,[],[f193,f160])).
fof(f193,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0'))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0')))))) | $spl8),
  inference(forward_demodulation,[],[f192,f137])).
fof(f192,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'T_0')))))) | $spl8),
  inference(forward_demodulation,[],[f74,f143])).
fof(f143,plain,(
  s(bool,'T_0') = s(bool,sK0) | $spl8),
  inference(backward_demodulation,[],[f137,f73])).
fof(f74,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK0))))))),
  inference(cnf_transformation,[],[f48])).
fof(f964,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X0))))))) )),
  inference(equality_resolution,[],[f117])).
fof(f117,plain,(
  ( ! [X0,X1] : (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1)))))) | s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1)))) & (s('type.closSem.v',X0) != s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))) | p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X0),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1)))))) <=> s('type.closSem.v',X0) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X1))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : (p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v',X7),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X8)))))) <=> s('type.closSem.v',X7) = s('type.closSem.v','const.closSem.Boolv_1'(s(bool,X8))))),
  file('thm.clos_numberProof.clos_numberProof_35_prove_ltb','thm.clos_numberProof.v_rel_Boolv|split|0')).
fof(f137,plain,(
  s(bool,'T_0') = s(bool,sK1) | $spl8),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s(bool,'T_0') = s(bool,sK1) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f139,plain,(
  s(bool,'F_0') = s(bool,sK0) | $spl10),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(bool,'F_0') = s(bool,sK0) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f130,plain,(
  s(bool,'F_0') = s(bool,sK0) | s(bool,'T_0') = s(bool,sK1)),
  inference(superposition,[],[f78,f73])).
fof(f78,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_numberProof.clos_numberProof_35_prove_ltb',aHL_BOOL_CASES)).
fof(f1261,plain,(
  $false | ($spl10 | $spl32)),
  inference(subsumption_resolution,[],[f1260,f964])).
fof(f1260,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0'))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0')))))) | ($spl10 | $spl32)),
  inference(forward_demodulation,[],[f1259,f1052])).
fof(f1259,plain,(
  ~p(s(bool,'const.clos_numberProof.v_rel_2'(s('type.closSem.v','const.closSem.Boolv_1'(s(bool,sK1))),s('type.closSem.v','const.closSem.Boolv_1'(s(bool,'F_0')))))) | $spl10),
  inference(forward_demodulation,[],[f74,f139])).
% SZS output end Proof for thm.clos_numberProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 28
% Passive clauses: 171
% Generated clauses: 1145
% Final active clauses: 20
% Final passive clauses: 102
% Input formulas: 21
% Initial clauses: 47
% 
% Duplicate literals: 133
% Trivial inequalities: 6
% Fw subsumption resolutions: 98
% Fw demodulations: 72
% Bw demodulations: 3
% 
% Simple tautologies: 60
% Equational tautologies: 33
% Forward subsumptions: 524
% 
% Binary resolution: 31
% Factoring: 1
% Forward superposition: 261
% Backward superposition: 372
% Self superposition: 48
% Equality factoring: 56
% Equality resolution: 1
% 
% Split clauses: 38
% Split components: 76
% SAT solver clauses: 63
% SAT solver unit clauses: 3
% SAT solver binary clauses: 33
% SAT solver learnt clauses: 1
% 
% Sat splits: 38
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 62
% TWLsolver calls for satisfiability: 16
% 
% Memory used [KB]: 1023
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
% Success in time 0.04 s
