Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.bvp_to_wordProps
% (101669)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 152
% Passive clauses: 3899
% Generated clauses: 25755
% Final active clauses: 132
% Final passive clauses: 3443
% Input formulas: 37
% Initial clauses: 68
% 
% Duplicate literals: 809
% Trivial inequalities: 292
% Fw subsumption resolutions: 2535
% Fw demodulations: 4348
% Bw demodulations: 14
% 
% Simple tautologies: 305
% Equational tautologies: 51
% Forward subsumptions: 12914
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 135
% Factoring: 1
% Forward superposition: 9476
% Backward superposition: 7571
% Self superposition: 253
% Equality factoring: 125
% Equality resolution: 39
% 
% Split clauses: 307
% Split components: 620
% SAT solver clauses: 731
% SAT solver unit clauses: 13
% SAT solver binary clauses: 118
% 
% Sat splits: 307
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 719
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 21108
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.bvp_to_wordProps
% (101772)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 1029
% Generated clauses: 18975
% Final active clauses: 128
% Final passive clauses: 520
% Input formulas: 37
% Initial clauses: 68
% 
% Duplicate literals: 141
% Trivial inequalities: 254
% Fw subsumption resolutions: 987
% Fw demodulations: 3337
% Bw demodulations: 325
% Condensations: 294
% Global subsumptions: 4
% 
% Simple tautologies: 79
% Equational tautologies: 14
% Forward subsumptions: 11276
% Bw demodulations to eq. taut.: 14
% 
% Binary resolution: 101
% Factoring: 1
% Forward superposition: 5949
% Backward superposition: 5971
% Self superposition: 106
% Equality factoring: 18
% Equality resolution: 41
% 
% Split clauses: 163
% Split components: 326
% SAT solver clauses: 1398
% SAT solver unit clauses: 109
% SAT solver binary clauses: 27
% SAT solver learnt clauses: 8
% 
% Sat splits: 163
% Sat splitting refutations: 58
% 
% TWLsolver clauses: 1382
% TWLsolver calls for satisfiability: 988
% 
% Memory used [KB]: 22771
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.bvp_to_wordProps
% (101845)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 843
% Passive clauses: 17071
% Generated clauses: 181833
% Final active clauses: 298
% Final passive clauses: 1715
% Input formulas: 37
% Initial clauses: 68
% Discarded non-redundant clauses: 6887
% 
% Duplicate literals: 4012
% Trivial inequalities: 1744
% Fw subsumption resolutions: 6074
% Fw demodulations: 18718
% Bw demodulations: 6821
% Condensations: 9664
% Global subsumptions: 55
% 
% Simple tautologies: 846
% Equational tautologies: 278
% Forward subsumptions: 122367
% Fw demodulations to eq. taut.: 15
% Bw demodulations to eq. taut.: 915
% 
% Unit resulting resolution: 4127
% Factoring: 3
% Forward superposition: 60001
% Backward superposition: 60277
% Self superposition: 469
% Equality factoring: 180
% Equality resolution: 73
% 
% Split clauses: 432
% Split components: 872
% SAT solver clauses: 18142
% SAT solver unit clauses: 1017
% SAT solver binary clauses: 4422
% 
% Sat splits: 432
% Sat splitting refutations: 150
% 
% Memory used [KB]: 106437
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.bvp_to_wordProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bvp_to_wordProps
% SZS output start Proof for thm.bvp_to_wordProps
fof(f3344,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f227,f236,f240,f241,f242,f243,f341,f421,f430,f433,f554,f565,f442,f533,f650,f655,f799,f960,f961,f231,f971,f248,f972,f249,f973,f250,f974,f251,f975,f969,f238,f992,f993,f432,f1004,f1005,f1006,f1116,f1007,f1008,f1117,f1114,f1009,f1010,f1011,f1121,f1012,f1125,f1148,f1158,f1245,f1251,f1152,f1150,f1299,f1302,f235,f1427,f1437,f1524,f1530,f1648,f1119,f1745,f1748,f1749,f1750,f1751,f1752,f1753,f1754,f1747,f1791,f1792,f1789,f1787,f1794,f1115,f1969,f2083,f2087,f2089,f2091,f2109,f2112,f2114,f2115,f2116,f2117,f2118,f2119,f2111,f2278,f2288,f2282,f2280,f2289,f2419,f2598,f2604,f1431,f1429,f2605,f2735,f2737,f2740,f229,f2762,f2765,f2766,f2767,f2768,f2769,f2770,f2771,f2764,f2902,f2908,f560,f558,f556,f2909,f2752,f3038,f3039,f3041,f1123,f233,f3059,f2086,f3060,f2088,f3061,f2090,f3062,f2092,f2085,f3204,f3214,f3208,f3206,f3215])).
fof(f3215,plain,(
  $false | ($spl6 | $spl37)),
  inference(subsumption_resolution,[],[f3082,f103])).
fof(f103,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3))))),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X4] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X4)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X3)))))) & ! [X2] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) & ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f81])).
fof(f81,plain,(
  ? [X0] : (! [X4] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X4)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X3] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X3)))))) & ! [X2] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) & ? [X1] : ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag',X1)))))),
  inference(rectify,[],[f73])).
fof(f73,plain,(
  ? [X0] : (! [X3] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X3)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X2)))))) & ! [X1] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) & ? [X4] : ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag',X4)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ? [X0] : ((! [X3] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X3)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X2)))))) & ! [X1] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'))))) & ? [X4] : ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag',X4)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ((! [X3] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X3)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X2] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X2)))))) & ! [X1] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'))))) => ! [X4] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X0),s('type.bvp_to_wordProps.tag',X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((! [X7] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X7)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X7)))))) & ! [X8] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X8)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'))))) => ! [X9] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag',X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((! [X7] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X7)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) & ! [X7] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X7)))))) & ! [X8] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X8)))))) & p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'))))) => ! [X9] : p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),X6),s('type.bvp_to_wordProps.tag',X9)))))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_induction')).
fof(f3082,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3)))) | ($spl6 | $spl37)),
  inference(superposition,[],[f101,f3044])).
fof(f3044,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | ($spl6 | $spl37)),
  inference(forward_demodulation,[],[f235,f1784])).
fof(f1784,plain,(
  s(bool,'T_0') = s(bool,sK10(sK3)) | $spl37),
  inference(trivial_inequality_removal,[],[f1783])).
fof(f1783,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,sK10(sK3)) | $spl37),
  inference(superposition,[],[f1120,f152])).
fof(f152,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bvp_to_wordProps.tag_induction',aHL_BOOL_CASES)).
fof(f1120,plain,(
  s(bool,'F_0') != s(bool,sK10(sK3)) | $spl37),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  s(bool,'F_0') != s(bool,sK10(sK3)) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f101,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2))))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f3206,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | $spl96),
  inference(cnf_transformation,[],[f3206_D])).
fof(f3206_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f3208,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | $spl98),
  inference(cnf_transformation,[],[f3208_D])).
fof(f3208_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f3214,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f3213,f3044])).
fof(f3213,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f3212,f3044])).
fof(f3212,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f3211,f237])).
fof(f237,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag',sK3) | $spl9),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag',sK3) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f3211,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f3210,f3044])).
fof(f3210,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl37)),
  inference(subsumption_resolution,[],[f3209,f124])).
fof(f124,plain,(
  ( ! [X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X48,X17] : ~s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X17))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X48)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_distinct|split|2')).
fof(f3209,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl37)),
  inference(subsumption_resolution,[],[f3081,f230])).
fof(f230,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag',sK3) | $spl3),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag',sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f3081,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl6 | $spl37)),
  inference(superposition,[],[f3044,f122])).
fof(f122,plain,(
  ( ! [X0] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X0) | s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(X0)))) | s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X0) | s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(X0))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X0) | s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(X0)))) | s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10])],[f46])).
fof(f46,plain,(
  ! [X0] : (? [X3] : s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X3))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X0) | ? [X2] : s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X2))) | ? [X1] : s('type.bvp_to_wordProps.tag',X0) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X0))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X33] : (? [X7] : s('type.bvp_to_wordProps.tag',X33) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X7))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X33) | ? [X7] : s('type.bvp_to_wordProps.tag',X33) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X7))) | ? [X8] : s('type.bvp_to_wordProps.tag',X33) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X8))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X33))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_nchotomy')).
fof(f3204,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl3 | $spl6 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f3203,f230])).
fof(f3203,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl6 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f3202,f3044])).
fof(f3202,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | ($spl6 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f3201,f3044])).
fof(f3201,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | ($spl6 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f3200,f3044])).
fof(f3200,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | ($spl6 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f3199,f124])).
fof(f3199,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl6 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f3080,f237])).
fof(f3080,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))))) | ($spl6 | $spl37)),
  inference(superposition,[],[f3044,f122])).
fof(f2085,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | $spl62),
  inference(cnf_transformation,[],[f2085_D])).
fof(f2085_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2092,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'T_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36))))) ) | $spl70),
  inference(cnf_transformation,[],[f2092_D])).
fof(f2092_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'T_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3062,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'T_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f3048,f125])).
fof(f125,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun(bool,X0),X5),s(bool,X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : ! [X17,X30,X31,X12,X13,X32] : s(X10,happ(s(fun(bool,X10),X13),s(bool,X17))) = s(X10,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X17))),s(fun('type.num.num',X10),X30),s(X10,X31),s(fun('type.num.num',X10),X12),s(fun(bool,X10),X13),s(X10,X32)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_case_def|split|3')).
fof(f3048,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f111])).
fof(f111,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) | s(X0,happ(s(fun(bool,X0),X3),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | sP1(X0,X2,X9,X7) | sP0(X0,X5,X8,X7) | s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7) | sP0(X0,X5,X8,X7) | sP1(X0,X2,X9,X7) | (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X0,happ(s(fun(bool,X0),X3),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X1) != s(X0,X12)) | s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7) | sP0(X0,X5,X8,X7) | sP1(X0,X2,X9,X7) | ? [X14] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X14))) & s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X14))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',X14)))) | ? [X13] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X13))) & s(X0,happ(s(fun(bool,X0),X3),s(bool,X13))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,X13)))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X1) != s(X0,X12)) | s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(definition_folding,[],[f75,f79,f78])).
fof(f78,plain,(
  ! [X0,X5,X8,X7] : (~sP0(X0,X5,X8,X7) | ? [X15] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X15))) & s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X15))) != s(X0,happ(s(fun('type.num.num',X0),X8),s('type.num.num',X15)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f79,plain,(
  ! [X0,X2,X9,X7] : (~sP1(X0,X2,X9,X7) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X2) != s(X0,X9)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f75,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7) | ? [X15] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X15))) & s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X15))) != s(X0,happ(s(fun('type.num.num',X0),X8),s('type.num.num',X15)))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X2) != s(X0,X9)) | ? [X14] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X14))) & s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X14))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',X14)))) | ? [X13] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X13))) & s(X0,happ(s(fun(bool,X0),X3),s(bool,X13))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,X13)))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X1) != s(X0,X12)) | s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7) | ? [X15] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X15))) & s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X15))) != s(X0,happ(s(fun('type.num.num',X0),X8),s('type.num.num',X15)))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X2) != s(X0,X9)) | ? [X14] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X14))) & s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X14))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',X14)))) | ? [X13] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X13))) & s(X0,happ(s(fun(bool,X0),X3),s(bool,X13))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,X13)))) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X1) != s(X0,X12))) | s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((s('type.bvp_to_wordProps.tag',X6) = s('type.bvp_to_wordProps.tag',X7) & ! [X15] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X15))) => s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X15))) = s(X0,happ(s(fun('type.num.num',X0),X8),s('type.num.num',X15)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) => s(X0,X2) = s(X0,X9)) & ! [X14] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X14))) => s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X14))) = s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',X14)))) & ! [X13] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X13))) => s(X0,happ(s(fun(bool,X0),X3),s(bool,X13))) = s(X0,happ(s(fun(bool,X0),X11),s(bool,X13)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) => s(X0,X1) = s(X0,X12))) => s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((s('type.bvp_to_wordProps.tag',X6) = s('type.bvp_to_wordProps.tag',X7) & ! [X15] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X15))) => s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X15))) = s(X0,happ(s(fun('type.num.num',X0),X8),s('type.num.num',X15)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) => s(X0,X2) = s(X0,X9)) & ! [X14] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X14))) => s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X14))) = s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',X14)))) & ! [X13] : (s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X13))) => s(X0,happ(s(fun(bool,X0),X3),s(bool,X13))) = s(X0,happ(s(fun(bool,X0),X11),s(bool,X13)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) => s(X0,X1) = s(X0,X12))) => s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : ! [X23,X24,X25,X26,X27,X28,X29,X30,X31,X12,X13,X32] : ((s('type.bvp_to_wordProps.tag',X28) = s('type.bvp_to_wordProps.tag',X29) & ! [X17] : (s('type.bvp_to_wordProps.tag',X29) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X17))) => s(X10,happ(s(fun('type.num.num',X10),X27),s('type.num.num',X17))) = s(X10,happ(s(fun('type.num.num',X10),X30),s('type.num.num',X17)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X29) => s(X10,X24) = s(X10,X31)) & ! [X17] : (s('type.bvp_to_wordProps.tag',X29) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X17))) => s(X10,happ(s(fun('type.num.num',X10),X12),s('type.num.num',X17))) = s(X10,happ(s(fun('type.num.num',X10),X26),s('type.num.num',X17)))) & ! [X17] : (s('type.bvp_to_wordProps.tag',X29) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X17))) => s(X10,happ(s(fun(bool,X10),X13),s(bool,X17))) = s(X10,happ(s(fun(bool,X10),X25),s(bool,X17)))) & (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X29) => s(X10,X23) = s(X10,X32))) => s(X10,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X28),s(fun('type.num.num',X10),X30),s(X10,X31),s(fun('type.num.num',X10),X12),s(fun(bool,X10),X13),s(X10,X32))) = s(X10,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X29),s(fun('type.num.num',X10),X27),s(X10,X24),s(fun('type.num.num',X10),X26),s(fun(bool,X10),X25),s(X10,X23))))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_case_cong')).
fof(f2068,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f1126,f1784])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1013,f230])).
fof(f1013,plain,(
  ( ! [X2,X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ~sP1(X0,X1,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))) ) | $spl6),
  inference(superposition,[],[f235,f104])).
fof(f104,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X3) | ~sP1(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : (~sP1(X0,X1,X2,X3) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X3) & s(X0,X1) != s(X0,X2)))),
  inference(rectify,[],[f83])).
fof(f83,plain,(
  ! [X0,X2,X9,X7] : (~sP1(X0,X2,X9,X7) | (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',X7) & s(X0,X2) != s(X0,X9)))),
  inference(nnf_transformation,[],[f79])).
fof(f2090,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'T_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24))))) ) | $spl68),
  inference(cnf_transformation,[],[f2090_D])).
fof(f2090_D,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'T_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3061,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'T_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f3047,f125])).
fof(f3047,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f110])).
fof(f110,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) | s(X0,happ(s(fun(bool,X0),X3),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun(bool,X0),X11),s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | sP1(X0,X2,X9,X7) | sP0(X0,X5,X8,X7) | s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7)) )),
  inference(cnf_transformation,[],[f88])).
fof(f2088,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'T_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12))))) ) | $spl66),
  inference(cnf_transformation,[],[f2088_D])).
fof(f2088_D,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'T_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3060,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'T_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f3046,f125])).
fof(f3046,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f109])).
fof(f109,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) | s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X0,happ(s(fun('type.num.num',X0),X10),s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | sP1(X0,X2,X9,X7) | sP0(X0,X5,X8,X7) | s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7)) )),
  inference(cnf_transformation,[],[f88])).
fof(f2086,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'T_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0))))) ) | $spl64),
  inference(cnf_transformation,[],[f2086_D])).
fof(f2086_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'T_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3059,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'T_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f3045,f125])).
fof(f3045,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f108])).
fof(f108,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X6),s(fun('type.num.num',X0),X8),s(X0,X9),s(fun('type.num.num',X0),X10),s(fun(bool,X0),X11),s(X0,X12))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X7),s(fun('type.num.num',X0),X5),s(X0,X2),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X3),s(X0,X1))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',X7) | s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.bvp_to_wordProps.tag',X7) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X12,X11,X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | sP1(X0,X2,X9,X7) | sP0(X0,X5,X8,X7) | s('type.bvp_to_wordProps.tag',X6) != s('type.bvp_to_wordProps.tag',X7)) )),
  inference(cnf_transformation,[],[f88])).
fof(f233,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | $spl4),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1123,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | $spl38),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3041,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | ($spl1 | $spl3 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f3040,f228])).
fof(f228,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl1),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f3040,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f1967,f230])).
fof(f1967,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl9 | $spl37)),
  inference(subsumption_resolution,[],[f1885,f237])).
fof(f1885,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl37),
  inference(superposition,[],[f122,f1784])).
fof(f3039,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | ($spl1 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f2736,f228])).
fof(f2736,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl9 | $spl37)),
  inference(forward_demodulation,[],[f1301,f1784])).
fof(f1301,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(subsumption_resolution,[],[f1300,f237])).
fof(f1300,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f1278,f99])).
fof(f99,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0'))))),
  inference(cnf_transformation,[],[f82])).
fof(f1278,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
fof(f3038,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | ($spl1 | $spl3 | $spl37)),
  inference(subsumption_resolution,[],[f2739,f228])).
fof(f2739,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl37)),
  inference(forward_demodulation,[],[f2738,f1784])).
fof(f2738,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl3),
  inference(subsumption_resolution,[],[f1297,f230])).
fof(f1297,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f1277,f102])).
fof(f102,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0'))))),
  inference(cnf_transformation,[],[f82])).
fof(f1277,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
fof(f2752,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | ($spl1 | $spl3 | $spl9 | $spl37)),
  inference(forward_demodulation,[],[f2751,f1784])).
fof(f2751,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | ($spl1 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f2750,f230])).
fof(f2750,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ($spl1 | $spl9)),
  inference(subsumption_resolution,[],[f2749,f237])).
fof(f2749,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | $spl1),
  inference(trivial_inequality_removal,[],[f2741])).
fof(f2741,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | $spl1),
  inference(superposition,[],[f228,f122])).
fof(f2909,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f2783,f103])).
fof(f2783,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3)))) | $spl4),
  inference(superposition,[],[f100,f233])).
fof(f100,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X3))))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f556,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl22),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f558,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl24),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f560,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl26),
  inference(cnf_transformation,[],[f560_D])).
fof(f560_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2908,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(forward_demodulation,[],[f2907,f233])).
fof(f2907,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(forward_demodulation,[],[f2906,f233])).
fof(f2906,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(forward_demodulation,[],[f2905,f233])).
fof(f2905,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f2904,f237])).
fof(f2904,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f2903,f233])).
fof(f2903,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2782,f230])).
fof(f2782,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f233,f122])).
fof(f2902,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(forward_demodulation,[],[f2901,f233])).
fof(f2901,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl3 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f2900,f230])).
fof(f2900,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f2899,f233])).
fof(f2899,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f2898,f233])).
fof(f2898,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f2897,f233])).
fof(f2897,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f2781,f237])).
fof(f2781,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f233,f122])).
fof(f2764,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | $spl86),
  inference(cnf_transformation,[],[f2764_D])).
fof(f2764_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2771,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK9(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36))))) ) | $spl94),
  inference(cnf_transformation,[],[f2771_D])).
fof(f2771_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK9(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2770,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK9(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(forward_demodulation,[],[f2757,f127])).
fof(f127,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X4),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10] : ! [X17,X30,X31,X12,X13,X32] : s(X10,happ(s(fun('type.num.num',X10),X12),s('type.num.num',X17))) = s(X10,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',X17))),s(fun('type.num.num',X10),X30),s(X10,X31),s(fun('type.num.num',X10),X12),s(fun(bool,X10),X13),s(X10,X32)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_case_def|split|2')).
fof(f2757,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(resolution,[],[f533,f111])).
fof(f2769,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X33),s('type.num.num',sK9(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24))))) ) | $spl92),
  inference(cnf_transformation,[],[f2769_D])).
fof(f2769_D,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X33),s('type.num.num',sK9(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2768,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X33),s('type.num.num',sK9(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(forward_demodulation,[],[f2756,f127])).
fof(f2756,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(resolution,[],[f533,f110])).
fof(f2767,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK9(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12))))) ) | $spl90),
  inference(cnf_transformation,[],[f2767_D])).
fof(f2767_D,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK9(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2766,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK9(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(forward_demodulation,[],[f2755,f127])).
fof(f2755,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(resolution,[],[f533,f109])).
fof(f2765,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X9),s('type.num.num',sK9(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0))))) ) | $spl88),
  inference(cnf_transformation,[],[f2765_D])).
fof(f2765_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X9),s('type.num.num',sK9(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2762,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X9),s('type.num.num',sK9(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(forward_demodulation,[],[f2754,f127])).
fof(f2754,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(resolution,[],[f533,f108])).
fof(f229,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl0),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2740,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl7 | $spl37)),
  inference(subsumption_resolution,[],[f2739,f2099])).
fof(f2099,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | ($spl7 | $spl37)),
  inference(forward_demodulation,[],[f234,f1784])).
fof(f234,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | $spl7),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f2737,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl7 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f2736,f2099])).
fof(f2735,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl7 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f2734,f230])).
fof(f2734,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl7 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f1967,f2099])).
fof(f2605,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f2478,f103])).
fof(f2478,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3)))) | $spl0),
  inference(superposition,[],[f98,f229])).
fof(f98,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X4))))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f1429,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl46),
  inference(cnf_transformation,[],[f1429_D])).
fof(f1429_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1431,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl48),
  inference(cnf_transformation,[],[f1431_D])).
fof(f1431_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2604,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f2603,f229])).
fof(f2603,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f2602,f229])).
fof(f2602,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f2601,f237])).
fof(f2601,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2600,f229])).
fof(f2600,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2599,f124])).
fof(f2599,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2477,f230])).
fof(f2477,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f229,f122])).
fof(f2598,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f2597,f229])).
fof(f2597,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f2596,f230])).
fof(f2596,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f2595,f229])).
fof(f2595,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f2594,f229])).
fof(f2594,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2593,f124])).
fof(f2593,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2476,f237])).
fof(f2476,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f229,f122])).
fof(f2419,plain,(
  $false | ($spl7 | $spl37 | $spl38)),
  inference(subsumption_resolution,[],[f1123,f2099])).
fof(f2289,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f2159,f103])).
fof(f2159,plain,(
  p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag',sK3)))) | $spl34),
  inference(superposition,[],[f101,f1119])).
fof(f2280,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | $spl82),
  inference(cnf_transformation,[],[f2280_D])).
fof(f2280_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2282,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | $spl84),
  inference(cnf_transformation,[],[f2282_D])).
fof(f2282_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2288,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl9 | $spl34)),
  inference(forward_demodulation,[],[f2287,f1119])).
fof(f2287,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl9 | $spl34)),
  inference(forward_demodulation,[],[f2286,f1119])).
fof(f2286,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl9 | $spl34)),
  inference(subsumption_resolution,[],[f2285,f237])).
fof(f2285,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl34)),
  inference(forward_demodulation,[],[f2284,f1119])).
fof(f2284,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f2283,f124])).
fof(f2283,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl34)),
  inference(subsumption_resolution,[],[f2158,f230])).
fof(f2158,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | $spl34),
  inference(superposition,[],[f1119,f122])).
fof(f2278,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl3 | $spl9 | $spl34)),
  inference(subsumption_resolution,[],[f2277,f230])).
fof(f2277,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl9 | $spl34)),
  inference(forward_demodulation,[],[f2276,f1119])).
fof(f2276,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl9 | $spl34)),
  inference(forward_demodulation,[],[f2275,f1119])).
fof(f2275,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl9 | $spl34)),
  inference(forward_demodulation,[],[f2274,f1119])).
fof(f2274,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | ($spl9 | $spl34)),
  inference(subsumption_resolution,[],[f2273,f124])).
fof(f2273,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | ($spl9 | $spl34)),
  inference(subsumption_resolution,[],[f2157,f237])).
fof(f2157,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))))) | $spl34),
  inference(superposition,[],[f1119,f122])).
fof(f2111,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl72),
  inference(cnf_transformation,[],[f2111_D])).
fof(f2111_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2119,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X43),s('type.num.num',sK8(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36))))) ) | $spl80),
  inference(cnf_transformation,[],[f2119_D])).
fof(f2119_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X43),s('type.num.num',sK8(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2118,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun('type.num.num',X36),X43),s('type.num.num',sK8(sK3)))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2103,f150])).
fof(f150,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X1))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X1))),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : ! [X17,X30,X31,X12,X13,X32] : s(X10,happ(s(fun('type.num.num',X10),X30),s('type.num.num',X17))) = s(X10,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',X17))),s(fun('type.num.num',X10),X30),s(X10,X31),s(fun('type.num.num',X10),X12),s(fun(bool,X10),X13),s(X10,X32)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bvp_to_wordProps.tag_case_def|split|0')).
fof(f2103,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f1405,f111])).
fof(f1405,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1304,f230])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ~sP1(X0,X1,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | $spl0),
  inference(superposition,[],[f229,f104])).
fof(f2117,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X31),s('type.num.num',sK8(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24))))) ) | $spl78),
  inference(cnf_transformation,[],[f2117_D])).
fof(f2117_D,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X31),s('type.num.num',sK8(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2116,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun('type.num.num',X24),X31),s('type.num.num',sK8(sK3)))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2102,f150])).
fof(f2102,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f1405,f110])).
fof(f2115,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X19),s('type.num.num',sK8(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12))))) ) | $spl76),
  inference(cnf_transformation,[],[f2115_D])).
fof(f2115_D,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X19),s('type.num.num',sK8(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2114,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun('type.num.num',X12),X19),s('type.num.num',sK8(sK3)))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2113,f150])).
fof(f2113,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2101,f124])).
fof(f2101,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f1405,f109])).
fof(f2112,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',sK8(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0))))) ) | $spl74),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',sK8(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2109,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun('type.num.num',X0),X7),s('type.num.num',sK8(sK3)))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f2108,f150])).
fof(f2108,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2100,f124])).
fof(f2100,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f1405,f108])).
fof(f2091,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'T_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f2072,f125])).
fof(f2072,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f111])).
fof(f2089,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'T_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f2071,f125])).
fof(f2071,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f110])).
fof(f2087,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'T_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f2070,f125])).
fof(f2070,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f109])).
fof(f2083,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'T_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(forward_demodulation,[],[f2069,f125])).
fof(f2069,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0')))) ) | ($spl3 | $spl6 | $spl37)),
  inference(resolution,[],[f2068,f108])).
fof(f1969,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl5 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f1968,f230])).
fof(f1968,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl5 | $spl9 | $spl37)),
  inference(subsumption_resolution,[],[f1967,f232])).
fof(f232,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | $spl5),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1115,plain,(
  p(s(bool,sK10(sK3))) | $spl28),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  p(s(bool,sK10(sK3))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1794,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1)) ) | $spl60),
  inference(cnf_transformation,[],[f1794_D])).
fof(f1794_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1787,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,sK10(sK3)))) ) | $spl37),
  inference(subsumption_resolution,[],[f1776,f139])).
fof(f139,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bool.EQ_CLAUSES|split|2')).
fof(f1776,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,sK10(sK3)))) ) | $spl37),
  inference(superposition,[],[f1120,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X19,X20] : ((p(s(bool,X19)) => p(s(bool,X20))) => ((p(s(bool,X20)) => p(s(bool,X19))) => s(bool,X19) = s(bool,X20)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bool.IMP_ANTISYM_AX')).
fof(f1789,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,sK10(sK3)))) ) | $spl37),
  inference(subsumption_resolution,[],[f1777,f139])).
fof(f1777,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,sK10(sK3))) | p(s(bool,X1))) ) | $spl37),
  inference(superposition,[],[f1120,f130])).
fof(f1792,plain,(
  $false | ($spl29 | $spl37)),
  inference(subsumption_resolution,[],[f1785,f1114])).
fof(f1785,plain,(
  p(s(bool,sK10(sK3))) | $spl37),
  inference(trivial_inequality_removal,[],[f1781])).
fof(f1781,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,sK10(sK3))) | $spl37),
  inference(superposition,[],[f1120,f142])).
fof(f142,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.bvp_to_wordProps.tag_induction','thm.bool.EQ_CLAUSES|split|3')).
fof(f1791,plain,(
  $false | ($spl29 | $spl37)),
  inference(subsumption_resolution,[],[f1786,f1114])).
fof(f1786,plain,(
  p(s(bool,sK10(sK3))) | $spl37),
  inference(trivial_inequality_removal,[],[f1780])).
fof(f1780,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,sK10(sK3))) | $spl37),
  inference(superposition,[],[f1120,f140])).
fof(f140,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f1747,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | $spl50),
  inference(cnf_transformation,[],[f1747_D])).
fof(f1747_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1754,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'F_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36))))) ) | $spl58),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'F_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1753,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,happ(s(fun(bool,X36),X46),s(bool,'F_0'))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(forward_demodulation,[],[f1729,f125])).
fof(f1729,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X37) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(resolution,[],[f1725,f111])).
fof(f1725,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(forward_demodulation,[],[f1126,f1121])).
fof(f1752,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'F_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24))))) ) | $spl56),
  inference(cnf_transformation,[],[f1752_D])).
fof(f1752_D,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'F_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1751,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,happ(s(fun(bool,X24),X34),s(bool,'F_0'))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(forward_demodulation,[],[f1728,f125])).
fof(f1728,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X25) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(resolution,[],[f1725,f110])).
fof(f1750,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'F_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12))))) ) | $spl54),
  inference(cnf_transformation,[],[f1750_D])).
fof(f1750_D,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'F_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1749,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,happ(s(fun(bool,X12),X22),s(bool,'F_0'))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(forward_demodulation,[],[f1727,f125])).
fof(f1727,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X13) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(resolution,[],[f1725,f109])).
fof(f1748,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'F_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0))))) ) | $spl52),
  inference(cnf_transformation,[],[f1748_D])).
fof(f1748_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'F_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1745,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,happ(s(fun(bool,X0),X10),s(bool,'F_0'))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(forward_demodulation,[],[f1726,f125])).
fof(f1726,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')),X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,'const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',X1) != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0')))) ) | ($spl3 | $spl6 | $spl36)),
  inference(resolution,[],[f1725,f108])).
fof(f1119,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | $spl34),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1648,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl5 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f1647,f230])).
fof(f1647,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl5 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f1646,f232])).
fof(f1646,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f1568,f237])).
fof(f1568,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl36),
  inference(superposition,[],[f122,f1121])).
fof(f1530,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1529,f229])).
fof(f1529,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1528,f229])).
fof(f1528,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1527,f237])).
fof(f1527,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f1526,f229])).
fof(f1526,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1525,f124])).
fof(f1525,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1352,f230])).
fof(f1352,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f122,f229])).
fof(f1524,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1523,f229])).
fof(f1523,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1522,f230])).
fof(f1522,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f1521,f229])).
fof(f1521,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f1520,f229])).
fof(f1520,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1519,f124])).
fof(f1519,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1351,f237])).
fof(f1351,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f122,f229])).
fof(f1437,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1436,f229])).
fof(f1436,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1435,f229])).
fof(f1435,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1434,f237])).
fof(f1434,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f1433,f229])).
fof(f1433,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1432,f124])).
fof(f1432,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1310,f230])).
fof(f1310,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f229,f122])).
fof(f1427,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1426,f229])).
fof(f1426,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1425,f230])).
fof(f1425,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f1424,f229])).
fof(f1424,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f1423,f229])).
fof(f1423,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1422,f124])).
fof(f1422,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1309,f237])).
fof(f1309,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))))) | $spl0),
  inference(superposition,[],[f229,f122])).
fof(f235,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | $spl6),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1302,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl5 | $spl9)),
  inference(subsumption_resolution,[],[f1301,f232])).
fof(f1299,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f1298,f230])).
fof(f1298,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl5),
  inference(subsumption_resolution,[],[f1297,f232])).
fof(f1150,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl42),
  inference(cnf_transformation,[],[f1150_D])).
fof(f1150_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1152,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl44),
  inference(cnf_transformation,[],[f1152_D])).
fof(f1152_D,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1251,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(forward_demodulation,[],[f1250,f235])).
fof(f1250,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(forward_demodulation,[],[f1249,f235])).
fof(f1249,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1248,f237])).
fof(f1248,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f1247,f235])).
fof(f1247,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1246,f124])).
fof(f1246,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1061,f230])).
fof(f1061,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl6),
  inference(superposition,[],[f122,f235])).
fof(f1245,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1244,f230])).
fof(f1244,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1243,f235])).
fof(f1243,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1242,f235])).
fof(f1242,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1241,f235])).
fof(f1241,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1240,f124])).
fof(f1240,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1060,f237])).
fof(f1060,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl6),
  inference(superposition,[],[f122,f235])).
fof(f1158,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(forward_demodulation,[],[f1157,f235])).
fof(f1157,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(forward_demodulation,[],[f1156,f235])).
fof(f1156,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1155,f237])).
fof(f1155,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f1154,f235])).
fof(f1154,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1153,f124])).
fof(f1153,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1019,f230])).
fof(f1019,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl6),
  inference(superposition,[],[f235,f122])).
fof(f1148,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl3 | $spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1147,f230])).
fof(f1147,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1146,f235])).
fof(f1146,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1145,f235])).
fof(f1145,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f1144,f235])).
fof(f1144,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1143,f124])).
fof(f1143,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | ($spl6 | $spl9)),
  inference(subsumption_resolution,[],[f1018,f237])).
fof(f1018,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3))))))) | $spl6),
  inference(superposition,[],[f235,f122])).
fof(f1125,plain,(
  s(bool,'T_0') = s(bool,sK10(sK3)) | $spl40),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  s(bool,'T_0') = s(bool,sK10(sK3)) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1012,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(sK3)) | $spl6),
  inference(superposition,[],[f235,f152])).
fof(f1121,plain,(
  s(bool,'F_0') = s(bool,sK10(sK3)) | $spl36),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  s(bool,'F_0') = s(bool,sK10(sK3)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1011,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,sK10(sK3)) | $spl6),
  inference(superposition,[],[f235,f152])).
fof(f1010,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | p(s(bool,sK10(sK3))) | $spl6),
  inference(superposition,[],[f235,f142])).
fof(f1009,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,'F_0'))) | p(s(bool,sK10(sK3))) | $spl6),
  inference(superposition,[],[f235,f140])).
fof(f1114,plain,(
  ~p(s(bool,sK10(sK3))) | $spl29),
  inference(cnf_transformation,[],[f1114_D])).
fof(f1114_D,plain,(
  ~p(s(bool,sK10(sK3))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1117,plain,(
  ( ! [X2] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2))) | ~p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ( ! [X2] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2))) | ~p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1008,plain,(
  ( ! [X3] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X3))) | ~p(s(bool,X3)) | ~p(s(bool,sK10(sK3)))) ) | $spl6),
  inference(superposition,[],[f235,f133])).
fof(f133,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f77])).
fof(f1007,plain,(
  ( ! [X2] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X2))) | ~p(s(bool,sK10(sK3))) | ~p(s(bool,X2))) ) | $spl6),
  inference(superposition,[],[f235,f133])).
fof(f1116,plain,(
  ( ! [X0] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) | p(s(bool,X0))) ) | $spl30),
  inference(cnf_transformation,[],[f1116_D])).
fof(f1116_D,plain,(
  ( ! [X0] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) | p(s(bool,X0))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1006,plain,(
  ( ! [X1] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X1))) | p(s(bool,sK10(sK3))) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f235,f130])).
fof(f1005,plain,(
  ( ! [X0] : (s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,X0))) | p(s(bool,X0)) | p(s(bool,sK10(sK3)))) ) | $spl6),
  inference(superposition,[],[f235,f130])).
fof(f1004,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl5 | $spl9)),
  inference(subsumption_resolution,[],[f1003,f230])).
fof(f1003,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl5 | $spl9)),
  inference(subsumption_resolution,[],[f1002,f237])).
fof(f1002,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl5),
  inference(trivial_inequality_removal,[],[f994])).
fof(f994,plain,(
  s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl5),
  inference(superposition,[],[f232,f122])).
fof(f432,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | $spl19),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f993,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(subsumption_resolution,[],[f982,f237])).
fof(f982,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(superposition,[],[f237,f122])).
fof(f992,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f983,f230])).
fof(f983,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(trivial_inequality_removal,[],[f981])).
fof(f981,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(superposition,[],[f237,f122])).
fof(f238,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | $spl8),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f969,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl3),
  inference(trivial_inequality_removal,[],[f968])).
fof(f968,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl3),
  inference(superposition,[],[f230,f122])).
fof(f975,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl3),
  inference(subsumption_resolution,[],[f967,f230])).
fof(f967,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl3),
  inference(superposition,[],[f230,f122])).
fof(f251,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X37) | sP0(X36,X43,X38,sK3) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36))))) ) | $spl16),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X37) | sP0(X36,X43,X38,sK3) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f974,plain,(
  ( ! [X39,X47,X45,X43,X41,X50,X48,X46,X44,X42,X40,X49] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X39,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X39),X46),s(X39,X47),s(fun('type.num.num',X39),X48),s(fun(bool,X39),X49),s(X39,X50))) = s(X39,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X40),s(fun('type.num.num',X39),X41),s(X39,X42),s(fun('type.num.num',X39),X43),s(fun(bool,X39),X44),s(X39,X45))) | s(X39,happ(s(fun(bool,X39),X44),s(bool,sK6(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) != s(X39,happ(s(fun(bool,X39),X49),s(bool,sK6(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) | s(X39,happ(s(fun('type.num.num',X39),X43),s('type.num.num',sK5(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) != s(X39,happ(s(fun('type.num.num',X39),X48),s('type.num.num',sK5(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) | sP0(X39,X46,X41,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X40)) ) | $spl3),
  inference(subsumption_resolution,[],[f966,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,sK3)) )),
  inference(subsumption_resolution,[],[f205,f99])).
fof(f205,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) | ~sP1(X0,X1,X2,sK3)) )),
  inference(superposition,[],[f103,f104])).
fof(f966,plain,(
  ( ! [X39,X47,X45,X43,X41,X50,X48,X46,X44,X42,X40,X49] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X39,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X39),X46),s(X39,X47),s(fun('type.num.num',X39),X48),s(fun(bool,X39),X49),s(X39,X50))) = s(X39,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X40),s(fun('type.num.num',X39),X41),s(X39,X42),s(fun('type.num.num',X39),X43),s(fun(bool,X39),X44),s(X39,X45))) | s(X39,happ(s(fun(bool,X39),X44),s(bool,sK6(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) != s(X39,happ(s(fun(bool,X39),X49),s(bool,sK6(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) | s(X39,happ(s(fun('type.num.num',X39),X43),s('type.num.num',sK5(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) != s(X39,happ(s(fun('type.num.num',X39),X48),s('type.num.num',sK5(X45,X44,X43,X42,X41,sK3,X40,X46,X48,X49,X47,X50,X39)))) | sP1(X39,X47,X42,sK3) | sP0(X39,X46,X41,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X40)) ) | $spl3),
  inference(superposition,[],[f230,f111])).
fof(f250,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X25) | sP0(X24,X31,X26,sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24))))) ) | $spl14),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X25) | sP0(X24,X31,X26,sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f973,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X27,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X27),X34),s(X27,X35),s(fun('type.num.num',X27),X36),s(fun(bool,X27),X37),s(X27,X38))) = s(X27,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X28),s(fun('type.num.num',X27),X29),s(X27,X30),s(fun('type.num.num',X27),X31),s(fun(bool,X27),X32),s(X27,X33))) | s(X27,happ(s(fun(bool,X27),X32),s(bool,sK6(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) != s(X27,happ(s(fun(bool,X27),X37),s(bool,sK6(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) | sP0(X27,X34,X29,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X28)) ) | $spl3),
  inference(subsumption_resolution,[],[f965,f222])).
fof(f965,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X38,X36,X34,X32] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X27,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X27),X34),s(X27,X35),s(fun('type.num.num',X27),X36),s(fun(bool,X27),X37),s(X27,X38))) = s(X27,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X28),s(fun('type.num.num',X27),X29),s(X27,X30),s(fun('type.num.num',X27),X31),s(fun(bool,X27),X32),s(X27,X33))) | s(X27,happ(s(fun(bool,X27),X32),s(bool,sK6(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) != s(X27,happ(s(fun(bool,X27),X37),s(bool,sK6(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X33,X32,X31,X30,X29,sK3,X28,X34,X36,X37,X35,X38,X27)))) | sP1(X27,X35,X30,sK3) | sP0(X27,X34,X29,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X28)) ) | $spl3),
  inference(superposition,[],[f230,f110])).
fof(f249,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X13) | sP0(X12,X19,X14,sK3) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12))))) ) | $spl12),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X13) | sP0(X12,X19,X14,sK3) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f972,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X15),X22),s(X15,X23),s(fun('type.num.num',X15),X24),s(fun(bool,X15),X25),s(X15,X26))) = s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X16),s(fun('type.num.num',X15),X17),s(X15,X18),s(fun('type.num.num',X15),X19),s(fun(bool,X15),X20),s(X15,X21))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) | s(X15,happ(s(fun('type.num.num',X15),X19),s('type.num.num',sK5(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) != s(X15,happ(s(fun('type.num.num',X15),X24),s('type.num.num',sK5(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) | sP0(X15,X22,X17,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X16)) ) | $spl3),
  inference(subsumption_resolution,[],[f964,f222])).
fof(f964,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X15,X22,X20,X18,X16] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X15),X22),s(X15,X23),s(fun('type.num.num',X15),X24),s(fun(bool,X15),X25),s(X15,X26))) = s(X15,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X16),s(fun('type.num.num',X15),X17),s(X15,X18),s(fun('type.num.num',X15),X19),s(fun(bool,X15),X20),s(X15,X21))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) | s(X15,happ(s(fun('type.num.num',X15),X19),s('type.num.num',sK5(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) != s(X15,happ(s(fun('type.num.num',X15),X24),s('type.num.num',sK5(X21,X20,X19,X18,X17,sK3,X16,X22,X24,X25,X23,X26,X15)))) | sP1(X15,X23,X18,sK3) | sP0(X15,X22,X17,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X16)) ) | $spl3),
  inference(superposition,[],[f230,f109])).
fof(f248,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X1) | sP0(X0,X7,X2,sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0))))) ) | $spl10),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X1) | sP0(X0,X7,X2,sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f971,plain,(
  ( ! [X6,X4,X14,X12,X10,X8,X7,X5,X3,X13,X11,X9] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X3,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X3),X10),s(X3,X11),s(fun('type.num.num',X3),X12),s(fun(bool,X3),X13),s(X3,X14))) = s(X3,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X4),s(fun('type.num.num',X3),X5),s(X3,X6),s(fun('type.num.num',X3),X7),s(fun(bool,X3),X8),s(X3,X9))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X9,X8,X7,X6,X5,sK3,X4,X10,X12,X13,X11,X14,X3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X9,X8,X7,X6,X5,sK3,X4,X10,X12,X13,X11,X14,X3)))) | sP0(X3,X10,X5,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f963,f222])).
fof(f963,plain,(
  ( ! [X6,X4,X14,X12,X10,X8,X7,X5,X3,X13,X11,X9] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s(X3,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X3),X10),s(X3,X11),s(fun('type.num.num',X3),X12),s(fun(bool,X3),X13),s(X3,X14))) = s(X3,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X4),s(fun('type.num.num',X3),X5),s(X3,X6),s(fun('type.num.num',X3),X7),s(fun(bool,X3),X8),s(X3,X9))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X9,X8,X7,X6,X5,sK3,X4,X10,X12,X13,X11,X14,X3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X9,X8,X7,X6,X5,sK3,X4,X10,X12,X13,X11,X14,X3)))) | sP1(X3,X11,X6,sK3) | sP0(X3,X10,X5,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X4)) ) | $spl3),
  inference(superposition,[],[f230,f108])).
fof(f231,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | $spl2),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f961,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f681,f102])).
fof(f681,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
fof(f960,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f682,f99])).
fof(f682,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
fof(f799,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f702,f99])).
fof(f702,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) | $spl2),
  inference(backward_demodulation,[],[f231,f103])).
fof(f655,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f654,f233])).
fof(f654,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f653,f233])).
fof(f653,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f652,f233])).
fof(f652,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f651,f237])).
fof(f651,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(forward_demodulation,[],[f480,f233])).
fof(f480,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f122,f233])).
fof(f650,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f649,f233])).
fof(f649,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f648,f233])).
fof(f648,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f647,f233])).
fof(f647,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f646,f233])).
fof(f646,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f479,f237])).
fof(f479,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f122,f233])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl20),
  inference(cnf_transformation,[],[f533_D])).
fof(f533_D,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f442,plain,(
  ( ! [X14,X12,X13] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ~sP1(X12,X13,X14,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl4),
  inference(superposition,[],[f104,f233])).
fof(f565,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f564,f233])).
fof(f564,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f563,f233])).
fof(f563,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f562,f233])).
fof(f562,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f561,f237])).
fof(f561,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(forward_demodulation,[],[f439,f233])).
fof(f439,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f233,f122])).
fof(f554,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f553,f233])).
fof(f553,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f552,f233])).
fof(f552,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f551,f233])).
fof(f551,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f550,f233])).
fof(f550,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f438,f237])).
fof(f438,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8('const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))))) | $spl4),
  inference(superposition,[],[f233,f122])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | ~sP1(X0,X1,X2,'const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3))))) ) | $spl4),
  inference(superposition,[],[f233,f104])).
fof(f430,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(subsumption_resolution,[],[f420,f237])).
fof(f420,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(superposition,[],[f237,f122])).
fof(f421,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(trivial_inequality_removal,[],[f419])).
fof(f419,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') != s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3)))) | $spl9),
  inference(superposition,[],[f237,f122])).
fof(f341,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f252,f102])).
fof(f252,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) | $spl8),
  inference(backward_demodulation,[],[f238,f103])).
fof(f243,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X36),X43),s(X36,X44),s(fun('type.num.num',X36),X45),s(fun(bool,X36),X46),s(X36,X47))) = s(X36,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X37),s(fun('type.num.num',X36),X38),s(X36,X39),s(fun('type.num.num',X36),X40),s(fun(bool,X36),X41),s(X36,X42))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s(X36,happ(s(fun(bool,X36),X41),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun(bool,X36),X46),s(bool,sK6(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) | s(X36,happ(s(fun('type.num.num',X36),X40),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) != s(X36,happ(s(fun('type.num.num',X36),X45),s('type.num.num',sK5(X42,X41,X40,X39,X38,sK3,X37,X43,X45,X46,X44,X47,X36)))) | sP0(X36,X43,X38,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X37)) )),
  inference(resolution,[],[f222,f111])).
fof(f242,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X24),X31),s(X24,X32),s(fun('type.num.num',X24),X33),s(fun(bool,X24),X34),s(X24,X35))) = s(X24,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X25),s(fun('type.num.num',X24),X26),s(X24,X27),s(fun('type.num.num',X24),X28),s(fun(bool,X24),X29),s(X24,X30))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s(X24,happ(s(fun(bool,X24),X29),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) != s(X24,happ(s(fun(bool,X24),X34),s(bool,sK6(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X30,X29,X28,X27,X26,sK3,X25,X31,X33,X34,X32,X35,X24)))) | sP0(X24,X31,X26,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X25)) )),
  inference(resolution,[],[f222,f110])).
fof(f241,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X12),X19),s(X12,X20),s(fun('type.num.num',X12),X21),s(fun(bool,X12),X22),s(X12,X23))) = s(X12,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X13),s(fun('type.num.num',X12),X14),s(X12,X15),s(fun('type.num.num',X12),X16),s(fun(bool,X12),X17),s(X12,X18))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) | s(X12,happ(s(fun('type.num.num',X12),X16),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) != s(X12,happ(s(fun('type.num.num',X12),X21),s('type.num.num',sK5(X18,X17,X16,X15,X14,sK3,X13,X19,X21,X22,X20,X23,X12)))) | sP0(X12,X19,X14,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X13)) )),
  inference(resolution,[],[f222,f109])).
fof(f240,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',sK3),s(fun('type.num.num',X0),X7),s(X0,X8),s(fun('type.num.num',X0),X9),s(fun(bool,X0),X10),s(X0,X11))) = s(X0,'const.bvp_to_wordProps.tag_CASE_6'(s('type.bvp_to_wordProps.tag',X1),s(fun('type.num.num',X0),X2),s(X0,X3),s(fun('type.num.num',X0),X4),s(fun(bool,X0),X5),s(X0,X6))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK6(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK5(X6,X5,X4,X3,X2,sK3,X1,X7,X9,X10,X8,X11,X0)))) | sP0(X0,X7,X2,sK3) | s('type.bvp_to_wordProps.tag',sK3) != s('type.bvp_to_wordProps.tag',X1)) )),
  inference(resolution,[],[f222,f108])).
fof(f236,plain,(
  s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f211,f99])).
fof(f211,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0')))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
fof(f227,plain,(
  s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(subsumption_resolution,[],[f210,f102])).
fof(f210,plain,(
  ~p(s(bool,happ(s(fun('type.bvp_to_wordProps.tag',bool),sK2),s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.Word64Tag_0')))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.NumTag_1'(s(bool,sK10(sK3)))) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BytesTag_1'(s('type.num.num',sK9(sK3)))) | s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.RefTag_0') = s('type.bvp_to_wordProps.tag',sK3) | s('type.bvp_to_wordProps.tag',sK3) = s('type.bvp_to_wordProps.tag','const.bvp_to_wordProps.BlockTag_1'(s('type.num.num',sK8(sK3))))),
  inference(superposition,[],[f103,f122])).
% SZS output end Proof for thm.bvp_to_wordProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 81
% Passive clauses: 588
% Generated clauses: 3192
% Final active clauses: 59
% Final passive clauses: 174
% Input formulas: 37
% Initial clauses: 55
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 23
% Trivial inequalities: 20
% Fw subsumption resolutions: 444
% Fw demodulations: 1045
% Bw demodulations: 7
% 
% Simple tautologies: 37
% Equational tautologies: 1
% Forward subsumptions: 488
% 
% Binary resolution: 92
% Forward superposition: 498
% Backward superposition: 981
% 
% Split clauses: 82
% Split components: 193
% SAT solver clauses: 92
% SAT solver unit clauses: 5
% SAT solver binary clauses: 7
% 
% Sat splits: 82
% Sat splitting refutations: 9
% 
% Memory used [KB]: 4221
% Time elapsed: 0.149 s
% ------------------------------
% ------------------------------
% Success in time 11.9 s
