Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.typeSysProps
% (69766)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 1925
% Generated clauses: 21974
% Final active clauses: 89
% Final passive clauses: 1631
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 1930
% Trivial inequalities: 262
% Fw subsumption resolutions: 2051
% Fw demodulations: 1934
% Bw demodulations: 12
% 
% Simple tautologies: 455
% Equational tautologies: 48
% Forward subsumptions: 9977
% 
% Binary resolution: 17
% Forward superposition: 7510
% Backward superposition: 7352
% Self superposition: 307
% Equality factoring: 177
% Equality resolution: 12
% 
% Split clauses: 1122
% Split components: 2265
% SAT solver clauses: 1893
% SAT solver unit clauses: 64
% SAT solver binary clauses: 426
% 
% Sat splits: 1122
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1887
% TWLsolver calls for satisfiability: 150
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.typeSysProps
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.typeSysProps
% SZS output start Proof for thm.typeSysProps
fof(f835,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f833,f155,f157,f83,f130,f82,f132,f767,f786,f694,f699,f679,f690,f667,f668,f81,f134,f84,f138,f122,f154,f167,f169,f80,f136,f86,f140,f88,f141,f768,f787])).
fof(f787,plain,(
  ( ! [X2] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X2),s('type.typeSystem.tenv_val',sK2)))) ) | $spl124),
  inference(cnf_transformation,[],[f787_D])).
fof(f787_D,plain,(
  ( ! [X2] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X2),s('type.typeSystem.tenv_val',sK2)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f768,plain,(
  ( ! [X2] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X2),s('type.typeSystem.tenv_val',sK2)))) ) | ($spl2 | $spl120)),
  inference(superposition,[],[f699,f132])).
fof(f141,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f88,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.typeSysProps.typeSysProps_77_prove_ltb',aHL_BOOL_CASES)).
fof(f140,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f86,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.typeSysProps.typeSysProps_77_prove_ltb',aHL_FALSITY)).
fof(f136,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2)))) | $spl6),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f80,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))))),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK2))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4))))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))) => s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X8)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X8))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10)))))) => s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num',X6),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X8)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X8))) & s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10)))))) => s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num',X6),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10))))))),
  file('thm.typeSysProps.typeSysProps_77_prove_ltb','thm.typeSysProps.typeSysProps_77_prove_ltb')).
fof(f169,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))) | $spl30),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f167,plain,(
  s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))) | ($spl6 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f162,f140])).
fof(f162,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))) | ($spl6 | $spl12)),
  inference(superposition,[],[f136,f141])).
fof(f154,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2))))) ) | $spl22),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f122,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) | ((s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))) & (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2))) | (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) <=> s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2)))) | (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) <=> s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X2)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X2)))) => (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4))))) <=> s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X3),s('type.num.num',X0),s('type.typeSystem.tenv_val',X2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X4)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',X8)))) & s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',X8)))) => (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10))))) <=> s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X9),s('type.num.num',X6),s('type.typeSystem.tenv_val',X8))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X7),s('type.ast.t',X10)))))))),
  file('thm.typeSysProps.typeSysProps_77_prove_ltb','thm.typeSysProps.lookup_tenv_val_inc_tvs')).
fof(f138,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f138_D])).
fof(f138_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f84,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.typeSysProps.typeSysProps_77_prove_ltb',aHL_TRUTH)).
fof(f134,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK2))) | $spl4),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f81,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.typeSystem.num_tvs_1'(s('type.typeSystem.tenv_val',sK2)))),
  inference(cnf_transformation,[],[f58])).
fof(f668,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2)))))) ) | $spl108),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2)))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f667,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2)))))) ) | ($spl4 | $spl8 | $spl22 | $spl30)),
  inference(subsumption_resolution,[],[f666,f138])).
fof(f666,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2)))))) ) | ($spl4 | $spl22 | $spl30)),
  inference(forward_demodulation,[],[f658,f169])).
fof(f658,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))))) ) | ($spl4 | $spl22)),
  inference(trivial_inequality_removal,[],[f657])).
fof(f657,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X3),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',X1),s('type.ast.t',X2))))) | ~p(s(bool,'const.typeSoundInvariants.tenv_val_ok_1'(s('type.typeSystem.tenv_val',sK2))))) ) | ($spl4 | $spl22)),
  inference(superposition,[],[f154,f134])).
fof(f690,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) | $spl116),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f679,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) | s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),X0),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) | ($spl2 | $spl108)),
  inference(superposition,[],[f668,f132])).
fof(f699,plain,(
  ( ! [X0] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) | $spl120),
  inference(cnf_transformation,[],[f699_D])).
fof(f699_D,plain,(
  ( ! [X0] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f694,plain,(
  ( ! [X0] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))) ) | $spl116),
  inference(equality_resolution,[],[f690])).
fof(f786,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2)))) ) | $spl122),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f767,plain,(
  ( ! [X0,X1] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X0),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X1),s('type.typeSystem.tenv_val',sK2)))) ) | $spl120),
  inference(superposition,[],[f699,f699])).
fof(f132,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4))))) | $spl2),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f82,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) = s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))),
  inference(cnf_transformation,[],[f58])).
fof(f130,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4))))) | $spl1),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f83,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.option.SOME_1'(s('type.pair.prod'('type.num.num','type.ast.t'),'const.pair.,_2'(s('type.num.num',sK1),s('type.ast.t',sK4)))))),
  inference(cnf_transformation,[],[f58])).
fof(f157,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) | $spl25),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f155,plain,(
  s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',sK0),s('type.typeSystem.tenv_val',sK2))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f130,f132])).
fof(f833,plain,(
  $false | ($spl25 | $spl122 | $spl124)),
  inference(subsumption_resolution,[],[f807,f787])).
fof(f807,plain,(
  ( ! [X7] : (s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num','const.num.|48|_0'),s('type.typeSystem.tenv_val',sK2))) != s('type.option.option'('type.pair.prod'('type.num.num','type.ast.t')),'const.typeSystem.lookup_tenv_val_3'(s('type.list.list'('type.string.char'),sK3),s('type.num.num',X7),s('type.typeSystem.tenv_val',sK2)))) ) | ($spl25 | $spl122)),
  inference(superposition,[],[f157,f786])).
% SZS output end Proof for thm.typeSysProps
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 25
% Passive clauses: 67
% Generated clauses: 730
% Final active clauses: 24
% Final passive clauses: 38
% Input formulas: 21
% Initial clauses: 44
% 
% Duplicate literals: 62
% Trivial inequalities: 6
% Fw subsumption resolutions: 6
% Fw demodulations: 40
% Bw demodulations: 4
% Condensations: 81
% Global subsumptions: 8
% 
% Simple tautologies: 38
% Equational tautologies: 14
% Forward subsumptions: 292
% 
% Binary resolution: 3
% Forward superposition: 197
% Backward superposition: 144
% Self superposition: 36
% Equality factoring: 26
% Equality resolution: 2
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 179
% SAT solver unit clauses: 62
% SAT solver binary clauses: 17
% SAT solver learnt clauses: 17
% 
% Sat splits: 1
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 145
% TWLsolver calls for satisfiability: 137
% 
% Memory used [KB]: 1535
% Time elapsed: 0.054 s
% ------------------------------
% ------------------------------
% Success in time 0.71 s
