Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.wordSem
% (10307)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 120
% Passive clauses: 2643
% Generated clauses: 16059
% Final active clauses: 110
% Final passive clauses: 2395
% Input formulas: 37
% Initial clauses: 69
% 
% Duplicate literals: 529
% Trivial inequalities: 278
% Fw subsumption resolutions: 1614
% Fw demodulations: 2130
% Bw demodulations: 9
% 
% Simple tautologies: 217
% Equational tautologies: 43
% Forward subsumptions: 8336
% 
% Binary resolution: 191
% Factoring: 1
% Forward superposition: 5630
% Backward superposition: 5240
% Self superposition: 186
% Equality factoring: 98
% Equality resolution: 27
% 
% Split clauses: 187
% Split components: 389
% SAT solver clauses: 392
% SAT solver unit clauses: 13
% SAT solver binary clauses: 81
% 
% Sat splits: 187
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 392
% TWLsolver calls for satisfiability: 110
% 
% Memory used [KB]: 12537
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.wordSem
% (10330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 792
% Generated clauses: 15865
% Final active clauses: 132
% Final passive clauses: 652
% Input formulas: 37
% Initial clauses: 69
% 
% Duplicate literals: 157
% Trivial inequalities: 214
% Fw subsumption resolutions: 783
% Fw demodulations: 1558
% Bw demodulations: 8
% Condensations: 602
% Global subsumptions: 4
% 
% Simple tautologies: 106
% Equational tautologies: 13
% Forward subsumptions: 10372
% 
% Binary resolution: 70
% Factoring: 1
% Forward superposition: 5782
% Backward superposition: 5640
% Self superposition: 156
% Equality factoring: 20
% Equality resolution: 6
% 
% Split clauses: 54
% Split components: 109
% SAT solver clauses: 1347
% SAT solver unit clauses: 109
% SAT solver binary clauses: 26
% SAT solver learnt clauses: 8
% 
% Sat splits: 54
% 
% TWLsolver clauses: 1331
% TWLsolver calls for satisfiability: 976
% 
% Memory used [KB]: 22771
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.wordSem
% (10364)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 107
% Passive clauses: 458
% Generated clauses: 8064
% Final active clauses: 101
% Final passive clauses: 351
% Input formulas: 37
% Initial clauses: 69
% 
% Duplicate literals: 13
% Trivial inequalities: 436
% Fw subsumption resolutions: 102
% Fw demodulations: 1603
% Bw demodulations: 6
% Condensations: 27
% 
% Simple tautologies: 38
% Equational tautologies: 23
% Forward subsumptions: 5364
% 
% Unit resulting resolution: 899
% Factoring: 1
% Forward superposition: 2409
% Backward superposition: 2401
% Self superposition: 68
% Equality factoring: 2
% Equality resolution: 28
% 
% SAT solver clauses: 460
% SAT solver unit clauses: 157
% SAT solver binary clauses: 268
% 
% Memory used [KB]: 200082
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.wordSem
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.wordSem
% SZS output start Proof for thm.wordSem
fof(f3601,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f289,f295,f298,f299,f301,f302,f304,f306,f308,f311,f316,f291,f317,f459,f722,f731,f734,f738,f745,f751,f755,f759,f727,f765,f770,f773,f724,f735,f779,f784,f787,f726,f739,f866,f867,f865,f849,f850,f868,f864,f853,f870,f863,f856,f872,f873,f861,f906,f908,f909,f910,f911,f914,f918,f922,f924,f925,f293,f948,f951,f952,f953,f945,f954,f933,f934,f869,f956,f938,f959,f960,f961,f297,f1104,f1248,f1250,f294,f1252,f300,f1254,f303,f1255,f1256,f307,f1235,f871,f1236,f957,f1258,f1260,f950,f1262,f723,f1286,f1288,f1287,f1289,f1290,f725,f1309,f1310,f949,f1369,f1372,f1375,f2316,f2320,f2324,f2504,f2509,f2514,f315,f313,f2873,f3360])).
fof(f3360,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3359,f150])).
fof(f150,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.wordSem.result_induction',aHL_FALSITY)).
fof(f3359,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(forward_demodulation,[],[f3129,f259])).
fof(f259,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4)))),
  inference(resolution,[],[f109,f147])).
fof(f147,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.wordSem.result_induction','thm.bool.EQ_CLAUSES|split|3')).
fof(f109,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4))))),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X5,X6] : p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),X5),s('type.wordSem.word_loc'(sK2),X6)))))) & ! [X3,X4] : p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),X3),s('type.wordSem.word_loc'(sK2),X4)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) & ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3,sK4])],[f86])).
fof(f86,plain,(
  ? [X0,X1] : (! [X5,X6] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X5),s('type.wordSem.word_loc'(X0),X6)))))) & ! [X3,X4] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X3),s('type.wordSem.word_loc'(X0),X4)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Error_0')))) & ? [X2] : ~p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),X2)))))),
  inference(rectify,[],[f78])).
fof(f78,plain,(
  ? [X0,X1] : (! [X4,X5] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Error_0')))) & ? [X6] : ~p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),X6)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ? [X0,X1] : ((! [X4,X5] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Error_0'))))) & ? [X6] : ~p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),X6)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1] : ((! [X4,X5] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Error_0'))))) => ! [X6] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),X6)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1] : ((! [X4,X5] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5)))))) & ! [X2,X3] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),'const.wordSem.Error_0'))))) => ! [X6] : p(s(bool,happ(s(fun('type.wordSem.result'(X0),bool),X1),s('type.wordSem.result'(X0),X6)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : ((! [X8,X9] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9)))))) & ! [X8,X9] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Error_0'))))) => ! [X10] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : ((! [X8,X9] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9)))))) & ! [X8,X9] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9)))))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0')))) & p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),'const.wordSem.Error_0'))))) => ! [X10] : p(s(bool,happ(s(fun('type.wordSem.result'(X6),bool),X7),s('type.wordSem.result'(X6),X10)))))),
  file('thm.wordSem.result_induction','thm.wordSem.result_induction')).
fof(f3129,plain,(
  p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4)))) | $spl14),
  inference(superposition,[],[f104,f313])).
fof(f104,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),X5),s('type.wordSem.word_loc'(sK2),X6))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f2873,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2872,f150])).
fof(f2872,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(forward_demodulation,[],[f2642,f259])).
fof(f2642,plain,(
  p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4)))) | $spl16),
  inference(superposition,[],[f105,f315])).
fof(f105,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),X3),s('type.wordSem.word_loc'(sK2),X4))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f313,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl14),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f315,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | $spl16),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2514,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2513,f292])).
fof(f292,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),sK4) | $spl3),
  inference(cnf_transformation,[],[f292_D])).
fof(f292_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),sK4) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f2513,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl32)),
  inference(forward_demodulation,[],[f2512,f949])).
fof(f2512,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f2511,f1759])).
fof(f1759,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f1758,f162])).
fof(f162,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))),
  inference(resolution,[],[f106,f149])).
fof(f149,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.wordSem.result_induction','thm.bool.EQ_CLAUSES|split|0')).
fof(f106,plain,(
  p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'))))),
  inference(cnf_transformation,[],[f87])).
fof(f1758,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))),
  inference(subsumption_resolution,[],[f1639,f160])).
fof(f160,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f139])).
fof(f139,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11] : ! [X22] : (s(X11,X22) = s(X11,X22) <=> p(s(bool,'T_0')))),
  file('thm.wordSem.result_induction','thm.bool.REFL_CLAUSE')).
fof(f1639,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))),
  inference(superposition,[],[f144,f162])).
fof(f144,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.wordSem.result_induction','thm.bool.EQ_CLAUSES|split|2')).
fof(f2511,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(forward_demodulation,[],[f2510,f162])).
fof(f2510,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f2340,f292])).
fof(f2340,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f259,f128])).
fof(f128,plain,(
  ( ! [X0,X1] : (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),sK12(X1,X0)),s('type.wordSem.word_loc'(X0),sK13(X1,X0)))) | s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),sK10(X1,X0)),s('type.wordSem.word_loc'(X0),sK11(X1,X0))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),sK10(X1,X0)),s('type.wordSem.word_loc'(X0),sK11(X1,X0)))) | s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),sK12(X1,X0)),s('type.wordSem.word_loc'(X0),sK13(X1,X0)))) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (? [X4,X5] : s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5))) | ? [X2,X3] : s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X1))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : (? [X4,X5] : s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X4),s('type.wordSem.word_loc'(X0),X5))) | ? [X2,X3] : s('type.wordSem.result'(X0),X1) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X2),s('type.wordSem.word_loc'(X0),X3))) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X1) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X1))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X38] : (? [X8,X9] : s('type.wordSem.result'(X6),X38) = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9))) | ? [X8,X9] : s('type.wordSem.result'(X6),X38) = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X8),s('type.wordSem.word_loc'(X6),X9))) | s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),X38) | s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),X38) | s('type.wordSem.result'(X6),'const.wordSem.Error_0') = s('type.wordSem.result'(X6),X38))),
  file('thm.wordSem.result_induction','thm.wordSem.result_nchotomy')).
fof(f2509,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2508,f292])).
fof(f2508,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f2507,f949])).
fof(f2507,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f2506,f1759])).
fof(f2506,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(forward_demodulation,[],[f2505,f189])).
fof(f189,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))),
  inference(resolution,[],[f107,f149])).
fof(f107,plain,(
  p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))))),
  inference(cnf_transformation,[],[f87])).
fof(f2505,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f2339,f290])).
fof(f290,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),sK4) | $spl1),
  inference(cnf_transformation,[],[f290_D])).
fof(f290_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),sK4) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2339,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f259,f128])).
fof(f2504,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2503,f1759])).
fof(f2503,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f2502,f189])).
fof(f2502,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f2501,f949])).
fof(f2501,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f2500,f290])).
fof(f2500,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f2338,f292])).
fof(f2338,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f259,f128])).
fof(f2324,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2323,f292])).
fof(f2323,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl32)),
  inference(forward_demodulation,[],[f2322,f949])).
fof(f2322,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f2321,f292])).
fof(f2321,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f2269,f162])).
fof(f2269,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f258,f128])).
fof(f258,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),sK4)))),
  inference(resolution,[],[f109,f148])).
fof(f148,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f2320,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2319,f292])).
fof(f2319,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl32)),
  inference(forward_demodulation,[],[f2318,f949])).
fof(f2318,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f2317,f290])).
fof(f2317,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f2268,f189])).
fof(f2268,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f258,f128])).
fof(f2316,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(subsumption_resolution,[],[f2315,f189])).
fof(f2315,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl32)),
  inference(forward_demodulation,[],[f2314,f949])).
fof(f2314,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f2313,f290])).
fof(f2313,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f2267,f292])).
fof(f2267,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f258,f128])).
fof(f1375,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1374,f292])).
fof(f1374,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(subsumption_resolution,[],[f1373,f296])).
fof(f296,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),sK4) | $spl7),
  inference(cnf_transformation,[],[f296_D])).
fof(f296_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),sK4) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1373,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f1331,f106])).
fof(f1331,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f1372,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f1371,f290])).
fof(f1371,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(subsumption_resolution,[],[f1370,f296])).
fof(f1370,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f1330,f107])).
fof(f1330,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f1369,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f1368,f290])).
fof(f1368,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f1367,f292])).
fof(f1367,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f1329,f108])).
fof(f108,plain,(
  p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0'))))),
  inference(cnf_transformation,[],[f87])).
fof(f1329,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f949,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | $spl32),
  inference(cnf_transformation,[],[f949_D])).
fof(f949_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1310,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X96] : (s(X89,X99) = s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X90),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X92),s(X89,X93),s(X89,X94),s(X89,X95))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),X90) | sP0(X89,sK2,X96,X91,'const.wordSem.NotEnoughSpace_0') | sP1(X89,sK2,X97,X92,'const.wordSem.NotEnoughSpace_0') | s(X89,X94) != s(X89,X99)) ) | $spl40),
  inference(cnf_transformation,[],[f1310_D])).
fof(f1310_D,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X96] : (s(X89,X99) = s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X90),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X92),s(X89,X93),s(X89,X94),s(X89,X95))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),X90) | sP0(X89,sK2,X96,X91,'const.wordSem.NotEnoughSpace_0') | sP1(X89,sK2,X97,X92,'const.wordSem.NotEnoughSpace_0') | s(X89,X94) != s(X89,X99)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1309,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X96] : (s(X89,X99) = s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X90),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X92),s(X89,X93),s(X89,X94),s(X89,X95))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X89,X94) != s(X89,X99) | sP1(X89,sK2,X97,X92,'const.wordSem.NotEnoughSpace_0') | sP0(X89,sK2,X96,X91,'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),X90)) ) | $spl21),
  inference(forward_demodulation,[],[f1308,f130])).
fof(f130,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X5) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X11] : ! [X34,X13,X35,X36,X37] : s(X11,X36) = s(X11,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X34),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X13),s(X11,X35),s(X11,X36),s(X11,X37)))),
  file('thm.wordSem.result_induction','thm.wordSem.result_case_def|split|3')).
fof(f1308,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X100,X98,X96] : (s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X96),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X97),s(X89,X98),s(X89,X99),s(X89,X100))) = s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X90),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X92),s(X89,X93),s(X89,X94),s(X89,X95))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X89,X94) != s(X89,X99) | sP1(X89,sK2,X97,X92,'const.wordSem.NotEnoughSpace_0') | sP0(X89,sK2,X96,X91,'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),X90)) ) | $spl21),
  inference(trivial_inequality_removal,[],[f1300])).
fof(f1300,plain,(
  ( ! [X94,X92,X90,X99,X97,X95,X93,X91,X89,X100,X98,X96] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X96),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X97),s(X89,X98),s(X89,X99),s(X89,X100))) = s(X89,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X90),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X89)),X92),s(X89,X93),s(X89,X94),s(X89,X95))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X89,X94) != s(X89,X99) | sP1(X89,sK2,X97,X92,'const.wordSem.NotEnoughSpace_0') | sP0(X89,sK2,X96,X91,'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),X90)) ) | $spl21),
  inference(superposition,[],[f725,f116])).
fof(f116,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) | s(X1,X3) != s(X1,X12) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8) | sP0(X1,X0,X6,X9,X8) | sP1(X1,X0,X5,X10,X8) | (s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) & s(X1,X4) != s(X1,X11)) | (s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) & s(X1,X3) != s(X1,X12)) | (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) & s(X1,X2) != s(X1,X13)) | s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))))),
  inference(definition_folding,[],[f80,f84,f83])).
fof(f83,plain,(
  ! [X1,X0,X6,X9,X8] : (~sP0(X1,X0,X6,X9,X8) | ? [X16,X17] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X16),s('type.wordSem.word_loc'(X0),X17))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f84,plain,(
  ! [X1,X0,X5,X10,X8] : (~sP1(X1,X0,X5,X10,X8) | ? [X14,X15] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X14),s('type.wordSem.word_loc'(X0),X15))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8) | ? [X16,X17] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X16),s('type.wordSem.word_loc'(X0),X17))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17)))) | ? [X14,X15] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X14),s('type.wordSem.word_loc'(X0),X15))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15)))) | (s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) & s(X1,X4) != s(X1,X11)) | (s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) & s(X1,X3) != s(X1,X12)) | (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) & s(X1,X2) != s(X1,X13)) | s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : ((s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8) | ? [X16,X17] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X16),s('type.wordSem.word_loc'(X0),X17))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17)))) | ? [X14,X15] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X14),s('type.wordSem.word_loc'(X0),X15))) & s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15))) != s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15)))) | (s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) & s(X1,X4) != s(X1,X11)) | (s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) & s(X1,X3) != s(X1,X12)) | (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) & s(X1,X2) != s(X1,X13))) | s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : ((s('type.wordSem.result'(X0),X7) = s('type.wordSem.result'(X0),X8) & ! [X16,X17] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X16),s('type.wordSem.word_loc'(X0),X17))) => s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17))) = s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17)))) & ! [X14,X15] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X14),s('type.wordSem.word_loc'(X0),X15))) => s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15))) = s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15)))) & (s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) => s(X1,X4) = s(X1,X11)) & (s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) => s(X1,X3) = s(X1,X12)) & (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) => s(X1,X2) = s(X1,X13))) => s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : ((s('type.wordSem.result'(X0),X7) = s('type.wordSem.result'(X0),X8) & ! [X16,X17] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X0),X16),s('type.wordSem.word_loc'(X0),X17))) => s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17))) = s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s('type.wordSem.word_loc'(X0),X16))),s('type.wordSem.word_loc'(X0),X17)))) & ! [X14,X15] : (s('type.wordSem.result'(X0),X8) = s('type.wordSem.result'(X0),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X0),X14),s('type.wordSem.word_loc'(X0),X15))) => s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15))) = s(X1,happ(s(fun('type.wordSem.word_loc'(X0),X1),happ(s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s('type.wordSem.word_loc'(X0),X14))),s('type.wordSem.word_loc'(X0),X15)))) & (s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) => s(X1,X4) = s(X1,X11)) & (s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) => s(X1,X3) = s(X1,X12)) & (s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) => s(X1,X2) = s(X1,X13))) => s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X11] : ! [X27,X28,X29,X30,X31,X32,X33,X34,X13,X35,X36,X37] : ((s('type.wordSem.result'(X6),X32) = s('type.wordSem.result'(X6),X33) & ! [X18,X19] : (s('type.wordSem.result'(X6),X33) = s('type.wordSem.result'(X6),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(X6),X18),s('type.wordSem.word_loc'(X6),X19))) => s(X11,happ(s(fun('type.wordSem.word_loc'(X6),X11),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X31),s('type.wordSem.word_loc'(X6),X18))),s('type.wordSem.word_loc'(X6),X19))) = s(X11,happ(s(fun('type.wordSem.word_loc'(X6),X11),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X34),s('type.wordSem.word_loc'(X6),X18))),s('type.wordSem.word_loc'(X6),X19)))) & ! [X18,X19] : (s('type.wordSem.result'(X6),X33) = s('type.wordSem.result'(X6),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(X6),X18),s('type.wordSem.word_loc'(X6),X19))) => s(X11,happ(s(fun('type.wordSem.word_loc'(X6),X11),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X13),s('type.wordSem.word_loc'(X6),X18))),s('type.wordSem.word_loc'(X6),X19))) = s(X11,happ(s(fun('type.wordSem.word_loc'(X6),X11),happ(s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X30),s('type.wordSem.word_loc'(X6),X18))),s('type.wordSem.word_loc'(X6),X19)))) & (s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X6),X33) => s(X11,X29) = s(X11,X35)) & (s('type.wordSem.result'(X6),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X6),X33) => s(X11,X28) = s(X11,X36)) & (s('type.wordSem.result'(X6),'const.wordSem.Error_0') = s('type.wordSem.result'(X6),X33) => s(X11,X27) = s(X11,X37))) => s(X11,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),X32),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X34),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X13),s(X11,X35),s(X11,X36),s(X11,X37))) = s(X11,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),X33),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X31),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X30),s(X11,X29),s(X11,X28),s(X11,X27))))),
  file('thm.wordSem.result_induction','thm.wordSem.result_case_cong')).
fof(f725,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | $spl21),
  inference(cnf_transformation,[],[f725_D])).
fof(f725_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1290,plain,(
  ( ! [X177,X175,X173,X182,X180,X178,X176,X184,X174,X181,X179] : (s(X173,X184) = s(X173,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X174),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X175),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X176),s(X173,X177),s(X173,X178),s(X173,X179))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X174) | sP0(X173,sK2,X180,X175,'const.wordSem.Error_0') | sP1(X173,sK2,X181,X176,'const.wordSem.Error_0') | s(X173,X177) != s(X173,X182) | s(X173,X179) != s(X173,X184)) ) | $spl38),
  inference(cnf_transformation,[],[f1290_D])).
fof(f1290_D,plain,(
  ( ! [X177,X175,X173,X182,X180,X178,X176,X184,X174,X181,X179] : (s(X173,X184) = s(X173,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X174),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X175),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X176),s(X173,X177),s(X173,X178),s(X173,X179))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X174) | sP0(X173,sK2,X180,X175,'const.wordSem.Error_0') | sP1(X173,sK2,X181,X176,'const.wordSem.Error_0') | s(X173,X177) != s(X173,X182) | s(X173,X179) != s(X173,X184)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1289,plain,(
  ( ! [X177,X175,X173,X182,X180,X178,X176,X184,X174,X181,X179] : (s(X173,X184) = s(X173,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X174),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X175),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X176),s(X173,X177),s(X173,X178),s(X173,X179))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X173,X179) != s(X173,X184) | s(X173,X177) != s(X173,X182) | sP1(X173,sK2,X181,X176,'const.wordSem.Error_0') | sP0(X173,sK2,X180,X175,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X174)) ) | $spl19),
  inference(forward_demodulation,[],[f1278,f127])).
fof(f127,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X6) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X11] : ! [X34,X13,X35,X36,X37] : s(X11,X37) = s(X11,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X34),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X13),s(X11,X35),s(X11,X36),s(X11,X37)))),
  file('thm.wordSem.result_induction','thm.wordSem.result_case_def|split|4')).
fof(f1278,plain,(
  ( ! [X177,X175,X173,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X173,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X180),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X181),s(X173,X182),s(X173,X183),s(X173,X184))) = s(X173,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X174),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X175),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X173)),X176),s(X173,X177),s(X173,X178),s(X173,X179))) | s(X173,X179) != s(X173,X184) | s(X173,X177) != s(X173,X182) | sP1(X173,sK2,X181,X176,'const.wordSem.Error_0') | sP0(X173,sK2,X180,X175,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X174)) ) | $spl19),
  inference(superposition,[],[f723,f119])).
fof(f119,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s(X1,X2) != s(X1,X13) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) | s(X1,X4) != s(X1,X11) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f1287,plain,(
  ( ! [X156,X154,X152,X151,X149,X157,X155,X153,X160,X150] : (s(X149,X160) = s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X150),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X151),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X152),s(X149,X153),s(X149,X154),s(X149,X155))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X150) | sP0(X149,sK2,X156,X151,'const.wordSem.Error_0') | sP1(X149,sK2,X157,X152,'const.wordSem.Error_0') | s(X149,X155) != s(X149,X160)) ) | $spl36),
  inference(cnf_transformation,[],[f1287_D])).
fof(f1287_D,plain,(
  ( ! [X156,X154,X152,X151,X149,X157,X155,X153,X160,X150] : (s(X149,X160) = s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X150),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X151),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X152),s(X149,X153),s(X149,X154),s(X149,X155))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X150) | sP0(X149,sK2,X156,X151,'const.wordSem.Error_0') | sP1(X149,sK2,X157,X152,'const.wordSem.Error_0') | s(X149,X155) != s(X149,X160)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1288,plain,(
  ( ! [X167,X165,X163,X161,X169,X166,X164,X162,X172,X168] : (s(X161,X172) = s(X161,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X162),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X163),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X164),s(X161,X165),s(X161,X166),s(X161,X167))) | s(X161,X167) != s(X161,X172) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | sP1(X161,sK2,X169,X164,'const.wordSem.Error_0') | sP0(X161,sK2,X168,X163,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X162)) ) | $spl19),
  inference(forward_demodulation,[],[f1284,f127])).
fof(f1284,plain,(
  ( ! [X167,X165,X163,X161,X171,X169,X166,X164,X162,X172,X170,X168] : (s(X161,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X168),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X169),s(X161,X170),s(X161,X171),s(X161,X172))) = s(X161,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X162),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X163),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X164),s(X161,X165),s(X161,X166),s(X161,X167))) | s(X161,X167) != s(X161,X172) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | sP1(X161,sK2,X169,X164,'const.wordSem.Error_0') | sP0(X161,sK2,X168,X163,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X162)) ) | $spl19),
  inference(trivial_inequality_removal,[],[f1277])).
fof(f1277,plain,(
  ( ! [X167,X165,X163,X161,X171,X169,X166,X164,X162,X172,X170,X168] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s(X161,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X168),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X169),s(X161,X170),s(X161,X171),s(X161,X172))) = s(X161,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X162),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X163),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X161)),X164),s(X161,X165),s(X161,X166),s(X161,X167))) | s(X161,X167) != s(X161,X172) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | sP1(X161,sK2,X169,X164,'const.wordSem.Error_0') | sP0(X161,sK2,X168,X163,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X162)) ) | $spl19),
  inference(superposition,[],[f723,f118])).
fof(f118,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s(X1,X2) != s(X1,X13) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f1286,plain,(
  ( ! [X156,X154,X152,X151,X149,X157,X155,X153,X160,X150] : (s(X149,X160) = s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X150),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X151),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X152),s(X149,X153),s(X149,X154),s(X149,X155))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X149,X155) != s(X149,X160) | sP1(X149,sK2,X157,X152,'const.wordSem.Error_0') | sP0(X149,sK2,X156,X151,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X150)) ) | $spl19),
  inference(forward_demodulation,[],[f1285,f127])).
fof(f1285,plain,(
  ( ! [X158,X156,X154,X152,X151,X149,X159,X157,X155,X153,X160,X150] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X156),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X157),s(X149,X158),s(X149,X159),s(X149,X160))) = s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X150),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X151),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X152),s(X149,X153),s(X149,X154),s(X149,X155))) | s(X149,X155) != s(X149,X160) | sP1(X149,sK2,X157,X152,'const.wordSem.Error_0') | sP0(X149,sK2,X156,X151,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X150)) ) | $spl19),
  inference(subsumption_resolution,[],[f1276,f723])).
fof(f1276,plain,(
  ( ! [X158,X156,X154,X152,X151,X149,X159,X157,X155,X153,X160,X150] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.Error_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X156),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X157),s(X149,X158),s(X149,X159),s(X149,X160))) = s(X149,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X150),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X151),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X149)),X152),s(X149,X153),s(X149,X154),s(X149,X155))) | s(X149,X155) != s(X149,X160) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | sP1(X149,sK2,X157,X152,'const.wordSem.Error_0') | sP0(X149,sK2,X156,X151,'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),X150)) ) | $spl19),
  inference(superposition,[],[f723,f118])).
fof(f723,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | $spl19),
  inference(cnf_transformation,[],[f723_D])).
fof(f723_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1262,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1261,f292])).
fof(f1261,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(subsumption_resolution,[],[f1239,f296])).
fof(f1239,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(superposition,[],[f296,f128])).
fof(f950,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | $spl33),
  inference(cnf_transformation,[],[f950_D])).
fof(f950_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1260,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f1259,f290])).
fof(f1259,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(subsumption_resolution,[],[f1238,f296])).
fof(f1238,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(superposition,[],[f296,f128])).
fof(f1258,plain,(
  s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1257,f290])).
fof(f1257,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1240,f292])).
fof(f1240,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(trivial_inequality_removal,[],[f1237])).
fof(f1237,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl7),
  inference(superposition,[],[f296,f128])).
fof(f957,plain,(
  ( ! [X134,X132,X142,X140,X138,X136,X135,X133,X143,X141,X139,X137] : (s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X139),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X140),s(X132,X141),s(X132,X142),s(X132,X143))) = s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X133),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X134),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X135),s(X132,X136),s(X132,X137),s(X132,X138))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X133) | sP0(X132,sK2,X139,X134,sK4) | sP1(X132,sK2,X140,X135,sK4) | s(X132,X137) != s(X132,X142) | s(X132,X138) != s(X132,X143)) ) | $spl34),
  inference(cnf_transformation,[],[f957_D])).
fof(f957_D,plain,(
  ( ! [X134,X132,X142,X140,X138,X136,X135,X133,X143,X141,X139,X137] : (s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X139),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X140),s(X132,X141),s(X132,X142),s(X132,X143))) = s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X133),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X134),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X135),s(X132,X136),s(X132,X137),s(X132,X138))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X133) | sP0(X132,sK2,X139,X134,sK4) | sP1(X132,sK2,X140,X135,sK4) | s(X132,X137) != s(X132,X142) | s(X132,X138) != s(X132,X143)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1236,plain,(
  ( ! [X134,X132,X142,X140,X138,X136,X135,X133,X143,X141,X139,X137] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X139),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X140),s(X132,X141),s(X132,X142),s(X132,X143))) = s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X133),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X134),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X135),s(X132,X136),s(X132,X137),s(X132,X138))) | s(X132,X138) != s(X132,X143) | s(X132,X137) != s(X132,X142) | sP1(X132,sK2,X140,X135,sK4) | sP0(X132,sK2,X139,X134,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X133)) ) | $spl7),
  inference(superposition,[],[f296,f120])).
fof(f120,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s(X1,X2) != s(X1,X13) | s(X1,X3) != s(X1,X12) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f871,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X128),s(X120,X129),s(X120,X130),s(X120,X131))) = s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X121),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X123),s(X120,X124),s(X120,X125),s(X120,X126))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X121) | sP0(X120,sK2,X127,X122,sK4) | sP1(X120,sK2,X128,X123,sK4) | s(X120,X124) != s(X120,X129) | s(X120,X126) != s(X120,X131)) ) | $spl30),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X128),s(X120,X129),s(X120,X130),s(X120,X131))) = s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X121),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X123),s(X120,X124),s(X120,X125),s(X120,X126))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X121) | sP0(X120,sK2,X127,X122,sK4) | sP1(X120,sK2,X128,X123,sK4) | s(X120,X124) != s(X120,X129) | s(X120,X126) != s(X120,X131)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1235,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X128),s(X120,X129),s(X120,X130),s(X120,X131))) = s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X121),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X123),s(X120,X124),s(X120,X125),s(X120,X126))) | s(X120,X126) != s(X120,X131) | s(X120,X124) != s(X120,X129) | sP1(X120,sK2,X128,X123,sK4) | sP0(X120,sK2,X127,X122,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X121)) ) | $spl7),
  inference(superposition,[],[f296,f119])).
fof(f307,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97) | sP0(X96,sK2,X103,X98,sK4) | sP1(X96,sK2,X104,X99,sK4) | s(X96,X102) != s(X96,X107)) ) | $spl12),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97) | sP0(X96,sK2,X103,X98,sK4) | sP1(X96,sK2,X104,X99,sK4) | s(X96,X102) != s(X96,X107)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1256,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1234,f292])).
fof(f1234,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl7),
  inference(superposition,[],[f296,f118])).
fof(f1255,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f1233,f290])).
fof(f1233,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl7),
  inference(superposition,[],[f296,f118])).
fof(f303,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61) | sP0(X60,sK2,X67,X62,sK4) | sP1(X60,sK2,X68,X63,sK4) | s(X60,X65) != s(X60,X70)) ) | $spl10),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61) | sP0(X60,sK2,X67,X62,sK4) | sP1(X60,sK2,X68,X63,sK4) | s(X60,X65) != s(X60,X70)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1254,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl7),
  inference(subsumption_resolution,[],[f1231,f296])).
fof(f1231,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl7),
  inference(superposition,[],[f296,f116])).
fof(f300,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37) | sP0(X36,sK2,X43,X38,sK4) | sP1(X36,sK2,X44,X39,sK4) | s(X36,X40) != s(X36,X45)) ) | $spl8),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37) | sP0(X36,sK2,X43,X38,sK4) | sP1(X36,sK2,X44,X39,sK4) | s(X36,X40) != s(X36,X45)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1252,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) ) | $spl7),
  inference(subsumption_resolution,[],[f1229,f296])).
fof(f1229,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) ) | $spl7),
  inference(superposition,[],[f296,f115])).
fof(f115,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) | s(X1,X4) != s(X1,X11) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f294,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1) | sP0(X0,sK2,X7,X2,sK4) | sP1(X0,sK2,X8,X3,sK4)) ) | $spl4),
  inference(cnf_transformation,[],[f294_D])).
fof(f294_D,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1) | sP0(X0,sK2,X7,X2,sK4) | sP1(X0,sK2,X8,X3,sK4)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1250,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1249,f292])).
fof(f1249,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl7),
  inference(subsumption_resolution,[],[f1227,f296])).
fof(f1227,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl7),
  inference(superposition,[],[f296,f114])).
fof(f114,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) | s('type.wordSem.result'(X0),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(X0),X8) | s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(X0),X8) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f1248,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f1247,f290])).
fof(f1247,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl7),
  inference(subsumption_resolution,[],[f1226,f296])).
fof(f1226,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl7),
  inference(superposition,[],[f296,f114])).
fof(f1104,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f962,f108])).
fof(f962,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | $spl6),
  inference(backward_demodulation,[],[f297,f109])).
fof(f297,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | $spl6),
  inference(cnf_transformation,[],[f297_D])).
fof(f297_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f961,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f941,f292])).
fof(f941,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(superposition,[],[f292,f128])).
fof(f960,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f942,f290])).
fof(f942,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(trivial_inequality_removal,[],[f940])).
fof(f940,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(superposition,[],[f292,f128])).
fof(f959,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f958,f290])).
fof(f958,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(subsumption_resolution,[],[f939,f292])).
fof(f939,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl3),
  inference(superposition,[],[f292,f128])).
fof(f938,plain,(
  ( ! [X134,X132,X142,X140,X138,X136,X135,X133,X143,X141,X139,X137] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X139),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X140),s(X132,X141),s(X132,X142),s(X132,X143))) = s(X132,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X133),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X134),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X132)),X135),s(X132,X136),s(X132,X137),s(X132,X138))) | s(X132,X138) != s(X132,X143) | s(X132,X137) != s(X132,X142) | sP1(X132,sK2,X140,X135,sK4) | sP0(X132,sK2,X139,X134,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X133)) ) | $spl3),
  inference(superposition,[],[f292,f120])).
fof(f956,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl3),
  inference(subsumption_resolution,[],[f936,f292])).
fof(f936,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl3),
  inference(superposition,[],[f292,f118])).
fof(f869,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85) | sP0(X84,sK2,X91,X86,sK4) | sP1(X84,sK2,X92,X87,sK4) | s(X84,X88) != s(X84,X93) | s(X84,X89) != s(X84,X94)) ) | $spl28),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85) | sP0(X84,sK2,X91,X86,sK4) | sP1(X84,sK2,X92,X87,sK4) | s(X84,X88) != s(X84,X93) | s(X84,X89) != s(X84,X94)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f934,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s(X84,X89) != s(X84,X94) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85)) ) | $spl3),
  inference(superposition,[],[f292,f117])).
fof(f117,plain,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X7),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X9),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X10),s(X1,X11),s(X1,X12),s(X1,X13))) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),X8),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X6),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X5),s(X1,X4),s(X1,X3),s(X1,X2))) | s('type.wordSem.result'(X0),'const.wordSem.Error_0') = s('type.wordSem.result'(X0),X8) | s(X1,X3) != s(X1,X12) | s(X1,X4) != s(X1,X11) | sP1(X1,X0,X5,X10,X8) | sP0(X1,X0,X6,X9,X8) | s('type.wordSem.result'(X0),X7) != s('type.wordSem.result'(X0),X8)) )),
  inference(cnf_transformation,[],[f85])).
fof(f933,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl3),
  inference(superposition,[],[f292,f116])).
fof(f954,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f932,f290])).
fof(f932,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl3),
  inference(superposition,[],[f292,f116])).
fof(f945,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f931])).
fof(f931,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) ) | $spl3),
  inference(superposition,[],[f292,f115])).
fof(f953,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl3),
  inference(subsumption_resolution,[],[f930,f292])).
fof(f930,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl3),
  inference(superposition,[],[f292,f115])).
fof(f952,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl3),
  inference(subsumption_resolution,[],[f929,f292])).
fof(f929,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl3),
  inference(superposition,[],[f292,f114])).
fof(f951,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f946,f290])).
fof(f946,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f928])).
fof(f928,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl3),
  inference(superposition,[],[f292,f114])).
fof(f948,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f947,f290])).
fof(f947,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | $spl3),
  inference(subsumption_resolution,[],[f927,f292])).
fof(f927,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | $spl3),
  inference(superposition,[],[f292,f114])).
fof(f293,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | $spl2),
  inference(cnf_transformation,[],[f293_D])).
fof(f293_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f925,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f893,f106])).
fof(f893,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f924,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f923,f290])).
fof(f923,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f892,f107])).
fof(f892,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f922,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f921,f290])).
fof(f921,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f891,f108])).
fof(f891,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f918,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) )),
  inference(subsumption_resolution,[],[f888,f106])).
fof(f888,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) )),
  inference(superposition,[],[f109,f118])).
fof(f914,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) )),
  inference(subsumption_resolution,[],[f885,f106])).
fof(f885,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) )),
  inference(superposition,[],[f109,f116])).
fof(f911,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) )),
  inference(subsumption_resolution,[],[f883,f107])).
fof(f883,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) )),
  inference(superposition,[],[f109,f115])).
fof(f910,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) )),
  inference(subsumption_resolution,[],[f882,f108])).
fof(f882,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) )),
  inference(superposition,[],[f109,f115])).
fof(f909,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) )),
  inference(subsumption_resolution,[],[f881,f106])).
fof(f881,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) )),
  inference(superposition,[],[f109,f114])).
fof(f908,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl1),
  inference(subsumption_resolution,[],[f907,f290])).
fof(f907,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) )),
  inference(subsumption_resolution,[],[f880,f107])).
fof(f880,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) )),
  inference(superposition,[],[f109,f114])).
fof(f906,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f905,f290])).
fof(f905,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) )),
  inference(subsumption_resolution,[],[f879,f108])).
fof(f879,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) )),
  inference(superposition,[],[f109,f114])).
fof(f861,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(trivial_inequality_removal,[],[f860])).
fof(f860,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(superposition,[],[f290,f128])).
fof(f873,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f859,f290])).
fof(f859,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(superposition,[],[f290,f128])).
fof(f872,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f858,f290])).
fof(f858,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2)))) | $spl1),
  inference(superposition,[],[f290,f128])).
fof(f856,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X128),s(X120,X129),s(X120,X130),s(X120,X131))) = s(X120,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X121),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X120)),X123),s(X120,X124),s(X120,X125),s(X120,X126))) | s(X120,X126) != s(X120,X131) | s(X120,X124) != s(X120,X129) | sP1(X120,sK2,X128,X123,sK4) | sP0(X120,sK2,X127,X122,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X121)) ) | $spl1),
  inference(superposition,[],[f290,f119])).
fof(f863,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl1),
  inference(trivial_inequality_removal,[],[f855])).
fof(f855,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl1),
  inference(superposition,[],[f290,f118])).
fof(f870,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl1),
  inference(subsumption_resolution,[],[f854,f290])).
fof(f854,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl1),
  inference(superposition,[],[f290,f118])).
fof(f853,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s(X84,X89) != s(X84,X94) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85)) ) | $spl1),
  inference(superposition,[],[f290,f117])).
fof(f864,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl1),
  inference(trivial_inequality_removal,[],[f852])).
fof(f852,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl1),
  inference(superposition,[],[f290,f116])).
fof(f868,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl1),
  inference(subsumption_resolution,[],[f851,f290])).
fof(f851,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl1),
  inference(superposition,[],[f290,f116])).
fof(f850,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) ) | $spl1),
  inference(superposition,[],[f290,f115])).
fof(f849,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl1),
  inference(superposition,[],[f290,f115])).
fof(f865,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl1),
  inference(trivial_inequality_removal,[],[f848])).
fof(f848,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl1),
  inference(superposition,[],[f290,f114])).
fof(f867,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl1),
  inference(subsumption_resolution,[],[f847,f290])).
fof(f847,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) ) | $spl1),
  inference(superposition,[],[f290,f114])).
fof(f866,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f846,f290])).
fof(f846,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) ) | $spl1),
  inference(superposition,[],[f290,f114])).
fof(f739,plain,(
  ( ! [X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X85) | sP0(X84,sK2,X91,X86,sK4) | sP1(X84,sK2,X92,X87,sK4) | s(X84,X88) != s(X84,X93) | s(X84,X90) != s(X84,X95) | s(X84,X93) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90)))) ) | $spl26),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  ( ! [X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X85) | sP0(X84,sK2,X91,X86,sK4) | sP1(X84,sK2,X92,X87,sK4) | s(X84,X88) != s(X84,X93) | s(X84,X90) != s(X84,X95) | s(X84,X93) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f726,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | $spl20),
  inference(cnf_transformation,[],[f726_D])).
fof(f726_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f787,plain,(
  ( ! [X364,X366,X360,X362,X368,X365,X367,X361,X363,X369,X371] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X361,X369) = s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X360),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X362),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X363),s(X361,X364),s(X361,X365),s(X361,X366))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X360) | s(X361,X366) != s(X361,X371) | s(X361,X364) != s(X361,X369) | sP1(X361,sK2,X368,X363,sK4) | sP0(X361,sK2,X367,X362,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f786,f291])).
fof(f786,plain,(
  ( ! [X364,X366,X360,X362,X368,X365,X367,X361,X363,X369,X371] : (s(X361,X369) = s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X360),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X362),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X363),s(X361,X364),s(X361,X365),s(X361,X366))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X360) | s(X361,X366) != s(X361,X371) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X361,X364) != s(X361,X369) | sP1(X361,sK2,X368,X363,sK4) | sP0(X361,sK2,X367,X362,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f785,f131])).
fof(f131,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X1,X4) = s(X1,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X0),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X2),s(fun('type.wordSem.word_loc'(X0),fun('type.wordSem.word_loc'(X0),X1)),X3),s(X1,X4),s(X1,X5),s(X1,X6)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X11] : ! [X34,X13,X35,X36,X37] : s(X11,X35) = s(X11,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(X6),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X34),s(fun('type.wordSem.word_loc'(X6),fun('type.wordSem.word_loc'(X6),X11)),X13),s(X11,X35),s(X11,X36),s(X11,X37)))),
  file('thm.wordSem.result_induction','thm.wordSem.result_case_def|split|2')).
fof(f785,plain,(
  ( ! [X364,X366,X360,X362,X368,X370,X365,X367,X361,X363,X369,X371] : (s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X367),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X368),s(X361,X369),s(X361,X370),s(X361,X371))) = s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X360),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X362),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X363),s(X361,X364),s(X361,X365),s(X361,X366))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X360) | s(X361,X366) != s(X361,X371) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X361,X364) != s(X361,X369) | sP1(X361,sK2,X368,X363,sK4) | sP0(X361,sK2,X367,X362,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f620,f291])).
fof(f620,plain,(
  ( ! [X364,X366,X360,X362,X368,X370,X365,X367,X361,X363,X369,X371] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X360) | s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X367),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X368),s(X361,X369),s(X361,X370),s(X361,X371))) = s(X361,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X360),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X362),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X361)),X363),s(X361,X364),s(X361,X365),s(X361,X366))) | s(X361,X366) != s(X361,X371) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X361,X364) != s(X361,X369) | sP1(X361,sK2,X368,X363,sK4) | sP0(X361,sK2,X367,X362,sK4)) ) | $spl0),
  inference(superposition,[],[f119,f291])).
fof(f784,plain,(
  ( ! [X350,X356,X352,X354,X349,X351,X357,X359,X353,X355,X348] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X349) | s(X348,X357) = s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X349),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X350),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X351),s(X348,X352),s(X348,X353),s(X348,X354))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X348,X354) != s(X348,X359) | s(X348,X352) != s(X348,X357) | sP1(X348,sK2,X356,X351,sK4) | sP0(X348,sK2,X355,X350,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f783,f291])).
fof(f783,plain,(
  ( ! [X350,X356,X352,X354,X349,X351,X357,X359,X353,X355,X348] : (s(X348,X357) = s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X349),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X350),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X351),s(X348,X352),s(X348,X353),s(X348,X354))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X348,X354) != s(X348,X359) | s(X348,X352) != s(X348,X357) | sP1(X348,sK2,X356,X351,sK4) | sP0(X348,sK2,X355,X350,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X349)) ) | $spl0),
  inference(forward_demodulation,[],[f782,f131])).
fof(f782,plain,(
  ( ! [X350,X356,X358,X352,X354,X349,X351,X357,X359,X353,X355,X348] : (s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X355),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X356),s(X348,X357),s(X348,X358),s(X348,X359))) = s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X349),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X350),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X351),s(X348,X352),s(X348,X353),s(X348,X354))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X348,X354) != s(X348,X359) | s(X348,X352) != s(X348,X357) | sP1(X348,sK2,X356,X351,sK4) | sP0(X348,sK2,X355,X350,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X349)) ) | $spl0),
  inference(forward_demodulation,[],[f619,f291])).
fof(f619,plain,(
  ( ! [X350,X356,X358,X352,X354,X349,X351,X357,X359,X353,X355,X348] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X355),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X356),s(X348,X357),s(X348,X358),s(X348,X359))) = s(X348,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X349),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X350),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X348)),X351),s(X348,X352),s(X348,X353),s(X348,X354))) | s(X348,X354) != s(X348,X359) | s(X348,X352) != s(X348,X357) | sP1(X348,sK2,X356,X351,sK4) | sP0(X348,sK2,X355,X350,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X349)) ) | $spl0),
  inference(superposition,[],[f119,f291])).
fof(f779,plain,(
  ( ! [X325,X327,X333,X335,X329,X331,X324,X326,X332,X328,X330] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X325) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X324,X333) = s(X324,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X325),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X326),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X327),s(X324,X328),s(X324,X329),s(X324,X330))) | s(X324,X330) != s(X324,X335) | s(X324,X328) != s(X324,X333) | sP1(X324,sK2,X332,X327,sK4) | sP0(X324,sK2,X331,X326,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f778,f291])).
fof(f778,plain,(
  ( ! [X325,X327,X333,X335,X329,X331,X324,X326,X332,X328,X330] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X324,X333) = s(X324,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X325),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X326),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X327),s(X324,X328),s(X324,X329),s(X324,X330))) | s(X324,X330) != s(X324,X335) | s(X324,X328) != s(X324,X333) | sP1(X324,sK2,X332,X327,sK4) | sP0(X324,sK2,X331,X326,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X325)) ) | $spl0),
  inference(forward_demodulation,[],[f777,f291])).
fof(f777,plain,(
  ( ! [X325,X327,X333,X335,X329,X331,X324,X326,X332,X328,X330] : (s(X324,X333) = s(X324,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X325),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X326),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X327),s(X324,X328),s(X324,X329),s(X324,X330))) | s(X324,X330) != s(X324,X335) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X324,X328) != s(X324,X333) | sP1(X324,sK2,X332,X327,sK4) | sP0(X324,sK2,X331,X326,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X325)) ) | $spl0),
  inference(forward_demodulation,[],[f617,f131])).
fof(f617,plain,(
  ( ! [X325,X327,X333,X335,X329,X331,X324,X326,X332,X334,X328,X330] : (s(X324,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X331),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X332),s(X324,X333),s(X324,X334),s(X324,X335))) = s(X324,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X325),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X326),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X324)),X327),s(X324,X328),s(X324,X329),s(X324,X330))) | s(X324,X330) != s(X324,X335) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X324,X328) != s(X324,X333) | sP1(X324,sK2,X332,X327,sK4) | sP0(X324,sK2,X331,X326,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X325)) ) | $spl0),
  inference(superposition,[],[f119,f291])).
fof(f735,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X61) | sP0(X60,sK2,X67,X62,sK4) | sP1(X60,sK2,X68,X63,sK4) | s(X60,X64) != s(X60,X69) | s(X60,X65) != s(X60,X70) | s(X60,X69) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66)))) ) | $spl24),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X61) | sP0(X60,sK2,X67,X62,sK4) | sP1(X60,sK2,X68,X63,sK4) | s(X60,X64) != s(X60,X69) | s(X60,X65) != s(X60,X70) | s(X60,X69) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f724,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | $spl18),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f773,plain,(
  ( ! [X243,X241,X249,X246,X244,X242,X240,X250,X248,X247,X245] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X241,X249) = s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X240),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X242),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X243),s(X241,X244),s(X241,X245),s(X241,X246))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X240) | s(X241,X245) != s(X241,X250) | s(X241,X244) != s(X241,X249) | sP1(X241,sK2,X248,X243,sK4) | sP0(X241,sK2,X247,X242,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f772,f291])).
fof(f772,plain,(
  ( ! [X243,X241,X249,X246,X244,X242,X240,X250,X248,X247,X245] : (s(X241,X249) = s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X240),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X242),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X243),s(X241,X244),s(X241,X245),s(X241,X246))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X240) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X241,X245) != s(X241,X250) | s(X241,X244) != s(X241,X249) | sP1(X241,sK2,X248,X243,sK4) | sP0(X241,sK2,X247,X242,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f771,f131])).
fof(f771,plain,(
  ( ! [X243,X241,X251,X249,X246,X244,X242,X240,X250,X248,X247,X245] : (s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X247),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X248),s(X241,X249),s(X241,X250),s(X241,X251))) = s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X240),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X242),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X243),s(X241,X244),s(X241,X245),s(X241,X246))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X240) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X241,X245) != s(X241,X250) | s(X241,X244) != s(X241,X249) | sP1(X241,sK2,X248,X243,sK4) | sP0(X241,sK2,X247,X242,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f610,f291])).
fof(f610,plain,(
  ( ! [X243,X241,X251,X249,X246,X244,X242,X240,X250,X248,X247,X245] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X240) | s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X247),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X248),s(X241,X249),s(X241,X250),s(X241,X251))) = s(X241,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X240),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X242),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X241)),X243),s(X241,X244),s(X241,X245),s(X241,X246))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X241,X245) != s(X241,X250) | s(X241,X244) != s(X241,X249) | sP1(X241,sK2,X248,X243,sK4) | sP0(X241,sK2,X247,X242,sK4)) ) | $spl0),
  inference(superposition,[],[f117,f291])).
fof(f770,plain,(
  ( ! [X231,X229,X237,X235,X233,X230,X228,X238,X236,X234,X232] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X229) | s(X228,X237) = s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X229),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X230),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X231),s(X228,X232),s(X228,X233),s(X228,X234))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X228,X233) != s(X228,X238) | s(X228,X232) != s(X228,X237) | sP1(X228,sK2,X236,X231,sK4) | sP0(X228,sK2,X235,X230,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f769,f291])).
fof(f769,plain,(
  ( ! [X231,X229,X237,X235,X233,X230,X228,X238,X236,X234,X232] : (s(X228,X237) = s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X229),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X230),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X231),s(X228,X232),s(X228,X233),s(X228,X234))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X228,X233) != s(X228,X238) | s(X228,X232) != s(X228,X237) | sP1(X228,sK2,X236,X231,sK4) | sP0(X228,sK2,X235,X230,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X229)) ) | $spl0),
  inference(forward_demodulation,[],[f768,f131])).
fof(f768,plain,(
  ( ! [X231,X229,X239,X237,X235,X233,X230,X228,X238,X236,X234,X232] : (s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X235),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X236),s(X228,X237),s(X228,X238),s(X228,X239))) = s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X229),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X230),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X231),s(X228,X232),s(X228,X233),s(X228,X234))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X228,X233) != s(X228,X238) | s(X228,X232) != s(X228,X237) | sP1(X228,sK2,X236,X231,sK4) | sP0(X228,sK2,X235,X230,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X229)) ) | $spl0),
  inference(forward_demodulation,[],[f609,f291])).
fof(f609,plain,(
  ( ! [X231,X229,X239,X237,X235,X233,X230,X228,X238,X236,X234,X232] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X235),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X236),s(X228,X237),s(X228,X238),s(X228,X239))) = s(X228,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X229),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X230),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X228)),X231),s(X228,X232),s(X228,X233),s(X228,X234))) | s(X228,X233) != s(X228,X238) | s(X228,X232) != s(X228,X237) | sP1(X228,sK2,X236,X231,sK4) | sP0(X228,sK2,X235,X230,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X229)) ) | $spl0),
  inference(superposition,[],[f117,f291])).
fof(f765,plain,(
  ( ! [X212,X210,X208,X206,X204,X213,X211,X209,X207,X205,X214] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X205) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X204,X213) = s(X204,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X205),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X206),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X207),s(X204,X208),s(X204,X209),s(X204,X210))) | s(X204,X209) != s(X204,X214) | s(X204,X208) != s(X204,X213) | sP1(X204,sK2,X212,X207,sK4) | sP0(X204,sK2,X211,X206,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f764,f291])).
fof(f764,plain,(
  ( ! [X212,X210,X208,X206,X204,X213,X211,X209,X207,X205,X214] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X204,X213) = s(X204,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X205),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X206),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X207),s(X204,X208),s(X204,X209),s(X204,X210))) | s(X204,X209) != s(X204,X214) | s(X204,X208) != s(X204,X213) | sP1(X204,sK2,X212,X207,sK4) | sP0(X204,sK2,X211,X206,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X205)) ) | $spl0),
  inference(forward_demodulation,[],[f763,f291])).
fof(f763,plain,(
  ( ! [X212,X210,X208,X206,X204,X213,X211,X209,X207,X205,X214] : (s(X204,X213) = s(X204,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X205),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X206),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X207),s(X204,X208),s(X204,X209),s(X204,X210))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X204,X209) != s(X204,X214) | s(X204,X208) != s(X204,X213) | sP1(X204,sK2,X212,X207,sK4) | sP0(X204,sK2,X211,X206,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X205)) ) | $spl0),
  inference(forward_demodulation,[],[f607,f131])).
fof(f607,plain,(
  ( ! [X212,X210,X208,X206,X204,X215,X213,X211,X209,X207,X205,X214] : (s(X204,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X211),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X212),s(X204,X213),s(X204,X214),s(X204,X215))) = s(X204,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X205),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X206),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X204)),X207),s(X204,X208),s(X204,X209),s(X204,X210))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X204,X209) != s(X204,X214) | s(X204,X208) != s(X204,X213) | sP1(X204,sK2,X212,X207,sK4) | sP0(X204,sK2,X211,X206,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X205)) ) | $spl0),
  inference(superposition,[],[f117,f291])).
fof(f727,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X25) | sP0(X24,sK2,X31,X26,sK4) | sP1(X24,sK2,X32,X27,sK4) | s(X24,X28) != s(X24,X33) | s(X24,X33) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30)))) ) | $spl22),
  inference(cnf_transformation,[],[f727_D])).
fof(f727_D,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X25) | sP0(X24,sK2,X31,X26,sK4) | sP1(X24,sK2,X32,X27,sK4) | s(X24,X28) != s(X24,X33) | s(X24,X33) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f759,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120,X129] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X121,X129) = s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X120),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X123),s(X121,X124),s(X121,X125),s(X121,X126))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X120) | s(X121,X124) != s(X121,X129) | sP1(X121,sK2,X128,X123,sK4) | sP0(X121,sK2,X127,X122,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f758,f291])).
fof(f758,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120,X129] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X121,X129) = s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X120),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X123),s(X121,X124),s(X121,X125),s(X121,X126))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X120) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X121,X124) != s(X121,X129) | sP1(X121,sK2,X128,X123,sK4) | sP0(X121,sK2,X127,X122,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f757,f291])).
fof(f757,plain,(
  ( ! [X127,X125,X123,X121,X128,X126,X124,X122,X120,X129] : (s(X121,X129) = s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X120),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X123),s(X121,X124),s(X121,X125),s(X121,X126))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X120) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X121,X124) != s(X121,X129) | sP1(X121,sK2,X128,X123,sK4) | sP0(X121,sK2,X127,X122,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f756,f131])).
fof(f756,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X128),s(X121,X129),s(X121,X130),s(X121,X131))) = s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X120),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X123),s(X121,X124),s(X121,X125),s(X121,X126))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X120) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X121,X124) != s(X121,X129) | sP1(X121,sK2,X128,X123,sK4) | sP0(X121,sK2,X127,X122,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f600,f291])).
fof(f600,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X120,X131,X129] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X120) | s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X127),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X128),s(X121,X129),s(X121,X130),s(X121,X131))) = s(X121,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X120),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X122),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X121)),X123),s(X121,X124),s(X121,X125),s(X121,X126))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X121,X124) != s(X121,X129) | sP1(X121,sK2,X128,X123,sK4) | sP0(X121,sK2,X127,X122,sK4)) ) | $spl0),
  inference(superposition,[],[f115,f291])).
fof(f755,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X109) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X108,X117) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X108,X112) != s(X108,X117) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f754,f291])).
fof(f754,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X108,X117) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X108,X112) != s(X108,X117) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl0),
  inference(forward_demodulation,[],[f753,f291])).
fof(f753,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X117,X115,X113] : (s(X108,X117) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X108,X112) != s(X108,X117) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl0),
  inference(forward_demodulation,[],[f752,f131])).
fof(f752,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X108,X112) != s(X108,X117) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl0),
  inference(forward_demodulation,[],[f599,f291])).
fof(f599,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X108,X112) != s(X108,X117) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) ) | $spl0),
  inference(superposition,[],[f115,f291])).
fof(f751,plain,(
  ( ! [X103,X101,X99,X97,X105,X102,X100,X98,X96,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X97) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X96,X105) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X96,X100) != s(X96,X105) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f750,f291])).
fof(f750,plain,(
  ( ! [X103,X101,X99,X97,X105,X102,X100,X98,X96,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X96,X105) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X96,X100) != s(X96,X105) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl0),
  inference(forward_demodulation,[],[f749,f291])).
fof(f749,plain,(
  ( ! [X103,X101,X99,X97,X105,X102,X100,X98,X96,X104] : (s(X96,X105) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X96,X100) != s(X96,X105) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl0),
  inference(forward_demodulation,[],[f748,f131])).
fof(f748,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X96,X100) != s(X96,X105) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl0),
  inference(forward_demodulation,[],[f598,f291])).
fof(f598,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X96,X100) != s(X96,X105) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) ) | $spl0),
  inference(superposition,[],[f115,f291])).
fof(f745,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X79,X77,X75,X73] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X73) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X72,X81) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s(X72,X76) != s(X72,X81) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f744,f291])).
fof(f744,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X79,X77,X75,X73] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X72,X81) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s(X72,X76) != s(X72,X81) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl0),
  inference(forward_demodulation,[],[f743,f291])).
fof(f743,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X79,X77,X75,X73] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X72,X81) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X76) != s(X72,X81) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl0),
  inference(forward_demodulation,[],[f742,f291])).
fof(f742,plain,(
  ( ! [X80,X78,X76,X74,X72,X81,X79,X77,X75,X73] : (s(X72,X81) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X76) != s(X72,X81) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl0),
  inference(forward_demodulation,[],[f596,f131])).
fof(f596,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X76) != s(X72,X81) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) ) | $spl0),
  inference(superposition,[],[f115,f291])).
fof(f738,plain,(
  ( ! [X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X85) | s(X84,X93) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X84,X90) != s(X84,X95) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f737,f291])).
fof(f737,plain,(
  ( ! [X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s(X84,X93) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X84,X90) != s(X84,X95) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85)) ) | $spl0),
  inference(forward_demodulation,[],[f736,f131])).
fof(f736,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X84,X90) != s(X84,X95) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85)) ) | $spl0),
  inference(forward_demodulation,[],[f587,f291])).
fof(f587,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X95,X93,X91,X89,X86,X84] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X91),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X92),s(X84,X93),s(X84,X94),s(X84,X95))) = s(X84,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X85),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X86),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X84)),X87),s(X84,X88),s(X84,X89),s(X84,X90))) | s(X84,X90) != s(X84,X95) | s(X84,X88) != s(X84,X93) | sP1(X84,sK2,X92,X87,sK4) | sP0(X84,sK2,X91,X86,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X85)) ) | $spl0),
  inference(superposition,[],[f291,f119])).
fof(f734,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X61) | s(X60,X69) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,X65) != s(X60,X70) | s(X60,X64) != s(X60,X69) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f733,f291])).
fof(f733,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X69,X67,X65,X63] : (s(X60,X69) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,X65) != s(X60,X70) | s(X60,X64) != s(X60,X69) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl0),
  inference(forward_demodulation,[],[f732,f131])).
fof(f732,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,X65) != s(X60,X70) | s(X60,X64) != s(X60,X69) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl0),
  inference(forward_demodulation,[],[f585,f291])).
fof(f585,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | s(X60,X64) != s(X60,X69) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) ) | $spl0),
  inference(superposition,[],[f291,f117])).
fof(f731,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X37) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X36,X45) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f730,f291])).
fof(f730,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X36,X45) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl0),
  inference(forward_demodulation,[],[f729,f291])).
fof(f729,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (s(X36,X45) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl0),
  inference(forward_demodulation,[],[f728,f131])).
fof(f728,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl0),
  inference(forward_demodulation,[],[f583,f291])).
fof(f583,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) ) | $spl0),
  inference(superposition,[],[f291,f115])).
fof(f722,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') != s('type.wordSem.result'(sK2),X25) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X24,X33) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,X28) != s(X24,X33) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4)) ) | $spl0),
  inference(forward_demodulation,[],[f721,f291])).
fof(f721,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') | s(X24,X33) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,X28) != s(X24,X33) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl0),
  inference(forward_demodulation,[],[f720,f291])).
fof(f720,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X24,X33) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X24,X28) != s(X24,X33) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl0),
  inference(forward_demodulation,[],[f719,f131])).
fof(f719,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0'),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X24,X28) != s(X24,X33) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl0),
  inference(forward_demodulation,[],[f582,f291])).
fof(f582,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),'const.wordSem.Error_0') | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X24,X28) != s(X24,X33) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) ) | $spl0),
  inference(superposition,[],[f291,f115])).
fof(f459,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f319,f107])).
fof(f319,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | $spl2),
  inference(backward_demodulation,[],[f293,f109])).
fof(f317,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f277,f106])).
fof(f277,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f291,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | $spl0),
  inference(cnf_transformation,[],[f291_D])).
fof(f291_D,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f316,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f276,f107])).
fof(f276,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f311,plain,(
  s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(subsumption_resolution,[],[f275,f108])).
fof(f275,plain,(
  ~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Exception_2'(s('type.wordSem.word_loc'(sK2),sK12(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK13(sK4,sK2)))) | s('type.wordSem.result'(sK2),sK4) = s('type.wordSem.result'(sK2),'const.wordSem.Result_2'(s('type.wordSem.word_loc'(sK2),sK10(sK4,sK2)),s('type.wordSem.word_loc'(sK2),sK11(sK4,sK2))))),
  inference(superposition,[],[f109,f128])).
fof(f308,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) )),
  inference(subsumption_resolution,[],[f272,f106])).
fof(f272,plain,(
  ( ! [X111,X109,X118,X116,X114,X112,X110,X108,X119,X117,X115,X113] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X115),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X116),s(X108,X117),s(X108,X118),s(X108,X119))) = s(X108,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X109),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X110),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X108)),X111),s(X108,X112),s(X108,X113),s(X108,X114))) | s(X108,X114) != s(X108,X119) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X108,sK2,X116,X111,sK4) | sP0(X108,sK2,X115,X110,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X109)) )),
  inference(superposition,[],[f109,f118])).
fof(f306,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) )),
  inference(subsumption_resolution,[],[f271,f107])).
fof(f271,plain,(
  ( ! [X103,X101,X99,X97,X107,X105,X102,X100,X98,X96,X106,X104] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X103),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X104),s(X96,X105),s(X96,X106),s(X96,X107))) = s(X96,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X97),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X98),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X96)),X99),s(X96,X100),s(X96,X101),s(X96,X102))) | s(X96,X102) != s(X96,X107) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X96,sK2,X104,X99,sK4) | sP0(X96,sK2,X103,X98,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X97)) )),
  inference(superposition,[],[f109,f118])).
fof(f304,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) )),
  inference(subsumption_resolution,[],[f269,f106])).
fof(f269,plain,(
  ( ! [X80,X78,X76,X74,X72,X83,X81,X79,X77,X75,X73,X82] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X79),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X80),s(X72,X81),s(X72,X82),s(X72,X83))) = s(X72,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X73),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X74),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X72)),X75),s(X72,X76),s(X72,X77),s(X72,X78))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X72,X77) != s(X72,X82) | sP1(X72,sK2,X80,X75,sK4) | sP0(X72,sK2,X79,X74,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X73)) )),
  inference(superposition,[],[f109,f116])).
fof(f302,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) )),
  inference(subsumption_resolution,[],[f268,f108])).
fof(f268,plain,(
  ( ! [X61,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X67),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X68),s(X60,X69),s(X60,X70),s(X60,X71))) = s(X60,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X61),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X62),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X60)),X63),s(X60,X64),s(X60,X65),s(X60,X66))) | s(X60,X65) != s(X60,X70) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X60,sK2,X68,X63,sK4) | sP0(X60,sK2,X67,X62,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X61)) )),
  inference(superposition,[],[f109,f116])).
fof(f301,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) )),
  inference(subsumption_resolution,[],[f267,f107])).
fof(f267,plain,(
  ( ! [X59,X57,X54,X52,X50,X48,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X55),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X56),s(X48,X57),s(X48,X58),s(X48,X59))) = s(X48,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X49),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X50),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X48)),X51),s(X48,X52),s(X48,X53),s(X48,X54))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s(X48,X52) != s(X48,X57) | sP1(X48,sK2,X56,X51,sK4) | sP0(X48,sK2,X55,X50,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X49)) )),
  inference(superposition,[],[f109,f115])).
fof(f299,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) )),
  inference(subsumption_resolution,[],[f266,f108])).
fof(f266,plain,(
  ( ! [X39,X37,X47,X45,X43,X41,X38,X36,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X43),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X44),s(X36,X45),s(X36,X46),s(X36,X47))) = s(X36,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X37),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X38),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X36)),X39),s(X36,X40),s(X36,X41),s(X36,X42))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s(X36,X40) != s(X36,X45) | sP1(X36,sK2,X44,X39,sK4) | sP0(X36,sK2,X43,X38,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X37)) )),
  inference(superposition,[],[f109,f115])).
fof(f298,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) )),
  inference(subsumption_resolution,[],[f265,f106])).
fof(f265,plain,(
  ( ! [X30,X28,X26,X24,X35,X33,X31,X29,X27,X25,X34,X32] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0')))) | s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X31),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X32),s(X24,X33),s(X24,X34),s(X24,X35))) = s(X24,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X25),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X26),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X24)),X27),s(X24,X28),s(X24,X29),s(X24,X30))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X24,sK2,X32,X27,sK4) | sP0(X24,sK2,X31,X26,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X25)) )),
  inference(superposition,[],[f109,f114])).
fof(f295,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) )),
  inference(subsumption_resolution,[],[f264,f107])).
fof(f264,plain,(
  ( ! [X14,X12,X23,X21,X19,X17,X15,X13,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0')))) | s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X19),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X20),s(X12,X21),s(X12,X22),s(X12,X23))) = s(X12,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X13),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X14),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X12)),X15),s(X12,X16),s(X12,X17),s(X12,X18))) | s('type.wordSem.result'(sK2),'const.wordSem.Error_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X12,sK2,X20,X15,sK4) | sP0(X12,sK2,X19,X14,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X13)) )),
  inference(superposition,[],[f109,f114])).
fof(f289,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) )),
  inference(subsumption_resolution,[],[f263,f108])).
fof(f263,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~p(s(bool,happ(s(fun('type.wordSem.result'(sK2),bool),sK3),s('type.wordSem.result'(sK2),'const.wordSem.Error_0')))) | s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),sK4),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X7),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X8),s(X0,X9),s(X0,X10),s(X0,X11))) = s(X0,'const.wordSem.result_CASE_6'(s('type.wordSem.result'(sK2),X1),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X2),s(fun('type.wordSem.word_loc'(sK2),fun('type.wordSem.word_loc'(sK2),X0)),X3),s(X0,X4),s(X0,X5),s(X0,X6))) | s('type.wordSem.result'(sK2),'const.wordSem.NotEnoughSpace_0') = s('type.wordSem.result'(sK2),sK4) | s('type.wordSem.result'(sK2),'const.wordSem.TimeOut_0') = s('type.wordSem.result'(sK2),sK4) | sP1(X0,sK2,X8,X3,sK4) | sP0(X0,sK2,X7,X2,sK4) | s('type.wordSem.result'(sK2),sK4) != s('type.wordSem.result'(sK2),X1)) )),
  inference(superposition,[],[f109,f114])).
% SZS output end Proof for thm.wordSem
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 70
% Passive clauses: 239
% Generated clauses: 3401
% Final active clauses: 59
% Final passive clauses: 161
% Input formulas: 37
% Initial clauses: 56
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 42
% Trivial inequalities: 27
% Fw subsumption resolutions: 433
% Fw demodulations: 1177
% Bw demodulations: 7
% 
% Simple tautologies: 71
% Equational tautologies: 1
% Forward subsumptions: 1224
% 
% Binary resolution: 56
% Forward superposition: 526
% Backward superposition: 1108
% 
% Split clauses: 87
% Split components: 213
% SAT solver clauses: 92
% SAT solver unit clauses: 4
% SAT solver binary clauses: 10
% 
% Sat splits: 87
% Sat splitting refutations: 4
% 
% Memory used [KB]: 3837
% Time elapsed: 0.139 s
% ------------------------------
% ------------------------------
% Success in time 11.93 s
