Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.clos_knownProof
% (86351)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 529
% Passive clauses: 2963
% Generated clauses: 30266
% Final active clauses: 451
% Final passive clauses: 1979
% Input formulas: 62
% Initial clauses: 157
% 
% Duplicate literals: 610
% Trivial inequalities: 66
% Fw subsumption resolutions: 3028
% Fw demodulations: 7359
% Bw demodulations: 75
% 
% Simple tautologies: 443
% Equational tautologies: 62
% Forward subsumptions: 15214
% 
% Binary resolution: 12479
% Factoring: 26
% Forward superposition: 2555
% Backward superposition: 3510
% Self superposition: 103
% Equality factoring: 76
% Equality resolution: 21
% 
% Split clauses: 431
% Split components: 865
% SAT solver clauses: 804
% SAT solver unit clauses: 27
% SAT solver binary clauses: 253
% 
% Sat splits: 431
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 797
% TWLsolver calls for satisfiability: 193
% 
% Memory used [KB]: 10490
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.clos_knownProof
% (86385)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 329
% Passive clauses: 867
% Generated clauses: 14682
% Final active clauses: 265
% Final passive clauses: 426
% Input formulas: 62
% Initial clauses: 157
% 
% Duplicate literals: 594
% Trivial inequalities: 72
% Fw subsumption resolutions: 1514
% Fw demodulations: 1926
% Bw demodulations: 122
% Condensations: 716
% Global subsumptions: 17
% 
% Simple tautologies: 263
% Equational tautologies: 64
% Forward subsumptions: 6945
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 968
% Factoring: 38
% Forward superposition: 3869
% Backward superposition: 3315
% Self superposition: 133
% Equality factoring: 78
% Equality resolution: 24
% 
% Split clauses: 184
% Split components: 371
% SAT solver clauses: 1739
% SAT solver unit clauses: 196
% SAT solver binary clauses: 81
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 14
% 
% Sat splits: 184
% Sat splitting refutations: 22
% 
% TWLsolver clauses: 1668
% TWLsolver calls for satisfiability: 1167
% 
% Memory used [KB]: 9978
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.clos_knownProof
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.clos_knownProof
% SZS output start Proof for thm.clos_knownProof
fof(f6499,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f246,f413,f411,f245,f412,f410,f1420,f1454,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3719,f3665,f3722,f3670,f3724,f3673,f3726,f3676,f3728,f3679,f3730,f3682,f3732,f3685,f3734,f3688,f3736,f3691,f3738,f3666,f3694,f3780,f3781,f3783,f3785,f3787,f3789,f3790,f3792,f3793,f3795,f3797,f3798,f3799,f4572,f4642,f4640,f4659,f4660,f4667,f4664,f4661,f4662,f4670,f4663,f4672,f4665,f4669,f4679,f4680,f4692,f4690,f4779,f4781,f4939,f4785,f4783,f4960,f4965,f4964,f4962,f4787,f5106,f5470,f5481,f5489,f5490,f5493,f5494,f5497,f5498,f5501,f5502,f5506,f5507,f5510,f5511,f5518,f5519,f5523,f5524,f5527,f5528,f5531,f5532,f5535,f5536,f5471,f5603,f5613,f5604,f5681,f5473,f5690,f5482,f5682,f5778,f6447,f6492])).
fof(f6492,plain,(
  $false | ($spl0 | $spl3 | $spl56)),
  inference(subsumption_resolution,[],[f6491,f413])).
fof(f6491,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | ($spl0 | $spl56)),
  inference(backward_demodulation,[],[f6474,f1456])).
fof(f1456,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(unit_resulting_resolution,[],[f410,f330])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) & (s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) <=> s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X46,X47] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X46),s('type.clos_known.val_approx',X47)))) <=> s('type.clos_known.val_approx',X47) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X46),s('type.clos_known.val_approx',X47))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_knownProof.subapprox_def')).
fof(f6474,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl56),
  inference(superposition,[],[f257,f4964])).
fof(f257,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X45] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X45)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|19')).
fof(f6447,plain,(
  $false | ($spl0 | $spl3 | $spl54)),
  inference(subsumption_resolution,[],[f6446,f413])).
fof(f6446,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | ($spl0 | $spl54)),
  inference(backward_demodulation,[],[f6430,f1456])).
fof(f6430,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl54),
  inference(superposition,[],[f259,f4962])).
fof(f259,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X44] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X44)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|18')).
fof(f5778,plain,(
  $false | ($spl0 | $spl3 | $spl52)),
  inference(subsumption_resolution,[],[f5194,f413])).
fof(f5194,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5105,f4787])).
fof(f5105,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.Impossible_0') | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f4966,f249])).
fof(f249,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(cnf_transformation,[],[f33])).
fof(f33,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|4')).
fof(f4966,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f4787,f1456])).
fof(f5682,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) ) | $spl88),
  inference(cnf_transformation,[],[f5682_D])).
fof(f5682_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5482,plain,(
  ( ! [X85] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85)) ) | $spl62),
  inference(cnf_transformation,[],[f5482_D])).
fof(f5482_D,plain,(
  ( ! [X85] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5690,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X101),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK15(X97,X98,X99,X100))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5689,f250])).
fof(f250,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(cnf_transformation,[],[f49])).
fof(f49,axiom,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|20')).
fof(f5689,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X101),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK15(X97,X98,X99,X100))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5675,f5621])).
fof(f5621,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5120,f316])).
fof(f316,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X30,X31] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X31),s('type.num.num',X30))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|10')).
fof(f5120,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5105,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X18,X19] : s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X19),s('type.num.num',X18))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X19),s('type.num.num',X18))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|1')).
fof(f5675,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X98),s('type.num.num',X100))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X101),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X98),s('type.num.num',X100))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK15(X97,X98,X99,X100))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5621,f3363])).
fof(f3363,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X101),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X98),s('type.num.num',X100))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X97),s('type.num.num',X99))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X98),s('type.num.num',X100))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,sK15(X97,X98,X99,X100))) )),
  inference(superposition,[],[f401,f524])).
fof(f524,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f263,f263])).
fof(f263,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.clos_knownProof.subapprox_Other|split|0',aHL_BOOL_CASES)).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK15(X3,X2,X1,X0)),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK15(X3,X2,X1,X0))) | (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1) | p(s(bool,sK15(X3,X2,X1,X0)))) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK15(X3,X2,X1,X0)),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1) | p(s(bool,X4))) & s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X4),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & ((s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,X4)))) & s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X4),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X4),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X25,X26,X27,X28] : ? [X24] : ((p(s(bool,X24)) <=> (s('type.num.num',X27) = s('type.num.num',X28) & s('type.num.num',X25) = s('type.num.num',X26))) & s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X24),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X27),s('type.num.num',X25))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X28),s('type.num.num',X26))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X27),s('type.num.num',X25))))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|6')).
fof(f5473,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl60),
  inference(cnf_transformation,[],[f5473_D])).
fof(f5473_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f5681,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5680,f250])).
fof(f5680,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5625,f5621])).
fof(f5625,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5621,f1762])).
fof(f1762,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X3))))) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,X3))) )),
  inference(superposition,[],[f401,f263])).
fof(f5604,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK14(X0,X1))) ) | $spl86),
  inference(cnf_transformation,[],[f5604_D])).
fof(f5604_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK14(X0,X1))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5613,plain,(
  ( ! [X92,X93,X91] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X93),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,sK14(X91,X92))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5612,f250])).
fof(f5612,plain,(
  ( ! [X92,X93,X91] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X93),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,sK14(X91,X92))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5611,f5537])).
fof(f5537,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5118,f260])).
fof(f260,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X34] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X34))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|13')).
fof(f5118,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5105,f262])).
fof(f262,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X20] : s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X20))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X20))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|2')).
fof(f5611,plain,(
  ( ! [X92,X93,X91] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X93),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X92))))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,sK14(X91,X92))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5593,f5537])).
fof(f5593,plain,(
  ( ! [X92,X93,X91] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X92))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X93),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X91),s('type.list.list'('type.clos_known.val_approx'),X92))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,sK14(X91,X92))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5537,f3361])).
fof(f3361,plain,(
  ( ! [X92,X93,X91] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X91))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X92))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X93),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X91),s('type.list.list'('type.clos_known.val_approx'),X92))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,sK14(X91,X92))) )),
  inference(superposition,[],[f416,f524])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK14(X1,X0)),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(forward_demodulation,[],[f353,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) = s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx',X0))),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) = s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx',X0))),s('type.clos_known.val_approx',X1)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X52,X53] : s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X52),s('type.clos_known.val_approx',X53))) = s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx',X52))),s('type.clos_known.val_approx',X53)))),
  file('thm.clos_knownProof.subapprox_Other|split|0','ARITY_0_2_const.clos_known.merge')).
fof(f353,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK14(X1,X0)),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((~p(s(bool,sK14(X1,X0))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1))) | p(s(bool,sK14(X1,X0)))) & s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK14(X1,X0)),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ? [X2] : ((~p(s(bool,X2)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1))) | p(s(bool,X2))) & s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ? [X2] : (((~p(s(bool,X2)) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1))) | p(s(bool,X2)))) & s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) & s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X22,X23] : ? [X24] : ((p(s(bool,X24)) <=> s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X22))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X23)))) & s('type.clos_known.val_approx',happ(s(fun('type.clos_known.val_approx','type.clos_known.val_approx'),happ(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X23))))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X22))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X24),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X23),s('type.list.list'('type.clos_known.val_approx'),X22))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|5')).
fof(f5603,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK14(X0,X1))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5602,f250])).
fof(f5602,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK14(X0,X1))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5601,f5537])).
fof(f5601,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))) | s(bool,'T_0') = s(bool,sK14(X0,X1))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5543,f5537])).
fof(f5543,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X0),s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK14(X0,X1))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5537,f1786])).
fof(f1786,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X0),s('type.list.list'('type.clos_known.val_approx'),X1))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK14(X0,X1))) )),
  inference(superposition,[],[f416,f263])).
fof(f5471,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,sK13(X1,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f5471_D])).
fof(f5471_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,sK13(X1,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5536,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) | $spl84),
  inference(cnf_transformation,[],[f5536_D])).
fof(f5536_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5535,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X25,X26))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5534,f250])).
fof(f5534,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X25,X26))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5533,f5386])).
fof(f5386,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5117,f258])).
fof(f258,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X40] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X40))),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|16')).
fof(f5117,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5105,f261])).
fof(f261,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X21] : s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X21))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X21))),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|3')).
fof(f5533,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X25))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X25,X26))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5465,f5386])).
fof(f5465,plain,(
  ( ! [X26,X27,X25] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X25))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X26))))))) | s(bool,'T_0') = s(bool,sK13(X25,X26))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4732])).
fof(f4732,plain,(
  ( ! [X26,X27,X25] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X26))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X26))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X27),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X25))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X26))))))) | s(bool,'T_0') = s(bool,sK13(X25,X26))) )),
  inference(superposition,[],[f1455,f1734])).
fof(f1734,plain,(
  ( ! [X2,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))))) | s(bool,'T_0') = s(bool,sK13(X1,X2))) )),
  inference(superposition,[],[f350,f263])).
fof(f350,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK13(X1,X0)),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,sK13(X1,X0))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | p(s(bool,sK13(X1,X0)))) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK13(X1,X0)),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ? [X2] : ((~p(s(bool,X2)) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | p(s(bool,X2))) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ? [X2] : (((~p(s(bool,X2)) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | p(s(bool,X2)))) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> s('type.integer.int',X0) = s('type.integer.int',X1)) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X2),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X29,X16] : ? [X24] : ((p(s(bool,X24)) <=> s('type.integer.int',X16) = s('type.integer.int',X29)) & s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X16))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X29))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X24),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X29))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|7')).
fof(f1455,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(subsumption_resolution,[],[f1418,f414])).
fof(f414,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f252,f247])).
fof(f247,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.clos_knownProof.subapprox_Other|split|0',aHL_TRUTH)).
fof(f252,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f1418,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(superposition,[],[f331,f263])).
fof(f331,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) | s('type.clos_known.val_approx',X1) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) )),
  inference(cnf_transformation,[],[f208])).
fof(f5532,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24)))) ) | $spl82),
  inference(cnf_transformation,[],[f5532_D])).
fof(f5532_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5531,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,sK13(X22,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5530,f250])).
fof(f5530,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,sK13(X22,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5529,f5386])).
fof(f5529,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X22))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,sK13(X22,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5464,f5386])).
fof(f5464,plain,(
  ( ! [X24,X23,X22] : (s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X22))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,sK13(X22,X23))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4731])).
fof(f4731,plain,(
  ( ! [X24,X23,X22] : (s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X22))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))),s('type.clos_known.val_approx',X24))) | s(bool,'T_0') = s(bool,sK13(X22,X23))) )),
  inference(superposition,[],[f1455,f1734])).
fof(f5528,plain,(
  ( ! [X21] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | $spl80),
  inference(cnf_transformation,[],[f5528_D])).
fof(f5528_D,plain,(
  ( ! [X21] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5527,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X19,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5526,f5386])).
fof(f5526,plain,(
  ( ! [X21,X19,X20] : (s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X19,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5525,f250])).
fof(f5525,plain,(
  ( ! [X21,X19,X20] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X19,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5463,f5386])).
fof(f5463,plain,(
  ( ! [X21,X19,X20] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X19))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X19))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X19,X20))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4730])).
fof(f4730,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X19))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X21),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X19))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))))))) | s(bool,'T_0') = s(bool,sK13(X19,X20))) )),
  inference(superposition,[],[f331,f1734])).
fof(f5524,plain,(
  ( ! [X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18))))) ) | $spl78),
  inference(cnf_transformation,[],[f5524_D])).
fof(f5524_D,plain,(
  ( ! [X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5523,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5522,f5386])).
fof(f5522,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) ) | ($spl0 | $spl52)),
  inference(subsumption_resolution,[],[f5521,f5116])).
fof(f5116,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5105,f256])).
fof(f256,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X17] : s('type.clos_known.val_approx',X17) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X17)))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|0')).
fof(f5521,plain,(
  ( ! [X17,X18,X16] : (s('type.clos_known.val_approx',X18) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X18))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5520,f5386])).
fof(f5520,plain,(
  ( ! [X17,X18,X16] : (s('type.clos_known.val_approx',X18) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx',X18))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5462,f5116])).
fof(f5462,plain,(
  ( ! [X17,X18,X16] : (s('type.clos_known.val_approx',X18) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))))),s('type.clos_known.val_approx',X18))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4729])).
fof(f4729,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X18)))) | s('type.clos_known.val_approx',X18) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X16))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))))),s('type.clos_known.val_approx',X18))) | s(bool,'T_0') = s(bool,sK13(X16,X17))) )),
  inference(superposition,[],[f331,f1734])).
fof(f5519,plain,(
  ( ! [X15] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f5519_D])).
fof(f5519_D,plain,(
  ( ! [X15] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5518,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X13,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5517,f5386])).
fof(f5517,plain,(
  ( ! [X14,X15,X13] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X13,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5516,f250])).
fof(f5516,plain,(
  ( ! [X14,X15,X13] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X13,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5461,f5386])).
fof(f5461,plain,(
  ( ! [X14,X15,X13] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X13))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X13))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s(bool,'T_0') = s(bool,sK13(X13,X14))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4728])).
fof(f4728,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X13))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X15),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X13))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))))))) | s(bool,'T_0') = s(bool,sK13(X13,X14))) )),
  inference(superposition,[],[f330,f1734])).
fof(f5511,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f5511_D])).
fof(f5511_D,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5510,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X24,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5509,f5386])).
fof(f5509,plain,(
  ( ! [X24,X23,X25] : (s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X24,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5508,f5386])).
fof(f5508,plain,(
  ( ! [X24,X23,X25] : (s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X24,X23))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5458,f5116])).
fof(f5458,plain,(
  ( ! [X24,X23,X25] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X24,X23))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4722])).
fof(f4722,plain,(
  ( ! [X24,X23,X25] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X24))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X24))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))))))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X25),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X23))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X24,X23))) )),
  inference(superposition,[],[f1455,f1734])).
fof(f5507,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22)))) ) | $spl72),
  inference(cnf_transformation,[],[f5507_D])).
fof(f5507_D,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5506,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5505,f5386])).
fof(f5505,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) ) | ($spl0 | $spl52)),
  inference(subsumption_resolution,[],[f5504,f5116])).
fof(f5504,plain,(
  ( ! [X21,X22,X20] : (s('type.clos_known.val_approx',X22) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5503,f5386])).
fof(f5503,plain,(
  ( ! [X21,X22,X20] : (s('type.clos_known.val_approx',X22) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5457,f5116])).
fof(f5457,plain,(
  ( ! [X21,X22,X20] : (s('type.clos_known.val_approx',X22) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4721])).
fof(f4721,plain,(
  ( ! [X21,X22,X20] : (s('type.clos_known.val_approx',X22) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X21))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X20))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X22))) | s(bool,'T_0') = s(bool,sK13(X21,X20))) )),
  inference(superposition,[],[f1455,f1734])).
fof(f5502,plain,(
  ( ! [X19] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) | $spl70),
  inference(cnf_transformation,[],[f5502_D])).
fof(f5502_D,plain,(
  ( ! [X19] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5501,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X18,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5500,f250])).
fof(f5500,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X18,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5499,f5386])).
fof(f5499,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X18))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X18,X17))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5456,f5386])).
fof(f5456,plain,(
  ( ! [X19,X17,X18] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X18))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17)))))))) | s(bool,'T_0') = s(bool,sK13(X18,X17))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4720])).
fof(f4720,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X18))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17)))))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X17))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X18,X17))) )),
  inference(superposition,[],[f331,f1734])).
fof(f5498,plain,(
  ( ! [X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16)))) ) | $spl68),
  inference(cnf_transformation,[],[f5498_D])).
fof(f5498_D,plain,(
  ( ! [X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5497,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16))) | s(bool,'T_0') = s(bool,sK13(X15,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5496,f250])).
fof(f5496,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16))) | s(bool,'T_0') = s(bool,sK13(X15,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5495,f5386])).
fof(f5495,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X15))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16))) | s(bool,'T_0') = s(bool,sK13(X15,X14))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5455,f5386])).
fof(f5455,plain,(
  ( ! [X14,X15,X16] : (s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16))) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X15))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))))),s('type.clos_known.val_approx',X16)))) | s(bool,'T_0') = s(bool,sK13(X15,X14))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4719])).
fof(f4719,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X15))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))))),s('type.clos_known.val_approx',X16)))) | s('type.clos_known.val_approx',X16) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X14))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X16))) | s(bool,'T_0') = s(bool,sK13(X15,X14))) )),
  inference(superposition,[],[f331,f1734])).
fof(f5494,plain,(
  ( ! [X13] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) | $spl66),
  inference(cnf_transformation,[],[f5494_D])).
fof(f5494_D,plain,(
  ( ! [X13] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5493,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X12,X11))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5492,f250])).
fof(f5492,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X12,X11))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5491,f5386])).
fof(f5491,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X12))),s('type.clos_known.val_approx','const.clos_known.Other_0')))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X12,X11))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5454,f5386])).
fof(f5454,plain,(
  ( ! [X12,X13,X11] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X12))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X11)))))))) | s(bool,'T_0') = s(bool,sK13(X12,X11))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4718])).
fof(f4718,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X12))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X11)))))))) | s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X11))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X13),s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X11))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) | s(bool,'T_0') = s(bool,sK13(X12,X11))) )),
  inference(superposition,[],[f330,f1734])).
fof(f5490,plain,(
  ( ! [X10] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10)))) ) | $spl64),
  inference(cnf_transformation,[],[f5490_D])).
fof(f5490_D,plain,(
  ( ! [X10] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5489,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10))) | s(bool,'T_0') = s(bool,sK13(X9,X8))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5488,f250])).
fof(f5488,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10))) | s(bool,'T_0') = s(bool,sK13(X9,X8))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5487,f5386])).
fof(f5487,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X9))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10))) | s(bool,'T_0') = s(bool,sK13(X9,X8))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5453,f5386])).
fof(f5453,plain,(
  ( ! [X10,X8,X9] : (s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10))) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X9))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X8))))),s('type.clos_known.val_approx',X10)))) | s(bool,'T_0') = s(bool,sK13(X9,X8))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f4717])).
fof(f4717,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X9))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X8))))),s('type.clos_known.val_approx',X10)))) | s('type.clos_known.val_approx',X10) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X8))),s('type.clos_known.val_approx','const.clos_known.Other_0'))),s('type.clos_known.val_approx',X10))) | s(bool,'T_0') = s(bool,sK13(X9,X8))) )),
  inference(superposition,[],[f330,f1734])).
fof(f5481,plain,(
  ( ! [X85,X83,X84] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,sK13(X83,X84))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5480,f250])).
fof(f5480,plain,(
  ( ! [X85,X83,X84] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,sK13(X83,X84))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5445,f5386])).
fof(f5445,plain,(
  ( ! [X85,X83,X84] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X83))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,sK13(X83,X84))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f3358])).
fof(f3358,plain,(
  ( ! [X85,X83,X84] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X83))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X84))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,X85),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X84))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,sK13(X83,X84))) )),
  inference(superposition,[],[f350,f524])).
fof(f5470,plain,(
  ( ! [X2,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X1,X2))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5469,f250])).
fof(f5469,plain,(
  ( ! [X2,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X1,X2))) ) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f5395,f5386])).
fof(f5395,plain,(
  ( ! [X2,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'F_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,'T_0') = s(bool,sK13(X1,X2))) ) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f5386,f1734])).
fof(f5106,plain,(
  $false | ($spl0 | $spl29 | $spl52)),
  inference(subsumption_resolution,[],[f5105,f4642])).
fof(f4787,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | $spl52),
  inference(cnf_transformation,[],[f4787_D])).
fof(f4787_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4962,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | $spl54),
  inference(cnf_transformation,[],[f4962_D])).
fof(f4962_D,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4964,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | $spl56),
  inference(cnf_transformation,[],[f4964_D])).
fof(f4964_D,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4965,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4931,f413])).
fof(f4931,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | $spl0),
  inference(duplicate_literal_removal,[],[f4919])).
fof(f4919,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | $spl0),
  inference(superposition,[],[f1456,f1816])).
fof(f1816,plain,(
  ( ! [X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X1) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X1)))) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X1)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X1)) )),
  inference(superposition,[],[f317,f264])).
fof(f264,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X0)),s('type.num.num',sK9(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : (s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X0)),s('type.num.num',sK9(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f75])).
fof(f75,plain,(
  ! [X0] : (? [X3,X4] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X4))) | ? [X2] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))) | ? [X1] : s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X12] : (? [X13,X14] : s('type.clos_known.val_approx',X12) = s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X13),s('type.num.num',X14))) | ? [X15] : s('type.clos_known.val_approx',X12) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X15))) | ? [X16] : s('type.clos_known.val_approx',X12) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X16))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X12) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X12))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.val_approx_nchotomy')).
fof(f317,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X42,X43] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X43),s('type.num.num',X42)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|17')).
fof(f4960,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4932,f413])).
fof(f4932,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | $spl0),
  inference(duplicate_literal_removal,[],[f4918])).
fof(f4918,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(sK7)))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(sK7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | $spl0),
  inference(superposition,[],[f1816,f1456])).
fof(f4783,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | $spl48),
  inference(cnf_transformation,[],[f4783_D])).
fof(f4783_D,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4785,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | $spl50),
  inference(cnf_transformation,[],[f4785_D])).
fof(f4785_D,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4939,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4938,f1456])).
fof(f4938,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4937,f1456])).
fof(f4937,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4936,f1456])).
fof(f4936,plain,(
  s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4935,f413])).
fof(f4935,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(forward_demodulation,[],[f4933,f1456])).
fof(f4933,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(duplicate_literal_removal,[],[f4897])).
fof(f4897,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(superposition,[],[f1816,f1456])).
fof(f4781,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | $spl46),
  inference(cnf_transformation,[],[f4781_D])).
fof(f4781_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4779,plain,(
  s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4778,f1456])).
fof(f4778,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4777,f1456])).
fof(f4777,plain,(
  s('type.clos_known.val_approx',sK7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4776,f1456])).
fof(f4776,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4775,f413])).
fof(f4775,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(forward_demodulation,[],[f4743,f1456])).
fof(f4743,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',sK7),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10('const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))) | $spl0),
  inference(superposition,[],[f1815,f1456])).
fof(f1815,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(superposition,[],[f316,f264])).
fof(f4690,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK12(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X0,X1,bool,bool))))) ) | $spl42),
  inference(cnf_transformation,[],[f4690_D])).
fof(f4690_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK12(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X0,X1,bool,bool))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4692,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl45),
  inference(cnf_transformation,[],[f4692_D])).
fof(f4692_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f4680,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK12(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1715,f524])).
fof(f1715,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK12(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK12(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f347,f263])).
fof(f347,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK12(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK12(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.clos_knownProof.subapprox_Other|split|0',aHL_EXT)).
fof(f4679,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK12(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK12(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1715,f524])).
fof(f4669,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl36),
  inference(cnf_transformation,[],[f4669_D])).
fof(f4669_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4665,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool))) ) | $spl32),
  inference(cnf_transformation,[],[f4665_D])).
fof(f4665_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4672,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl41),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f4663,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK12(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1714,f263])).
fof(f1714,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,X2,bool))) )),
  inference(superposition,[],[f347,f263])).
fof(f4670,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl38),
  inference(cnf_transformation,[],[f4670_D])).
fof(f4670_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4662,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK12(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1714,f524])).
fof(f4661,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK12(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1714,f524])).
fof(f4664,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f4664_D])).
fof(f4664_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4667,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl35),
  inference(cnf_transformation,[],[f4667_D])).
fof(f4667_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f4660,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK12(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1714,f524])).
fof(f4659,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK12(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1714,f524])).
fof(f4640,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | $spl26),
  inference(cnf_transformation,[],[f4640_D])).
fof(f4640_D,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4642,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.Impossible_0') | $spl29),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.Impossible_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f4572,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.Impossible_0') | s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Impossible_0')))),
  inference(superposition,[],[f1455,f249])).
fof(f3799,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791,f3763,f3794,f3796,f3768,f3772])).
fof(f3772,plain,(
  ( ! [X12,X13,X11] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx',X11),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.clos_known.val_approx',X11))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X11) | s('type.clos_known.val_approx',X11) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X11)))) | s('type.clos_known.val_approx',X11) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X11)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X11)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f1822])).
fof(f3768,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X0)))))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f401])).
fof(f3796,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f3766,f247])).
fof(f3766,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f399])).
fof(f3794,plain,(
  ( ! [X2,X3] : (s('type.num.num',X2) = s('type.num.num',X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f3765,f247])).
fof(f3765,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.num.num',X2) = s('type.num.num',X3)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f398])).
fof(f3763,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X1),s('type.list.list'('type.clos_known.val_approx'),X0))))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f416])).
fof(f3791,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3761,f247])).
fof(f3761,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f351])).
fof(f3759,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f350])).
fof(f3788,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl6),
  inference(subsumption_resolution,[],[f3757,f247])).
fof(f3757,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) = s('type.integer.int',X1)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f348])).
fof(f3786,plain,(
  ( ! [X21,X22] : (s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X22),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X21) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X21)))) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X21)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f3752,f247])).
fof(f3752,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X22),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X21) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X21)))) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X21)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X21)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f1827])).
fof(f3784,plain,(
  ( ! [X19,X20] : (s('type.clos_known.val_approx',X20) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X19)),s('type.num.num',sK9(X19)))),s('type.clos_known.val_approx',X20))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X19) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X19)))) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X19)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X19)) ) | $spl6),
  inference(subsumption_resolution,[],[f3751,f247])).
fof(f3751,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | s('type.clos_known.val_approx',X20) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X19)),s('type.num.num',sK9(X19)))),s('type.clos_known.val_approx',X20))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X19) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X19)))) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X19)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X19)) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f1826])).
fof(f3782,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) ) | $spl6),
  inference(subsumption_resolution,[],[f3749,f247])).
fof(f3749,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.clos_known.val_approx',X1) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx',X1)))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f330])).
fof(f3779,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f3739,f3666])).
fof(f3739,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f3666,f275])).
fof(f3465,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') != s(bool,f(s(X116,sK12(X117,X118,bool,X116)))) | s(fun(X116,bool),X117) = s(fun(X116,bool),X118) | s(bool,X119) = s(bool,g(s(X116,sK12(X117,X118,bool,X116)))) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f347,f524])).
fof(f3464,plain,(
  ( ! [X114,X112,X115,X113] : (s(X114,f(s(bool,'T_0'))) != s(X114,g(s(bool,'T_0'))) | s(fun(bool,X114),X112) = s(fun(bool,X114),X113) | s(bool,X115) = s(bool,sK12(X112,X113,X114,bool)) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f347,f524])).
fof(f3713,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s('type.num.num',X104) = s('type.num.num',X105) | s(bool,X108) = s(bool,sK15(X104,X105,X106,X107)) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f3463,f247])).
fof(f3463,plain,(
  ( ! [X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | s('type.num.num',X104) = s('type.num.num',X105) | s(bool,X108) = s(bool,sK15(X104,X105,X106,X107)) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f398,f524])).
fof(f3712,plain,(
  ( ! [X103,X101,X99,X102,X100] : (s('type.num.num',X101) = s('type.num.num',X102) | s(bool,X103) = s(bool,sK15(X99,X100,X101,X102)) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f3462,f247])).
fof(f3462,plain,(
  ( ! [X103,X101,X99,X102,X100] : (~p(s(bool,'T_0')) | s('type.num.num',X101) = s('type.num.num',X102) | s(bool,X103) = s(bool,sK15(X99,X100,X101,X102)) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f399,f524])).
fof(f3461,plain,(
  ( ! [X94,X97,X95,X98,X96] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X95),s('type.num.num',X97))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X94),s('type.num.num',X96))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X95),s('type.num.num',X97))))) | s(bool,X98) = s(bool,sK15(X94,X95,X96,X97)) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f401,f524])).
fof(f3711,plain,(
  ( ! [X92,X93,X91] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X91))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X92))) | s(bool,X93) = s(bool,sK14(X91,X92)) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f3460,f247])).
fof(f3460,plain,(
  ( ! [X92,X93,X91] : (~p(s(bool,'T_0')) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X91))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X92))) | s(bool,X93) = s(bool,sK14(X91,X92)) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f351,f524])).
fof(f3459,plain,(
  ( ! [X90,X88,X89] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X88))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X89))))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),'const.list.MAP2_3'(s(fun('type.clos_known.val_approx',fun('type.clos_known.val_approx','type.clos_known.val_approx')),'const.clos_known.merge_0'),s('type.list.list'('type.clos_known.val_approx'),X88),s('type.list.list'('type.clos_known.val_approx'),X89))))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s(bool,X90) = s(bool,sK14(X88,X89)) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f416,f524])).
fof(f3710,plain,(
  ( ! [X85,X86,X84] : (s('type.integer.int',X84) = s('type.integer.int',X85) | s(bool,X86) = s(bool,sK13(X84,X85)) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f3458,f247])).
fof(f3458,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'T_0')) | s('type.integer.int',X84) = s('type.integer.int',X85) | s(bool,X86) = s(bool,sK13(X84,X85)) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f348,f524])).
fof(f3457,plain,(
  ( ! [X83,X81,X82] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X82))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X81))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X82))))) | s(bool,X83) = s(bool,sK13(X81,X82)) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f350,f524])).
fof(f3709,plain,(
  ( ! [X80,X78,X79] : (s('type.clos_known.val_approx',X79) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X78),s('type.clos_known.val_approx',X79))) | s(bool,X80) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X78),s('type.clos_known.val_approx',X79))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f3456,f247])).
fof(f3456,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'T_0')) | s('type.clos_known.val_approx',X79) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X78),s('type.clos_known.val_approx',X79))) | s(bool,X80) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X78),s('type.clos_known.val_approx',X79))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f330,f524])).
fof(f3708,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~sP6(sK13(X66,X66),X65,X64) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f3449,f247])).
fof(f3449,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~sP6(sK13(X66,X66),X65,X64) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f839,f524])).
fof(f3707,plain,(
  ( ! [X59,X57,X60,X58] : (~sP5(X57,X58,X59) | sP6(X57,X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3447,f247])).
fof(f3447,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~sP5(X57,X58,X59) | sP6(X57,X58,X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f402,f524])).
fof(f3706,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6(X53,X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f3446,f247])).
fof(f3446,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP6(X53,X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f395,f524])).
fof(f3705,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | sP6(X49,X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f3445,f247])).
fof(f3445,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP6(X49,X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f394,f524])).
fof(f3704,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | sP6(X45,X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f3444,f247])).
fof(f3444,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP6(X45,X47,X46) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f393,f524])).
fof(f3703,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,X41) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f3443,f247])).
fof(f3443,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | ~p(s(bool,X43)) | ~sP5(X43,X42,X41) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f388,f524])).
fof(f3702,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,X37,X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f3442,f247])).
fof(f3442,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP2(X38,X37,X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f369,f524])).
fof(f3701,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,X33,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f3441,f247])).
fof(f3441,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP3(X34,X33,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f365,f524])).
fof(f3700,plain,(
  ( ! [X30,X31,X29,X32] : (sP4(X30,X31,X29) | ~sP3(X30,X29,X31) | ~sP2(X31,X29,X30) | ~sP1(X29,X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f3440,f247])).
fof(f3440,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | sP4(X30,X31,X29) | ~sP3(X30,X29,X31) | ~sP2(X31,X29,X30) | ~sP1(X29,X31,X30) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f362,f524])).
fof(f3699,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | ~sP4(X26,X25,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f3439,f247])).
fof(f3439,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP4(X26,X25,X27) | s(bool,X25) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f361,f524])).
fof(f3698,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3438,f247])).
fof(f3438,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f334,f524])).
fof(f3697,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f3437,f247])).
fof(f3437,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f333,f524])).
fof(f3696,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f3436,f247])).
fof(f3436,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f325,f524])).
fof(f3695,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3432,f247])).
fof(f3432,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f275,f524])).
fof(f3407,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f271,f524])).
fof(f3367,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) != s(bool,f(s(X120,sK12(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,g(s(X120,sK12(X121,X122,bool,X120))))) )),
  inference(superposition,[],[f347,f524])).
fof(f3366,plain,(
  ( ! [X118,X116,X119,X117] : (s(X118,f(s(bool,X119))) != s(X118,g(s(bool,X119))) | s(fun(bool,X118),X116) = s(fun(bool,X118),X117) | s(bool,'T_0') = s(bool,X119) | s(bool,'T_0') = s(bool,sK12(X116,X117,X118,bool))) )),
  inference(superposition,[],[f347,f524])).
fof(f3309,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f524,f524])).
fof(f3200,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X0),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X0)))) | s('type.clos_known.val_approx',X0) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X0)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X0)) )),
  inference(superposition,[],[f1382,f264])).
fof(f3028,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f1382,f325])).
fof(f3029,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f1382,f325])).
fof(f3030,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f333])).
fof(f3031,plain,(
  ( ! [X2,X0,X1] : (sP0('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1382,f334])).
fof(f3034,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f361])).
fof(f3035,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1382,f414,f1382,f361])).
fof(f3037,plain,(
  ( ! [X0,X1] : (sP4('T_0','T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f420,f1382,f362])).
fof(f3041,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1382,f362])).
fof(f3043,plain,(
  ( ! [X2,X0,X1] : (sP4(sK14(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f420,f1382,f362])).
fof(f3053,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1382,f362])).
fof(f3057,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1382,f362])).
fof(f3059,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK14(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f501,f1382,f362])).
fof(f3061,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',sK14(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f1132,f1382,f362])).
fof(f3065,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f1132,f1382,f362])).
fof(f3067,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(sK14(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f1132,f1382,f362])).
fof(f3068,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f365])).
fof(f3077,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f365])).
fof(f3078,plain,(
  ( ! [X2,X0,X3,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1382,f414,f1382,f365])).
fof(f3086,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f366])).
fof(f3087,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f369])).
fof(f3096,plain,(
  ( ! [X0,X1] : (~sP2('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f369])).
fof(f3097,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1382,f414,f1382,f369])).
fof(f3105,plain,(
  ( ! [X2,X0,X3,X1] : (sP2(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f370])).
fof(f3106,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f374])).
fof(f3107,plain,(
  ( ! [X2,X0,X3,X1] : (sP1(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1382,f375])).
fof(f3108,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f1382,f376])).
fof(f3109,plain,(
  ( ! [X2,X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X2,X2),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1382,f378])).
fof(f3111,plain,(
  ( ! [X2,X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',sK13(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f497,f1382,f378])).
fof(f3113,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1382,f379])).
fof(f3114,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1382,f379])).
fof(f3115,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f382])).
fof(f3144,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1382,f388])).
fof(f3151,plain,(
  ( ! [X0,X1] : (~sP5('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1382,f388])).
fof(f3152,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f1382,f388])).
fof(f3158,plain,(
  ( ! [X2,X0,X3,X1] : (sP5(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1382,f390])).
fof(f3159,plain,(
  ( ! [X0,X1] : (sP6('F_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f414,f1382,f392])).
fof(f3165,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f247,f1382,f394])).
fof(f3169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X4)) )),
  inference(unit_resulting_resolution,[],[f1382,f1382,f394])).
fof(f3170,plain,(
  ( ! [X0,X1] : (~sP6('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f395])).
fof(f3175,plain,(
  ( ! [X0,X1] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1382,f395])).
fof(f3176,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1382,f414,f1382,f395])).
fof(f3177,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1382,f402])).
fof(f3178,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f435,f1382,f402])).
fof(f3180,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X2,X2),X3)) )),
  inference(unit_resulting_resolution,[],[f509,f1382,f402])).
fof(f3181,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X2,X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f757])).
fof(f3182,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X2,X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f758])).
fof(f3187,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X2,X2),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f759])).
fof(f3192,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f1382,f761])).
fof(f3193,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0',sK13(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f799])).
fof(f3197,plain,(
  ( ! [X2,X0,X1] : (~sP6(sK13(X0,X0),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1382,f839])).
fof(f3010,plain,(
  ( ! [X0,X1] : (p(s(bool,sK15(X0,X0,X1,X1)))) )),
  inference(equality_resolution,[],[f1296])).
fof(f2754,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f1383,f325])).
fof(f2755,plain,(
  ( ! [X0,X1] : (s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f1383,f325])).
fof(f2756,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f333])).
fof(f2757,plain,(
  ( ! [X0,X1] : (sP0('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1383,f334])).
fof(f2760,plain,(
  ( ! [X0] : (~sP4('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f361])).
fof(f2761,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1383,f414,f1383,f361])).
fof(f2763,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f420,f1383,f362])).
fof(f2767,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1383,f362])).
fof(f2769,plain,(
  ( ! [X0,X1] : (sP4(sK14(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f420,f1383,f362])).
fof(f2779,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1383,f362])).
fof(f2783,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1383,f362])).
fof(f2785,plain,(
  ( ! [X2,X0,X1] : (sP4(sK14(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f501,f1383,f362])).
fof(f2787,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK14(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f1132,f1383,f362])).
fof(f2791,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f1132,f1383,f362])).
fof(f2793,plain,(
  ( ! [X2,X0,X1] : (sP4(sK14(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f1132,f1383,f362])).
fof(f2794,plain,(
  ( ! [X0] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f365])).
fof(f2803,plain,(
  ( ! [X0] : (~sP3('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f365])).
fof(f2804,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1383,f414,f1383,f365])).
fof(f2812,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f366])).
fof(f2813,plain,(
  ( ! [X0] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f369])).
fof(f2822,plain,(
  ( ! [X0] : (~sP2('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f369])).
fof(f2823,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1383,f414,f1383,f369])).
fof(f2831,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f370])).
fof(f2832,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f374])).
fof(f2833,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1383,f375])).
fof(f2834,plain,(
  ( ! [X2,X0,X1] : (sP1('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1383,f376])).
fof(f2835,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1383,f378])).
fof(f2837,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f497,f1383,f378])).
fof(f2839,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1383,f379])).
fof(f2840,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1383,f379])).
fof(f2841,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f382])).
fof(f2869,plain,(
  ( ! [X0] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1383,f388])).
fof(f2876,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1383,f388])).
fof(f2877,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f1383,f388])).
fof(f2883,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1383,f390])).
fof(f2884,plain,(
  ( ! [X0] : (sP6('F_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f414,f1383,f392])).
fof(f2889,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f247,f1383,f394])).
fof(f2893,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1383,f1383,f394])).
fof(f2894,plain,(
  ( ! [X0] : (~sP6('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f395])).
fof(f2899,plain,(
  ( ! [X0] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1383,f395])).
fof(f2900,plain,(
  ( ! [X0,X1] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1383,f414,f1383,f395])).
fof(f2901,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1383,f402])).
fof(f2902,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f435,f1383,f402])).
fof(f2904,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f509,f1383,f402])).
fof(f2905,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f757])).
fof(f2906,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f758])).
fof(f2911,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f759])).
fof(f2916,plain,(
  ( ! [X2,X0,X1] : (sP6(sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1383,f761])).
fof(f2917,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f799])).
fof(f2921,plain,(
  ( ! [X0,X1] : (~sP6(sK13(X0,X0),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1383,f839])).
fof(f2746,plain,(
  ( ! [X0] : (sP4('F_0','F_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f423,f428,f414,f1152,f363])).
fof(f2728,plain,(
  ( ! [X0] : (sP4('F_0',sK14(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1132,f426,f414,f1151,f364])).
fof(f2508,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f1384,f325])).
fof(f2509,plain,(
  ( ! [X0,X1] : (s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f1384,f325])).
fof(f2510,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f333])).
fof(f2511,plain,(
  ( ! [X0,X1] : (sP0('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1384,f334])).
fof(f2514,plain,(
  ( ! [X0] : (~sP4('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f361])).
fof(f2515,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1384,f414,f1384,f361])).
fof(f2517,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f420,f1384,f362])).
fof(f2521,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1384,f362])).
fof(f2523,plain,(
  ( ! [X0,X1] : (sP4(sK14(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f420,f1384,f362])).
fof(f2533,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1384,f362])).
fof(f2537,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1384,f362])).
fof(f2539,plain,(
  ( ! [X2,X0,X1] : (sP4(sK14(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f501,f1384,f362])).
fof(f2540,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK14(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f1132,f1384,f362])).
fof(f2542,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f1132,f1384,f362])).
fof(f2543,plain,(
  ( ! [X2,X0,X1] : (sP4(sK14(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1150,f1149,f1132,f1384,f362])).
fof(f2544,plain,(
  ( ! [X0] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f365])).
fof(f2553,plain,(
  ( ! [X0] : (~sP3('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f365])).
fof(f2554,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1384,f414,f1384,f365])).
fof(f2562,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f366])).
fof(f2563,plain,(
  ( ! [X0] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f369])).
fof(f2572,plain,(
  ( ! [X0] : (~sP2('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f369])).
fof(f2573,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1384,f414,f1384,f369])).
fof(f2581,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f370])).
fof(f2582,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f374])).
fof(f2583,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1384,f375])).
fof(f2584,plain,(
  ( ! [X2,X0,X1] : (sP1('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1384,f376])).
fof(f2585,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1384,f378])).
fof(f2587,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f497,f1384,f378])).
fof(f2589,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1384,f379])).
fof(f2590,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1384,f379])).
fof(f2591,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f382])).
fof(f2619,plain,(
  ( ! [X0] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1384,f388])).
fof(f2626,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1384,f388])).
fof(f2627,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f1384,f388])).
fof(f2633,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1384,f390])).
fof(f2634,plain,(
  ( ! [X0] : (sP6('F_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f414,f1384,f392])).
fof(f2639,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f247,f1384,f394])).
fof(f2643,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1384,f1384,f394])).
fof(f2644,plain,(
  ( ! [X0] : (~sP6('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f395])).
fof(f2649,plain,(
  ( ! [X0] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1384,f395])).
fof(f2650,plain,(
  ( ! [X0,X1] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1384,f414,f1384,f395])).
fof(f2651,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1384,f402])).
fof(f2652,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f435,f1384,f402])).
fof(f2654,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f509,f1384,f402])).
fof(f2655,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f757])).
fof(f2656,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f758])).
fof(f2661,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f759])).
fof(f2666,plain,(
  ( ! [X2,X0,X1] : (sP6(sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f1384,f761])).
fof(f2667,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f799])).
fof(f2671,plain,(
  ( ! [X0,X1] : (~sP6(sK13(X0,X0),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1384,f839])).
fof(f2496,plain,(
  ( ! [X0] : (sP4(sK14(X0,X0),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f421,f426,f414,f1150,f363])).
fof(f2484,plain,(
  ( ! [X0] : (sP4(sK14(X0,X0),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f432,f420,f1149,f362])).
fof(f2489,plain,(
  ( ! [X0,X1] : (sP4(sK14(X0,X0),sK14(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f432,f1132,f1149,f362])).
fof(f2290,plain,(
  s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f247,f1381,f325])).
fof(f2292,plain,(
  ( ! [X0] : (sP0(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1381,f333])).
fof(f2293,plain,(
  ( ! [X0] : (sP0('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1381,f334])).
fof(f2296,plain,(
  ~sP4('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f361])).
fof(f2297,plain,(
  ~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1381,f414,f1381,f361])).
fof(f2298,plain,(
  sP4('T_0','T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f431,f425,f420,f1381,f362])).
fof(f2300,plain,(
  ( ! [X0] : (sP4(sK13(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1381,f362])).
fof(f2307,plain,(
  ( ! [X0] : (sP4('T_0',sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1381,f362])).
fof(f2309,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1381,f362])).
fof(f2312,plain,(
  ( ! [X0] : (sP4('T_0',sK14(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f1132,f1381,f362])).
fof(f2313,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f1132,f1381,f362])).
fof(f2315,plain,(
  ~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f365])).
fof(f2324,plain,(
  ~sP3('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f365])).
fof(f2325,plain,(
  ~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1381,f414,f1381,f365])).
fof(f2333,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1381,f366])).
fof(f2334,plain,(
  ~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f369])).
fof(f2341,plain,(
  ~sP2('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f369])).
fof(f2342,plain,(
  ~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1381,f414,f1381,f369])).
fof(f2348,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1381,f370])).
fof(f2349,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f1381,f374])).
fof(f2350,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1381,f375])).
fof(f2351,plain,(
  ( ! [X0,X1] : (sP1('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1381,f376])).
fof(f2352,plain,(
  ( ! [X0] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X0,X0),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1381,f378])).
fof(f2354,plain,(
  ( ! [X0] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f497,f1381,f378])).
fof(f2356,plain,(
  ~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f438,f1381,f379])).
fof(f2357,plain,(
  ~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f438,f1381,f379])).
fof(f2358,plain,(
  ( ! [X0] : (~sP6('F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f382])).
fof(f2386,plain,(
  ~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f414,f247,f1381,f388])).
fof(f2393,plain,(
  ~sP5('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f414,f247,f1381,f388])).
fof(f2394,plain,(
  ~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f414,f1381,f1381,f388])).
fof(f2400,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1381,f390])).
fof(f2401,plain,(
  sP6('F_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f414,f414,f1381,f392])).
fof(f2406,plain,(
  ( ! [X0] : (sP6('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f247,f1381,f394])).
fof(f2410,plain,(
  ( ! [X0] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1381,f1381,f394])).
fof(f2411,plain,(
  ~sP6('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f395])).
fof(f2416,plain,(
  ~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f1381,f395])).
fof(f2417,plain,(
  ~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))),
  inference(unit_resulting_resolution,[],[f1381,f414,f1381,f395])).
fof(f2418,plain,(
  ( ! [X0] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1381,f402])).
fof(f2419,plain,(
  ( ! [X0] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f435,f1381,f402])).
fof(f2421,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f509,f1381,f402])).
fof(f2422,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f757])).
fof(f2423,plain,(
  ( ! [X0] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f758])).
fof(f2428,plain,(
  ( ! [X0] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f759])).
fof(f2432,plain,(
  ( ! [X0,X1] : (sP6(sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f1381,f761])).
fof(f2433,plain,(
  ( ! [X0] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')),'F_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f799])).
fof(f2437,plain,(
  ( ! [X0] : (~sP6(sK13(X0,X0),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) )),
  inference(unit_resulting_resolution,[],[f414,f1381,f839])).
fof(f2283,plain,(
  ( ! [X0] : (~sP4('F_0','T_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f1143,f359])).
fof(f2278,plain,(
  ( ! [X0] : (~sP4('F_0',sK14(X0,X0),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1135,f359])).
fof(f2028,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(unit_resulting_resolution,[],[f247,f1380,f325])).
fof(f2029,plain,(
  ( ! [X0,X1] : (s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0))) = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f1380,f1380,f325])).
fof(f2030,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f1380,f333])).
fof(f2031,plain,(
  ( ! [X0,X1] : (sP0('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f1380,f334])).
fof(f2034,plain,(
  ( ! [X0] : (~sP4('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f361])).
fof(f2035,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1380,f414,f1380,f361])).
fof(f2037,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f420,f1380,f362])).
fof(f2039,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),'T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1380,f362])).
fof(f2049,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1380,f362])).
fof(f2051,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X2)))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1380,f362])).
fof(f2054,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK14(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f1132,f1380,f362])).
fof(f2055,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X2)))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f1132,f1380,f362])).
fof(f2057,plain,(
  ( ! [X0] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f365])).
fof(f2066,plain,(
  ( ! [X0] : (~sP3('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f365])).
fof(f2067,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1380,f414,f1380,f365])).
fof(f2075,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X2)))) )),
  inference(unit_resulting_resolution,[],[f1380,f366])).
fof(f2076,plain,(
  ( ! [X0] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f369])).
fof(f2083,plain,(
  ( ! [X0] : (~sP2('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f369])).
fof(f2084,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1380,f414,f1380,f369])).
fof(f2090,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X2)))) )),
  inference(unit_resulting_resolution,[],[f1380,f370])).
fof(f2091,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X2)))) )),
  inference(unit_resulting_resolution,[],[f1380,f374])).
fof(f2092,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1380,f375])).
fof(f2093,plain,(
  ( ! [X2,X0,X1] : (sP1('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1380,f376])).
fof(f2094,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),sK13(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1380,f378])).
fof(f2096,plain,(
  ( ! [X0,X1] : (sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'T_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f497,f1380,f378])).
fof(f2098,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1380,f379])).
fof(f2099,plain,(
  ( ! [X0] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1380,f379])).
fof(f2100,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f382])).
fof(f2128,plain,(
  ( ! [X0] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1380,f388])).
fof(f2135,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1380,f388])).
fof(f2136,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f1380,f388])).
fof(f2142,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1380,f390])).
fof(f2143,plain,(
  ( ! [X0] : (sP6('F_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(unit_resulting_resolution,[],[f414,f414,f1380,f392])).
fof(f2148,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f247,f1380,f394])).
fof(f2152,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1380,f1380,f394])).
fof(f2153,plain,(
  ( ! [X0] : (~sP6('T_0','F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)))) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f395])).
fof(f2158,plain,(
  ( ! [X0] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1380,f395])).
fof(f2159,plain,(
  ( ! [X0,X1] : (~sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f1380,f414,f1380,f395])).
fof(f2160,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1380,f402])).
fof(f2161,plain,(
  ( ! [X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f435,f1380,f402])).
fof(f2163,plain,(
  ( ! [X2,X0,X1] : (sP6('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f509,f1380,f402])).
fof(f2164,plain,(
  ( ! [X0,X1] : (~sP4('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f757])).
fof(f2165,plain,(
  ( ! [X0,X1] : (~sP3('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f758])).
fof(f2170,plain,(
  ( ! [X0,X1] : (~sP2('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f759])).
fof(f2174,plain,(
  ( ! [X2,X0,X1] : (sP6(sK13(X0,X0),'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1380,f761])).
fof(f2175,plain,(
  ( ! [X0,X1] : (~sP5('const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0)),'F_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f799])).
fof(f2179,plain,(
  ( ! [X0,X1] : (~sP6(sK13(X0,X0),'F_0','const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X1)))) )),
  inference(unit_resulting_resolution,[],[f414,f1380,f839])).
fof(f2010,plain,(
  ( ! [X0] : (sP4('T_0',sK14(X0,X0),'T_0')) )),
  inference(unit_resulting_resolution,[],[f430,f425,f247,f1132,f362])).
fof(f2012,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),sK14(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f503,f502,f247,f1132,f362])).
fof(f2003,plain,(
  ( ! [X0] : (~sP4('T_0','F_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f1126,f360])).
fof(f1829,plain,(
  ( ! [X26,X25] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X26),s('type.clos_known.val_approx',X25)))) | s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X25)),s('type.num.num',sK9(X25)))) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X26),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X25)),s('type.num.num',sK9(X25)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X25) | s('type.clos_known.val_approx',X25) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X25)))) | s('type.clos_known.val_approx',X25) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X25)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X25)) )),
  inference(superposition,[],[f331,f264])).
fof(f1828,plain,(
  ( ! [X24,X23] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X23),s('type.clos_known.val_approx',X24)))) | s('type.clos_known.val_approx',X24) != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X23)),s('type.num.num',sK9(X23)))),s('type.clos_known.val_approx',X24))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X23) | s('type.clos_known.val_approx',X23) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X23)))) | s('type.clos_known.val_approx',X23) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X23)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X23)) )),
  inference(superposition,[],[f331,f264])).
fof(f1827,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X22),s('type.clos_known.val_approx',X21)))) | s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X22),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X21)),s('type.num.num',sK9(X21)))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X21) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X21)))) | s('type.clos_known.val_approx',X21) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X21)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X21)) )),
  inference(superposition,[],[f330,f264])).
fof(f1826,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx',X19),s('type.clos_known.val_approx',X20)))) | s('type.clos_known.val_approx',X20) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',sK8(X19)),s('type.num.num',sK9(X19)))),s('type.clos_known.val_approx',X20))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X19) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X19)))) | s('type.clos_known.val_approx',X19) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X19)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X19)) )),
  inference(superposition,[],[f330,f264])).
fof(f1822,plain,(
  ( ! [X12,X13,X11] : (s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.clos_known.val_approx',X11))) = s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,sK15(X12,sK8(X11),X13,sK9(X11))),s('type.clos_known.val_approx',X11),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X11) | s('type.clos_known.val_approx',X11) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X11)))) | s('type.clos_known.val_approx',X11) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X11)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X11)) )),
  inference(superposition,[],[f401,f264])).
fof(f1821,plain,(
  ( ! [X10,X9] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X10))),s('type.clos_known.val_approx',X9))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X9) | s('type.clos_known.val_approx',X9) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X9)))) | s('type.clos_known.val_approx',X9) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X9)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X9)) )),
  inference(superposition,[],[f357,f264])).
fof(f1820,plain,(
  ( ! [X8,X7] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X7),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X8))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X7) | s('type.clos_known.val_approx',X7) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X7)))) | s('type.clos_known.val_approx',X7) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X7)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X7)) )),
  inference(superposition,[],[f356,f264])).
fof(f1819,plain,(
  ( ! [X6,X5] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X6))),s('type.clos_known.val_approx',X5))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X5) | s('type.clos_known.val_approx',X5) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X5)))) | s('type.clos_known.val_approx',X5) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X5)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X5)) )),
  inference(superposition,[],[f355,f264])).
fof(f1818,plain,(
  ( ! [X4,X3] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X3),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X4))))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X3) | s('type.clos_known.val_approx',X3) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X3)))) | s('type.clos_known.val_approx',X3) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X3)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X3)) )),
  inference(superposition,[],[f354,f264])).
fof(f1817,plain,(
  ( ! [X2] : (s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx',X2),s('type.clos_known.val_approx','const.clos_known.Impossible_0'))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X2) | s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',sK11(X2)))) | s('type.clos_known.val_approx',X2) = s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),sK10(X2)))) | s('type.clos_known.val_approx','const.clos_known.Impossible_0') = s('type.clos_known.val_approx',X2)) )),
  inference(superposition,[],[f320,f264])).
fof(f1733,plain,(
  ( ! [X0] : (s('type.clos_known.val_approx','const.bool.COND_3'(s(bool,'T_0'),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))) = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0)))))) )),
  inference(superposition,[],[f350,f497])).
fof(f1381,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))))),
  inference(unit_resulting_resolution,[],[f250,f331])).
fof(f1380,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Impossible_0'),s('type.clos_known.val_approx',X0))))) )),
  inference(unit_resulting_resolution,[],[f256,f331])).
fof(f1382,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) )),
  inference(unit_resulting_resolution,[],[f316,f331])).
fof(f1383,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) )),
  inference(unit_resulting_resolution,[],[f260,f331])).
fof(f1384,plain,(
  ( ! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X0))),s('type.clos_known.val_approx','const.clos_known.Other_0'))))) )),
  inference(unit_resulting_resolution,[],[f258,f331])).
fof(f1296,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | p(s(bool,sK15(X2,X2,X1,X0)))) )),
  inference(equality_resolution,[],[f400])).
fof(f1100,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f247,f1094,f325])).
fof(f1101,plain,(
  ( ! [X0,X1] : (s(bool,sK14(X0,X0)) = s(bool,sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f1094,f1094,f325])).
fof(f1102,plain,(
  ( ! [X0,X1] : (sP0(X0,sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f1094,f333])).
fof(f1103,plain,(
  ( ! [X0,X1] : (sP0(sK14(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f1094,f334])).
fof(f1106,plain,(
  ( ! [X0] : (~sP4('T_0',sK14(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f361])).
fof(f1107,plain,(
  ( ! [X0,X1] : (~sP4(sK14(X0,X0),sK14(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f1094,f361])).
fof(f1109,plain,(
  ( ! [X0] : (sP4('T_0','T_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f420,f1094,f362])).
fof(f1111,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),'T_0',sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f1094,f362])).
fof(f1113,plain,(
  ( ! [X0,X1] : (sP4('T_0',sK13(X0,X0),sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f430,f425,f501,f1094,f362])).
fof(f1115,plain,(
  ( ! [X2,X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),sK14(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f1094,f362])).
fof(f1118,plain,(
  ( ! [X0] : (~sP3(sK14(X0,X0),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f365])).
fof(f1126,plain,(
  ( ! [X0] : (~sP3('T_0',sK14(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f365])).
fof(f1129,plain,(
  ( ! [X0,X1] : (~sP3(sK14(X0,X0),sK14(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f1094,f365])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,sK14(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f1094,f366])).
fof(f1135,plain,(
  ( ! [X0] : (~sP2(sK14(X0,X0),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f369])).
fof(f1143,plain,(
  ( ! [X0] : (~sP2('T_0',sK14(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f369])).
fof(f1146,plain,(
  ( ! [X0,X1] : (~sP2(sK14(X0,X0),sK14(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f1094,f369])).
fof(f1149,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,sK14(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f1094,f370])).
fof(f1150,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,sK14(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f1094,f374])).
fof(f1151,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,sK14(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f1094,f375])).
fof(f1152,plain,(
  ( ! [X2,X0,X1] : (sP1(sK14(X0,X0),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1094,f376])).
fof(f1154,plain,(
  ( ! [X0,X1] : (sP4(sK14(X0,X0),sK13(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f497,f1094,f378])).
fof(f1156,plain,(
  ( ! [X0,X1] : (sP4(sK14(X0,X0),'T_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f497,f1094,f378])).
fof(f1158,plain,(
  ( ! [X0] : (~sP4(sK14(X0,X0),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1094,f379])).
fof(f1160,plain,(
  ( ! [X0] : (~sP4(sK14(X0,X0),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f438,f1094,f379])).
fof(f1165,plain,(
  ( ! [X0,X1] : (~sP6('F_0',sK14(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f382])).
fof(f1189,plain,(
  ( ! [X0] : (~sP5(sK14(X0,X0),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1094,f388])).
fof(f1197,plain,(
  ( ! [X0] : (~sP5('T_0','F_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f247,f414,f1094,f388])).
fof(f1198,plain,(
  ( ! [X0,X1] : (~sP5(sK14(X0,X0),'F_0',sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f1094,f414,f1094,f388])).
fof(f1201,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,sK14(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f1094,f390])).
fof(f1204,plain,(
  ( ! [X0] : (sP6('F_0','F_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f414,f414,f1094,f392])).
fof(f1207,plain,(
  ( ! [X0,X1] : (sP6('T_0',sK14(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f247,f1094,f394])).
fof(f1210,plain,(
  ( ! [X2,X0,X1] : (sP6(sK14(X0,X0),sK14(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f1094,f1094,f394])).
fof(f1211,plain,(
  ( ! [X0] : (~sP6('T_0','F_0',sK14(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f395])).
fof(f1215,plain,(
  ( ! [X0] : (~sP6(sK14(X0,X0),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f1094,f395])).
fof(f1216,plain,(
  ( ! [X0,X1] : (~sP6(sK14(X0,X0),'F_0',sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f1094,f395])).
fof(f1217,plain,(
  ( ! [X0,X1] : (sP6(sK14(X0,X0),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f1094,f402])).
fof(f1219,plain,(
  ( ! [X0,X1] : (sP6(sK14(X0,X0),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f435,f1094,f402])).
fof(f1220,plain,(
  ( ! [X2,X0,X1] : (sP6(sK14(X0,X0),sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f509,f1094,f402])).
fof(f1221,plain,(
  ( ! [X0,X1] : (~sP4(sK14(X0,X0),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f757])).
fof(f1222,plain,(
  ( ! [X0,X1] : (~sP3(sK14(X0,X0),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f758])).
fof(f1225,plain,(
  ( ! [X0,X1] : (~sP2(sK14(X0,X0),sK13(X1,X1),'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f759])).
fof(f1228,plain,(
  ( ! [X2,X0,X1] : (sP6(sK13(X0,X0),sK14(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f1094,f761])).
fof(f1229,plain,(
  ( ! [X0,X1] : (~sP5(sK14(X0,X0),'F_0',sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f799])).
fof(f1232,plain,(
  ( ! [X0,X1] : (~sP6(sK13(X0,X0),'F_0',sK14(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f414,f1094,f839])).
fof(f1094,plain,(
  ( ! [X0] : (p(s(bool,sK14(X0,X0)))) )),
  inference(equality_resolution,[],[f352])).
fof(f1056,plain,(
  ( ! [X0] : (sP4('F_0',sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f504,f426,f421,f414,f364])).
fof(f988,plain,(
  sP4('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f430,f425,f421,f414,f363])).
fof(f989,plain,(
  ( ! [X0] : (sP4(sK13(X0,X0),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f503,f502,f421,f414,f363])).
fof(f999,plain,(
  ( ! [X0] : (sP4('F_0','F_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f505,f428,f423,f414,f363])).
fof(f956,plain,(
  ( ! [X10,X8,X9] : (sP4(X9,X10,sK13(X8,X8)) | ~sP3(X9,sK13(X8,X8),X10) | ~sP2(X10,sK13(X8,X8),X9)) )),
  inference(subsumption_resolution,[],[f955,f505])).
fof(f955,plain,(
  ( ! [X10,X8,X9] : (sP4(X9,X10,sK13(X8,X8)) | ~sP3(X9,sK13(X8,X8),X10) | ~sP2(X10,sK13(X8,X8),X9) | ~sP1(sK13(X8,X8),X10,X9)) )),
  inference(subsumption_resolution,[],[f953,f247])).
fof(f953,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | sP4(X9,X10,sK13(X8,X8)) | ~sP3(X9,sK13(X8,X8),X10) | ~sP2(X10,sK13(X8,X8),X9) | ~sP1(sK13(X8,X8),X10,X9)) )),
  inference(superposition,[],[f362,f497])).
fof(f929,plain,(
  sP4('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f432,f425,f420,f247,f362])).
fof(f932,plain,(
  ( ! [X0] : (sP4(sK13(X0,X0),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f503,f502,f420,f247,f362])).
fof(f938,plain,(
  ( ! [X0,X1] : (sP4(sK13(X0,X0),sK13(X1,X1),'T_0')) )),
  inference(unit_resulting_resolution,[],[f503,f502,f501,f247,f362])).
fof(f949,plain,(
  ( ! [X0] : (~sP2(sK13(X0,X0),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f504,f501,f728,f247,f362])).
fof(f915,plain,(
  ( ! [X0] : (~sP5(sK13(X0,X0),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f901,f387])).
fof(f910,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X9)) | ~sP6(sK13(X10,X10),X9,sK13(X8,X8))) )),
  inference(subsumption_resolution,[],[f907,f247])).
fof(f907,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | ~sP6(sK13(X10,X10),X9,sK13(X8,X8))) )),
  inference(superposition,[],[f839,f497])).
fof(f901,plain,(
  ( ! [X0] : (~sP6(sK13(X0,X0),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f839])).
fof(f891,plain,(
  ( ! [X0] : (~sP6('T_0','F_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f881,f383])).
fof(f881,plain,(
  ( ! [X0] : (~sP5('T_0','F_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f414,f247,f799])).
fof(f868,plain,(
  ( ! [X0] : (~sP2('T_0',sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f759])).
fof(f863,plain,(
  ( ! [X0] : (~sP4('T_0','F_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f852,f360])).
fof(f852,plain,(
  ( ! [X0] : (~sP3('T_0',sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f758])).
fof(f847,plain,(
  ( ! [X0] : (~sP4('T_0',sK13(X0,X0),'F_0')) )),
  inference(unit_resulting_resolution,[],[f247,f414,f757])).
fof(f839,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~sP6(sK13(X8,X8),X10,X9)) )),
  inference(subsumption_resolution,[],[f835,f247])).
fof(f835,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP6(sK13(X8,X8),X10,X9)) )),
  inference(superposition,[],[f395,f497])).
fof(f803,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f791,f383])).
fof(f799,plain,(
  ( ! [X10,X8,X9] : (~sP5(X10,X9,sK13(X8,X8)) | ~p(s(bool,X10)) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f795,f247])).
fof(f795,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | ~p(s(bool,X10)) | ~sP5(X10,X9,sK13(X8,X8))) )),
  inference(superposition,[],[f388,f497])).
fof(f791,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f414,f247,f388])).
fof(f777,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f436,f247,f414,f387])).
fof(f762,plain,(
  ( ! [X33,X31,X32] : (sP6(sK13(X31,X31),X32,X33) | ~sP5(sK13(X31,X31),X32,X33)) )),
  inference(subsumption_resolution,[],[f749,f247])).
fof(f749,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | ~sP5(sK13(X31,X31),X32,X33) | sP6(sK13(X31,X31),X32,X33)) )),
  inference(superposition,[],[f402,f497])).
fof(f761,plain,(
  ( ! [X30,X28,X29] : (sP6(sK13(X28,X28),X29,X30) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f748,f247])).
fof(f748,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP6(sK13(X28,X28),X29,X30)) )),
  inference(superposition,[],[f394,f497])).
fof(f760,plain,(
  ( ! [X26,X27,X25] : (sP6(sK13(X25,X25),X27,X26) | p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f747,f247])).
fof(f747,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP6(sK13(X25,X25),X27,X26)) )),
  inference(superposition,[],[f393,f497])).
fof(f759,plain,(
  ( ! [X24,X23,X22] : (~sP2(X23,sK13(X22,X22),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f746,f247])).
fof(f746,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP2(X23,sK13(X22,X22),X24)) )),
  inference(superposition,[],[f369,f497])).
fof(f758,plain,(
  ( ! [X21,X19,X20] : (~sP3(X20,sK13(X19,X19),X21) | p(s(bool,X21)) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f745,f247])).
fof(f745,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP3(X20,sK13(X19,X19),X21)) )),
  inference(superposition,[],[f365,f497])).
fof(f757,plain,(
  ( ! [X17,X18,X16] : (~sP4(X17,sK13(X16,X16),X18) | p(s(bool,X18)) | ~p(s(bool,X17))) )),
  inference(subsumption_resolution,[],[f744,f247])).
fof(f744,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | p(s(bool,X18)) | ~sP4(X17,sK13(X16,X16),X18)) )),
  inference(superposition,[],[f361,f497])).
fof(f728,plain,(
  ( ! [X0] : (~sP4('F_0',sK13(X0,X0),'T_0')) )),
  inference(unit_resulting_resolution,[],[f414,f497,f380])).
fof(f729,plain,(
  ( ! [X0] : (sP4('T_0',sK13(X0,X0),'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f497,f378])).
fof(f730,plain,(
  ( ! [X0] : (~sP4('F_0','T_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f414,f497,f380])).
fof(f731,plain,(
  ( ! [X0] : (sP4('T_0','T_0',sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f247,f497,f378])).
fof(f718,plain,(
  ~sP4('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f709,f358])).
fof(f709,plain,(
  ~sP1('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f414,f414,f414,f373])).
fof(f703,plain,(
  ~sP4('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f693,f359])).
fof(f693,plain,(
  ~sP2('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f414,f247,f247,f369])).
fof(f671,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f414,f247,f247,f365])).
fof(f609,plain,(
  ~sP4('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f247,f438,f379])).
fof(f613,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f247,f438,f379])).
fof(f589,plain,(
  sP4('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f414,f438,f377])).
fof(f591,plain,(
  sP4('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f414,f438,f377])).
fof(f570,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f564,f414])).
fof(f564,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f322,f263])).
fof(f497,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f495,f272])).
fof(f498,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,sK13(X0,X0))) )),
  inference(unit_resulting_resolution,[],[f495,f275])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP0(X0,sK13(X1,X1))) )),
  inference(unit_resulting_resolution,[],[f495,f333])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP0(sK13(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f495,f334])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,sK13(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f495,f366])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,sK13(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f495,f370])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,sK13(X2,X2))) )),
  inference(unit_resulting_resolution,[],[f495,f374])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f495,f375])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP1(sK13(X0,X0),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f495,f376])).
fof(f508,plain,(
  ( ! [X0,X1] : (~sP6('F_0',sK13(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f414,f495,f382])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f495,f390])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP6('T_0',sK13(X0,X0),X1)) )),
  inference(unit_resulting_resolution,[],[f247,f495,f394])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP6(sK13(X0,X0),sK13(X1,X1),X2)) )),
  inference(unit_resulting_resolution,[],[f495,f495,f394])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP6(sK13(X0,X0),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f433,f495,f402])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP6(sK13(X0,X0),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f435,f495,f402])).
fof(f495,plain,(
  ( ! [X0] : (p(s(bool,sK13(X0,X0)))) )),
  inference(equality_resolution,[],[f349])).
fof(f474,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f247,f394])).
fof(f467,plain,(
  ( ! [X0] : (sP6('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f247,f393])).
fof(f464,plain,(
  ( ! [X0] : (~sP6('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f414,f382])).
fof(f463,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f414,f381])).
fof(f460,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f414,f414,f415])).
fof(f438,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f414,f271])).
fof(f436,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f414,f391])).
fof(f435,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f390])).
fof(f433,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f414,f389])).
fof(f432,plain,(
  ( ! [X0,X1] : (sP1('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f247,f376])).
fof(f431,plain,(
  ( ! [X0,X1] : (sP1(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f247,f375])).
fof(f430,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f374])).
fof(f428,plain,(
  ( ! [X0,X1] : (sP2('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f414,f372])).
fof(f426,plain,(
  ( ! [X0,X1] : (sP2(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f414,f371])).
fof(f425,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f370])).
fof(f423,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f414,f368])).
fof(f421,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f414,f367])).
fof(f420,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f366])).
fof(f419,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f247,f334])).
fof(f418,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f247,f333])).
fof(f398,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK15(X3,X2,X1,X0))) | s('type.num.num',X2) = s('type.num.num',X3)) )),
  inference(cnf_transformation,[],[f244])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK15(X3,X2,X1,X0))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1) | p(s(bool,sK15(X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f244])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP6(X2,X0,X1)) & (~sP6(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP6(X2,X0,X1))),
  inference(definition_folding,[],[f165,f174,f173])).
fof(f173,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f174,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X49,X50,X51] : ((p(s(bool,X51)) <=> (p(s(bool,X50)) => p(s(bool,X49)))) <=> ((p(s(bool,X51)) | p(s(bool,X50))) & (p(s(bool,X51)) | ~p(s(bool,X49))) & (~p(s(bool,X50)) | p(s(bool,X49)) | ~p(s(bool,X51)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.sat.sat_27_prove_ltb')).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP6(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f173])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP5(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP5(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP5(X2,X0,X1)) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f174])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP5(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f236])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~sP5(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP4(X2,X1,X0)) & (~sP4(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP4(X2,X1,X0))),
  inference(definition_folding,[],[f135,f171,f170,f169,f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (sP1(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f169,plain,(
  ! [X1,X0,X2] : (sP2(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f170,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f171,plain,(
  ! [X2,X1,X0] : (sP4(X2,X1,X0) <=> (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X49,X50,X51] : ((p(s(bool,X51)) <=> s(bool,X49) = s(bool,X50)) <=> ((p(s(bool,X51)) | p(s(bool,X50)) | p(s(bool,X49))) & (p(s(bool,X51)) | ~p(s(bool,X49)) | ~p(s(bool,X50))) & (p(s(bool,X50)) | ~p(s(bool,X49)) | ~p(s(bool,X51))) & (p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.sat.sat_2|49|_prove_ltb')).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f233])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP4(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~sP1(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(nnf_transformation,[],[f168])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f232])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f232])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X1,X0,X2] : ((~sP2(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0,X2)))),
  inference(nnf_transformation,[],[f169])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f230])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f170])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f227])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f227])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (sP1(X2,X1,X0) & sP2(X1,X2,X0) & sP3(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP1(X2,X1,X0) | ~sP2(X1,X2,X0) | ~sP3(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f223])).
fof(f223,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP4(X2,X1,X0)))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X2,X1,X0] : ((~sP4(X2,X1,X0) | (sP1(X0,X1,X2) & sP2(X1,X0,X2) & sP3(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP1(X0,X1,X2) | ~sP2(X1,X0,X2) | ~sP3(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP4(X2,X1,X0)))),
  inference(nnf_transformation,[],[f171])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP2(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f224])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | sP3(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f224])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f224])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2) | ~sP3(X0,X2,X1) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f224])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X2,X1) | p(s(bool,X1)) | sP4(X0,X1,X2) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f224])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X2,X1) | p(s(bool,X0)) | sP4(X0,X1,X2) | ~sP2(X1,X2,X0) | ~sP1(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f224])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X34,X35,X36] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X34))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X36),s('type.num.num',X35)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|11')).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X30,X31,X32] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X31),s('type.num.num',X30))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X32)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|8')).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X38,X39,X40] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X40))),s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X39),s('type.num.num',X38)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|14')).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X30,X31,X33] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Clos_2'(s('type.num.num',X31),s('type.num.num',X30))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X33)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|9')).
fof(f351,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK14(X1,X0))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1)))) )),
  inference(cnf_transformation,[],[f221])).
fof(f352,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X0))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.clos_known.val_approx'),X1))) | p(s(bool,sK14(X1,X0)))) )),
  inference(cnf_transformation,[],[f221])).
fof(f348,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK13(X1,X0))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f349,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | p(s(bool,sK13(X1,X0)))) )),
  inference(cnf_transformation,[],[f218])).
fof(f415,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f332,f414])).
fof(f332,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f210])).
fof(f322,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f325,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f151])).
fof(f319,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X41,X40] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X40))),s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X41)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|15')).
fof(f318,plain,(
  ( ! [X0,X1] : (s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X0))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X34,X37] : s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Tuple_1'(s('type.list.list'('type.clos_known.val_approx'),X34))),s('type.clos_known.val_approx','const.clos_known.Int_1'(s('type.integer.int',X37)))))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_known.merge_def|split|12')).
fof(f275,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f276,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f183])).
fof(f271,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f272,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f181])).
fof(f3798,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791,f3763,f3794,f3796,f3768])).
fof(f3797,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791,f3763,f3794,f3796])).
fof(f3795,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791,f3763,f3794])).
fof(f3793,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791,f3763])).
fof(f3792,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759,f3791])).
fof(f3790,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788,f3759])).
fof(f3789,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786,f3788])).
fof(f3787,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784,f3786])).
fof(f3785,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782,f3784])).
fof(f3783,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779,f3782])).
fof(f3781,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f3741,f247])).
fof(f3741,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f3666,f414])).
fof(f3780,plain,(
  $false | $spl6),
  inference(global_subsumption,[],[f246,f245,f247,f249,f250,f414,f256,f257,f258,f259,f260,f261,f262,f263,f264,f272,f271,f276,f275,f316,f317,f318,f319,f320,f321,f325,f322,f331,f330,f334,f333,f415,f347,f350,f349,f348,f416,f352,f351,f354,f355,f356,f357,f364,f363,f362,f361,f360,f359,f358,f368,f367,f366,f365,f372,f371,f370,f369,f376,f375,f374,f373,f380,f379,f378,f377,f387,f386,f385,f402,f383,f382,f381,f391,f390,f389,f388,f395,f394,f393,f392,f401,f400,f399,f398,f418,f419,f420,f421,f423,f425,f426,f428,f430,f431,f432,f433,f435,f436,f438,f460,f463,f464,f467,f474,f495,f518,f516,f515,f512,f509,f508,f505,f504,f503,f502,f501,f500,f499,f498,f497,f524,f570,f591,f589,f613,f609,f671,f693,f703,f709,f718,f731,f730,f729,f728,f757,f758,f759,f760,f761,f762,f777,f791,f799,f803,f839,f847,f852,f863,f868,f881,f891,f901,f910,f915,f949,f938,f932,f929,f956,f999,f989,f988,f1056,f1094,f1232,f1229,f1228,f1225,f1222,f1221,f1220,f1219,f1217,f1216,f1215,f1211,f1210,f1207,f1204,f1201,f1198,f1197,f1189,f1165,f1160,f1158,f1156,f1154,f1152,f1151,f1150,f1149,f1146,f1143,f1135,f1132,f1129,f1126,f1118,f1115,f1113,f1111,f1109,f1107,f1106,f1103,f1102,f1101,f1100,f1296,f1384,f1383,f1382,f1380,f1455,f1381,f1714,f1715,f1733,f1734,f1762,f1786,f1815,f1816,f1817,f1818,f1819,f1820,f1821,f1822,f1826,f1827,f1828,f1829,f2003,f2012,f2010,f2179,f2175,f2174,f2170,f2165,f2164,f2163,f2161,f2160,f2159,f2158,f2153,f2152,f2148,f2143,f2142,f2136,f2135,f2128,f2100,f2099,f2098,f2096,f2094,f2093,f2092,f2091,f2090,f2084,f2083,f2076,f2075,f2067,f2066,f2057,f2055,f2054,f2051,f2049,f2039,f2037,f2035,f2034,f2031,f2030,f2029,f2028,f2278,f2283,f2437,f2433,f2432,f2428,f2423,f2422,f2421,f2419,f2418,f2417,f2416,f2411,f2410,f2406,f2401,f2400,f2394,f2393,f2386,f2358,f2357,f2356,f2354,f2352,f2351,f2350,f2349,f2348,f2342,f2341,f2334,f2333,f2325,f2324,f2315,f2313,f2312,f2309,f2307,f2300,f2298,f2297,f2296,f2293,f2292,f2290,f2489,f2484,f2496,f2671,f2667,f2666,f2661,f2656,f2655,f2654,f2652,f2651,f2650,f2649,f2644,f2643,f2639,f2634,f2633,f2627,f2626,f2619,f2591,f2590,f2589,f2587,f2585,f2584,f2583,f2582,f2581,f2573,f2572,f2563,f2562,f2554,f2553,f2544,f2543,f2542,f2540,f2539,f2537,f2533,f2523,f2521,f2517,f2515,f2514,f2511,f2510,f2509,f2508,f2728,f2746,f2921,f2917,f2916,f2911,f2906,f2905,f2904,f2902,f2901,f2900,f2899,f2894,f2893,f2889,f2884,f2883,f2877,f2876,f2869,f2841,f2840,f2839,f2837,f2835,f2834,f2833,f2832,f2831,f2823,f2822,f2813,f2812,f2804,f2803,f2794,f2793,f2791,f2787,f2785,f2783,f2779,f2769,f2767,f2763,f2761,f2760,f2757,f2756,f2755,f2754,f3010,f3197,f3193,f3192,f3187,f3182,f3181,f3180,f3178,f3177,f3176,f3175,f3170,f3169,f3165,f3159,f3158,f3152,f3151,f3144,f3115,f3114,f3113,f3111,f3109,f3108,f3107,f3106,f3105,f3097,f3096,f3087,f3086,f3078,f3077,f3068,f3067,f3065,f3061,f3059,f3057,f3053,f3043,f3041,f3037,f3035,f3034,f3031,f3030,f3029,f3028,f3200,f3309,f3358,f3361,f3363,f3366,f3367,f3407,f3664,f3669,f3672,f3675,f3678,f3681,f3684,f3687,f3690,f3693,f3695,f3696,f3697,f3698,f3699,f3700,f3701,f3702,f3703,f3704,f3705,f3706,f3707,f3708,f3709,f3457,f3710,f3459,f3711,f3461,f3712,f3713,f3464,f3465,f3719,f3722,f3724,f3726,f3728,f3730,f3732,f3734,f3736,f3738,f3666,f3779])).
fof(f3694,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~sP6(sK13(X2,X2),X1,'T_0')) ) | $spl24),
  inference(cnf_transformation,[],[f3694_D])).
fof(f3694_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~sP6(sK13(X2,X2),X1,'T_0')) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3666,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f3666_D])).
fof(f3666_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3738,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6(sK13(X2,X2),X1,'T_0')) )),
  inference(condensation,[],[f3737])).
fof(f3737,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | ~sP6(sK13(X70,X70),X69,'T_0') | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f3487,f570])).
fof(f3487,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP6(sK13(X70,X70),X69,'T_0') | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f839,f524])).
fof(f3691,plain,(
  ( ! [X2,X1] : (~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f3691_D])).
fof(f3691_D,plain,(
  ( ! [X2,X1] : (~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3736,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f3735])).
fof(f3735,plain,(
  ( ! [X61,X64,X62,X63] : (~sP5('T_0',X62,X63) | sP6('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f3485,f570])).
fof(f3485,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~sP5('T_0',X62,X63) | sP6('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f402,f524])).
fof(f3688,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) ) | $spl20),
  inference(cnf_transformation,[],[f3688_D])).
fof(f3688_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3734,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f3733])).
fof(f3733,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP6('T_0',X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3484,f570])).
fof(f3484,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP6('T_0',X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f395,f524])).
fof(f3685,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | sP6('T_0',X1,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f3685_D])).
fof(f3685_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | sP6('T_0',X1,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3732,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f3731])).
fof(f3731,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | sP6('T_0',X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f3483,f570])).
fof(f3483,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | sP6('T_0',X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f394,f524])).
fof(f3682,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP6('T_0',X2,X1)) ) | $spl16),
  inference(cnf_transformation,[],[f3682_D])).
fof(f3682_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | sP6('T_0',X2,X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3730,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f3729])).
fof(f3729,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,X50)) | sP6('T_0',X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f3482,f570])).
fof(f3482,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | sP6('T_0',X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f393,f524])).
fof(f3679,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f3679_D])).
fof(f3679_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3728,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f3727])).
fof(f3727,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~sP5(X47,X46,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f3481,f570])).
fof(f3481,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,X47)) | ~sP5(X47,X46,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f388,f524])).
fof(f3676,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP2(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f3676_D])).
fof(f3676_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP2(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3726,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f3725])).
fof(f3725,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f3480,f570])).
fof(f3480,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f369,f524])).
fof(f3673,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP3(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f3673_D])).
fof(f3673_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP3(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3724,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f3723])).
fof(f3723,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f3479,f570])).
fof(f3479,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f365,f524])).
fof(f3670,plain,(
  ( ! [X2,X1] : (sP4(X1,X2,'T_0') | ~sP2(X2,'T_0',X1) | ~sP3(X1,'T_0',X2)) ) | $spl8),
  inference(cnf_transformation,[],[f3670_D])).
fof(f3670_D,plain,(
  ( ! [X2,X1] : (sP4(X1,X2,'T_0') | ~sP2(X2,'T_0',X1) | ~sP3(X1,'T_0',X2)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3722,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f3721,f432])).
fof(f3721,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | ~sP1('T_0',X2,X1)) )),
  inference(condensation,[],[f3720])).
fof(f3720,plain,(
  ( ! [X35,X33,X36,X34] : (sP4(X34,X35,'T_0') | ~sP3(X34,'T_0',X35) | ~sP2(X35,'T_0',X34) | ~sP1('T_0',X35,X34) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f3478,f570])).
fof(f3478,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | sP4(X34,X35,'T_0') | ~sP3(X34,'T_0',X35) | ~sP2(X35,'T_0',X34) | ~sP1('T_0',X35,X34) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f362,f524])).
fof(f3665,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f3665_D])).
fof(f3665_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3719,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f3718])).
fof(f3718,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4(X30,'T_0',X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f3477,f570])).
fof(f3477,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4(X30,'T_0',X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f361,f524])).
fof(f3693,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6(sK13(X2,X2),X1,'T_0')) )),
  inference(condensation,[],[f3692])).
fof(f3692,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | ~sP6(sK13(X70,X70),X69,'T_0') | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f3428,f570])).
fof(f3428,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP6(sK13(X70,X70),X69,'T_0') | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f839,f524])).
fof(f3690,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP5('T_0',X1,X2) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f3689])).
fof(f3689,plain,(
  ( ! [X61,X64,X62,X63] : (~sP5('T_0',X62,X63) | sP6('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f3426,f570])).
fof(f3426,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~sP5('T_0',X62,X63) | sP6('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f402,f524])).
fof(f3687,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f3686])).
fof(f3686,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP6('T_0',X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3425,f570])).
fof(f3425,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP6('T_0',X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f395,f524])).
fof(f3684,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f3683])).
fof(f3683,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | sP6('T_0',X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f3424,f570])).
fof(f3424,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | sP6('T_0',X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f394,f524])).
fof(f3681,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f3680])).
fof(f3680,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,X50)) | sP6('T_0',X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f3423,f570])).
fof(f3423,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | sP6('T_0',X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f393,f524])).
fof(f3678,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X1,'T_0')) )),
  inference(condensation,[],[f3677])).
fof(f3677,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | ~sP5(X47,X46,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f3422,f570])).
fof(f3422,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,X47)) | ~sP5(X47,X46,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f388,f524])).
fof(f3675,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X1,'T_0',X2)) )),
  inference(condensation,[],[f3674])).
fof(f3674,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f3421,f570])).
fof(f3421,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP2(X42,'T_0',X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f369,f524])).
fof(f3672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f3671])).
fof(f3671,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f3420,f570])).
fof(f3420,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'T_0',X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f365,f524])).
fof(f3669,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f3668,f432])).
fof(f3668,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP4(X1,X2,'T_0') | ~sP3(X1,'T_0',X2) | ~sP2(X2,'T_0',X1) | ~sP1('T_0',X2,X1)) )),
  inference(condensation,[],[f3667])).
fof(f3667,plain,(
  ( ! [X35,X33,X36,X34] : (sP4(X34,X35,'T_0') | ~sP3(X34,'T_0',X35) | ~sP2(X35,'T_0',X34) | ~sP1('T_0',X35,X34) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f3419,f570])).
fof(f3419,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,X33)) | sP4(X34,X35,'T_0') | ~sP3(X34,'T_0',X35) | ~sP2(X35,'T_0',X34) | ~sP1('T_0',X35,X34) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f362,f524])).
fof(f3664,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f3663])).
fof(f3663,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4(X30,'T_0',X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f3418,f570])).
fof(f3418,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4(X30,'T_0',X31) | s(bool,X29) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f361,f524])).
fof(f1454,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1453,f250])).
fof(f1453,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1417,f414])).
fof(f1417,plain,(
  p(s(bool,'F_0')) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx','const.clos_known.merge_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | ($spl1 | $spl2)),
  inference(superposition,[],[f331,f446])).
fof(f446,plain,(
  s(bool,'F_0') = s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0'))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f417,f276])).
fof(f417,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx','const.clos_known.Other_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f411,f412])).
fof(f1420,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f1419,f414])).
fof(f1419,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1381,f446])).
fof(f410,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) | $spl0),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f412,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | $spl2),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f245,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7) | p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))))),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',sK7)) & (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx',sK7))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f176])).
fof(f176,plain,(
  ? [X0] : ((p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0)) & (~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))) | s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx',X0)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ? [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))) <~> s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ~! [X0] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X0)))) <=> s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X6)))) <=> s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',X6)))) <=> s('type.clos_known.val_approx','const.clos_known.Other_0') = s('type.clos_known.val_approx',X6))),
  file('thm.clos_knownProof.subapprox_Other|split|0','thm.clos_knownProof.subapprox_Other|split|0')).
fof(f411,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) | $spl1),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f413,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx',sK7) | $spl3),
  inference(cnf_transformation,[],[f413_D])).
fof(f413_D,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx',sK7) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f246,plain,(
  s('type.clos_known.val_approx','const.clos_known.Other_0') != s('type.clos_known.val_approx',sK7) | ~p(s(bool,'const.clos_knownProof.subapprox_2'(s('type.clos_known.val_approx','const.clos_known.Other_0'),s('type.clos_known.val_approx',sK7))))),
  inference(cnf_transformation,[],[f177])).
% SZS output end Proof for thm.clos_knownProof
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 390
% Passive clauses: 1510
% Generated clauses: 6675
% Final active clauses: 240
% Final passive clauses: 489
% Input formulas: 62
% Initial clauses: 157
% 
% Duplicate literals: 146
% Trivial inequalities: 27
% Fw subsumption resolutions: 140
% Fw demodulations: 1264
% Bw demodulations: 617
% Condensations: 83
% Global subsumptions: 12
% 
% Simple tautologies: 62
% Equational tautologies: 44
% Forward subsumptions: 3316
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 2447
% Factoring: 1
% Forward superposition: 532
% Backward superposition: 686
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 51
% Split components: 112
% SAT solver clauses: 1645
% SAT solver unit clauses: 630
% SAT solver binary clauses: 391
% 
% Sat splits: 51
% Sat splitting refutations: 19
% 
% Memory used [KB]: 4605
% Time elapsed: 0.142 s
% ------------------------------
% ------------------------------
% Success in time 1.46 s
